[1561747135] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177048 kB) free!
[1561747156] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197440 kB) free!
[1561747217] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (211296 kB) free!
[1561747255] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188280 kB) free!
[1561747276] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180748 kB) free!
[1561747337] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202952 kB) free!
[1561747375] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201528 kB) free!
[1561747396] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203820 kB) free!
[1561747457] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201120 kB) free!
[1561747708] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.38, 4.05
[1561748077] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561748452] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8713276 kB) free!
[1561748572] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12386388 kB) free.
[1561748733] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11245360 kB) free!
[1561748853] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (19551612 kB) free.
[1561749172] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8510072 kB) free!
[1561749292] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12384672 kB) free.
[1561749453] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (11419604 kB) free!
[1561749573] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (19542820 kB) free.
[1561749610] SERVICE ALERT: smscapp1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1561749877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561749892] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8459224 kB) free!
[1561749997] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561750012] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12376120 kB) free.
[1561750173] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11838060 kB) free!
[1561750293] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (19828000 kB) free.
[1561750544] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 12.1% (2961468 kB) free!
[1561750597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561750604] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 12.1% (2962836 kB) free!
[1561750612] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8421548 kB) free!
[1561750664] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 12.1% (2964832 kB) free!
[1561750717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561750725] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 12.1% (2964788 kB) free!
[1561750732] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12376080 kB) free.
[1561750893] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (12613264 kB) free!
[1561751014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (20214352 kB) free.
[1561751317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561751332] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8423184 kB) free!
[1561751438] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561751452] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12369456 kB) free.
[1561751613] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (12976420 kB) free!
[1561751733] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20224168 kB) free.
[1561752037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561752052] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8341836 kB) free!
[1561752157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561752173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12375636 kB) free.
[1561752333] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (12772216 kB) free!
[1561752453] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20236752 kB) free.
[1561752757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561752772] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8277504 kB) free!
[1561752877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561752892] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12377024 kB) free.
[1561753053] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (13031484 kB) free!
[1561753173] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20249428 kB) free.
[1561753477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561753492] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8104188 kB) free!
[1561753597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561753612] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12372328 kB) free.
[1561753708] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 71.21, 33.61, 13.24
[1561753828] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 75.36, 47.55, 20.81
[1561753948] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1561753948] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 75.59, 56.85, 27.46
[1561754197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561754212] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7980560 kB) free!
[1561754318] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561754332] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12375332 kB) free.
[1561754457] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (197880 kB) free!
[1561754548] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 12.52, 19.55
[1561754577] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (195212 kB) free!
[1561754657] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2714556 kB) free!
[1561754697] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (386080 kB) free!
[1561754917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561754932] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7705716 kB) free!
[1561755037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561755052] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12364484 kB) free.
[1561755573] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (13034160 kB) free!
[1561755637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561755693] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20224592 kB) free.
[1561755757] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561756252] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9055808 kB) free!
[1561756293] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12184684 kB) free!
[1561756348] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.12, 0.31, 3.09
[1561756357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561756373] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12375424 kB) free.
[1561756413] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (19911328 kB) free.
[1561756477] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561756972] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8968592 kB) free!
[1561757013] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12825888 kB) free!
[1561757092] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12376852 kB) free.
[1561757133] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20241240 kB) free.
[1561757678] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561757692] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8875292 kB) free!
[1561757734] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (12778564 kB) free!
[1561757798] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561757812] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12373648 kB) free.
[1561757854] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20249056 kB) free.
[1561758398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561758412] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8847332 kB) free!
[1561758454] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12849320 kB) free!
[1561758518] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561758532] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12375288 kB) free.
[1561758574] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20247468 kB) free.
[1561759118] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561759132] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8748648 kB) free!
[1561759174] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12544236 kB) free!
[1561759238] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561759252] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12366508 kB) free.
[1561759294] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (20072596 kB) free.
[1561759838] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561759852] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8614956 kB) free!
[1561759894] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12277412 kB) free!
[1561759958] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561759972] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12372000 kB) free.
[1561760014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (19778676 kB) free.
[1561760057] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2433968 kB) free!
[1561760097] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (1780232 kB) free!
[1561760559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561760572] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8374524 kB) free!
[1561760614] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11820860 kB) free!
[1561760679] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1561760692] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12374968 kB) free.
[1561760734] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (19316344 kB) free.
[1561761279] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561761293] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8115532 kB) free!
[1561761297] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1620776 kB) free!
[1561761334] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11232492 kB) free!
[1561761399] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561761413] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12372368 kB) free.
[1561761454] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (18795112 kB) free.
[1561761748] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 110.14, 212.88, 126.64
[1561761857] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (2956252 kB) free!
[1561761868] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.05, 142.52, 111.31
[1561761897] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1561761897] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (1634656 kB) free!
[1561761988] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.34, 95.49, 97.87
[1561761999] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561762013] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7830036 kB) free!
[1561762119] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561762133] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12376432 kB) free.
[1561762497] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1618056 kB) free!
[1561762699] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 68.18, 38.49, 16.08
[1561762719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561762733] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7620900 kB) free!
[1561762819] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 69.27, 48.71, 22.54
[1561762840] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561762853] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12366316 kB) free.
[1561762939] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.25, 46.02, 25.00
[1561763098] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.2% (1656596 kB) free!
[1561763367] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 60.25, 34.92, 15.81
[1561763439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561763487] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 61.29, 43.73, 21.36
[1561763539] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.11, 6.34, 13.20
[1561763559] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561763607] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 61.51, 49.63, 26.23
[1561763788] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 0.57, 14.37
[1561764053] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8900312 kB) free!
[1561764173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12377140 kB) free.
[1561764207] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.44, 12.83, 19.32
[1561764739] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.56, 3.95
[1561764759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561764773] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8729776 kB) free!
[1561764879] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561764893] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12377740 kB) free.
[1561764989] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.41, 4.21
[1561765479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561765493] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8516848 kB) free!
[1561765599] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561765614] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12373744 kB) free.
[1561766007] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.32, 2.34, 4.80
[1561766199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561766213] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8180392 kB) free!
[1561766320] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561766333] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12375380 kB) free.
[1561766685] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 478 processes
[1561766732] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22649132 kB) free.
[1561766745] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 483 processes
[1561766813] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.0% (22788760 kB) free.
[1561766854] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (12452352 kB) free!
[1561766919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561766933] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8027868 kB) free!
[1561766974] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (19731620 kB) free.
[1561767039] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561767053] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12366316 kB) free.
[1561767574] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11577884 kB) free!
[1561767639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561767653] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7760068 kB) free!
[1561767694] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (19192452 kB) free.
[1561767759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561767773] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11645140 kB) free.
[1561767879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561767898] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1521820 kB) free!
[1561768458] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2395588 kB) free!
[1561768894] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12324432 kB) free!
[1561768973] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8676724 kB) free!
[1561769014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (20210572 kB) free.
[1561769079] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561769093] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12348676 kB) free.
[1561769679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561769799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561769919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561770214] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12240020 kB) free!
[1561770294] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9652548 kB) free!
[1561770298] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 62.3% (10127512 kB) free.
[1561770334] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (19883768 kB) free.
[1561770414] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8282688 kB) free!
[1561770519] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1561770519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561770534] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12322924 kB) free.
[1561771119] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561771239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561771359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561771535] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (11540272 kB) free!
[1561771655] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (19809120 kB) free.
[1561771734] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9158456 kB) free!
[1561771735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1561771854] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7603772 kB) free!
[1561771855] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1561771974] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11650540 kB) free.
[1561772455] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561772560] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561772575] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1561773760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561773880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561773948] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.6% (21726632 kB) free.
[1561774374] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8839720 kB) free!
[1561774494] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12295304 kB) free.
[1561775080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561775095] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8623440 kB) free!
[1561775200] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561775215] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12292256 kB) free.
[1561775255] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12180700 kB) free!
[1561775375] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (19606068 kB) free.
[1561775789] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561775800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561775815] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8258836 kB) free!
[1561775920] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561775935] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12096932 kB) free.
[1561776139] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 88.86, 130.83, 71.11
[1561776225] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 34.5% (8453360 kB) free.
[1561776259] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.37, 87.67, 62.54
[1561776379] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.18, 58.83, 55.02
[1561776521] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561776575] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12338104 kB) free!
[1561776641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561776695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (19564964 kB) free.
[1561776761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561776989] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 74.91, 49.20, 21.91
[1561777109] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 75.46, 57.90, 28.41
[1561777229] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.71, 38.90, 25.01
[1561777245] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 488 processes
[1561777305] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 484 processes
[1561777361] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561777579] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 1.67, 15.53
[1561777829] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.21, 5.58, 13.35
[1561777895] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (12897016 kB) free!
[1561777961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561778015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (19393916 kB) free.
[1561778081] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561778681] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561778779] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.66, 0.48, 4.56
[1561778801] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561779029] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 0.39, 3.91
[1561779216] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10832056 kB) free!
[1561779336] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (19224908 kB) free.
[1561779402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561779522] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561779898] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1561779979] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 161.38, 128.23, 65.70
[1561780099] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 31.93, 93.86, 61.03
[1561780122] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561780135] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7683828 kB) free!
[1561780219] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1561780219] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.94, 63.42, 53.86
[1561780242] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561780255] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12276120 kB) free.
[1561780536] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (8260100 kB) free!
[1561780656] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (16939276 kB) free.
[1561780842] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561780962] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561781256] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (10626784 kB) free!
[1561781376] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (18138832 kB) free.
[1561781419] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 2.16, 15.74
[1561781455] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8697844 kB) free!
[1561781575] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12287084 kB) free.
[1561782162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561782175] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8403816 kB) free!
[1561782283] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561782295] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12284888 kB) free.
[1561782419] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.7% (1004332 kB) free!
[1561782539] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.0% (881456 kB) free!
[1561782576] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (7839788 kB) free!
[1561782620] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.21, 1.01, 4.91
[1561782659] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.1% (6698684 kB) free.
[1561782696] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (16550612 kB) free.
[1561782882] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561782895] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7800924 kB) free!
[1561783002] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561783015] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12280856 kB) free.
[1561783296] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (10017600 kB) free!
[1561783416] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (6063692 kB) free!
[1561783536] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (15488120 kB) free.
[1561783602] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561783669] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.18, 7.24, 5.97
[1561783722] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561783789] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.31, 6.90, 6.02
[1561783797] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.5% (2349140 kB) free!
[1561783842] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561783909] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.26, 6.22, 5.88
[1561783917] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 13.1% (2130932 kB) free!
[1561784037] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.8% (1925848 kB) free!
[1561784136] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (9043940 kB) free!
[1561784215] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9021024 kB) free!
[1561784256] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (16608776 kB) free.
[1561784335] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12272500 kB) free.
[1561784637] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (455840 kB) free!
[1561784857] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11646776 kB) free!
[1561784977] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (7620444 kB) free!
[1561785097] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (15879536 kB) free.
[1561785298] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3034160 kB) free!
[1561785418] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161288 kB) free!
[1561785539] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178204 kB) free!
[1561785697] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10936856 kB) free!
[1561785817] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (6930336 kB) free!
[1561785937] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (15288560 kB) free.
[1561786135] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8560684 kB) free!
[1561786243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561786255] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12254268 kB) free.
[1561786360] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.9% (14619752 kB) free.
[1561786407] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 69.8% (11340648 kB) free.
[1561786537] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (10294684 kB) free!
[1561786657] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (6048036 kB) free!
[1561786777] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (14648028 kB) free.
[1561786843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561786855] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8048260 kB) free!
[1561786963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561786975] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12005552 kB) free.
[1561787083] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561787377] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (9708816 kB) free!
[1561787497] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (5601012 kB) free!
[1561787617] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (14060500 kB) free.
[1561787824] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 362 processes
[1561787945] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 4 warn out of 365 processes
[1561788038] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (183664 kB) free!
[1561788065] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 370 processes
[1561788159] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171308 kB) free!
[1561788207] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 33.23, 76.99, 44.21
[1561788217] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (9307076 kB) free!
[1561788279] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175988 kB) free!
[1561788327] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.41, 52.73, 39.30
[1561788337] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (5217280 kB) free!
[1561788447] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.70, 36.32, 34.92
[1561788457] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (8071972 kB) free!
[1561788665] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1561788775] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9538868 kB) free!
[1561788884] HOST ALERT: sig4-vm2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1561788895] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8071020 kB) free!
[1561788906] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1561788966] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4-vm2' was out of bounds)
[1561789015] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12249408 kB) free.
[1561789048] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.79, 7.32, 19.63
[1561789088] SERVICE ALERT: sig4-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig4-vm2' was out of bounds)
[1561789111] SERVICE ALERT: sig4-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4-vm2' was out of bounds)
[1561789178] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4-vm2' was out of bounds)
[1561789234] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig4-vm2' was out of bounds)
[1561789259] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.2% (1043548 kB) free!
[1561789291] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1561789379] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 11.7% (859888 kB) free!
[1561789451] HOST ALERT: sig4-vm2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.191)
[1561789500] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (654072 kB) free!
[1561789615] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7614300 kB) free!
[1561789657] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (23770028 kB) free.
[1561789735] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12253880 kB) free.
[1561789909] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 6.46, 4.67, 4.80
[1561790100] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.0% (6768696 kB) free.
[1561790138] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180660 kB) free!
[1561790158] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.8% (3616136 kB) free!
[1561790224] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197324 kB) free!
[1561790258] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184732 kB) free!
[1561790279] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185000 kB) free!
[1561790345] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182636 kB) free!
[1561790355] HOST ALERT: sig4-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1561790379] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176052 kB) free!
[1561790399] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180204 kB) free!
[1561790465] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182576 kB) free!
[1561790857] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (12727520 kB) free!
[1561790977] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.8% (22855516 kB) free.
[1561791448] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.25, 3.74, 4.71
[1561792135] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8581296 kB) free!
[1561792177] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (10445500 kB) free!
[1561792255] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12237072 kB) free.
[1561792297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (21027168 kB) free.
[1561792309] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.77, 6.13, 5.57
[1561792368] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.9% (2591092 kB) free!
[1561792429] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.62, 5.99, 5.59
[1561792488] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.2% (2310176 kB) free!
[1561792550] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.28, 5.72, 5.54
[1561792608] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.7% (2068476 kB) free!
[1561792702] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.9% (23164300 kB) free.
[1561792759] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.4% (21415444 kB) free.
[1561792781] HOST ALERT: sig4-vm2;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1561792807] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.5% (21703692 kB) free.
[1561792845] HOST ALERT: sig4-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1561792907] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.4% (14692660 kB) free.
[1561793144] HOST ALERT: sig4-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.191)
[1561793150] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1561793150] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.55, 4.39, 4.89
[1561793208] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1561793208] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (791884 kB) free!
[1561793415] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172724 kB) free!
[1561793420] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561793461] HOST ALERT: sig4-vm2;UP;SOFT;1;PING OK - Packet loss = 40%, RTA = 0.60 ms
[1561793535] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170080 kB) free!
[1561793656] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (160952 kB) free!
[1561793750] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.50, 5.01, 5.07
[1561793848] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.55, 7.86, 6.30
[1561793870] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.99, 4.60, 4.91
[1561793968] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.52, 7.48, 6.36
[1561794088] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1561794088] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.02, 6.45, 6.13
[1561794688] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.98, 4.55, 5.00
[1561794697] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11655968 kB) free!
[1561794817] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (20577896 kB) free.
[1561795070] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.06, 5.24, 5.03
[1561795190] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.26, 5.58, 5.17
[1561795288] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.89, 5.56, 5.32
[1561795310] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.98, 6.64, 5.62
[1561795408] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.80, 4.92, 5.12
[1561795528] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.85, 4.18, 4.82
[1561796017] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (10029752 kB) free!
[1561796137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (19711932 kB) free.
[1561796455] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8872620 kB) free!
[1561796576] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7502672 kB) free!
[1561796683] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561796696] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12233488 kB) free.
[1561796700] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 12.1% (888176 kB) free!
[1561796820] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (687040 kB) free!
[1561796940] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (515408 kB) free!
[1561797029] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 108.31, 72.83, 32.98
[1561797149] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 108.72, 84.71, 42.16
[1561797167] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 892 processes
[1561797269] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.11, 56.84, 37.11
[1561797287] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 891 processes
[1561797337] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (8822580 kB) free!
[1561797407] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 879 processes
[1561797457] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (18832948 kB) free.
[1561797869] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.59, 8.07, 19.67
[1561797883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561797896] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9044160 kB) free!
[1561797928] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.87, 5.48, 5.58
[1561798003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561798007] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 316 processes
[1561798016] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7529032 kB) free!
[1561798049] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.50, 4.46, 5.19
[1561798123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561798136] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12234772 kB) free.
[1561798169] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.02, 3.97, 4.92
[1561798415] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171360 kB) free!
[1561798536] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (147880 kB) free!
[1561798656] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178224 kB) free!
[1561798657] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (7818968 kB) free!
[1561798724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561798736] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7771360 kB) free!
[1561798777] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (18113152 kB) free.
[1561798844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561798856] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11643616 kB) free.
[1561798910] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.35, 4.16, 4.85
[1561798964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561799251] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.7% (21989544 kB) free.
[1561799291] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.2% (14177364 kB) free.
[1561799346] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 64.0% (10407488 kB) free.
[1561799448] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1561799448] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.3% (6859828 kB) free.
[1561799669] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.71, 0.59, 3.29
[1561799958] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (803248 kB) free!
[1561799977] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (7204952 kB) free!
[1561800078] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (683072 kB) free!
[1561800097] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (17539984 kB) free.
[1561800198] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (828948 kB) free!
[1561800506] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194804 kB) free!
[1561800627] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190780 kB) free!
[1561800697] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (12427208 kB) free!
[1561800747] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179492 kB) free!
[1561800764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561800817] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (8252804 kB) free!
[1561800937] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (17182652 kB) free.
[1561801055] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195272 kB) free!
[1561801166] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194604 kB) free!
[1561801175] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185812 kB) free!
[1561801286] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192700 kB) free!
[1561801295] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1561801295] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187936 kB) free!
[1561801407] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190076 kB) free!
[1561801499] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1561801910] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.46, 5.35, 5.01
[1561801964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561802030] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.24, 4.95, 4.90
[1561802084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561802099] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 8620 MB (16% inode=98%):
[1561802137] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (8200796 kB) free!
[1561802205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561802257] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (16679864 kB) free.
[1561802858] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13067344 kB) free!
[1561802969] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 222.47, 181.43, 94.58
[1561802978] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (9046608 kB) free!
[1561803056] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7572348 kB) free!
[1561803089] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 189.22, 187.95, 107.80
[1561803098] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (5188540 kB) free!
[1561803115] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169472 kB) free!
[1561803176] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12220232 kB) free.
[1561803209] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 27.89, 126.72, 95.08
[1561803235] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169192 kB) free!
[1561803355] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169152 kB) free!
[1561803698] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (14199164 kB) free.
[1561803776] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7642436 kB) free!
[1561803869] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 203.93, 98.97, 39.18
[1561803896] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12220392 kB) free.
[1561803989] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 213.02, 137.11, 60.41
[1561804005] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561804110] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1561804110] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 72.31, 120.71, 64.56
[1561804298] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (8282064 kB) free!
[1561804418] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (4499228 kB) free!
[1561804538] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (15204680 kB) free.
[1561804696] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1561805009] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.04, 3.21, 16.19
[1561805096] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9027420 kB) free!
[1561805138] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (10170512 kB) free!
[1561805205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561805216] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12216556 kB) free.
[1561805258] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (6214944 kB) free!
[1561805310] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 2.76, 18.11
[1561805325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561805378] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (15708892 kB) free.
[1561805465] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1561805585] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1561805706] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1561805816] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8908024 kB) free!
[1561805925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561805936] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12214096 kB) free.
[1561805978] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11278756 kB) free!
[1561805998] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.3% (22536948 kB) free.
[1561806035] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.8% (22496156 kB) free.
[1561806046] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561806078] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.7% (22466212 kB) free.
[1561806098] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (7510692 kB) free!
[1561806124] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.7% (14747384 kB) free.
[1561806157] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 65.5% (10651976 kB) free.
[1561806218] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (3765076 kB) free!
[1561806306] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1561806768] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (942536 kB) free!
[1561806809] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.65, 3.19, 4.95
[1561806818] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1561806818] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (12454860 kB) free!
[1561806889] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (380820 kB) free!
[1561807009] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1561807009] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (956188 kB) free!
[1561807110] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.52, 3.03
[1561807137] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8244748 kB) free!
[1561807257] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12205692 kB) free.
[1561807258] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1561807328] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (180696 kB) free!
[1561807430] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1561807448] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (162308 kB) free!
[1561807506] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 358 processes
[1561807568] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166052 kB) free!
[1561807626] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 355 processes
[1561807746] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 356 processes
[1561808019] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (15081600 kB) free.
[1561808445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561808457] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8450396 kB) free!
[1561808565] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561808577] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12207524 kB) free.
[1561808619] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11587592 kB) free!
[1561808739] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (7704620 kB) free!
[1561808859] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (3859320 kB) free!
[1561808947] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1561808947] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1561809165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561809177] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7955216 kB) free!
[1561809285] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561809297] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12206788 kB) free.
[1561809459] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (18265044 kB) free.
[1561809496] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.5% (2848920 kB) free!
[1561809616] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (2715160 kB) free!
[1561809736] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.0% (2597500 kB) free!
[1561809886] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561809897] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7445688 kB) free!
[1561810006] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561810017] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12202628 kB) free.
[1561810147] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 352 processes
[1561810267] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1561810937] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1528180 kB) free!
[1561811009] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.79, 11.04, 6.18
[1561811129] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.96, 8.23, 5.74
[1561811206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561811217] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8670392 kB) free!
[1561811249] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.11, 6.50, 5.41
[1561811326] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561811337] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12205712 kB) free.
[1561811849] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.05, 3.26, 4.13
[1561811926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561811938] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8479632 kB) free!
[1561812046] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561812058] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12194708 kB) free.
[1561812067] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 358 processes
[1561812188] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 355 processes
[1561812308] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1561812646] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561812658] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8479352 kB) free!
[1561812767] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561812778] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12204264 kB) free.
[1561812908] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1561813235] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1561813366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561813378] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8044080 kB) free!
[1561813486] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561813498] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12203652 kB) free.
[1561814086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561814098] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7420064 kB) free!
[1561814206] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561814218] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12200624 kB) free.
[1561815406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561815418] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8430292 kB) free!
[1561815526] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561815526] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561815538] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12190392 kB) free.
[1561816127] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561816138] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7872064 kB) free!
[1561816247] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561816259] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12181036 kB) free.
[1561816609] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1561816659] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13092976 kB) free!
[1561816779] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (20420372 kB) free.
[1561816847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561816967] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561817249] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 111.08, 90.31, 41.68
[1561817310] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561817369] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.72, 61.43, 36.99
[1561817458] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8781268 kB) free!
[1561817489] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.86, 42.02, 32.84
[1561817578] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12189092 kB) free.
[1561818089] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.14, 8.23, 18.62
[1561818167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561818178] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8565740 kB) free!
[1561818287] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561818298] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12188208 kB) free.
[1561818887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561818898] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8291296 kB) free!
[1561819008] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561819018] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12185360 kB) free.
[1561819179] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11267720 kB) free!
[1561819299] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (21091760 kB) free.
[1561819607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561819618] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7894384 kB) free!
[1561819727] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561819738] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12185988 kB) free.
[1561820327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561820447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561820489] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.70, 2.83, 4.01
[1561820499] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (10037832 kB) free!
[1561820619] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (19646888 kB) free.
[1561820939] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8754384 kB) free!
[1561821020] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 254.28, 189.09, 90.28
[1561821059] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12179540 kB) free.
[1561821140] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 48.56, 138.58, 84.30
[1561821219] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561821260] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7.04, 92.89, 74.14
[1561821647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561821659] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8601424 kB) free!
[1561821767] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561821779] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12184832 kB) free.
[1561822368] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561822379] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8207704 kB) free!
[1561822488] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561822499] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12186496 kB) free.
[1561822710] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4.15, 44.78, 33.07
[1561822830] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.78, 30.04, 29.09
[1561822950] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.32, 20.16, 25.59
[1561823019] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (11154196 kB) free!
[1561823060] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 0.67, 11.05
[1561823088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561823099] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7552160 kB) free!
[1561823139] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (20466660 kB) free.
[1561823208] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561823219] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12181836 kB) free.
[1561823550] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 2.99, 13.58
[1561823808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561823819] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7467172 kB) free!
[1561823928] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561823939] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183168 kB) free.
[1561824260] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.49, 0.49, 3.39
[1561824308] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561824339] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (10228340 kB) free!
[1561824459] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (19437276 kB) free.
[1561824528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561824539] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7438192 kB) free!
[1561824648] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561824659] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12174348 kB) free.
[1561824750] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 0.46, 4.04
[1561825375] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1561825495] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1561825628] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 67.7% (11002492 kB) free.
[1561825659] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (11927292 kB) free!
[1561825743] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.2% (14010980 kB) free.
[1561825779] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (20677684 kB) free.
[1561825848] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561825860] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8379888 kB) free!
[1561825879] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.3% (14679132 kB) free.
[1561825955] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 35.6% (17540500 kB) free.
[1561825968] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561826568] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561826688] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561827288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561827408] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561828608] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561828728] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561829328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561829448] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561830048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561830089] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1561830168] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561830755] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8150400 kB) free!
[1561830875] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12852148 kB) free.
[1561831368] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561831489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561832075] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9577944 kB) free!
[1561832088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561832108] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1561832195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12843060 kB) free.
[1561832208] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561832228] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 352 processes
[1561832348] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1561832486] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (148960 kB) free!
[1561832606] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (150040 kB) free!
[1561832726] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180936 kB) free!
[1561832795] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9382288 kB) free!
[1561832808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561832915] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12849656 kB) free.
[1561832928] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561832949] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 363 processes