[1569695504] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 27.5% (6743764 kB) free.
[1569695587] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 33.8% (5493708 kB) free.
[1569695730] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 579 processes
[1569695765] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (3334584 kB) free!
[1569695856] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187356 kB) free!
[1569695885] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171172 kB) free!
[1569695976] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195456 kB) free!
[1569696005] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198320 kB) free!
[1569696102] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190688 kB) free!
[1569697230] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 587 processes
[1569697292] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 591 processes
[1569697310] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (1271696 kB) free!
[1569697370] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185412 kB) free!
[1569697430] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.9% (225156 kB) free!
[1569697490] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 42.4% (10389028 kB) free.
[1569697929] SERVICE FLAPPING ALERT: smscapp1;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1569698557] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1569700073] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 449 processes
[1569700892] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8493680 kB) free!
[1569701012] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.4% (7564980 kB) free!
[1569701132] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (6770804 kB) free!
[1569704343] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238520 kB) free!
[1569704463] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242932 kB) free!
[1569704584] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224712 kB) free!
[1569705784] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.9% (47260328 kB) free.
[1569706205] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.8% (2883804 kB) free!
[1569706805] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (1099220 kB) free!
[1569707140] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1569707260] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1569707405] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1569707405] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 57.8% (14160448 kB) free.
[1569707502] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.6% (12395236 kB) free.
[1569718805] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569718893] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 592 processes
[1569718953] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 585 processes
[1569719881] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569720001] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1569720757] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.9% (14787636 kB) free.
[1569721205] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (4201576 kB) free!
[1569721304] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201952 kB) free!
[1569721331] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (203588 kB) free!
[1569721353] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 595 processes
[1569721357] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182200 kB) free!
[1569721413] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 596 processes
[1569721424] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170136 kB) free!
[1569721460] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187968 kB) free!
[1569721477] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186060 kB) free!
[1569721544] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175452 kB) free!
[1569721597] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187560 kB) free!
[1569723813] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 602 processes
[1569723874] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 600 processes
[1569723934] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 602 processes
[1569723994] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 2 crit, 0 warn out of 608 processes
[1569724594] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 596 processes
[1569725794] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 604 processes
[1569725854] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 590 processes
[1569725873] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 596 processes
[1569725993] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 594 processes
[1569726113] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 594 processes
[1569728512] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.92, 10.37, 6.70
[1569728632] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.87, 11.39, 7.51
[1569728752] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 12.50, 11.44, 7.98
[1569729932] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4942880 kB) free!
[1569730552] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.67, 3.74, 4.77
[1569731752] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 15.86, 10.67, 6.96
[1569731872] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.46, 10.31, 7.30
[1569731990] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 334 processes
[1569731992] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.34, 7.67, 6.70
[1569732110] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1569732188] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3064236 kB) free!
[1569732308] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.8% (3048024 kB) free!
[1569732428] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (3047896 kB) free!
[1569732796] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.2% (21774756 kB) free.
[1569732878] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569732886] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 69.3% (16977124 kB) free.
[1569732965] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39615684 kB) free.
[1569732999] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1569733025] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.1% (22076380 kB) free.
[1569733088] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569733208] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1569733405] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172536 kB) free!
[1569733493] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177968 kB) free!
[1569733525] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189756 kB) free!
[1569733613] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189136 kB) free!
[1569733638] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196668 kB) free!
[1569733645] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197496 kB) free!
[1569733693] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.5% (4536628 kB) free!
[1569733733] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189336 kB) free!
[1569733758] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187108 kB) free!
[1569733813] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.2% (4454968 kB) free!
[1569733878] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179896 kB) free!
[1569733933] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.8% (4372296 kB) free!
[1569735113] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 632 processes
[1569736193] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1569736193] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.37, 25.09, 20.30
[1569736793] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 14.14, 18.98, 19.52
[1569737393] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.77, 21.56, 20.21
[1569737584] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9415420 kB) free!
[1569737704] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8882724 kB) free!
[1569737824] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.0% (8364228 kB) free!
[1569738133] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2309236 kB) free!
[1569739024] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3207568 kB) free!
[1569741463] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 24760 MB (17% inode=99%):
[1569741583] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 24760 MB (17% inode=99%):
[1569741593] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 15.94, 17.82, 19.71
[1569741703] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 24760 MB (17% inode=99%):
[1569745793] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.27, 2.89, 4.40
[1569749503] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 53405 MB (38% inode=99%):
[1569749699] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 93.6% (22930600 kB) free.
[1569749727] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.9% (22666032 kB) free.
[1569749778] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (260520 kB) free!
[1569749898] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213680 kB) free!
[1569749933] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.3% (22870888 kB) free.
[1569749968] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 67.8% (33414588 kB) free.
[1569750004] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 35.6% (17550152 kB) free.
[1569750045] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.6% (22944080 kB) free.
[1569750335] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198408 kB) free!
[1569750455] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190260 kB) free!
[1569750575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175284 kB) free!
[1569750577] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239684 kB) free!
[1569750697] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (222968 kB) free!
[1569750817] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (221200 kB) free!
[1569753001] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1569753011] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 98.82, 45.70, 18.62
[1569753121] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1569753132] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 104.76, 65.59, 29.19
[1569753252] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 105.60, 78.83, 38.45
[1569753575] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (2500160 kB) free!
[1569753593] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1569754176] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1569754176] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (404096 kB) free!
[1569754452] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 2.46, 12.85
[1569756252] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.64, 0.77, 2.88
[1569761235] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.2% (21774416 kB) free.
[1569761273] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 35.3% (17409884 kB) free.
[1569762121] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1569762241] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1569762445] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181628 kB) free!
[1569762565] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176568 kB) free!
[1569762685] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188988 kB) free!
[1569771252] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.49, 19.62, 9.35
[1569771372] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.91, 13.31, 8.28
[1569771492] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.40, 9.37, 7.45
[1569771685] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1569772692] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.58, 1.86, 4.78
[1569775235] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.4% (23553664 kB) free.
[1569775799] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.30, 10.96, 6.90
[1569775919] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.02, 7.60, 6.16
[1569776039] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.89, 5.45, 5.55
[1569776639] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.21, 1.79, 3.47
[1569779041] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569779161] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569780460] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 304 processes
[1569780580] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes