[1587148665] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10498284 kB) free!
[1587148785] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6801096 kB) free!
[1587148905] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3261844 kB) free!
[1587150080] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3027532 kB) free!
[1587150105] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23968628 kB) free!
[1587150119] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (1547536 kB) free!
[1587150179] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187756 kB) free!
[1587150206] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.5% (4462476 kB) free.
[1587150240] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (190820 kB) free!
[1587150301] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (191780 kB) free!
[1587150705] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4429516 kB) free!
[1587150806] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (1716408 kB) free!
[1587150926] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.5% (1707844 kB) free!
[1587151047] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.8% (1758488 kB) free!
[1587151201] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 80.3% (19678276 kB) free.
[1587151305] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14228504 kB) free!
[1587151647] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1544312 kB) free!
[1587152847] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (1832848 kB) free!
[1587153105] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (13008120 kB) free!
[1587153286] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.5% (13873820 kB) free!
[1587153406] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3575396 kB) free!
[1587153482] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234904 kB) free!
[1587153526] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (504532 kB) free!
[1587153602] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242544 kB) free!
[1587153705] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22439748 kB) free!
[1587153722] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219836 kB) free!
[1587154047] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (527916 kB) free!
[1587154306] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34172672 kB) free.
[1587154647] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (1819544 kB) free!
[1587154906] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11590024 kB) free!
[1587155026] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8046024 kB) free!
[1587155146] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34138016 kB) free.
[1587155522] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.4% (6134916 kB) free!
[1587155746] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11171224 kB) free!
[1587155866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7608060 kB) free!
[1587155986] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27790648 kB) free.
[1587156586] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11655340 kB) free!
[1587156706] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7858536 kB) free!
[1587156826] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4402132 kB) free!
[1587157674] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 66.2% (32623272 kB) free.
[1587158026] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20174204 kB) free!
[1587158216] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.5% (45095920 kB) free.
[1587158626] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30088196 kB) free.
[1587158876] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 36.4% (8920432 kB) free.
[1587159088] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2476) < 172.105.59.173 [18/Apr/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1587159123] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.0% (44357504 kB) free.
[1587159209] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(167) < 172.105.59.173 [18/Apr/2020:03:02:32 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1587159226] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8483284 kB) free!
[1587159329] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587159346] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33738100 kB) free.
[1587159425] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(307) < 172.105.57.57 [18/Apr/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1587159434] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(373) < 2020/04/18 03:01:32 [error] 22134#0: *84646322 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1587159447] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.3% (4106556 kB) free.
[1587159476] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (239120 kB) free!
[1587159545] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587159555] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587159596] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (210828 kB) free!
[1587159716] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (300296 kB) free!
[1587159947] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11421560 kB) free!
[1587160067] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (7977356 kB) free!
[1587160187] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31484108 kB) free.
[1587160787] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2691500 kB) free!
[1587160907] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20980460 kB) free!
[1587161027] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28472236 kB) free.
[1587161627] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2564628 kB) free!
[1587161747] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33318556 kB) free.
[1587162116] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.7% (9963952 kB) free.
[1587162347] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6731632 kB) free!
[1587162467] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2278488 kB) free!
[1587162587] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31356692 kB) free.
[1587163187] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7069168 kB) free!
[1587163307] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33268328 kB) free.
[1587163316] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (257788 kB) free!
[1587163437] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (299092 kB) free!
[1587163557] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1587163557] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (247940 kB) free!
[1587163907] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12437360 kB) free!
[1587164027] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9051552 kB) free!
[1587164147] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33348872 kB) free.
[1587164748] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13074060 kB) free!
[1587164757] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 36.8% (9009952 kB) free.
[1587164868] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34194732 kB) free.
[1587165468] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14237424 kB) free!
[1587165588] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35262364 kB) free.
[1587166188] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15909108 kB) free!
[1587166308] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35263608 kB) free.
[1587166557] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188972 kB) free!
[1587166647] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1587166677] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (231172 kB) free!
[1587166797] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185436 kB) free!
[1587166908] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17107304 kB) free!
[1587167028] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37323396 kB) free.
[1587167382] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 57.1% (14095068 kB) free.
[1587167536] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.9% (19092620 kB) free.
[1587167628] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19530300 kB) free!
[1587167748] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37282336 kB) free.
[1587167871] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 73.2% (17948280 kB) free.
[1587168348] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19027568 kB) free!
[1587168468] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37219172 kB) free.
[1587169068] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17382940 kB) free!
[1587169188] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37083708 kB) free.
[1587169788] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16547672 kB) free!
[1587169909] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36852244 kB) free.
[1587170508] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15532824 kB) free!
[1587170628] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35889840 kB) free.
[1587171228] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15761820 kB) free!
[1587171348] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35452288 kB) free.
[1587171901] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4685076 kB) free!
[1587171948] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16156580 kB) free!
[1587171962] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 36.3% (8896048 kB) free.
[1587172068] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34908164 kB) free.
[1587172127] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 47.4% (62525472 kB) free.
[1587172182] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (450120 kB) free!
[1587172302] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178788 kB) free!
[1587172422] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170080 kB) free!
[1587172668] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16151488 kB) free!
[1587172788] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587172788] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34537908 kB) free.
[1587173388] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16049820 kB) free!
[1587173508] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34009992 kB) free.
[1587174708] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16854312 kB) free!
[1587174828] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35751224 kB) free.
[1587175428] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16673572 kB) free!
[1587175548] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35566476 kB) free.
[1587176149] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15534252 kB) free!
[1587176269] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34775304 kB) free.
[1587176869] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15117492 kB) free!
[1587176989] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33980680 kB) free.
[1587176997] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1587177589] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12858928 kB) free!
[1587177709] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32971692 kB) free.
[1587178909] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15360744 kB) free!
[1587179029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34109776 kB) free.
[1587179629] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13493308 kB) free!
[1587179749] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33569800 kB) free.
[1587180349] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12981560 kB) free!
[1587180469] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33034256 kB) free.
[1587180471] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3391808 kB) free!
[1587180591] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (2181700 kB) free!
[1587180711] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1617748 kB) free!
[1587181058] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.1% (4684832 kB) free!
[1587181069] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10735660 kB) free!
[1587181170] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.13, 6.73, 6.04
[1587181178] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.1% (4681848 kB) free!
[1587181189] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32217320 kB) free.
[1587181290] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.81, 6.37, 5.98
[1587181299] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.1% (4673968 kB) free!
[1587181410] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.93, 7.71, 6.54
[1587181790] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9392756 kB) free!
[1587181910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31121724 kB) free.
[1587182510] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8383820 kB) free!
[1587182630] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4180164 kB) free!
[1587182750] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27330248 kB) free.
[1587183350] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5177280 kB) free!
[1587183470] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27632604 kB) free.
[1587184070] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4744500 kB) free!
[1587184190] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27815024 kB) free.
[1587184790] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4612308 kB) free!
[1587184910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26786932 kB) free.
[1587185510] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5079272 kB) free!
[1587185630] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (623868 kB) free!
[1587185750] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24360680 kB) free!
[1587185904] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1587186024] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1587186136] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179672 kB) free!
[1587186144] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1587186256] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163428 kB) free!
[1587186277] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1587186351] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414808 kB) free!
[1587186376] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185796 kB) free!
[1587186397] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1587186517] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1587186557] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 334 processes
[1587186677] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1587186744] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1587187117] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1587187529] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(397) < 103.255.146.164 [18/Apr/2020:10:47:05 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587187551] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22634656 kB) free!
[1587187649] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587187933] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4333500 kB) free!
[1587188053] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (3689828 kB) free!
[1587188173] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3038020 kB) free!
[1587188751] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1587188751] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (573848 kB) free!
[1587188773] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (195432 kB) free!
[1587189447] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (2496456 kB) free!
[1587189449] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(288) < 103.255.146.164 [18/Apr/2020:11:20:12 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587189562] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 50.3% (12325856 kB) free.
[1587189567] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3253028 kB) free!
[1587189569] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587189687] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3731368 kB) free.
[1587189744] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1587189864] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1587189951] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16451896 kB) free!
[1587189984] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 346 processes
[1587190551] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (445776 kB) free!
[1587190887] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3152916 kB) free!
[1587191007] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.8% (2565572 kB) free!
[1587191128] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3001104 kB) free!
[1587191369] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(137) < 103.255.146.164 [18/Apr/2020:11:54:43 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587191489] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587191751] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31843560 kB) free.
[1587191773] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16257056 kB) free.
[1587192351] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5489564 kB) free!
[1587192471] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1335468 kB) free!
[1587192474] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9196836 kB) free!
[1587192562] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.9% (3646040 kB) free!
[1587192591] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34749696 kB) free.
[1587192594] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (8874840 kB) free!
[1587192682] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 13.2% (3226528 kB) free!
[1587192714] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (8545196 kB) free!
[1587192802] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 11.5% (2825316 kB) free!
[1587193191] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7843328 kB) free!
[1587193312] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3422080 kB) free!
[1587193403] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (795956 kB) free!
[1587193431] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395928 kB) free!
[1587194010] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 35.06, 30.94, 24.64
[1587194128] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (920344 kB) free!
[1587194489] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(128) < 103.255.146.164 [18/Apr/2020:12:47:12 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587194514] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3609536 kB) free!
[1587194609] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587194631] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (25076180 kB) free!
[1587195210] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 16.86, 17.63, 19.83
[1587195231] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389668 kB) free!
[1587195328] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3591676 kB) free.
[1587195416] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8405060 kB) free!
[1587195536] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.4% (8088012 kB) free!
[1587195656] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.7% (7759312 kB) free!
[1587195809] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(96) < 103.255.146.164 [18/Apr/2020:13:11:50 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587195928] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3230316 kB) free!
[1587195929] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587195976] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 21.3% (5209732 kB) free.
[1587196048] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (1600704 kB) free!
[1587196168] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3748356 kB) free.
[1587196431] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26393780 kB) free!
[1587196576] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4730760 kB) free!
[1587196696] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4653272 kB) free!
[1587196768] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2826952 kB) free!
[1587196816] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4586184 kB) free!
[1587196888] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2664088 kB) free!
[1587197008] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1587197008] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2532808 kB) free!
[1587197031] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2401844 kB) free!
[1587197129] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(232) < 103.255.146.164 [18/Apr/2020:13:31:07 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587197249] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587197456] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4704440 kB) free!
[1587198209] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1247512 kB) free!
[1587198231] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32089188 kB) free.
[1587198809] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (2712144 kB) free!
[1587198831] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6217184 kB) free!
[1587198951] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1671644 kB) free!
[1587198973] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (3889336 kB) free!
[1587199071] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35603520 kB) free.
[1587199093] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3590868 kB) free!
[1587199213] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.5% (3251196 kB) free!
[1587199216] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2279688 kB) free!
[1587199409] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (760244 kB) free!
[1587199671] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10164316 kB) free!
[1587199791] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5663056 kB) free!
[1587199813] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1587199813] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1781116 kB) free!
[1587199912] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1317640 kB) free!
[1587200009] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.3% (1993636 kB) free!
[1587200249] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(101) < 103.255.146.164 [18/Apr/2020:14:18:09 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587200369] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587200512] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16396776 kB) free!
[1587201569] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(89) < 103.255.146.164 [18/Apr/2020:14:41:59 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587201689] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587201712] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412860 kB) free!
[1587202213] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.8% (16337412 kB) free.
[1587202290] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(117) < 185.255.8.58 [18/Apr/2020:15:00:10 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1587202312] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23576160 kB) free!
[1587202410] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587202912] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374268 kB) free!
[1587203009] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (920344 kB) free!
[1587203010] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(72) < 103.255.146.164 [18/Apr/2020:15:11:47 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587203130] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587203609] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (2067768 kB) free!
[1587204099] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 31.0% (7610204 kB) free.
[1587204112] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40214388 kB) free.
[1587204712] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7208000 kB) free!
[1587204832] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2442992 kB) free!
[1587204930] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(49) < 103.255.146.164 [18/Apr/2020:15:44:16 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587204952] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414780 kB) free!
[1587205032] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 619 processes
[1587205051] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587205152] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 619 processes
[1587205272] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 368 processes
[1587205552] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15285296 kB) free!
[1587205856] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.9% (9827812 kB) free!
[1587206152] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39436812 kB) free.
[1587206514] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 20.4% (10052584 kB) free.
[1587206752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4320096 kB) free!
[1587206850] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(483) < 103.255.146.164 [18/Apr/2020:16:14:00 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587206872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386008 kB) free!
[1587206970] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587206993] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399208 kB) free!
[1587207210] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1521588 kB) free!
[1587207593] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15504552 kB) free!
[1587207810] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.5% (2361376 kB) free!
[1587208127] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (25179688 kB) free!
[1587208193] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (33002984 kB) free.
[1587208247] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.7% (40479092 kB) free.
[1587208793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6061456 kB) free!
[1587208913] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1816212 kB) free!
[1587209010] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (667816 kB) free!
[1587209033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40018608 kB) free.
[1587209370] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(307) < 103.255.146.164 [18/Apr/2020:16:53:05 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587209490] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587209610] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.1% (2456488 kB) free!
[1587209633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10241340 kB) free!
[1587209753] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5097148 kB) free!
[1587209873] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (516452 kB) free!
[1587210057] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3991856 kB) free!
[1587210210] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (216216 kB) free!
[1587210473] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15267624 kB) free!
[1587210690] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(328) < 103.255.146.164 [18/Apr/2020:17:13:53 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587210810] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (2204972 kB) free!
[1587210811] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587211073] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37871848 kB) free.
[1587211213] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.1% (3385952 kB) free!
[1587211333] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3118636 kB) free!
[1587211410] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (907188 kB) free!
[1587211453] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (2868380 kB) free!
[1587211673] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9418724 kB) free!
[1587211793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4885324 kB) free!
[1587211913] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (650828 kB) free!
[1587212010] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2805432 kB) free!
[1587212053] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1766144 kB) free!
[1587212398] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2499072 kB) free!
[1587212513] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17175180 kB) free!
[1587212998] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198608 kB) free!
[1587213114] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36896916 kB) free.
[1587213210] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(379) < 103.255.146.164 [18/Apr/2020:18:00:10 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587213259] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.6% (16296128 kB) free.
[1587213330] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587213361] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (580848 kB) free!
[1587213421] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (579284 kB) free!
[1587213481] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 2.4% (580512 kB) free!
[1587213541] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 2.3% (575572 kB) free!
[1587213598] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (2553176 kB) free!
[1587213647] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (25923140 kB) free!
[1587213714] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8151468 kB) free!
[1587213767] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.4% (34785100 kB) free.
[1587213810] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1399984 kB) free!
[1587213834] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3323448 kB) free!
[1587213954] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403476 kB) free!
[1587214554] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13414148 kB) free!
[1587214968] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25559328 kB) free!
[1587215088] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.1% (34461168 kB) free.
[1587215130] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(232) < 103.255.146.164 [18/Apr/2020:18:31:16 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587215154] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (12053348 kB) free!
[1587215250] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587215341] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 23.9% (5855832 kB) free.
[1587215610] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (2899724 kB) free!
[1587215998] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1587215998] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2312572 kB) free!
[1587216954] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34341288 kB) free.
[1587217051] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(166) < 103.255.146.164 [18/Apr/2020:19:01:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587217171] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587217410] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.52, 4.15, 4.98
[1587217554] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7280112 kB) free!
[1587217674] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2655516 kB) free!
[1587217794] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391624 kB) free!
[1587218010] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4028796 kB) free.
[1587218394] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13493608 kB) free!
[1587218610] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3174844 kB) free!
[1587218730] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.9% (3877128 kB) free.
[1587218994] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32053988 kB) free.
[1587219330] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2585984 kB) free!
[1587219450] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3340968 kB) free.
[1587219594] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6423868 kB) free!
[1587219715] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2119212 kB) free!
[1587219835] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37496940 kB) free.
[1587220050] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.0% (1782020 kB) free!
[1587220170] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.7% (2382744 kB) free!
[1587220290] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (2993996 kB) free!
[1587220435] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11814976 kB) free!
[1587220555] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7440416 kB) free!
[1587220675] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3294444 kB) free!
[1587220890] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1056504 kB) free!
[1587221184] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1587221275] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14696320 kB) free!
[1587221490] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3471024 kB) free.
[1587221875] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35955536 kB) free.
[1587222090] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (1197988 kB) free!
[1587222210] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (1609392 kB) free!
[1587222331] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (1898404 kB) free!
[1587222475] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7989616 kB) free!
[1587222595] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3789088 kB) free!
[1587222715] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409640 kB) free!
[1587222931] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3294764 kB) free.
[1587223531] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3226952 kB) free!
[1587223651] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3223456 kB) free!
[1587223771] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3205276 kB) free!
[1587223915] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28465912 kB) free.
[1587224060] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1587224515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4283632 kB) free!
[1587224635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (560804 kB) free!
[1587224755] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37821336 kB) free.
[1587225355] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9529104 kB) free!
[1587225476] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5428988 kB) free!
[1587225596] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1592592 kB) free!
[1587225598] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1587226196] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14704972 kB) free!
[1587226796] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28360548 kB) free.
[1587227126] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1674 processes
[1587227396] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5971824 kB) free!
[1587227516] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37131176 kB) free.
[1587227726] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1843 processes
[1587227971] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.4% (5430168 kB) free.
[1587228116] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12306500 kB) free!
[1587228236] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8254072 kB) free!
[1587228356] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4465224 kB) free!
[1587228956] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14756644 kB) free!
[1587229556] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27242844 kB) free.
[1587230156] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5944584 kB) free!
[1587230276] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35481492 kB) free.
[1587230876] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10698880 kB) free!
[1587230996] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6970284 kB) free!
[1587231117] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36494828 kB) free.
[1587231716] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12312220 kB) free!
[1587231836] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8527228 kB) free!
[1587231956] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36421892 kB) free.
[1587232556] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13371472 kB) free!
[1587232676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9375920 kB) free!
[1587232796] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5898008 kB) free!
[1587233396] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14755908 kB) free!
[1587234596] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35621768 kB) free.