[1547318466] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547318922] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547322522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1547322642] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547328309] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547337066] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1547337186] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547338242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547338362] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547339563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547339683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547339803] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1547343683] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.60, 10.61, 10.26
[1547343803] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.38, 10.85, 10.38
[1547343923] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.25, 10.59, 10.34
[1547349797] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 27.5% (6729608 kB) free.
[1547351123] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.41, 7.76, 9.09
[1547352403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547352523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1547352643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547353386] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547353506] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547353626] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547353997] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4653052 kB) free!
[1547354117] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4535284 kB) free!
[1547354237] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (4421352 kB) free!
[1547354443] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547354826] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547355309] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9874016 kB) free!
[1547355429] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12680304 kB) free.
[1547355643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547355763] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547356637] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 21.8% (5332332 kB) free.
[1547358163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547358283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547358403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547358437] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4675516 kB) free!
[1547358557] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4405448 kB) free!
[1547358677] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1547358677] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4314176 kB) free!
[1547359029] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9476608 kB) free!
[1547359149] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12179064 kB) free.
[1547360203] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1547360203] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1547360803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547360923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547361044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547361427] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547361547] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547362277] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2215780 kB) free!
[1547362843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547362963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547363083] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1547363347] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547363349] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9559992 kB) free!
[1547363467] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547363469] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12144296 kB) free.
[1547363587] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547363683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547363723] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.52, 11.20, 10.37
[1547363803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547363843] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.44, 12.05, 10.79
[1547363923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547363963] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.28, 11.66, 10.80
[1547364187] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1547364307] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547364428] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547365269] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9143508 kB) free!
[1547365389] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11764512 kB) free.
[1547365606] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.2% (7998792 kB) free!
[1547365628] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1547365628] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547365726] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8340404 kB) free!
[1547365846] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.6% (9182068 kB) free!
[1547366835] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5545 MB (10% inode=97%):
[1547366944] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 15.33, 11.17, 5.25
[1547367064] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.16, 7.88, 4.76
[1547367077] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (3875780 kB) free!
[1547367428] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547367435] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 8091 MB (15% inode=98%):
[1547367523] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547367548] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547367668] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1547367789] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9006364 kB) free!
[1547367909] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11534260 kB) free.
[1547368123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547368243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547368277] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1199952 kB) free!
[1547368291] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1547368363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547368411] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1547368531] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1547368869] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547368989] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547369109] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547369131] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 13 users currently logged in
[1547369382] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 447 processes
[1547369442] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 446 processes
[1547369502] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 1 warn out of 448 processes
[1547369562] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 1 warn out of 445 processes
[1547369709] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547369709] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9619404 kB) free!
[1547369830] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12277684 kB) free.
[1547369863] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 447 processes
[1547370163] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 428 processes
[1547370223] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 425 processes
[1547370283] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 2 warn out of 424 processes
[1547370309] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547370343] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 2 warn out of 422 processes
[1547370429] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547371247] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4324584 kB) free!
[1547371363] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547371629] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8917276 kB) free!
[1547371629] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547371749] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11594696 kB) free.
[1547371749] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547371869] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547373069] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547373163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547373284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547373404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547373669] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547373789] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547374149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8514208 kB) free!
[1547374266] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.30, 10.27, 4.29
[1547374269] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11130560 kB) free.
[1547374386] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 47.54, 23.39, 9.83
[1547374389] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547374507] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 18.32, 24.35, 12.13
[1547374509] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547374877] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.2% (22589920 kB) free.
[1547375204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547375470] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8217332 kB) free!
[1547375590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10854056 kB) free.
[1547375707] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.09, 1.53, 4.19
[1547375709] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547375830] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547376404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547376429] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547376524] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547376549] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547376644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547377149] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547377269] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547377389] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547377390] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8060264 kB) free!
[1547377510] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10826440 kB) free.
[1547377989] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547378109] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547378229] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547378830] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547378950] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547379044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547379070] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547379310] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8395940 kB) free!
[1547379430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10893616 kB) free.
[1547379563] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 10.50, 9.40, 9.90
[1547379644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547379670] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547379764] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547379790] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547379884] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547379910] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547380510] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547380630] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547380751] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1547380877] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1547381107] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.61, 20.22, 20.39
[1547381227] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.13, 13.87, 18.04
[1547381231] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8221116 kB) free!
[1547381347] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.92, 9.64, 15.99
[1547381350] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547381351] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10724412 kB) free.
[1547381470] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547382048] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.3% (5554904 kB) free!
[1547382070] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547382190] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547382648] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (2420984 kB) free!
[1547382678] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193688 kB) free!
[1547382790] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547382798] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185844 kB) free!
[1547382884] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547382910] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547382918] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189148 kB) free!
[1547383030] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547383147] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1547383147] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.08, 1.20, 3.28
[1547383151] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8131316 kB) free!
[1547383248] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.4% (5599508 kB) free!
[1547383271] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10694508 kB) free.
[1547383484] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547383604] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547383630] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547383724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547383751] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547383763] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.99, 10.53, 10.14
[1547383848] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1547383848] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (1728668 kB) free!
[1547383871] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547383883] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.07, 10.21, 10.06
[1547384003] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.36, 10.27, 10.10
[1547384131] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1547384251] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1547384371] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1547384448] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.5% (5158300 kB) free!
[1547384471] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547384591] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547385048] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (1959124 kB) free!
[1547385071] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8132528 kB) free!
[1547385191] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547385191] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10683084 kB) free.
[1547385311] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547385571] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1547385648] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5102904 kB) free!
[1547385911] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547386031] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547386151] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547386243] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 435 processes
[1547386248] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (2318052 kB) free!
[1547386391] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7910160 kB) free!
[1547386512] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10561220 kB) free.
[1547386724] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1547386752] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547386872] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1547386872] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547387324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1547387444] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547387472] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547387564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547387592] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547387711] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8670884 kB) free!
[1547387831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11035756 kB) free.
[1547388192] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1547388312] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1547388431] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8001628 kB) free!
[1547388551] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10612044 kB) free.
[1547388764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547388912] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1547389032] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1547389092] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 4 warn out of 337 processes
[1547389364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547389485] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547389633] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547389751] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7933236 kB) free!
[1547389753] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547389871] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10459576 kB) free.
[1547390353] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547390473] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547390593] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547390684] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547390804] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547391071] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8040920 kB) free!
[1547391191] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10666696 kB) free.
[1547391193] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547391313] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547391404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547391524] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547391644] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 449 processes
[1547391704] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 449 processes
[1547391764] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 2 warn out of 448 processes
[1547391792] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8197972 kB) free!
[1547391824] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 2 warn out of 447 processes
[1547391912] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10786080 kB) free.
[1547392124] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547392244] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547392365] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547392512] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7778964 kB) free!
[1547392514] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547392632] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10501164 kB) free.
[1547392634] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547393234] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547393347] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1547393354] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547393565] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547393614] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3177800 kB) free!
[1547393734] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 19.4% (3154168 kB) free!
[1547393832] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8187508 kB) free!
[1547393854] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 19.3% (3141364 kB) free!
[1547393952] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10677928 kB) free.
[1547393954] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547393971] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 17 users currently logged in
[1547394074] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547394091] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 17 users currently logged in
[1547394165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547394211] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 17 users currently logged in
[1547394224] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 446 processes
[1547394285] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547394405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547394547] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 158.56, 34.68, 12.01
[1547394667] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 76.87, 70.91, 30.24
[1547394674] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547394788] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.85, 47.64, 26.64
[1547394794] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547394811] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1547394811] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 13 users currently logged in
[1547394914] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547395388] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.25, 7.00, 14.34
[1547396003] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 10.06, 10.03, 9.99
[1547396114] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547396352] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7885220 kB) free!
[1547396472] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10429816 kB) free.
[1547396588] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.57, 0.78, 4.44
[1547396603] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.87, 10.03, 10.03
[1547396714] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547396723] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.64, 10.15, 10.07
[1547396834] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547396843] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.47, 10.50, 10.20
[1547397648] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547398034] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547398154] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547398754] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547398875] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547399205] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1547399472] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8783376 kB) free!
[1547399474] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547399592] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11268896 kB) free.
[1547399594] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547399843] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.43, 8.80, 9.48
[1547401606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547401726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547401846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547403194] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547403315] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes