[1584211045] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197488 kB) free!
[1584211165] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189968 kB) free!
[1584211285] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193696 kB) free!
[1584211425] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177828 kB) free!
[1584211545] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190572 kB) free!
[1584211607] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15605280 kB) free!
[1584211665] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169684 kB) free!
[1584212207] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396432 kB) free!
[1584212807] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (22004072 kB) free!
[1584212923] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180516 kB) free!
[1584212984] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173032 kB) free!
[1584212988] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584213044] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (186296 kB) free!
[1584213104] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (180724 kB) free!
[1584213404] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 76.8% (18832340 kB) free.
[1584213407] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408148 kB) free!
[1584213915] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584215115] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 24.9% (6138144 kB) free.
[1584215715] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187204 kB) free!
[1584215835] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169344 kB) free!
[1584215955] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178352 kB) free!
[1584215958] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237248 kB) free!
[1584215988] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1584216078] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (248196 kB) free!
[1584216198] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213088 kB) free!
[1584216436] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584216689] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(119) < 92.42.104.166 [15/Mar/2020:01:40:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584216801] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (402660 kB) free!
[1584216809] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584216921] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6085968 kB) free!
[1584217041] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6093428 kB) free!
[1584217850] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1393.84, 300.82, 99.87
[1584217971] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8018.01, 3195.97, 1185.17
[1584217998] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.3% (30695016 kB) free.
[1584218009] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(10) < 49.50.86.152 [15/Mar/2020:02:02:59 +0530] TCP 502 0 0 120.003 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.002, 0.000"
[1584218091] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8910.85, 5132.93, 2138.78
[1584218130] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584218313] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2020/03/15 02:00:56 [error] 12254#0: *70754012 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584218388] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 12409.56, 10409.50, 5511.48
[1584218433] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584218508] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 5385.00, 9158.02, 5697.09
[1584218628] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 727.35, 6128.23, 5006.46
[1584219322] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1584219407] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14766652 kB) free!
[1584219425] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9506804 kB) free!
[1584219443] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1584219545] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4281784 kB) free!
[1584219563] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1584219666] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230680 kB) free!
[1584219744] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3612580 kB) free!
[1584219864] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215744 kB) free!
[1584219984] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242136 kB) free!
[1584220007] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396680 kB) free!
[1584220428] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.20, 16.09, 721.44
[1584220491] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.35, 10.01, 301.06
[1584220584] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 89.8% (44230792 kB) free.
[1584220866] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.0% (41902724 kB) free.
[1584221028] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.51, 3.45, 378.73
[1584221363] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1584221433] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(333) < 2020/03/15 03:00:32 [error] 12253#0: *70771964 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584221485] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (2618124 kB) free!
[1584221553] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2020/03/15 03:01:32 [error] 12257#0: *70771995 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584221673] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584221729] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1954) < 203.212.70.170 [15/Mar/2020:03:02:12 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584221807] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17614332 kB) free!
[1584221849] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584221986] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(277) < 103.16.101.72 [15/Mar/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1584222085] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2292324 kB) free!
[1584222108] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584222408] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412144 kB) free!
[1584222866] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584222987] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584223008] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16589692 kB) free!
[1584223065] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 45.1% (11045656 kB) free.
[1584223608] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448928 kB) free!
[1584224208] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23954976 kB) free!
[1584224808] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397764 kB) free!
[1584224963] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584225083] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1584225203] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1584225408] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20718252 kB) free!
[1584226008] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384288 kB) free!
[1584226379] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1584226499] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 337 processes
[1584226608] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23276264 kB) free!
[1584226619] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1584227208] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2681944 kB) free!
[1584227808] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21803908 kB) free!
[1584228408] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3009108 kB) free!
[1584229008] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19552128 kB) free!
[1584229065] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189744 kB) free!
[1584229185] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177048 kB) free!
[1584229305] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184988 kB) free!
[1584229609] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1890028 kB) free!
[1584230209] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17562100 kB) free!
[1584230355] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 49.2% (12135900 kB) free.
[1584230485] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.4% (20197796 kB) free.
[1584230809] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36265876 kB) free.
[1584231409] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16584696 kB) free!
[1584231529] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (13124252 kB) free!
[1584231649] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (9837220 kB) free!
[1584232249] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20105492 kB) free!
[1584232473] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/03/15 06:02:53 [error] 12255#0: *70827740 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584232492] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 920.35, 538.49, 211.54
[1584232594] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584232612] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 125.15, 360.71, 186.04
[1584232649] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(206) < 203.212.70.170 [15/Mar/2020:06:06:58 +0530] TCP 502 0 0 180.001 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 60.000, 0.000"
[1584232769] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584232849] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36102352 kB) free.
[1584233241] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.9% (30175064 kB) free.
[1584233449] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15663532 kB) free!
[1584233569] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (11931836 kB) free!
[1584233689] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (8286560 kB) free!
[1584233841] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1580424 kB) free!
[1584233961] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378380 kB) free!
[1584234081] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (31755976 kB) free.
[1584234289] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24481888 kB) free!
[1584234681] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3346932 kB) free!
[1584234801] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397644 kB) free!
[1584234889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36617064 kB) free.
[1584234921] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399600 kB) free!
[1584235182] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 89.0% (117263356 kB) free.
[1584235229] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39276032 kB) free.
[1584235277] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.7% (21741608 kB) free.
[1584235489] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19116492 kB) free!
[1584235609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15591800 kB) free!
[1584235729] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (12114388 kB) free!
[1584236330] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24553140 kB) free!
[1584236930] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36627756 kB) free.
[1584236955] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177840 kB) free!
[1584237075] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178040 kB) free!
[1584237195] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186464 kB) free!
[1584237530] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19704036 kB) free!
[1584237650] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16408476 kB) free!
[1584237770] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (13128800 kB) free!
[1584238370] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24069308 kB) free!
[1584238970] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34551632 kB) free.
[1584239570] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16888600 kB) free!
[1584239690] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13336528 kB) free!
[1584239804] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 10.2% (2489228 kB) free!
[1584239810] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12390252 kB) free!
[1584239864] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (903300 kB) free!
[1584239924] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 26.7% (6545304 kB) free.
[1584240196] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.3% (20063200 kB) free.
[1584240224] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.7% (3357864 kB) free!
[1584240236] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.0% (50136008 kB) free.
[1584240284] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3358192 kB) free!
[1584240344] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 13.7% (3363060 kB) free!
[1584240404] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 13.7% (3362832 kB) free!
[1584242646] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13388780 kB) free!
[1584242766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37920036 kB) free.
[1584243366] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17014100 kB) free!
[1584243486] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12566304 kB) free!
[1584243606] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (8302772 kB) free!
[1584243803] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 730 processes
[1584243923] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 730 processes
[1584244043] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 728 processes
[1584244206] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19807408 kB) free!
[1584244408] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584244513] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584244513] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584244513] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584244513] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584244537] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584244538] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584244633] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584244633] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2313 processes
[1584244635] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584244635] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 194881 MB (86% inode=99%):
[1584244659] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 21.92, 23.74, 19.17
[1584244668] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584244807] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409128 kB) free!
[1584244839] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9476704 kB) free!
[1584244959] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9474464 kB) free!
[1584245079] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (9474008 kB) free!
[1584245268] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.5% (30778264 kB) free.
[1584246007] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29202716 kB) free.
[1584246607] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8480588 kB) free!
[1584246727] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4220440 kB) free!
[1584246814] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44690608 kB) free.
[1584247020] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1584247140] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1584247285] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4363424 kB) free!
[1584247307] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187956 kB) free!
[1584247405] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163732 kB) free!
[1584247428] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175252 kB) free!
[1584247525] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186088 kB) free!
[1584247548] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193224 kB) free!
[1584247769] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1584247788] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584247889] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584247908] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584248009] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584248028] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584248609] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1584248618] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12197192 kB) free!
[1584248679] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3724060 kB) free!
[1584248738] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7918956 kB) free!
[1584248858] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3642704 kB) free!
[1584249205] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4533132 kB) free!
[1584249325] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4437308 kB) free!
[1584249445] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (4348036 kB) free!
[1584249458] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17666144 kB) free!
[1584250058] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34533424 kB) free.
[1584250658] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11710700 kB) free!
[1584250778] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7487416 kB) free!
[1584250898] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3372000 kB) free!
[1584251245] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1272812 kB) free!
[1584251498] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20144876 kB) free!
[1584252098] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374660 kB) free!
[1584252698] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14400484 kB) free!
[1584252824] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21672156 kB) free.
[1584252895] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.4% (20584128 kB) free.
[1584252944] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39250768 kB) free.
[1584252996] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.6% (20006272 kB) free.
[1584253298] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33993600 kB) free.
[1584253741] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1584253861] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1584253899] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8220796 kB) free!
[1584254019] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3698644 kB) free!
[1584254139] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387468 kB) free!
[1584254205] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195856 kB) free!
[1584254211] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254227] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254243] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254243] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254243] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254243] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254252] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254278] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254325] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165760 kB) free!
[1584254341] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254357] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254370] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584254370] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 175737 MB (78% inode=99%):
[1584254371] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2322 processes
[1584254371] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584254382] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584254404] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.9% (30482412 kB) free.
[1584254445] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175924 kB) free!
[1584254461] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38778 MB (75% inode=98%):
[1584254487] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254628] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1584254695] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180956 kB) free!
[1584254815] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167736 kB) free!
[1584254867] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254935] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1584254935] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185512 kB) free!
[1584254973] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254973] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254982] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584254991] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 18.90, 21.51, 20.33
[1584255008] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255088] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4646232 kB) free!
[1584255095] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584255095] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2325 processes
[1584255095] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2330 processes
[1584255112] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255138] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255208] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4407000 kB) free!
[1584255232] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584255263] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.7% (30394292 kB) free.
[1584255328] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4209820 kB) free!
[1584255573] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255573] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255645] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1584255645] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 44.0% (10781256 kB) free.
[1584255671] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255700] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584255701] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 174404 MB (77% inode=99%):
[1584255703] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255703] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255705] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255801] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255833] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255833] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255835] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2337 processes
[1584255868] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584255922] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38777 MB (75% inode=98%):
[1584255963] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584255963] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2334 processes
[1584255998] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584256128] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.8% (30443044 kB) free.
[1584256303] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584256303] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584256427] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584256427] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 173645 MB (77% inode=99%):
[1584256531] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584256661] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584256728] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584256791] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584256850] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.7% (30407744 kB) free.
[1584256959] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.9% (20456508 kB) free.
[1584257014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.7% (47095832 kB) free.
[1584257128] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2001168 kB) free!
[1584257294] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.7% (20521120 kB) free.
[1584257392] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38777 MB (75% inode=98%):
[1584257617] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23243996 kB) free!
[1584257737] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (19122208 kB) free!
[1584257843] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1584257857] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35332748 kB) free.
[1584258349] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9726492 kB) free!
[1584258457] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7057988 kB) free!
[1584258469] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9725372 kB) free!
[1584258577] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2327748 kB) free!
[1584258589] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (9725196 kB) free!
[1584258698] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390772 kB) free!
[1584259942] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33726220 kB) free.
[1584260445] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4547440 kB) free!
[1584260565] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4444900 kB) free!
[1584260685] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4323000 kB) free!
[1584261146] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7885752 kB) free!
[1584261266] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31747000 kB) free.
[1584261569] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(167) < 92.42.104.166 [15/Mar/2020:14:05:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584261690] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584261767] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4564224 kB) free!
[1584261866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1241232 kB) free!
[1584261887] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4435852 kB) free!
[1584261986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (388140 kB) free!
[1584262007] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 20.4% (5023128 kB) free.
[1584262107] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (440036 kB) free!
[1584262189] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4511268 kB) free!
[1584262485] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2324664 kB) free!
[1584262607] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.5% (3820492 kB) free!
[1584262707] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (21989620 kB) free!
[1584262727] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.7% (2634148 kB) free!
[1584262847] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1975600 kB) free!
[1584262861] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1584262981] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1584263029] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.2% (3961548 kB) free!
[1584263149] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.3% (3739016 kB) free!
[1584263269] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.3% (3510788 kB) free!
[1584263307] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394936 kB) free!
[1584263907] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17940388 kB) free!
[1584264066] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9249224 kB) free!
[1584264186] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9101076 kB) free!
[1584264306] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8927848 kB) free!
[1584264460] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4163008 kB) free!
[1584264469] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1793892 kB) free!
[1584264507] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396708 kB) free!
[1584264580] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (3988312 kB) free!
[1584264601] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.5% (20847720 kB) free.
[1584264658] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584264658] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.8% (21763912 kB) free.
[1584264690] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(184) < 92.42.104.166 [15/Mar/2020:15:00:15 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584264700] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.8% (15442376 kB) free.
[1584264721] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33447624 kB) free.
[1584264763] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39246848 kB) free.
[1584264781] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1584264788] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584264810] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584264835] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584264842] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584264901] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1584264913] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.0% (30055280 kB) free.
[1584264963] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584264963] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584264965] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584264972] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584265093] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584265093] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584265095] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584265097] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (4227800 kB) free!
[1584265102] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584265147] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.1% (20613816 kB) free.
[1584265217] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.4% (4029784 kB) free!
[1584265223] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584265223] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584265337] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (3835660 kB) free!
[1584265518] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584265648] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584265657] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (233952 kB) free!
[1584265700] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584265700] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2362 processes
[1584265733] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.7% (14085120 kB) free.
[1584265778] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584265805] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185748 kB) free!
[1584265823] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2362 processes
[1584265925] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185464 kB) free!
[1584265937] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1584265937] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169444 kB) free!
[1584266033] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584266033] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584266045] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181348 kB) free!
[1584266163] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584266163] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584266293] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584266293] SERVICE ALERT: smscapp2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584266310] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584266430] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2371 processes
[1584266526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4376788 kB) free!
[1584266646] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389988 kB) free!
[1584266766] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375272 kB) free!
[1584266997] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584267033] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584267040] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584267127] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584267163] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584267168] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9235992 kB) free!
[1584267170] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584267257] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584267288] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9171596 kB) free!
[1584267293] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584267299] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2384 processes
[1584267307] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4771024 kB) free!
[1584267366] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26200380 kB) free!
[1584267408] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1584267408] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (8701424 kB) free!
[1584267501] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 145444 MB (64% inode=99%):
[1584267501] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1584267585] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.9% (30005280 kB) free.
[1584267861] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 19.53, 21.89, 23.45
[1584267899] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1584267899] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2388 processes
[1584267966] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393692 kB) free!
[1584268111] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584268111] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584268231] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 148057 MB (65% inode=99%):
[1584268234] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1584268461] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 17.95, 37.12, 31.76
[1584268566] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20557372 kB) free!
[1584268581] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 13.78, 29.08, 29.45
[1584268702] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584269166] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400364 kB) free!
[1584269766] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19509084 kB) free!
[1584270366] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401976 kB) free!
[1584270408] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4120592 kB) free!
[1584270966] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19613452 kB) free!
[1584271566] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (426500 kB) free!
[1584271876] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.8% (20190356 kB) free.
[1584271924] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.4% (57206104 kB) free.
[1584271960] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39246140 kB) free.
[1584272008] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.0% (44352780 kB) free.
[1584272117] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.3% (20652380 kB) free.
[1584272184] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9561684 kB) free!
[1584272204] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 90.4% (22147848 kB) free.
[1584272304] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9441164 kB) free!
[1584272424] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9350088 kB) free!
[1584272955] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584273088] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169240 kB) free!
[1584273208] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183272 kB) free!
[1584273328] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187812 kB) free!
[1584273725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (2997604 kB) free!
[1584273765] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235540 kB) free!
[1584273845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395060 kB) free!
[1584273885] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227520 kB) free!
[1584273965] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392572 kB) free!
[1584274005] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230980 kB) free!
[1584274565] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20808184 kB) free!
[1584275165] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399072 kB) free!
[1584275765] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17890964 kB) free!
[1584275791] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584275912] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584275916] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 13.84, 14.47, 13.74
[1584276036] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584276329] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584276365] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387916 kB) free!
[1584276521] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584276556] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173312 kB) free!
[1584276642] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 18.00, 16.32, 14.70
[1584276676] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180460 kB) free!
[1584276796] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190504 kB) free!
[1584276965] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19928260 kB) free!
[1584277220] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.7% (19924280 kB) free.
[1584277258] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.4% (38612116 kB) free.
[1584277298] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.6% (17558524 kB) free.
[1584277428] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1584277491] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584277548] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1584277565] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398152 kB) free!
[1584277795] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584277824] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4818828 kB) free!
[1584277825] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171264 kB) free!
[1584277905] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175056 kB) free!
[1584277916] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.4% (29766924 kB) free.
[1584277945] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178352 kB) free!
[1584278025] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188336 kB) free!
[1584278065] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175908 kB) free!
[1584278145] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181868 kB) free!
[1584278148] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1584278165] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20011968 kB) free!
[1584278268] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1584278402] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584278522] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38776 MB (75% inode=98%):
[1584278765] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377904 kB) free!
[1584279299] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584279365] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19306260 kB) free!
[1584279665] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.6% (4566724 kB) free!
[1584279666] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9522912 kB) free!
[1584279785] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.3% (4490072 kB) free!
[1584279786] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8928952 kB) free!
[1584279905] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.9% (4389112 kB) free!
[1584279906] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (9285772 kB) free!
[1584279965] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409432 kB) free!
[1584280565] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18492072 kB) free!
[1584280658] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.2% (20033152 kB) free.
[1584280688] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 30.5% (15047204 kB) free.
[1584280715] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1584280715] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.2% (21118696 kB) free.
[1584281165] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404100 kB) free!
[1584281766] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19136720 kB) free!
[1584282305] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2345148 kB) free!
[1584282366] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403328 kB) free!
[1584282966] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17164384 kB) free!
[1584283566] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386960 kB) free!
[1584284706] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3739588 kB) free!
[1584285301] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1584285421] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1584285966] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23425884 kB) free!
[1584286566] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379196 kB) free!
[1584287166] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27669180 kB) free.
[1584287728] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (4594208 kB) free!
[1584287766] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (2971136 kB) free!
[1584287848] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (4508220 kB) free!
[1584287869] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584287886] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (381684 kB) free!
[1584287968] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4438088 kB) free!
[1584288006] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (427768 kB) free!
[1584288396] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584288399] HOST ALERT: rq-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.203)
[1584288519] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;HARD;1;(Return code of 255 for service 'rq1-rabbitmqerror' on host 'rq-1' was out of bounds)
[1584288606] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15712752 kB) free!
[1584288614] SERVICE ALERT: rq-1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'rq-1' was out of bounds)
[1584288614] SERVICE ALERT: rq-1;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'rq-1' was out of bounds)
[1584288624] HOST ALERT: rq-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1584288650] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 7695.08, 3309.79, 1251.93
[1584288707] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(134) < 49.50.86.152 [15/Mar/2020:21:40:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1584288734] SERVICE ALERT: rq-1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 400998 MB (99% inode=99%):
[1584288734] SERVICE ALERT: rq-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /rq1vns1 1068147 MB (99% inode=99%):
[1584288770] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8287.31, 4987.59, 2116.30
[1584288810] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3533) < 203.212.70.170 [15/Mar/2020:21:42:52 +0530] TCP 502 0 0 180.000 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1584288828] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584288890] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2820.54, 4471.51, 2309.26
[1584288930] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584288993] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(507) < 2020/03/15 21:41:25 [error] 12249#0: *71206041 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584289113] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584289116] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584289168] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 24.4% (5976052 kB) free.
[1584289206] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394344 kB) free!
[1584289241] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 114.30, 98.17, 56.90
[1584289344] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 361 processes
[1584289364] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 103.62, 98.96, 62.22
[1584289464] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1584289476] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.96, 8.04, 5.16
[1584289484] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 84.54, 94.19, 65.01
[1584289530] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 172.105.59.173 [15/Mar/2020:21:53:44 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1584289597] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.26, 7.96, 5.47
[1584289650] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584289715] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584289717] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.30, 8.07, 5.81
[1584289807] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16613740 kB) free!
[1584289969] SERVICE ALERT: rq-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1584290089] SERVICE ALERT: rq-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 353 processes
[1584290090] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.96, 83.73, 638.88
[1584290209] SERVICE ALERT: rq-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1584290407] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395100 kB) free!
[1584290690] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.49, 13.19, 335.83
[1584290843] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584290963] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1584291563] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584291568] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4857664 kB) free!
[1584291607] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35287324 kB) free.
[1584291683] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1584291688] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4777520 kB) free!
[1584291804] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584291808] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (4720976 kB) free!
[1584291884] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.65, 8.75, 29.48
[1584291906] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1584292207] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11705908 kB) free!
[1584292327] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7450208 kB) free!
[1584292351] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21782016 kB) free.
[1584292380] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27009152 kB) free.
[1584292404] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 60.0% (29581200 kB) free.
[1584292668] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182844 kB) free!
[1584292788] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (203644 kB) free!
[1584292908] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193956 kB) free!
[1584293585] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20609992 kB) free!
[1584293705] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (16957996 kB) free!
[1584293825] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35958852 kB) free.
[1584294203] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584294323] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584294425] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12703796 kB) free!
[1584294545] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8614256 kB) free!
[1584294665] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4715344 kB) free!
[1584295266] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22501284 kB) free!
[1584295866] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2850800 kB) free!
[1584296007] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9660624 kB) free!
[1584296127] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9660660 kB) free!
[1584296247] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (9659184 kB) free!
[1584296466] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18394140 kB) free!
[1584296809] SERVICE ALERT: rq-1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes