[1551724211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551724233] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 338 processes
[1551724284] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (2888360 kB) free!
[1551724304] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47457 MB (21% inode=99%):
[1551724307] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9640624 kB) free!
[1551724331] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551724404] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.0% (9346740 kB) free!
[1551724427] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6785852 kB) free!
[1551724547] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13143796 kB) free.
[1551725044] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.18, 6.86, 5.95
[1551725147] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.1% (6474428 kB) free!
[1551725164] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.43, 6.70, 6.00
[1551725267] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.9% (12255912 kB) free.
[1551725284] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.68, 6.58, 6.03
[1551725433] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 337 processes
[1551725867] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8288056 kB) free!
[1551725987] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.5% (5662056 kB) free!
[1551726107] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.6% (11628188 kB) free.
[1551726204] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10184864 kB) free.
[1551726484] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.42, 3.93, 4.87
[1551726704] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42066 MB (18% inode=99%):
[1551726707] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8965736 kB) free!
[1551726732] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551726804] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4272476 kB) free!
[1551726824] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40018 MB (17% inode=99%):
[1551726827] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.2% (6528172 kB) free!
[1551726852] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551726924] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10185256 kB) free.
[1551726944] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47414 MB (21% inode=99%):
[1551726947] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.6% (12129708 kB) free.
[1551726972] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551727524] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5187580 kB) free!
[1551727544] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39222 MB (17% inode=99%):
[1551727547] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8893928 kB) free!
[1551727572] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551727644] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3503956 kB) free!
[1551727664] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1551727664] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47414 MB (21% inode=99%):
[1551727667] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.6% (6222964 kB) free!
[1551727692] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551727764] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10199444 kB) free.
[1551727788] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.6% (12132912 kB) free.
[1551728264] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43318 MB (19% inode=99%):
[1551728292] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551728364] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4776980 kB) free!
[1551728384] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39222 MB (17% inode=99%):
[1551728412] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551728468] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1551728484] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10189628 kB) free.
[1551728504] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47414 MB (21% inode=99%):
[1551728532] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551728987] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8628464 kB) free!
[1551729084] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5590096 kB) free!
[1551729104] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39222 MB (17% inode=99%):
[1551729107] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5948856 kB) free!
[1551729133] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551729204] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4309740 kB) free!
[1551729224] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39222 MB (17% inode=99%):
[1551729227] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.5% (11576012 kB) free.
[1551729253] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551729324] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10191984 kB) free.
[1551729344] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47414 MB (21% inode=99%):
[1551729373] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551729827] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.1% (7457204 kB) free!
[1551729925] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5171916 kB) free!
[1551729944] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39222 MB (17% inode=99%):
[1551729947] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.6% (12137152 kB) free.
[1551729973] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551730045] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (3454848 kB) free!
[1551730064] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47414 MB (21% inode=99%):
[1551730093] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551730165] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10193916 kB) free.
[1551730547] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.2% (7500324 kB) free!
[1551730664] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39206 MB (17% inode=99%):
[1551730667] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.6% (12138164 kB) free.
[1551730684] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.34, 5.29, 5.02
[1551730693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551730765] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4200984 kB) free!
[1551730785] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39190 MB (17% inode=99%):
[1551730804] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.81, 5.17, 5.02
[1551730813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551730885] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10141464 kB) free.
[1551730905] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47350 MB (21% inode=99%):
[1551730924] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.37, 4.91, 4.94
[1551730933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551731267] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8652404 kB) free!
[1551731387] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5979240 kB) free!
[1551731505] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47018 MB (20% inode=99%):
[1551731507] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11713836 kB) free.
[1551731625] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43178 MB (19% inode=99%):
[1551731745] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43178 MB (19% inode=99%):
[1551732085] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9542940 kB) free!
[1551732107] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.4% (6117104 kB) free!
[1551732205] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7391596 kB) free!
[1551732228] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3561512 kB) free!
[1551732325] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5915152 kB) free!
[1551732348] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10308908 kB) free.
[1551732888] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2446328 kB) free!
[1551732948] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.9% (6335956 kB) free!
[1551733007] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1551733068] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.6% (11630904 kB) free.
[1551733668] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7551516 kB) free!
[1551733788] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4935368 kB) free!
[1551733908] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10524292 kB) free.
[1551733933] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551734508] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.2% (5539388 kB) free!
[1551734533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551734628] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.6% (11132252 kB) free.
[1551734653] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551734725] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9968364 kB) free.
[1551734745] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47250 MB (21% inode=99%):
[1551735253] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551735325] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3413696 kB) free!
[1551735345] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39058 MB (17% inode=99%):
[1551735373] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551735445] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9963720 kB) free.
[1551735465] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47250 MB (21% inode=99%):
[1551735724] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.75, 5.80, 5.26
[1551735844] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.72, 5.41, 5.18
[1551735964] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.04, 4.92, 5.03
[1551735974] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551736045] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4747060 kB) free!
[1551736065] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38950 MB (17% inode=99%):
[1551736094] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551736165] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2954804 kB) free!
[1551736185] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47142 MB (20% inode=99%):
[1551736214] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551736285] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9968728 kB) free.
[1551736305] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47142 MB (20% inode=99%):
[1551736564] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.75, 3.87, 4.40
[1551736814] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551736885] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4520876 kB) free!
[1551736934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551737005] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9964864 kB) free.
[1551737054] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551737605] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5518500 kB) free!
[1551737654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551737725] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4282584 kB) free!
[1551737774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551737845] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9971284 kB) free.
[1551737895] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551738446] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.7% (5293284 kB) free!
[1551738494] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551738566] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (4025188 kB) free!
[1551738614] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551738686] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9960800 kB) free.
[1551738705] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551738705] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47270 MB (21% inode=99%):
[1551738734] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551739286] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4988496 kB) free!
[1551739305] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38950 MB (17% inode=99%):
[1551739334] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551739406] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3401588 kB) free!
[1551739425] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 38950 MB (17% inode=99%):
[1551739454] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551739526] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9965024 kB) free.
[1551739545] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47142 MB (20% inode=99%):
[1551740054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551740126] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4655176 kB) free!
[1551740174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551740246] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9962064 kB) free.
[1551740294] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551740764] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.28, 5.37, 5.05
[1551740846] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5497596 kB) free!
[1551740884] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.63, 5.06, 4.97
[1551740895] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551740966] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (4243188 kB) free!
[1551741015] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551741086] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9965344 kB) free.
[1551741135] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551741345] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47237 MB (21% inode=99%):
[1551741687] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5098536 kB) free!
[1551741735] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551741807] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3548256 kB) free!
[1551741855] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551741927] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9955632 kB) free.
[1551741945] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47141 MB (20% inode=99%):
[1551742065] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47141 MB (20% inode=99%):
[1551742185] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46117 MB (20% inode=99%):
[1551742455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551742527] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4679624 kB) free!
[1551742575] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551742647] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9969240 kB) free.
[1551742695] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551743247] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5602136 kB) free!
[1551743296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551743367] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4350080 kB) free!
[1551743416] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551743433] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 336 processes
[1551743487] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9959600 kB) free.
[1551743536] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551744087] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5222444 kB) free!
[1551744136] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551744207] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (3814076 kB) free!
[1551744256] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551744327] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9963852 kB) free.
[1551744856] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551744927] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4853308 kB) free!
[1551744976] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551745048] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (3105860 kB) free!
[1551745096] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551745168] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9960244 kB) free.
[1551745185] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47269 MB (21% inode=99%):
[1551745697] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551745768] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4506824 kB) free!
[1551745785] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39077 MB (17% inode=99%):
[1551745817] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551745888] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9962340 kB) free.
[1551745905] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47269 MB (21% inode=99%):
[1551745937] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1551745937] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551746488] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5364304 kB) free!
[1551746505] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39077 MB (17% inode=99%):
[1551746537] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551746608] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (4074680 kB) free!
[1551746625] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39077 MB (17% inode=99%):
[1551746628] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.2% (7510140 kB) free!
[1551746657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551746728] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9955380 kB) free.
[1551746745] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47268 MB (21% inode=99%):
[1551746748] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.4% (5635936 kB) free!
[1551746777] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551746868] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11236064 kB) free.
[1551747328] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4957912 kB) free!
[1551747345] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39076 MB (17% inode=99%):
[1551747377] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551747448] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (3313764 kB) free!
[1551747465] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39076 MB (17% inode=99%):
[1551747468] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.8% (5808484 kB) free!
[1551747497] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551747568] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1551747568] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9964392 kB) free.
[1551747585] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47267 MB (21% inode=99%):
[1551747588] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.6% (11151600 kB) free.
[1551748084] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551748098] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551748168] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4581620 kB) free!
[1551748185] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41507 MB (18% inode=99%):
[1551748188] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5764268 kB) free!
[1551748218] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551748289] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9954124 kB) free.
[1551748305] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47267 MB (21% inode=99%):
[1551748308] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11254188 kB) free.
[1551748338] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551748888] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5485892 kB) free!
[1551748905] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39075 MB (17% inode=99%):
[1551748908] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.9% (6380552 kB) free!
[1551748938] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551749008] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (4224972 kB) free!
[1551749026] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39075 MB (17% inode=99%):
[1551749029] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.6% (11640868 kB) free.
[1551749058] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551749128] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9957888 kB) free.
[1551749146] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47267 MB (21% inode=99%):
[1551749178] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551749628] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.3% (6551452 kB) free!
[1551749728] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5103184 kB) free!
[1551749746] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39075 MB (17% inode=99%):
[1551749748] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11940988 kB) free.
[1551749778] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551749848] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3564888 kB) free!
[1551749866] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39075 MB (17% inode=99%):
[1551749898] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551749968] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9958684 kB) free.
[1551749986] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47267 MB (21% inode=99%):
[1551750348] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.6% (5741012 kB) free!
[1551750468] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.6% (11161384 kB) free.
[1551750498] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551750568] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4682544 kB) free!
[1551750586] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41507 MB (18% inode=99%):
[1551750619] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551750688] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9959500 kB) free.
[1551750706] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47267 MB (21% inode=99%):
[1551750739] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551751068] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.7% (6727228 kB) free!
[1551751188] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.5% (12066472 kB) free.
[1551751288] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5563780 kB) free!
[1551751306] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39075 MB (17% inode=99%):
[1551751339] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551751408] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4296552 kB) free!
[1551751426] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39073 MB (17% inode=99%):
[1551751459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551751528] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9951252 kB) free.
[1551751546] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47265 MB (21% inode=99%):
[1551751579] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551751788] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7866944 kB) free!
[1551751908] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.5% (12079472 kB) free.
[1551752129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5093164 kB) free!
[1551752146] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39071 MB (17% inode=99%):
[1551752179] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551752249] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3536384 kB) free!
[1551752266] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39071 MB (17% inode=99%):
[1551752299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551752369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9963328 kB) free.
[1551752386] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47263 MB (21% inode=99%):
[1551752508] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.3% (6061272 kB) free!
[1551752628] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 32.3% (15941956 kB) free.
[1551752899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551752969] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4656204 kB) free!
[1551752986] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39071 MB (17% inode=99%):
[1551753020] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551753089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9951368 kB) free.
[1551753106] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47263 MB (21% inode=99%):
[1551753140] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551753668] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 45540 MB (21% inode=99%):
[1551753689] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5528132 kB) free!
[1551753706] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39071 MB (17% inode=99%):
[1551753740] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551753809] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4272460 kB) free!
[1551753826] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39071 MB (17% inode=99%):
[1551753860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551753929] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9955412 kB) free.
[1551753946] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47263 MB (21% inode=99%):
[1551753980] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551754529] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5137208 kB) free!
[1551754546] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39070 MB (17% inode=99%):
[1551754580] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551754649] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (3659512 kB) free!
[1551754666] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39070 MB (17% inode=99%):
[1551754700] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551754769] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9954004 kB) free.
[1551754786] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47262 MB (21% inode=99%):
[1551754868] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39446 MB (18% inode=99%):
[1551754988] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46622 MB (22% inode=99%):
[1551755300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551755369] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4727396 kB) free!
[1551755387] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39070 MB (17% inode=99%):
[1551755421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551755490] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 18.5% (9109776 kB) free!
[1551755507] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47262 MB (21% inode=99%):
[1551755541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551755610] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9958900 kB) free.
[1551755884] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.76, 5.31, 5.18
[1551756004] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.87, 5.21, 5.16
[1551756107] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39068 MB (17% inode=99%):
[1551756124] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.25, 4.91, 5.05
[1551756141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551756210] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4310896 kB) free!
[1551756227] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39070 MB (17% inode=99%):
[1551756261] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551756330] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9946764 kB) free.
[1551756347] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47262 MB (21% inode=99%):
[1551756381] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551756724] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.06, 4.26, 4.65
[1551756930] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9627080 kB) free!
[1551756947] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46749 MB (20% inode=99%):
[1551757050] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7660084 kB) free!
[1551757067] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43165 MB (19% inode=99%):
[1551757170] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6280816 kB) free!
[1551757187] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43165 MB (19% inode=99%):
[1551757988] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42651 MB (20% inode=99%):
[1551758108] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48763 MB (23% inode=99%):
[1551759570] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4906440 kB) free!
[1551760324] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.47, 6.19, 5.60
[1551760444] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.64, 6.29, 5.70
[1551760508] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40315 MB (19% inode=99%):
[1551760564] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.71, 6.43, 5.82
[1551760628] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47599 MB (22% inode=99%):
[1551761228] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40414 MB (19% inode=99%):
[1551761348] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46434 MB (22% inode=99%):
[1551761387] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47201 MB (21% inode=99%):
[1551761764] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1551761764] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.99, 4.64, 5.00
[1551761948] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41753 MB (20% inode=99%):
[1551761970] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9877896 kB) free!
[1551762068] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46217 MB (22% inode=99%):
[1551762174] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2357716 kB) free!
[1551762587] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46668 MB (20% inode=99%):
[1551762707] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44748 MB (19% inode=99%):
[1551762827] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42861 MB (19% inode=99%):
[1551763269] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42257 MB (20% inode=99%):
[1551763389] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46252 MB (22% inode=99%):
[1551763989] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40381 MB (19% inode=99%):
[1551764028] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4804696 kB) free!
[1551764109] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44457 MB (21% inode=99%):
[1551764148] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10197504 kB) free.
[1551764164] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.28, 5.15, 5.02
[1551764284] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.15, 5.13, 5.03
[1551764404] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.58, 4.96, 4.98
[1551764709] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39635 MB (18% inode=99%):
[1551764748] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227272 kB) free!
[1551764829] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44183 MB (21% inode=99%):
[1551764868] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.9% (5865848 kB) free!
[1551764988] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1554636 kB) free!
[1551765004] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.04, 5.18, 5.11
[1551765124] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.51, 5.32, 5.17
[1551765244] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.67, 5.51, 5.26
[1551765429] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39885 MB (19% inode=99%):
[1551765549] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 35532 MB (17% inode=99%):
[1551765570] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1551765570] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222992 kB) free!
[1551765669] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43201 MB (20% inode=99%):
[1551766444] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.57, 4.40, 4.80
[1551767388] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.3% (5547360 kB) free!
[1551767749] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48337 MB (21% inode=99%):
[1551767807] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.4% (2333972 kB) free!
[1551767989] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4433524 kB) free!
[1551768408] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 40.8% (6638364 kB) free.
[1551768589] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1551768589] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.3% (6569232 kB) free!
[1551768844] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.71, 5.85, 5.12
[1551768964] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.49, 6.02, 5.28
[1551769084] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.74, 6.02, 5.38
[1551769170] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.6% (5714364 kB) free!
[1551769189] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1775516 kB) free!
[1551769608] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (404960 kB) free!
[1551769728] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 29.7% (4824704 kB) free.
[1551769770] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4656012 kB) free!
[1551771893] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 633 processes
[1551771953] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 633 processes
[1551772876] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.5% (1868296 kB) free!
[1551773328] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2848376 kB) free!
[1551773448] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2537436 kB) free!
[1551773476] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1216536 kB) free!
[1551773568] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.6% (2216536 kB) free!
[1551773627] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1551773636] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 19.3% (1418092 kB) free!
[1551773756] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 18.7% (1376908 kB) free!
[1551773876] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 18.2% (1337608 kB) free!
[1551774169] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1551774169] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (623348 kB) free!
[1551774769] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.9% (2104096 kB) free!
[1551775369] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1444340 kB) free!
[1551776877] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (688636 kB) free!
[1551776969] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (3074868 kB) free!
[1551777082] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 19.63, 191.43, 141.02
[1551777202] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.88, 128.19, 123.96
[1551777322] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.44, 85.80, 108.94
[1551778170] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.2% (8497856 kB) free!
[1551778336] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (1414008 kB) free!
[1551778369] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.7% (1736172 kB) free!
[1551778456] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.8% (1384124 kB) free!
[1551778576] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 18.4% (1353672 kB) free!
[1551778684] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1551778969] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1284680 kB) free!
[1551779122] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 0.60, 15.94
[1551779989] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.9% (10799752 kB) free.
[1551780069] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44030 MB (21% inode=99%):
[1551780076] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.1% (13512352 kB) free.
[1551780322] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.34, 4.63
[1551780589] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.4% (6092652 kB) free!
[1551780669] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 35197 MB (16% inode=99%):
[1551780709] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (3060092 kB) free!
[1551780789] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43318 MB (20% inode=99%):
[1551780829] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.8% (9747380 kB) free!
[1551780909] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41270 MB (19% inode=99%):
[1551781429] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3730964 kB) free!
[1551781577] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (735160 kB) free!
[1551781739] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 81.29, 30.04, 11.15
[1551781859] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.89, 20.42, 9.91
[1551781970] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 11.2% (1817000 kB) free!
[1551781979] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.93, 13.79, 8.76
[1551782029] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.6% (8696120 kB) free!
[1551782570] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1123624 kB) free!
[1551782579] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.40, 2.39, 4.91
[1551782629] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4283632 kB) free!
[1551783229] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.7% (8225068 kB) free!
[1551783391] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.75, 93.40, 75.27
[1551783511] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.17, 62.80, 66.26
[1551783631] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.63, 42.19, 58.29
[1551784170] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1551784429] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (603772 kB) free!
[1551784831] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.96, 1.48, 16.54
[1551785029] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.5% (5185056 kB) free!
[1551785370] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2341040 kB) free!
[1551785630] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242516 kB) free!
[1551786230] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.7% (5256796 kB) free!
[1551786631] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.27, 0.91, 3.09
[1551786830] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (2625928 kB) free!
[1551787084] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.54, 4.76, 4.96
[1551787581] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551787770] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2653836 kB) free!
[1551787831] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.14, 29.78, 20.25
[1551787951] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.75, 20.08, 17.85
[1551788071] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1551788071] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.48, 13.57, 15.73
[1551788122] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.38, 11.61, 12.28
[1551788181] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551788242] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.34, 7.90, 10.84
[1551788250] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (155068 kB) free!
[1551788261] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 320 processes
[1551788301] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551788362] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.06, 5.30, 9.53
[1551788370] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (149200 kB) free!
[1551788422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551788490] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166816 kB) free!
[1551789090] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.4% (14526104 kB) free.
[1551789171] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.7% (2881216 kB) free!
[1551789272] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.74, 4.73
[1551789562] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1551789562] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.48, 2.89
[1551790084] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.07, 6.21, 5.49
[1551790204] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.27, 6.12, 5.54
[1551790324] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.34, 6.22, 5.65
[1551791370] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461192 kB) free!
[1551792124] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.47, 4.45, 4.81
[1551792774] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4308112 kB) free!
[1551793170] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4930652 kB) free!
[1551793971] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1559032 kB) free!
[1551794261] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 322 processes
[1551795124] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.27, 7.62, 6.15
[1551795244] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.80, 8.01, 6.47
[1551795364] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1551795364] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.81, 8.03, 6.67
[1551796430] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.7% (5763348 kB) free!
[1551797474] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 19.6% (1442936 kB) free!
[1551797594] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.2% (1415548 kB) free!
[1551797630] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4884616 kB) free!
[1551797714] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 19.1% (1407108 kB) free!
[1551798230] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11687312 kB) free.
[1551798364] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.94, 4.23, 4.97
[1551798830] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.8% (7283652 kB) free!
[1551798950] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4485020 kB) free!
[1551798964] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.17, 5.54, 5.17
[1551799070] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.1% (11390664 kB) free.
[1551799085] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.76, 5.56, 5.22
[1551799205] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.39, 5.51, 5.25
[1551799670] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7926312 kB) free!
[1551799790] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.3% (5085320 kB) free!
[1551799876] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3116888 kB) free!
[1551799910] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2193260 kB) free!
[1551799996] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (3070992 kB) free!
[1551800072] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1551800117] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (3019584 kB) free!
[1551800362] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1551800510] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.6% (7665068 kB) free!
[1551801110] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4228920 kB) free!
[1551801171] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1551801914] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (691396 kB) free!
[1551802311] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.7% (9206288 kB) free!
[1551803511] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3101640 kB) free!
[1551804005] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.74, 4.41, 4.71
[1551804111] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.6% (9638340 kB) free!
[1551805311] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3259248 kB) free!
[1551805805] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.34, 5.89, 5.39
[1551805911] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11426012 kB) free.
[1551805925] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.73, 5.43, 5.28
[1551806045] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.97, 4.89, 5.10
[1551806422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551806511] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9586664 kB) free!
[1551806574] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461776 kB) free!
[1551806631] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.9% (7350504 kB) free!
[1551806645] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.43, 3.79, 4.45
[1551806751] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4874428 kB) free!
[1551807022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551807142] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551807951] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9847160 kB) free!
[1551808170] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5907672 kB) free!
[1551808342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551808462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551808553] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 581 processes
[1551808582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551808613] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 579 processes
[1551808674] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 2 warn out of 574 processes
[1551808734] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 2 warn out of 578 processes
[1551809034] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 571 processes
[1551809045] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.80, 6.11, 5.14
[1551809117] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1624192 kB) free!
[1551809151] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (1000600 kB) free!
[1551809165] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.19, 6.49, 5.40
[1551809285] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.77, 6.53, 5.54
[1551809752] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.5% (8134596 kB) free!
[1551810352] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4381524 kB) free!
[1551810382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551810502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes