[1562610605] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12192068 kB) free.
[1562610948] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.81, 5.32, 5.06
[1562611068] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.42, 5.35, 5.10
[1562611115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562611188] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.36, 5.00, 5.01
[1562611235] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562611324] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (3645564 kB) free!
[1562611445] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3623728 kB) free!
[1562611516] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (208768 kB) free!
[1562611565] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (3594584 kB) free!
[1562611636] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201656 kB) free!
[1562611756] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191292 kB) free!
[1562611788] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.22, 4.94, 4.93
[1562611805] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8806236 kB) free!
[1562611836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562611926] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12204908 kB) free.
[1562611956] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562611970] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (1884116 kB) free!
[1562612525] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8649032 kB) free!
[1562612556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562612645] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12211720 kB) free.
[1562612676] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562612831] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (5113540 kB) free!
[1562613245] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8389908 kB) free!
[1562613276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562613365] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12198940 kB) free.
[1562613396] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1562613396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562613651] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.9% (19972088 kB) free.
[1562613665] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.7% (21249224 kB) free.
[1562613675] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1562613675] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 68.4% (33689892 kB) free.
[1562613686] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.7% (21987368 kB) free.
[1562613706] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 22.1% (3597108 kB) free.
[1562613710] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (2875984 kB) free!
[1562613715] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.4% (2020508 kB) free!
[1562613721] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.0% (1956880 kB) free!
[1562613732] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (1476880 kB) free!
[1562613744] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 68.0% (11054096 kB) free.
[1562613819] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 75.3% (12234408 kB) free.
[1562613965] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8337416 kB) free!
[1562613996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562614085] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12208224 kB) free.
[1562614116] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562614275] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.1% (5982960 kB) free!
[1562614388] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 570 processes
[1562614395] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2932264 kB) free!
[1562614429] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (163332 kB) free!
[1562614448] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 568 processes
[1562614497] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194564 kB) free!
[1562614515] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2885340 kB) free!
[1562614549] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173448 kB) free!
[1562614557] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188520 kB) free!
[1562614617] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (188840 kB) free!
[1562614669] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164308 kB) free!
[1562614677] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (192212 kB) free!
[1562614685] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8382916 kB) free!
[1562614717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1562614757] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1562614805] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12195872 kB) free.
[1562614837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562614855] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189020 kB) free!
[1562614899] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1562614975] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174316 kB) free!
[1562615095] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196788 kB) free!
[1562615388] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.07, 5.14, 5.02
[1562615405] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8233448 kB) free!
[1562615437] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562615508] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.36, 4.89, 4.94
[1562615525] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12206700 kB) free.
[1562615557] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562616125] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8122560 kB) free!
[1562616157] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562616246] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12205936 kB) free.
[1562616277] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562616557] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1562616677] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1562616748] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (201684 kB) free!
[1562616797] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1562616845] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8336408 kB) free!
[1562616868] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (194836 kB) free!
[1562616877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562616965] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12199564 kB) free.
[1562616988] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1562616988] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (283300 kB) free!
[1562616997] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562617299] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1562617565] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8444868 kB) free!
[1562617598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562617615] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1562617685] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12204680 kB) free.
[1562617718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562617735] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562618071] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.6% (2705192 kB) free!
[1562618191] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (153520 kB) free!
[1562618285] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8301988 kB) free!
[1562618311] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 42.8% (6950768 kB) free.
[1562618318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562618405] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12197512 kB) free.
[1562618438] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562618508] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.85, 5.08, 5.06
[1562618628] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.64, 5.01, 5.04
[1562618748] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.81, 4.94, 5.01
[1562619005] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8041948 kB) free!
[1562619038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562619125] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12203472 kB) free.
[1562619158] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562619348] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.65, 4.88, 4.99
[1562619725] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7689588 kB) free!
[1562619845] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12207796 kB) free.
[1562620358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562620397] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1562620445] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7321988 kB) free!
[1562620478] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562620548] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.12, 5.12, 5.01
[1562620566] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12196776 kB) free.
[1562620668] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.54, 4.92, 4.95
[1562621079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562621199] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562621765] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8575572 kB) free!
[1562621799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562621885] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12203240 kB) free.
[1562621919] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562622485] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8257448 kB) free!
[1562622519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562622605] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12194748 kB) free.
[1562622639] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562623068] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.36, 5.20, 5.02
[1562623069] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.6% (14404320 kB) free.
[1562623077] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 62.0% (15185460 kB) free.
[1562623114] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.0% (22461936 kB) free.
[1562623129] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 70.4% (34669668 kB) free.
[1562623152] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 70.5% (11466832 kB) free.
[1562623188] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.16, 5.10, 5.00
[1562623205] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8045472 kB) free!
[1562623239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562623325] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12199780 kB) free.
[1562623359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562623516] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1562623669] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173968 kB) free!
[1562623737] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2517148 kB) free!
[1562623789] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169528 kB) free!
[1562623857] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (221816 kB) free!
[1562623910] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171200 kB) free!
[1562623925] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7943040 kB) free!
[1562623960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562623977] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (211496 kB) free!
[1562624045] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12206528 kB) free.
[1562624080] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1562624645] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8006780 kB) free!
[1562624680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562624765] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12186776 kB) free.
[1562624800] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562625365] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7978936 kB) free!
[1562625400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562625486] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12197284 kB) free.
[1562625520] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562626085] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7669892 kB) free!
[1562626205] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183596 kB) free.
[1562626720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562626805] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7739756 kB) free!
[1562626840] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562626925] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12196076 kB) free.
[1562627441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562627561] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562628125] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8428136 kB) free!
[1562628161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562628245] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12196076 kB) free.
[1562628281] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562628518] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562628555] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562628845] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8047396 kB) free!
[1562628881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562628965] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12186568 kB) free.
[1562629001] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562629565] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7743496 kB) free!
[1562629686] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12197256 kB) free.
[1562630201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562630285] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7770140 kB) free!
[1562630322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562630405] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11842524 kB) free.
[1562630442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562630988] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562631605] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8460600 kB) free!
[1562631725] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12166680 kB) free.
[1562632242] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562632842] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562632925] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9581872 kB) free!
[1562632962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562633045] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8212820 kB) free!
[1562633082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562633165] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12129880 kB) free.
[1562633355] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171024 kB) free!
[1562633475] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 47.4% (7700552 kB) free.
[1562633935] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1562633988] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.87, 5.39, 5.12
[1562634055] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562634108] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.91, 5.18, 5.08
[1562634177] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562634228] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.58, 4.98, 5.02
[1562634365] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9313152 kB) free!
[1562634486] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7960692 kB) free!
[1562634606] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12111728 kB) free.
[1562634828] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.95, 4.96, 4.97
[1562634848] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 549 processes
[1562634859] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (3054052 kB) free!
[1562634908] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 550 processes
[1562634979] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (3020512 kB) free!
[1562635099] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (2991192 kB) free!
[1562635206] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7964692 kB) free!
[1562635310] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562635326] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12116508 kB) free.
[1562636082] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1562636082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562636455] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1562636526] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8286456 kB) free!
[1562636575] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562636646] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12118292 kB) free.
[1562637246] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8335436 kB) free!
[1562637282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562637366] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12110604 kB) free.
[1562637402] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562637966] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8228408 kB) free!
[1562638003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562638086] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12115348 kB) free.
[1562638123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562638243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562640643] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562641086] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8245540 kB) free!
[1562641099] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1554300 kB) free!
[1562641206] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12101020 kB) free.
[1562641806] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8231008 kB) free!
[1562641844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562641926] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12109284 kB) free.
[1562641964] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562642526] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8371524 kB) free!
[1562642564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562642646] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12111448 kB) free.
[1562642684] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562643246] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8289392 kB) free!
[1562643284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562643366] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12098668 kB) free.
[1562643404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562643967] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8199704 kB) free!
[1562644004] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562644087] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12107496 kB) free.
[1562644125] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562644245] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562644845] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562645287] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8627796 kB) free!
[1562645407] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7253752 kB) free!
[1562645527] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12101940 kB) free.
[1562645575] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562645695] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1562645716] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1562645836] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1562645956] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1562646045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562646127] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8435608 kB) free!
[1562646165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562646228] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.96, 5.27, 5.08
[1562646247] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12096908 kB) free.
[1562646285] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562646348] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.74, 5.44, 5.17
[1562646468] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.77, 5.44, 5.19
[1562646518] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163908 kB) free!
[1562646638] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181592 kB) free!
[1562646758] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202824 kB) free!
[1562646847] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9104852 kB) free!
[1562646885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562646967] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7621256 kB) free!
[1562647005] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562647087] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12102696 kB) free.
[1562647606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562647668] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.31, 4.68, 4.94
[1562647688] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7726484 kB) free!
[1562647726] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562647808] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12089884 kB) free.
[1562648268] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.59, 5.11, 5.01
[1562648326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562648388] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.26, 5.05, 4.99
[1562648400] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.47, 6.64, 5.47
[1562648408] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7452856 kB) free!
[1562648446] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562648520] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.94, 6.02, 5.38
[1562648528] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12099396 kB) free.
[1562648640] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.40, 6.44, 5.60
[1562649046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562649128] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7153528 kB) free!
[1562649166] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562649248] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12093972 kB) free.
[1562649766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562649768] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15073.91, 12250.52, 6279.58
[1562649887] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562649888] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15076.77, 13185.81, 7346.29
[1562650008] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15078.19, 13811.93, 8283.76
[1562650440] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.90, 4.81, 4.95
[1562650448] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8133028 kB) free!
[1562650486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562650568] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12077108 kB) free.
[1562650606] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562650786] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (12896416 kB) free!
[1562650906] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (20546836 kB) free.
[1562651506] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (12771304 kB) free!
[1562651626] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (20447644 kB) free.
[1562651640] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.73, 6.00, 5.55
[1562651760] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.66, 6.21, 5.67
[1562651768] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8497172 kB) free!
[1562651806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562651880] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.16, 7.68, 6.30
[1562651889] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12078840 kB) free.
[1562651926] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562652226] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12489964 kB) free!
[1562652346] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (20086824 kB) free.
[1562652588] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.43, 5.13, 5.06
[1562652708] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.50, 4.95, 5.00
[1562652942] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21782096 kB) free.
[1562652946] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (12617952 kB) free!
[1562653029] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.5% (39181208 kB) free.
[1562653066] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (19622244 kB) free.
[1562653088] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7499836 kB) free!
[1562653100] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.8% (14109876 kB) free.
[1562653168] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.5% (14551804 kB) free.
[1562653208] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12068252 kB) free.
[1562653308] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.71, 4.99, 5.01
[1562653397] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1562653430] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.20, 5.06, 5.03
[1562653495] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562653517] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1562653547] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195460 kB) free!
[1562653550] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.96, 5.04, 5.03
[1562653570] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.50, 5.34, 5.13
[1562653615] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562653638] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1562653666] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12521328 kB) free!
[1562653667] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (237256 kB) free!
[1562653674] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (563228 kB) free!
[1562653690] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.00, 5.33, 5.16
[1562653777] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157456 kB) free!
[1562653786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (19129456 kB) free.
[1562653787] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (237644 kB) free!
[1562653794] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (322420 kB) free!
[1562653810] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.43, 5.35, 5.18
[1562653865] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174712 kB) free!
[1562653895] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186136 kB) free!
[1562653897] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (151264 kB) free!
[1562653915] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.1% (6773836 kB) free.
[1562653985] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (250332 kB) free!
[1562654015] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190360 kB) free!
[1562654017] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175572 kB) free!
[1562654105] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191236 kB) free!
[1562654135] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166192 kB) free!
[1562654150] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1562654150] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.36, 5.02, 4.99
[1562654215] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562654238] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.3% (5544360 kB) free!
[1562654335] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562654358] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.0% (5431376 kB) free!
[1562654386] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (11000668 kB) free!
[1562654410] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.63, 4.39, 4.84
[1562654478] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4139156 kB) free!
[1562654506] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.1% (24428872 kB) free.
[1562655950] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 80.99, 58.40, 29.54
[1562656008] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.31, 3.07, 18.82
[1562656070] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 80.25, 65.65, 35.70
[1562656190] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 43.64, 61.20, 37.99
[1562656726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562656846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562656966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562657390] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.27, 6.57, 14.45
[1562657846] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.0% (18515608 kB) free.
[1562658131] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.7% (21257504 kB) free.
[1562658259] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39447156 kB) free.
[1562658336] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.2% (21134780 kB) free.
[1562658408] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.99, 2.25, 4.08
[1562658410] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.1% (14486512 kB) free.
[1562658426] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.7% (4823372 kB) free!
[1562658546] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.6% (4801168 kB) free!
[1562658666] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.5% (4776148 kB) free!
[1562659018] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (149608 kB) free!
[1562659055] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167688 kB) free!
[1562659110] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (3017636 kB) free!
[1562659138] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166912 kB) free!
[1562659175] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177336 kB) free!
[1562659230] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.3% (2813256 kB) free!
[1562659258] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1562659258] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (150064 kB) free!
[1562659295] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1562659295] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186988 kB) free!
[1562659350] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.0% (2595340 kB) free!
[1562659680] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.59, 4.00, 4.81
[1562659966] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562660408] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8700608 kB) free!
[1562660528] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7310312 kB) free!
[1562660550] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (758412 kB) free!
[1562660648] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12056768 kB) free.
[1562661106] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12682076 kB) free!
[1562661166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562661226] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (20776028 kB) free.
[1562661248] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7910188 kB) free!
[1562661286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562661368] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12056032 kB) free.
[1562661407] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562661714] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.0% (1246940 kB) free!
[1562661834] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.4% (1131868 kB) free!
[1562661955] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.7% (1009612 kB) free!
[1562661968] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8295784 kB) free!
[1562662006] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562662088] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12059288 kB) free.
[1562662126] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562662555] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (378940 kB) free!
[1562662688] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8318780 kB) free!
[1562662726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562662809] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12048640 kB) free.
[1562662846] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562663026] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12148700 kB) free!
[1562663146] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20256932 kB) free.
[1562663408] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8365340 kB) free!
[1562663446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562663528] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12055240 kB) free.
[1562663538] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (2171984 kB) free!
[1562663566] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562663658] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (2137132 kB) free!
[1562663744] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173640 kB) free!
[1562663778] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1759516 kB) free!
[1562663838] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1562663864] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163896 kB) free!
[1562663884] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.0% (22464020 kB) free.
[1562663931] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2591.57, 2426.07, 1085.93
[1562663932] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1562663932] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.7% (21497612 kB) free.
[1562663984] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2008656 kB) free!
[1562663991] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1039.67, 2048.86, 1046.31
[1562664012] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1562664012] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.3% (14030612 kB) free.
[1562664051] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 349.97, 1623.45, 954.29
[1562664099] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1562664099] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.9% (22272652 kB) free.
[1562664128] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8299604 kB) free!
[1562664157] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.6% (14564452 kB) free.
[1562664166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562664171] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 48.16, 1086.56, 838.68
[1562664223] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6833108 kB) free.
[1562664248] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12056220 kB) free.
[1562664287] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562664484] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.1% (3731432 kB) free!
[1562664604] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177896 kB) free!
[1562664724] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172564 kB) free!
[1562664766] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175224 kB) free!
[1562664848] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7198968 kB) free!
[1562664886] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165900 kB) free!
[1562664946] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (12731992 kB) free!
[1562664968] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12046980 kB) free.
[1562665006] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159684 kB) free!
[1562665066] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (20725128 kB) free.
[1562665466] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9301280 kB) free!
[1562665486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562665586] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8535504 kB) free!
[1562665606] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562665706] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (7843932 kB) free!
[1562666168] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8954652 kB) free!
[1562666206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562666288] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7776412 kB) free!
[1562666326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562666408] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12037088 kB) free.
[1562666446] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562666866] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12335644 kB) free!
[1562666880] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.82, 5.91, 5.13
[1562666906] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241736 kB) free!
[1562666986] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (19820228 kB) free.
[1562667000] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.98, 6.03, 5.25
[1562667008] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8380328 kB) free!
[1562667046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562667120] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.40, 6.27, 5.44
[1562667129] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12046008 kB) free.
[1562667167] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562667308] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 557 processes
[1562667368] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 554 processes
[1562667610] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.12, 5.25, 5.07
[1562667730] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.90, 5.19, 5.07
[1562667771] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.75, 0.51, 17.94
[1562667850] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.19, 5.17, 5.08
[1562668191] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.49, 4.84, 16.33
[1562668267] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.96, 6.49, 5.10
[1562668320] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.00, 4.24, 4.80
[1562668328] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8152852 kB) free!
[1562668366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562668387] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.74, 6.55, 5.29
[1562668448] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12048592 kB) free.
[1562668486] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562668507] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.55, 6.23, 5.33
[1562668786] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11841788 kB) free!
[1562668906] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (18887384 kB) free.
[1562669086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562669206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562669326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562669506] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (12940208 kB) free!
[1562669571] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.65, 2.03, 4.07
[1562669626] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.9% (9125028 kB) free!
[1562669650] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.71, 4.13, 4.66
[1562669746] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (17373308 kB) free.
[1562669926] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562670248] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8860132 kB) free!
[1562670346] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (10748468 kB) free!
[1562670368] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7162804 kB) free!
[1562670467] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (18120140 kB) free.
[1562670488] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12036344 kB) free.
[1562670907] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.44, 2.22, 3.84
[1562671219] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3120172 kB) free!
[1562671339] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (2970364 kB) free!
[1562671459] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (2811956 kB) free!
[1562671735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1562671855] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1562671975] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562672050] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.16, 5.88, 5.20
[1562672170] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.29, 5.93, 5.30
[1562672288] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8345956 kB) free!
[1562672290] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.43, 5.51, 5.23
[1562672326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562672409] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12022648 kB) free.
[1562672446] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562672939] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180436 kB) free!
[1562673059] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (290640 kB) free!
[1562673113] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181980 kB) free!
[1562673179] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (236712 kB) free!
[1562673230] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (1090600 kB) free!
[1562673233] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180552 kB) free!
[1562673259] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (984680 kB) free!
[1562673350] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.3% (975512 kB) free!
[1562673353] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181252 kB) free!
[1562673470] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.9% (872716 kB) free!
[1562674070] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (320684 kB) free!
[1562674208] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7211212 kB) free!
[1562674328] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12030360 kB) free.
[1562674847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562674866] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2416648 kB) free!
[1562674928] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7326112 kB) free!
[1562674967] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562675048] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12027064 kB) free.
[1562675567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562675648] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7396460 kB) free!
[1562675687] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562675768] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12018680 kB) free.
[1562675806] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562675871] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6835472 kB) free.
[1562675991] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.79, 4.78, 4.92
[1562676124] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562676287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562676407] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562676591] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 28.55, 21.89, 11.73
[1562676711] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.37, 16.28, 10.90
[1562676831] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.98, 12.41, 10.14
[1562676968] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8574788 kB) free!
[1562677007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562677008] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 37.16, 1003.53, 800.03
[1562677088] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12022016 kB) free.
[1562677128] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562677128] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8.31, 672.75, 703.49
[1562677248] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.18, 451.62, 618.73
[1562677690] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1562677690] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.31, 3.25, 4.20
[1562677975] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1562678095] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562678288] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8283660 kB) free!
[1562678327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562678408] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12012828 kB) free.
[1562678447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562678457] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.8% (21511504 kB) free.
[1562678520] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.5% (20615872 kB) free.
[1562678570] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.4% (21185340 kB) free.
[1562678625] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.6% (37772716 kB) free.
[1562678659] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 31.9% (5181784 kB) free.
[1562678713] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.6% (21949628 kB) free.
[1562678773] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.8% (14270304 kB) free.
[1562679008] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7907196 kB) free!
[1562679047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562679075] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169564 kB) free!
[1562679128] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11871180 kB) free.
[1562679167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562679195] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155400 kB) free!
[1562679225] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224004 kB) free!
[1562679270] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (177748 kB) free!
[1562679287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562679315] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174216 kB) free!
[1562679345] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239020 kB) free!
[1562679378] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176316 kB) free!
[1562679390] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174592 kB) free!
[1562679465] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222952 kB) free!
[1562679498] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167964 kB) free!
[1562679510] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (177204 kB) free!
[1562679618] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169980 kB) free!
[1562679728] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4686096 kB) free!
[1562679832] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.48, 4.63, 4.96
[1562679848] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (786748 kB) free!
[1562679968] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177796 kB) free!
[1562680038] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1562680110] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 79.5% (12913496 kB) free.
[1562680158] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1562680279] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1562680488] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562680848] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.80, 2.37, 14.97
[1562680928] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8475776 kB) free!
[1562681048] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12017444 kB) free.
[1562681290] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.54, 8.27, 6.34
[1562681410] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.27, 8.30, 6.59
[1562681530] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.36, 8.33, 6.81
[1562682288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562682408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562682528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562682648] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.18, 3.14, 4.67
[1562682848] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9401624 kB) free!
[1562682907] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.93, 5.80, 5.22
[1562682969] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7990184 kB) free!
[1562683028] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.61, 5.72, 5.26
[1562683089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12003384 kB) free.
[1562683148] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.74, 5.73, 5.32
[1562683432] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.61, 5.25, 5.08
[1562683552] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.80, 5.08, 5.03
[1562683672] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.95, 5.04, 5.03
[1562683920] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.88, 5.85, 5.08
[1562684040] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.18, 5.33, 4.99
[1562684272] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.78, 4.70, 4.85
[1562684289] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8784028 kB) free!
[1562684409] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7069684 kB) free!
[1562684529] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11346480 kB) free.
[1562684574] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1562684640] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.54, 5.96, 5.31
[1562684760] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.48, 6.04, 5.41
[1562684880] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.20, 6.59, 5.69
[1562684928] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562684952] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.0% (18516508 kB) free.
[1562685008] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39316340 kB) free.
[1562685078] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 84.7% (13770844 kB) free.
[1562685136] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 68.7% (11174072 kB) free.
[1562685313] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1562685472] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.48, 5.13, 5.03
[1562685548] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.86, 2.23, 3.65
[1562685558] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (1706216 kB) free!
[1562685592] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.89, 5.01, 5.00
[1562685678] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181760 kB) free!
[1562685679] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1562685689] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176396 kB) free!
[1562685731] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.36, 3.12, 4.94
[1562685798] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1562685798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195512 kB) free!
[1562685809] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (152344 kB) free!
[1562685929] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1562685929] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (149956 kB) free!
[1562686080] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.68, 4.38, 4.70
[1562686220] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.1% (4981232 kB) free!
[1562686248] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 69.28, 26.34, 11.57
[1562686340] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237576 kB) free!
[1562686369] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 82.77, 45.68, 20.44
[1562686460] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1562686460] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218352 kB) free!
[1562686489] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1562686489] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.06, 32.62, 18.77
[1562686728] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562686848] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562687089] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.43, 7.41, 11.50
[1562687448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562687568] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562688129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8041320 kB) free!
[1562688169] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562688250] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11968876 kB) free.
[1562688289] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.92, 1.19, 4.13
[1562688289] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562688466] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (10395920 kB) free!
[1562688548] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.59, 5.72, 5.33
[1562688587] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (18405252 kB) free.
[1562688668] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.58, 5.36, 5.25
[1562688788] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.94, 4.65, 5.02
[1562688849] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7667408 kB) free!
[1562688889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1562688969] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11960312 kB) free.
[1562689009] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562689110] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3124820 kB) free!
[1562689230] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (3044852 kB) free!
[1562689350] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (2968604 kB) free!
[1562689388] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1562689388] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.42, 1.77, 3.24
[1562689569] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7257668 kB) free!
[1562689689] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11964428 kB) free.
[1562690209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562690289] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8067148 kB) free!
[1562690329] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562690409] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11956588 kB) free.
[1562690929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562690986] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12575624 kB) free!
[1562691049] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562691106] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (9520180 kB) free!
[1562691226] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (17459928 kB) free.
[1562691592] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.23, 5.19, 5.01
[1562691609] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8183636 kB) free!
[1562691649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562691712] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.38, 5.18, 5.02
[1562691729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11964600 kB) free.
[1562691769] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562691832] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.66, 4.94, 4.96
[1562692329] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7932648 kB) free!
[1562692369] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562692426] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10835092 kB) free!
[1562692449] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11963020 kB) free.
[1562692489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562692546] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (18439628 kB) free.
[1562693050] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7653884 kB) free!
[1562693170] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11956808 kB) free.
[1562693531] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1562693689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562693770] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7216292 kB) free!
[1562693810] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562693890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11965568 kB) free.
[1562694346] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (10203996 kB) free!
[1562694385] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1562694409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562694466] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (17451712 kB) free.
[1562694505] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1562694529] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562695066] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (11198160 kB) free!
[1562695090] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8317080 kB) free!
[1562695129] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562695186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (18222728 kB) free.
[1562695210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11953412 kB) free.
[1562695249] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562695350] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562695787] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10893344 kB) free!
[1562695810] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8053828 kB) free!
[1562695849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562695907] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (18057852 kB) free.
[1562695930] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11964808 kB) free.
[1562695969] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562695969] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562696032] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562696060] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562696129] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562696507] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (10229892 kB) free!
[1562696530] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8044800 kB) free!
[1562696569] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562696598] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562696627] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (17449600 kB) free.
[1562696650] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11963256 kB) free.
[1562696689] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562696956] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 369 processes