[1551637900] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43850 MB (19% inode=99%):
[1551638824] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40265 MB (19% inode=99%):
[1551638944] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44264 MB (21% inode=99%):
[1551640196] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551640275] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11143368 kB) free.
[1551640300] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48194 MB (21% inode=99%):
[1551640796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551640916] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551641475] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6209944 kB) free!
[1551641500] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40004 MB (17% inode=99%):
[1551641516] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551641595] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4581068 kB) free!
[1551641620] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551641637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551641715] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11140308 kB) free.
[1551642220] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551642236] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551642315] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5686792 kB) free!
[1551642340] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551642356] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551642435] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11146788 kB) free.
[1551642460] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551642476] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551643035] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6502820 kB) free!
[1551643060] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40004 MB (17% inode=99%):
[1551643076] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551643155] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.3% (5096816 kB) free!
[1551643180] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40004 MB (17% inode=99%):
[1551643196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551643275] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11141092 kB) free.
[1551643300] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48196 MB (21% inode=99%):
[1551643316] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551643744] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42820 MB (20% inode=99%):
[1551643864] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 36676 MB (17% inode=99%):
[1551643875] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5984360 kB) free!
[1551643900] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42116 MB (18% inode=99%):
[1551643916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551643984] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44875 MB (21% inode=99%):
[1551643995] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.6% (5733012 kB) free!
[1551644020] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48196 MB (21% inode=99%):
[1551644036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551644115] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11142160 kB) free.
[1551644620] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40004 MB (17% inode=99%):
[1551644637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551644716] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5496320 kB) free!
[1551644740] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551644757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551644836] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11144092 kB) free.
[1551644860] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48196 MB (21% inode=99%):
[1551644877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551645436] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6389116 kB) free!
[1551645460] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551645477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551645556] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4927128 kB) free!
[1551645581] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551645597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551645676] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11139380 kB) free.
[1551645701] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551646197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551646276] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (5999312 kB) free!
[1551646301] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551646317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551646396] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (4253080 kB) free!
[1551646421] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47683 MB (21% inode=99%):
[1551646437] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551646516] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11137412 kB) free.
[1551647021] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40004 MB (17% inode=99%):
[1551647038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551647116] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5621860 kB) free!
[1551647141] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551647158] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551647236] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11145420 kB) free.
[1551647261] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551647278] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551647584] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42864 MB (20% inode=99%):
[1551647704] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40816 MB (19% inode=99%):
[1551647824] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44918 MB (21% inode=99%):
[1551647836] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6496740 kB) free!
[1551647861] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551647878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551647957] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5135288 kB) free!
[1551647981] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551647998] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551648077] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11139648 kB) free.
[1551648101] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551648118] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551648677] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6199776 kB) free!
[1551648701] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42403 MB (18% inode=99%):
[1551648718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551648797] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (4612520 kB) free!
[1551648821] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551648838] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551648917] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11138068 kB) free.
[1551649421] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44195 MB (19% inode=99%):
[1551649439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551649517] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5902804 kB) free!
[1551649541] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42147 MB (18% inode=99%):
[1551649559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551649637] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11143536 kB) free.
[1551649661] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551649679] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551650237] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6860712 kB) free!
[1551650261] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42179 MB (18% inode=99%):
[1551650279] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551650357] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5592880 kB) free!
[1551650381] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42179 MB (18% inode=99%):
[1551650399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551650477] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11139256 kB) free.
[1551650501] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551650519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551651077] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6496652 kB) free!
[1551651101] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42083 MB (18% inode=99%):
[1551651119] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551651197] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.3% (5091684 kB) free!
[1551651221] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 37987 MB (16% inode=99%):
[1551651239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551651318] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1551651318] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11135012 kB) free.
[1551651341] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551651360] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1551651360] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551651918] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6182016 kB) free!
[1551651942] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551651959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551652024] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42898 MB (20% inode=99%):
[1551652035] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (524268 kB) free!
[1551652038] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4573544 kB) free!
[1551652062] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48194 MB (21% inode=99%):
[1551652079] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551652144] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40850 MB (19% inode=99%):
[1551652158] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11142976 kB) free.
[1551652264] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44950 MB (21% inode=99%):
[1551652662] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43170 MB (19% inode=99%):
[1551652679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551652758] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5856836 kB) free!
[1551652782] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39074 MB (17% inode=99%):
[1551652799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551652878] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11138036 kB) free.
[1551652902] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1551652902] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48194 MB (21% inode=99%):
[1551652919] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551653235] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 66.1% (10748828 kB) free.
[1551653478] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6822820 kB) free!
[1551653502] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40002 MB (17% inode=99%):
[1551653519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551653573] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 36485 MB (16% inode=99%):
[1551653598] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5536692 kB) free!
[1551653622] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40002 MB (17% inode=99%):
[1551653639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551653718] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11138232 kB) free.
[1551653742] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48194 MB (21% inode=99%):
[1551653759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551653984] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 495.46, 333.14, 149.89
[1551654104] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 497.41, 387.59, 192.06
[1551654224] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 497.49, 423.98, 229.10
[1551654318] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6505348 kB) free!
[1551654325] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.2% (2489856 kB) free!
[1551654342] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39490 MB (17% inode=99%):
[1551654360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551654438] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5147668 kB) free!
[1551654462] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39490 MB (17% inode=99%):
[1551654480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551654559] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11147380 kB) free.
[1551654582] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551654600] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551654925] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2444976 kB) free!
[1551655158] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6345872 kB) free!
[1551655182] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40001 MB (17% inode=99%):
[1551655200] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551655264] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42904 MB (20% inode=99%):
[1551655278] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (4855840 kB) free!
[1551655302] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40001 MB (17% inode=99%):
[1551655320] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551655384] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40856 MB (19% inode=99%):
[1551655398] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11148236 kB) free.
[1551655422] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48193 MB (21% inode=99%):
[1551655504] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44957 MB (21% inode=99%):
[1551655920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551655998] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6030084 kB) free!
[1551656022] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40003 MB (17% inode=99%):
[1551656040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551656118] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4270568 kB) free!
[1551656142] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1551656160] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551656238] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11145452 kB) free.
[1551656624] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.76, 1.73, 18.94
[1551656742] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42178 MB (18% inode=99%):
[1551656760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551656838] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5726320 kB) free!
[1551656862] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42178 MB (18% inode=99%):
[1551656881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551656958] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11147628 kB) free.
[1551656982] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48194 MB (21% inode=99%):
[1551657001] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551657558] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6679008 kB) free!
[1551657582] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40002 MB (17% inode=99%):
[1551657601] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551657678] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.0% (5412516 kB) free!
[1551657703] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40643 MB (18% inode=99%):
[1551657721] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551657799] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11145140 kB) free.
[1551657823] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48194 MB (21% inode=99%):
[1551657841] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551658398] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6400040 kB) free!
[1551658423] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40002 MB (17% inode=99%):
[1551658424] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.22, 1.80, 4.12
[1551658441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551658518] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.1% (4953092 kB) free!
[1551658543] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40002 MB (17% inode=99%):
[1551658561] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551658638] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11145932 kB) free.
[1551658663] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48194 MB (21% inode=99%):
[1551659104] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42933 MB (20% inode=99%):
[1551659161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551659224] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44932 MB (21% inode=99%):
[1551659238] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6420548 kB) free!
[1551659263] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40002 MB (17% inode=99%):
[1551659282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551659358] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4918584 kB) free!
[1551659383] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40002 MB (17% inode=99%):
[1551659402] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551659478] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11141568 kB) free.
[1551659503] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48194 MB (21% inode=99%):
[1551659824] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40866 MB (19% inode=99%):
[1551659945] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48806 MB (23% inode=99%):
[1551660002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551660078] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6035368 kB) free!
[1551660103] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42433 MB (18% inode=99%):
[1551660122] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551660198] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4322576 kB) free!
[1551660223] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48193 MB (21% inode=99%):
[1551660242] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551660318] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11142556 kB) free.
[1551660823] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39489 MB (17% inode=99%):
[1551660842] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551660918] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5728336 kB) free!
[1551660943] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39489 MB (17% inode=99%):
[1551660962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551661038] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11145204 kB) free.
[1551661063] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48193 MB (21% inode=99%):
[1551661082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551661639] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6661900 kB) free!
[1551661663] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42241 MB (18% inode=99%):
[1551661683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551661759] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5377100 kB) free!
[1551661783] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42241 MB (18% inode=99%):
[1551661803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551661879] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11142392 kB) free.
[1551661903] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48193 MB (21% inode=99%):
[1551661923] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551662479] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6305484 kB) free!
[1551662503] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40000 MB (17% inode=99%):
[1551662523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551662599] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4759672 kB) free!
[1551662623] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40000 MB (17% inode=99%):
[1551662643] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551662719] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11140736 kB) free.
[1551662743] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48192 MB (21% inode=99%):
[1551663243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551663319] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5913556 kB) free!
[1551663343] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42304 MB (18% inode=99%):
[1551663363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551663439] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11146136 kB) free.
[1551663463] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48192 MB (21% inode=99%):
[1551663483] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551664039] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6836604 kB) free!
[1551664064] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42336 MB (18% inode=99%):
[1551664084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551664159] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5554808 kB) free!
[1551664184] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42336 MB (18% inode=99%):
[1551664204] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551664279] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11138844 kB) free.
[1551664304] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48191 MB (21% inode=99%):
[1551664324] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551664880] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6439540 kB) free!
[1551664904] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39999 MB (17% inode=99%):
[1551664924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551665000] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.2% (5002388 kB) free!
[1551665024] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39999 MB (17% inode=99%):
[1551665044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551665120] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11136692 kB) free.
[1551665144] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48191 MB (21% inode=99%):
[1551665164] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551666320] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6468964 kB) free!
[1551666344] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39999 MB (17% inode=99%):
[1551666364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551666440] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5134280 kB) free!
[1551666464] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39999 MB (17% inode=99%):
[1551666485] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551666560] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11140584 kB) free.
[1551666584] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48190 MB (21% inode=99%):
[1551666605] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1551667160] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6116892 kB) free!
[1551667184] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42781 MB (19% inode=99%):
[1551667205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551667280] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4478948 kB) free!
[1551667304] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48189 MB (21% inode=99%):
[1551667325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551667400] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10444288 kB) free.
[1551667925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551668000] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9820620 kB) free!
[1551668045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551668121] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7974016 kB) free!
[1551668165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551668241] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.5% (6660840 kB) free!
[1551669104] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47158 MB (20% inode=99%):
[1551669224] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44086 MB (19% inode=99%):
[1551669344] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43574 MB (19% inode=99%):
[1551669657] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 436.68, 396.22, 195.63
[1551669777] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 59.44, 265.35, 172.00
[1551669898] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.62, 177.86, 151.26
[1551670641] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4793684 kB) free!
[1551671241] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11122348 kB) free.
[1551672298] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.89, 0.87, 12.08
[1551672365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551672944] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47650 MB (21% inode=99%):
[1551673498] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.59, 0.64, 3.81
[1551673565] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551673685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551673745] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43727 MB (20% inode=99%):
[1551673806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551673865] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46277 MB (22% inode=99%):
[1551674241] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9798992 kB) free!
[1551674362] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7870540 kB) free!
[1551674482] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551674482] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6389572 kB) free!
[1551675944] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47033 MB (20% inode=99%):
[1551676064] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43961 MB (19% inode=99%):
[1551676184] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43961 MB (19% inode=99%):
[1551677465] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 37737 MB (18% inode=99%):
[1551677482] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4580036 kB) free!
[1551677585] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 45010 MB (21% inode=99%):
[1551679282] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10868164 kB) free.
[1551679385] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40186 MB (19% inode=99%):
[1551679505] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44816 MB (21% inode=99%):
[1551680105] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 37633 MB (18% inode=99%):
[1551680225] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44917 MB (21% inode=99%):
[1551680406] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1551681425] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39754 MB (19% inode=99%):
[1551681545] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43710 MB (20% inode=99%):
[1551681584] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47399 MB (21% inode=99%):
[1551681665] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41662 MB (19% inode=99%):
[1551682206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551682282] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9404512 kB) free!
[1551682326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551682402] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7535000 kB) free!
[1551682446] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1551682446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551682522] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6173440 kB) free!
[1551682865] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47505 MB (22% inode=99%):
[1551683033] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 19.8% (1452916 kB) free!
[1551683153] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 19.7% (1447012 kB) free!
[1551683273] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 19.6% (1439176 kB) free!
[1551683984] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46882 MB (20% inode=99%):
[1551684048] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4854016 kB) free!
[1551684104] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43298 MB (19% inode=99%):
[1551684168] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4819548 kB) free!
[1551684224] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1551684224] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43296 MB (19% inode=99%):
[1551684288] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4805188 kB) free!
[1551685522] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4661424 kB) free!
[1551687322] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10665464 kB) free.
[1551687665] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40164 MB (19% inode=99%):
[1551687785] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48088 MB (23% inode=99%):
[1551688385] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40140 MB (19% inode=99%):
[1551688447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551688505] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47285 MB (22% inode=99%):
[1551689024] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47787 MB (21% inode=99%):
[1551689122] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9690672 kB) free!
[1551689235] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.0% (2918176 kB) free!
[1551689242] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7572176 kB) free!
[1551689355] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (2862248 kB) free!
[1551689362] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6189004 kB) free!
[1551689476] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (2805556 kB) free!
[1551689705] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43117 MB (20% inode=99%):
[1551689825] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47201 MB (22% inode=99%):
[1551690247] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551690367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551690425] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38995 MB (18% inode=99%):
[1551690483] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.33, 5.40, 5.13
[1551690488] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551690545] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47044 MB (22% inode=99%):
[1551690603] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.41, 5.40, 5.16
[1551690723] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.22, 5.65, 5.28
[1551691145] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38970 MB (18% inode=99%):
[1551691265] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46126 MB (22% inode=99%):
[1551691323] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.77, 4.20, 4.73
[1551691424] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46628 MB (20% inode=99%):
[1551691544] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43556 MB (19% inode=99%):
[1551691664] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43556 MB (19% inode=99%):
[1551691865] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38947 MB (18% inode=99%):
[1551691985] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46746 MB (22% inode=99%):
[1551692363] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4752716 kB) free!
[1551692585] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38925 MB (18% inode=99%):
[1551692705] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44548 MB (21% inode=99%):
[1551693076] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1547448 kB) free!
[1551693305] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36858 MB (17% inode=99%):
[1551693425] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44015 MB (21% inode=99%):
[1551694026] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36837 MB (17% inode=99%):
[1551694146] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 43996 MB (21% inode=99%):
[1551694673] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (724208 kB) free!
[1551695346] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41986 MB (20% inode=99%):
[1551695466] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 35842 MB (17% inode=99%):
[1551695567] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 25.2% (4093912 kB) free.
[1551695586] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44090 MB (21% inode=99%):
[1551695963] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10494892 kB) free.
[1551695989] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54254 MB (24% inode=99%):
[1551696186] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39443 MB (18% inode=99%):
[1551696306] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44075 MB (21% inode=99%):
[1551696906] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 34051 MB (16% inode=99%):
[1551697026] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41211 MB (19% inode=99%):
[1551697146] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 38139 MB (18% inode=99%):
[1551697688] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551698264] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47472 MB (21% inode=99%):
[1551698860] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 321 processes
[1551698864] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39280 MB (17% inode=99%):
[1551698984] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46447 MB (20% inode=99%):
[1551699104] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43375 MB (19% inode=99%):
[1551699150] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.56, 9.98, 5.87
[1551699271] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.46, 6.74, 5.19
[1551699391] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.45, 4.66, 4.61
[1551699563] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9843056 kB) free!
[1551699683] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7685744 kB) free!
[1551699803] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6352752 kB) free!
[1551699826] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9196924 kB) free!
[1551699946] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.3% (14420032 kB) free.
[1551700324] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.08, 5.90, 5.06
[1551700444] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.85, 5.89, 5.15
[1551700546] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9871180 kB) free!
[1551700564] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.91, 5.87, 5.23
[1551700666] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.7% (15143576 kB) free.
[1551700688] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551700808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551700928] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551701164] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1551701164] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.44, 4.09, 4.64
[1551702461] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 320 processes
[1551702466] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8209488 kB) free!
[1551702586] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.9% (13734776 kB) free.
[1551702803] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4732600 kB) free!
[1551703187] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8292340 kB) free!
[1551703307] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.3% (12463148 kB) free.
[1551704346] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46154 MB (22% inode=99%):
[1551704603] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10445412 kB) free.
[1551704947] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41380 MB (19% inode=99%):
[1551705067] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46146 MB (22% inode=99%):
[1551705667] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38974 MB (18% inode=99%):
[1551705787] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46906 MB (22% inode=99%):
[1551706133] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4845492 kB) free!
[1551706253] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4808644 kB) free!
[1551706304] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47630 MB (21% inode=99%):
[1551706328] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551706373] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (4776996 kB) free!
[1551706387] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41427 MB (19% inode=99%):
[1551706403] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9570044 kB) free!
[1551706507] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47093 MB (22% inode=99%):
[1551706523] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7499484 kB) free!
[1551706643] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6214068 kB) free!
[1551706928] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551706967] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169468 kB) free!
[1551707048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551707087] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155140 kB) free!
[1551707107] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39974 MB (19% inode=99%):
[1551707168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551707207] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165276 kB) free!
[1551707227] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44581 MB (21% inode=99%):
[1551707504] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46606 MB (20% inode=99%):
[1551707624] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43534 MB (19% inode=99%):
[1551707744] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42478 MB (18% inode=99%):
[1551707827] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40484 MB (19% inode=99%):
[1551707947] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 36388 MB (17% inode=99%):
[1551708067] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44058 MB (21% inode=99%):
[1551708173] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24523 MB (10% inode=99%):
[1551708443] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4741808 kB) free!
[1551708667] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36365 MB (17% inode=99%):
[1551708787] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44018 MB (21% inode=99%):
[1551709387] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36327 MB (17% inode=99%):
[1551709507] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44268 MB (21% inode=99%):
[1551709643] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10428600 kB) free.
[1551709923] SERVICE ALERT: sig3-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 277 processes
[1551709936] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 75.1% (5525560 kB) free.
[1551710043] SERVICE ALERT: sig3-vm1;CPU_Procs;OK;HARD;1;CPU OK: 272 processes
[1551710107] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36334 MB (17% inode=99%):
[1551710227] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44528 MB (21% inode=99%):
[1551710744] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47618 MB (21% inode=99%):
[1551710768] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551710827] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40434 MB (19% inode=99%):
[1551710844] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9674516 kB) free!
[1551710947] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44534 MB (21% inode=99%):
[1551710964] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7540872 kB) free!
[1551711084] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6264620 kB) free!
[1551711547] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39642 MB (19% inode=99%):
[1551711667] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44543 MB (21% inode=99%):
[1551712247] SERVICE ALERT: sig3-vm2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig3-vm2' was out of bounds)
[1551712247] SERVICE ALERT: sig3-vm2;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'sig3-vm2' was out of bounds)
[1551712250] HOST ALERT: sig3-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.181)
[1551712267] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40452 MB (19% inode=99%):
[1551712314] HOST ALERT: sig3-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1551712367] SERVICE ALERT: sig3-vm2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1551712367] SERVICE ALERT: sig3-vm2;Total Processes;OK;HARD;1;PROCS OK: 272 processes
[1551712387] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 36356 MB (17% inode=99%):
[1551712508] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43528 MB (20% inode=99%):
[1551712564] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551712568] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551712689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551712809] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551713474] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 68.9% (5065844 kB) free.
[1551713636] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 75.6% (5560864 kB) free.
[1551713744] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46577 MB (20% inode=99%):
[1551713864] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43504 MB (19% inode=99%):
[1551713873] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 76.3% (5614832 kB) free.
[1551713984] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43472 MB (19% inode=99%):
[1551714084] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4708820 kB) free!
[1551714684] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5093332 kB) free!
[1551715784] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47591 MB (21% inode=99%):
[1551716108] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44544 MB (21% inode=99%):
[1551716164] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.83, 5.32, 5.11
[1551716284] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.55, 5.38, 5.16
[1551716404] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.49, 5.44, 5.21
[1551716708] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39493 MB (18% inode=99%):
[1551716828] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44467 MB (21% inode=99%):
[1551716989] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45530 MB (20% inode=99%):
[1551717107] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8771440 kB) free!
[1551717109] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48461 MB (21% inode=99%):
[1551717227] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.2% (14408896 kB) free.
[1551717288] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2474420 kB) free!
[1551717407] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.7% (1745048 kB) free!
[1551717428] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40326 MB (19% inode=99%):
[1551717548] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 36102 MB (17% inode=99%):
[1551717668] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1551717668] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43270 MB (20% inode=99%):
[1551717773] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (916656 kB) free!
[1551717898] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 19.60, 79.86, 52.46
[1551718007] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1580512 kB) free!
[1551718018] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.75, 53.47, 46.12
[1551718138] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.43, 35.82, 40.54
[1551718204] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.57, 4.61, 4.83
[1551718209] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551718232] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 335 processes
[1551718284] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10294460 kB) free.
[1551718373] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (4012968 kB) free!
[1551718784] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39272 MB (17% inode=99%):
[1551718809] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551718884] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3465508 kB) free!
[1551718904] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47392 MB (21% inode=99%):
[1551718929] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551719004] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10318084 kB) free.
[1551719339] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 0.90, 11.32
[1551719504] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39328 MB (17% inode=99%):
[1551719529] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551719604] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4880264 kB) free!
[1551719624] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39328 MB (17% inode=99%):
[1551719649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551719724] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10308440 kB) free.
[1551719744] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47520 MB (21% inode=99%):
[1551719769] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551720227] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8139680 kB) free!
[1551720324] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5729408 kB) free!
[1551720344] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39328 MB (17% inode=99%):
[1551720347] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.8% (13705140 kB) free.
[1551720369] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551720444] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (4455956 kB) free!
[1551720464] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39328 MB (17% inode=99%):
[1551720490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551720539] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.15, 0.23, 3.29
[1551720564] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10318544 kB) free.
[1551720584] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47520 MB (21% inode=99%):
[1551720604] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.29, 5.44, 5.20
[1551720610] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551720724] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.16, 5.33, 5.19
[1551720844] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.64, 5.18, 5.16
[1551720947] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9165800 kB) free!
[1551721007] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.4% (1696524 kB) free!
[1551721067] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.9% (14244132 kB) free.
[1551721164] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5374004 kB) free!
[1551721184] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39328 MB (17% inode=99%):
[1551721210] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551721284] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3715144 kB) free!
[1551721304] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39328 MB (17% inode=99%):
[1551721330] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551721405] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10309292 kB) free.
[1551721424] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47520 MB (21% inode=99%):
[1551721444] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1551721444] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.87, 4.14, 4.62
[1551721667] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9593068 kB) free!
[1551721787] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.9% (14239240 kB) free.
[1551721930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551722004] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4951024 kB) free!
[1551722024] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39328 MB (17% inode=99%):
[1551722050] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551722088] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (3162592 kB) free!
[1551722124] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10306140 kB) free.
[1551722144] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47520 MB (21% inode=99%):
[1551722170] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1551722724] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5614240 kB) free!
[1551722744] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41400 MB (18% inode=99%):
[1551722770] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551722807] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1551722807] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1148536 kB) free!
[1551722844] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (4138484 kB) free!
[1551722864] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41393 MB (18% inode=99%):
[1551722891] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551722964] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10287904 kB) free.
[1551722984] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47377 MB (21% inode=99%):
[1551723011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551723587] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9067080 kB) free!
[1551723611] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551723708] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.4% (13972700 kB) free.
[1551723709] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47163 MB (20% inode=99%):
[1551723829] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46549 MB (20% inode=99%):
[1551723949] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46333 MB (20% inode=99%):
[1551724164] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4746024 kB) free!
[1551724184] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39260 MB (17% inode=99%):