[1566585572] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183132 kB) free!
[1566585692] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (212508 kB) free!
[1566585794] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200912 kB) free!
[1566585812] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (210932 kB) free!
[1566585902] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177828 kB) free!
[1566585914] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175064 kB) free!
[1566586022] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163924 kB) free!
[1566586034] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176052 kB) free!
[1566586142] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196576 kB) free!
[1566587744] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 210.26, 164.03, 81.38
[1566587765] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 85.22, 31.48, 12.84
[1566587864] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 164.10, 169.18, 93.69
[1566587886] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.70, 21.91, 11.60
[1566587942] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (2557532 kB) free!
[1566587984] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.83, 113.44, 82.41
[1566588006] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.22, 15.54, 10.51
[1566588212] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.6% (11246912 kB) free.
[1566588617] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 594 processes
[1566588677] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 593 processes
[1566588737] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 596 processes
[1566588775] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1566588797] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 592 processes
[1566588842] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 292 processes
[1566588962] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 291 processes
[1566589082] SERVICE ALERT: sig5-vm2;CPU_Procs;OK;HARD;1;CPU OK: 220 processes
[1566589142] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2099304 kB) free!
[1566589184] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.3% (17594704 kB) free!
[1566589206] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.44, 2.77, 4.74
[1566589304] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6893312 kB) free!
[1566589424] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (431688 kB) free!
[1566589784] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.70, 0.96, 12.46
[1566590984] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.76, 0.83, 4.06
[1566593612] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190904 kB) free!
[1566593672] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232404 kB) free!
[1566593733] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198144 kB) free!
[1566593792] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (258168 kB) free!
[1566593853] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184216 kB) free!
[1566593912] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226148 kB) free!
[1566594513] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (45247188 kB) free.
[1566596342] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.3% (10620688 kB) free.
[1566596834] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 59.7% (14632344 kB) free.
[1566606606] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.81, 11.02, 7.72
[1566606726] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.09, 8.13, 7.06
[1566606846] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.98, 6.46, 6.58
[1566607227] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 408 processes
[1566607347] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 409 processes
[1566607424] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.4% (22990232 kB) free!
[1566607446] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.23, 3.03, 4.69
[1566608024] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 44.4% (58499544 kB) free.
[1566608853] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 65.3% (16104872 kB) free.
[1566608984] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 211.14, 159.76, 77.35
[1566609104] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 84.66, 141.35, 81.46
[1566609224] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.81, 94.73, 71.64
[1566611024] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.84, 1.10, 11.01
[1566611342] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173960 kB) free!
[1566611462] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184160 kB) free!
[1566611582] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175904 kB) free!
[1566611837] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193056 kB) free!
[1566611854] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186712 kB) free!
[1566611962] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187184 kB) free!
[1566611974] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193712 kB) free!
[1566612082] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177768 kB) free!
[1566612094] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173488 kB) free!
[1566612224] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.11, 1.65, 4.40
[1566616347] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1566616468] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566620400] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.6% (21954136 kB) free.
[1566620456] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1566620461] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.2% (14496680 kB) free.
[1566620503] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.6% (22688356 kB) free.
[1566620559] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.9% (40832716 kB) free.
[1566620576] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1566620604] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.2% (23239040 kB) free.
[1566622848] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.05, 5.77, 5.46
[1566622968] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.81, 5.56, 5.43
[1566623088] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.30, 5.28, 5.35
[1566625135] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566625255] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566625375] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566625861] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (231480 kB) free!
[1566625981] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (236828 kB) free!
[1566626101] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (225140 kB) free!
[1566628663] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 253.70, 169.55, 76.16
[1566628783] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 254.57, 197.74, 97.81
[1566628903] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 215.98, 208.27, 114.12
[1566629667] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566629787] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1566630288] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.92, 4.35, 4.60
[1566630703] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 1.23, 17.06
[1566630759] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7196960 kB) free!
[1566630879] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.2% (6509292 kB) free!
[1566630999] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.9% (5851564 kB) free!
[1566631599] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2641756 kB) free!
[1566632503] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.58, 2.98
[1566632575] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566633113] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.2% (14504464 kB) free.
[1566633174] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.9% (40831572 kB) free.
[1566633224] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 284.40, 152.89, 62.82
[1566633344] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 290.53, 198.75, 90.55
[1566633464] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 291.79, 229.65, 114.99
[1566633513] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9865708 kB) free!
[1566633633] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9607960 kB) free!
[1566633753] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9331728 kB) free!
[1566635520] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (2357648 kB) free!
[1566635640] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.7% (1745948 kB) free!
[1566635760] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1124952 kB) free!
[1566635864] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.25, 1.25, 11.40
[1566636154] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4937756 kB) free!
[1566659013] SERVICE ALERT: webservercluster2;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'webservercluster2' was out of bounds)
[1566659013] SERVICE ALERT: primary-server;Sensors Status;CRITICAL;SOFT;1;(Return code of 255 for service 'Sensors Status' on host 'primary-server' was out of bounds)
[1566659013] SERVICE ALERT: primary-server;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'primary-server' was out of bounds)
[1566659013] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;SOFT;1;(Return code of 255 for service 'Sensors Status' on host 'webservercluster2' was out of bounds)
[1566659013] SERVICE ALERT: webserver2;HTTP;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 80: Connection refused
[1566659013] SERVICE ALERT: primary-server;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'primary-server' was out of bounds)
[1566659013] SERVICE ALERT: primary-server;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'primary-server' was out of bounds)
[1566659013] SERVICE ALERT: primary-server;Nginx Status;CRITICAL;SOFT;1;(Return code of 255 for service 'Nginx Status' on host 'primary-server' was out of bounds)
[1566659013] SERVICE ALERT: webservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'webservercluster2' was out of bounds)
[1566659013] SERVICE ALERT: primary-server;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'primary-server' was out of bounds)
[1566659013] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3217784 kB) free!
[1566659013] SERVICE ALERT: webservercluster2;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'webservercluster2' was out of bounds)
[1566659013] SERVICE ALERT: webservercluster2;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'webservercluster2' was out of bounds)
[1566659013] SERVICE ALERT: webservercluster2;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'webservercluster2' was out of bounds)
[1566659013] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1566659013] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1566659013] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1566659013] SERVICE ALERT: sigqscluster2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1566659013] SERVICE ALERT: webserver2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 288 processes with STATE = RSZDT
[1566659013] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1566659013] SERVICE ALERT: qservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566659013] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 469 processes
[1566659013] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1566659013] SERVICE ALERT: webserver1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1566659013] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8154500 kB) free!
[1566659013] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.9% (47740156 kB) free.
[1566659013] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8152100 kB) free!
[1566659014] SERVICE ALERT: primary-server;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'primary-server' was out of bounds)
[1566659014] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'webservercluster2' was out of bounds)
[1566659015] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;(Return code of 255 for service 'Memory' on host 'webservercluster2' was out of bounds)
[1566659015] HOST ALERT: pservercluster2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.38)
[1566659015] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'pservercluster2' was out of bounds)
[1566659015] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'pservercluster2' was out of bounds)
[1566659015] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'pservercluster2' was out of bounds)
[1566659015] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'pservercluster2' was out of bounds)
[1566659015] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'pservercluster2' was out of bounds)
[1566659015] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'pservercluster2' was out of bounds)
[1566659015] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'pservercluster2' was out of bounds)
[1566659015] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'pservercluster2' was out of bounds)
[1566659015] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;(Return code of 255 for service 'HDD_Root status' on host 'pservercluster2' was out of bounds)
[1566659016] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sservercluster1' was out of bounds)
[1566659016] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sservercluster1' was out of bounds)
[1566659016] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sservercluster1' was out of bounds)
[1566659016] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;1;(Return code of 255 for service 'Sensors Status' on host 'sservercluster1' was out of bounds)
[1566659016] HOST ALERT: sservercluster1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.49)
[1566659016] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sservercluster1' was out of bounds)
[1566659016] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sservercluster1' was out of bounds)
[1566659016] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'primary-server' was out of bounds)
[1566659016] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'webservercluster2' was out of bounds)
[1566659016] SERVICE ALERT: primary-server;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'primary-server' was out of bounds)
[1566659016] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sservercluster1' was out of bounds)
[1566659016] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sservercluster1' was out of bounds)
[1566659016] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sservercluster1' was out of bounds)
[1566659016] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sservercluster1' was out of bounds)
[1566659020] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'webservercluster2' was out of bounds)
[1566659020] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'primary-server' was out of bounds)
[1566659020] SERVICE ALERT: primary-server;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'primary-server' was out of bounds)
[1566659020] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;(Return code of 255 for service 'Memory' on host 'primary-server' was out of bounds)
[1566659020] SERVICE ALERT: primary-server;HDD_Home status;CRITICAL;SOFT;2;(Return code of 255 for service 'HDD_Home status' on host 'primary-server' was out of bounds)
[1566659020] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;SOFT;2;(Return code of 255 for service 'Current Users' on host 'webservercluster2' was out of bounds)
[1566659020] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;(Return code of 255 for service 'CPU_Procs' on host 'webservercluster2' was out of bounds)
[1566659023] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1566659073] SERVICE ALERT: webserver2;HTTP;OK;HARD;1;HTTP OK: HTTP/1.1 200 OK - 355 bytes in 0.016 second response time
[1566659073] SERVICE ALERT: webserver2;Total Processes;OK;HARD;1;PROCS OK: 178 processes with STATE = RSZDT
[1566659073] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 357 processes
[1566659133] SERVICE ALERT: primary-server;Nginx Status;CRITICAL;SOFT;2;(Return code of 255 for service 'Nginx Status' on host 'primary-server' was out of bounds)
[1566659133] SERVICE ALERT: primary-server;CPU Load;CRITICAL;SOFT;2;(Return code of 255 for service 'CPU Load' on host 'primary-server' was out of bounds)
[1566659133] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;SOFT;2;(Return code of 255 for service 'Sensors Status' on host 'webservercluster2' was out of bounds)
[1566659133] SERVICE ALERT: primary-server;Current Users;CRITICAL;SOFT;2;(Return code of 255 for service 'Current Users' on host 'primary-server' was out of bounds)
[1566659133] SERVICE ALERT: primary-server;Total Processes;CRITICAL;SOFT;2;(Return code of 255 for service 'Total Processes' on host 'primary-server' was out of bounds)
[1566659133] SERVICE ALERT: webservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'webservercluster2' was out of bounds)
[1566659133] SERVICE ALERT: primary-server;Sensors Status;CRITICAL;SOFT;2;(Return code of 255 for service 'Sensors Status' on host 'primary-server' was out of bounds)
[1566659133] SERVICE ALERT: webservercluster2;HDD_Root status;CRITICAL;SOFT;2;(Return code of 255 for service 'HDD_Root status' on host 'webservercluster2' was out of bounds)
[1566659133] SERVICE ALERT: primary-server;HDD_Root status;CRITICAL;SOFT;2;(Return code of 255 for service 'HDD_Root status' on host 'primary-server' was out of bounds)
[1566659133] SERVICE ALERT: webservercluster2;HDD_Home status;CRITICAL;SOFT;2;(Return code of 255 for service 'HDD_Home status' on host 'webservercluster2' was out of bounds)
[1566659133] SERVICE ALERT: webservercluster2;Total Processes;CRITICAL;SOFT;2;(Return code of 255 for service 'Total Processes' on host 'webservercluster2' was out of bounds)
[1566659133] SERVICE ALERT: webservercluster2;CPU Load;CRITICAL;SOFT;2;(Return code of 255 for service 'CPU Load' on host 'webservercluster2' was out of bounds)
[1566659133] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3222404 kB) free!
[1566659133] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1566659133] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1566659133] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.5% (8147424 kB) free!
[1566659134] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;(Return code of 255 for service 'Memory' on host 'webservercluster2' was out of bounds)
[1566659135] SERVICE ALERT: primary-server;HDD_Home status;CRITICAL;HARD;3;(Return code of 255 for service 'HDD_Home status' on host 'primary-server' was out of bounds)
[1566659136] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;(Return code of 255 for service 'CPU_Procs' on host 'primary-server' was out of bounds)
[1566659136] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;(Return code of 255 for service 'CPU_Procs' on host 'webservercluster2' was out of bounds)
[1566659136] SERVICE ALERT: primary-server;LocalService-NTPD SERVICE;CRITICAL;HARD;3;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'primary-server' was out of bounds)
[1566659139] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;2;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sservercluster1' was out of bounds)
[1566659140] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;(Return code of 255 for service 'Current Users' on host 'webservercluster2' was out of bounds)
[1566659140] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;(Return code of 255 for service 'Memory' on host 'primary-server' was out of bounds)
[1566659253] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;HARD;3;(Return code of 255 for service 'Sensors Status' on host 'webservercluster2' was out of bounds)
[1566659253] SERVICE ALERT: primary-server;Nginx Status;CRITICAL;HARD;3;(Return code of 255 for service 'Nginx Status' on host 'primary-server' was out of bounds)
[1566659253] SERVICE ALERT: primary-server;CPU Load;CRITICAL;HARD;3;(Return code of 255 for service 'CPU Load' on host 'primary-server' was out of bounds)
[1566659253] SERVICE ALERT: primary-server;Current Users;CRITICAL;HARD;3;(Return code of 255 for service 'Current Users' on host 'primary-server' was out of bounds)
[1566659253] SERVICE ALERT: primary-server;Total Processes;CRITICAL;HARD;3;(Return code of 255 for service 'Total Processes' on host 'primary-server' was out of bounds)
[1566659253] SERVICE ALERT: primary-server;Sensors Status;CRITICAL;HARD;3;(Return code of 255 for service 'Sensors Status' on host 'primary-server' was out of bounds)
[1566659253] SERVICE ALERT: webservercluster2;HDD_Root status;CRITICAL;HARD;3;(Return code of 255 for service 'HDD_Root status' on host 'webservercluster2' was out of bounds)
[1566659253] SERVICE ALERT: webservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'webservercluster2' was out of bounds)
[1566659253] SERVICE ALERT: primary-server;HDD_Root status;CRITICAL;HARD;3;(Return code of 255 for service 'HDD_Root status' on host 'primary-server' was out of bounds)
[1566659253] SERVICE ALERT: webservercluster2;HDD_Home status;CRITICAL;HARD;3;(Return code of 255 for service 'HDD_Home status' on host 'webservercluster2' was out of bounds)
[1566659253] SERVICE ALERT: webservercluster2;CPU Load;CRITICAL;HARD;3;(Return code of 255 for service 'CPU Load' on host 'webservercluster2' was out of bounds)
[1566659253] SERVICE ALERT: webservercluster2;Total Processes;CRITICAL;HARD;3;(Return code of 255 for service 'Total Processes' on host 'webservercluster2' was out of bounds)
[1566659253] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3217640 kB) free!
[1566659253] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1566659397] HOST ALERT: pservercluster2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1566659520] HOST ALERT: sservercluster1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.49)
[1566659612] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.27, 0.60, 0.23
[1566659613] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1566659614] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7009 MB (13% inode=98%):
[1566659616] HOST ALERT: sservercluster1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1566659736] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 124930 MB (55% inode=99%):
[1566659736] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1566659736] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1566659736] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.01, 0.22, 0.15
[1566659736] SERVICE ALERT: sservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 41424 MB (80% inode=99%):
[1566659736] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 146390 MB (66% inode=99%):
[1566659736] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1566659736] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 312 processes
[1566659736] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1566659736] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1566659736] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1566659736] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1566659736] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 305 processes
[1566659736] SERVICE ALERT: sservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1566659736] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.7% (23456784 kB) free.
[1566659736] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 96.1% (23560720 kB) free.
[1566660334] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 82.2% (13365824 kB) free.
[1566660336] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1566660340] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1566660453] SERVICE ALERT: webservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 223136 MB (99% inode=99%):
[1566660453] SERVICE ALERT: webservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 40380 MB (78% inode=99%):
[1566660453] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.08, 0.08, 0.10
[1566660453] SERVICE ALERT: webservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1566660453] SERVICE ALERT: webservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1566660453] SERVICE ALERT: webservercluster2;Total Processes;OK;HARD;1;PROCS OK: 293 processes
[1566661412] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.17, 6.33, 5.71
[1566661532] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.40, 5.20, 5.37
[1566661652] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.64, 4.72, 5.17
[1566662134] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4894168 kB) free!
[1566662852] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.38, 4.52, 4.89
[1566663213] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566663335] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 130073 MB (57% inode=99%):
[1566663336] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566663336] SERVICE ALERT: primary-server;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1566663340] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 92.7% (22717028 kB) free.
[1566663453] SERVICE ALERT: primary-server;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39101 MB (76% inode=99%):
[1566663453] SERVICE ALERT: primary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1566663453] SERVICE ALERT: primary-server;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.07, 0.06
[1566663453] SERVICE ALERT: primary-server;Sensors Status;OK;HARD;1;SENSORS OK
[1566663453] SERVICE ALERT: primary-server;Total Processes;OK;HARD;1;PROCS OK: 343 processes
[1566663453] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1566663454] SERVICE ALERT: primary-server;Nginx Status;OK;HARD;1;NGINX OK -  0.437 sec. response time, Active: 0 (Writing: 0 Reading: 0 Waiting: 0) ReqPerSec: 0.000 ConnPerSec: 0.000 ReqPerConn: 0.000
[1566663812] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566663813] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 294 processes
[1566663813] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 319 processes
[1566663932] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1566663933] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 293 processes
[1566663933] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 321 processes
[1566664053] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 293 processes
[1566664054] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 323 processes
[1566670536] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)