[1565116758] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565116774] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7196440 kB) free!
[1565116780] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181112 kB) free!
[1565116878] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565116881] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.6% (3098420 kB) free!
[1565116894] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11338992 kB) free.
[1565116900] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194828 kB) free!
[1565116918] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 578 processes
[1565117001] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.6% (3091644 kB) free!
[1565117020] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189148 kB) free!
[1565117038] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 584 processes
[1565117121] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.6% (3086180 kB) free!
[1565117158] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;HARD;3;PROCS WARNING: 590 processes
[1565117178] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565117298] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565117479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565117494] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7393072 kB) free!
[1565117599] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565117614] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11345572 kB) free.
[1565117758] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 625 processes
[1565118199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565118214] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7477628 kB) free!
[1565118319] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565118334] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11338460 kB) free.
[1565118919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565118934] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7527332 kB) free!
[1565119039] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565119054] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11345360 kB) free.
[1565119355] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.0% (9880816 kB) free.
[1565119639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565119654] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7609716 kB) free!
[1565119760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565119775] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11344796 kB) free.
[1565119880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565120374] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7767432 kB) free!
[1565120480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565120494] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11234708 kB) free.
[1565120600] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565121094] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7970408 kB) free!
[1565121200] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565121214] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6213952 kB) free!
[1565121320] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565121334] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11336112 kB) free.
[1565121920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565121934] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6500628 kB) free!
[1565122040] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565122054] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11344128 kB) free.
[1565122641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565122654] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6595152 kB) free!
[1565122761] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565122774] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11345140 kB) free.
[1565123361] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565123374] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6728044 kB) free!
[1565123481] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1565123494] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11337148 kB) free.
[1565124081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565124094] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6790728 kB) free!
[1565124201] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565124215] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11344080 kB) free.
[1565124692] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9124080 kB) free!
[1565124755] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185628 kB) free!
[1565124801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565124812] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (1789388 kB) free!
[1565124814] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6863684 kB) free!
[1565124876] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176036 kB) free!
[1565124921] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565124932] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212476 kB) free!
[1565124934] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11325900 kB) free.
[1565124996] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194904 kB) free!
[1565125521] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565125532] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (44827824 kB) free.
[1565125534] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6973224 kB) free!
[1565125642] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565125654] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332924 kB) free.
[1565126241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565126254] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7167660 kB) free!
[1565126361] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565126374] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11337196 kB) free.
[1565126961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565126974] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7381108 kB) free!
[1565127081] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565127094] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11325832 kB) free.
[1565127322] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169992 kB) free!
[1565127681] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565127694] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7598372 kB) free!
[1565127801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565127814] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11334736 kB) free.
[1565127820] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 45.4% (11137392 kB) free.
[1565127921] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565127922] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.9% (10993700 kB) free.
[1565128414] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7886880 kB) free!
[1565128522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565128535] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6130696 kB) free!
[1565128642] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565128655] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332564 kB) free.
[1565129242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565129255] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6627272 kB) free!
[1565129362] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565129375] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11334860 kB) free.
[1565129962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565129975] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7079392 kB) free!
[1565130082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565130095] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11326836 kB) free.
[1565130682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1565130695] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7449852 kB) free!
[1565130802] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565130815] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11330876 kB) free.
[1565131403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565131415] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7962056 kB) free!
[1565131523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565131535] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6323128 kB) free!
[1565131643] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565131655] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11331512 kB) free.
[1565132243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565132255] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6822160 kB) free!
[1565132363] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565132376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11331756 kB) free.
[1565132745] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1300.36, 874.17, 392.51
[1565132865] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1305.21, 1017.07, 503.28
[1565132963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565132975] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7087512 kB) free!
[1565132985] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1306.18, 1112.73, 600.62
[1565133083] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565133095] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11325456 kB) free.
[1565133683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565133695] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7443024 kB) free!
[1565133804] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565133815] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332628 kB) free.
[1565134403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565134415] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7723188 kB) free!
[1565134523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565134535] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11221984 kB) free.
[1565134596] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2752796 kB) free!
[1565134643] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565134645] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 167.49, 636.36, 407.98
[1565134765] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 22.83, 425.92, 358.56
[1565134885] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.24, 285.07, 315.11
[1565135135] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8013348 kB) free!
[1565135243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565135255] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6350880 kB) free!
[1565135363] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565135375] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11334780 kB) free.
[1565135963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565135975] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6881732 kB) free!
[1565136083] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565136095] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11325388 kB) free.
[1565136585] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.45, 2.60, 16.00
[1565136683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565136696] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7417132 kB) free!
[1565136804] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565136816] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11335800 kB) free.
[1565137403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565137416] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7689124 kB) free!
[1565137523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565137536] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11326552 kB) free.
[1565137643] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565137885] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 0.54, 13.01
[1565138136] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7906676 kB) free!
[1565138243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565138256] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6202256 kB) free!
[1565138363] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565138376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11335920 kB) free.
[1565138385] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.51, 2.63, 4.58
[1565138796] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2472260 kB) free!
[1565138963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565138976] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6723808 kB) free!
[1565139083] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565139085] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.07, 0.98, 4.21
[1565139096] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323852 kB) free.
[1565139684] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565139696] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7187012 kB) free!
[1565139804] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565139816] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332356 kB) free.
[1565140404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565140416] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7507324 kB) free!
[1565140524] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565140537] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323748 kB) free.
[1565141124] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565141136] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7773468 kB) free!
[1565141244] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565141256] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5967168 kB) free!
[1565141364] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565141376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332952 kB) free.
[1565141964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565141976] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6396260 kB) free!
[1565142085] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565142096] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323876 kB) free.
[1565142684] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565142696] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7105180 kB) free!
[1565142804] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565142816] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11331696 kB) free.
[1565143404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565143416] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7508212 kB) free!
[1565143524] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565143536] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11321808 kB) free.
[1565144124] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565144136] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7802228 kB) free!
[1565144244] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565144256] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6272096 kB) free!
[1565144364] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565144377] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332392 kB) free.
[1565144819] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 1951 processes
[1565144879] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 1952 processes
[1565144964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565144976] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6928896 kB) free!
[1565145085] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565145096] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11322368 kB) free.
[1565145684] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565145696] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7289384 kB) free!
[1565145804] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565145816] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11329400 kB) free.
[1565145821] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197384 kB) free!
[1565145922] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (2062156 kB) free!
[1565145941] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192192 kB) free!
[1565146042] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (2046028 kB) free!
[1565146061] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179620 kB) free!
[1565146098] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565146162] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2032820 kB) free!
[1565146218] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1565146404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565146416] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7613284 kB) free!
[1565146524] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565146536] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11322656 kB) free.
[1565146644] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565147136] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7923096 kB) free!
[1565147244] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565147256] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6240632 kB) free!
[1565147364] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565147376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11329712 kB) free.
[1565147965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565147976] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6630176 kB) free!
[1565148085] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565148096] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11318512 kB) free.
[1565148685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565148697] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7182872 kB) free!
[1565148805] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565148817] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11329336 kB) free.
[1565149405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565149417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7566872 kB) free!
[1565149525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565149537] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11317620 kB) free.
[1565149645] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565150137] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7957064 kB) free!
[1565150245] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565150257] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6330284 kB) free!
[1565150366] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565150377] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11326912 kB) free.
[1565150965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565150977] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6900436 kB) free!
[1565151085] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565151097] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11318788 kB) free.
[1565151685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565151697] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7163888 kB) free!
[1565151805] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565151817] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323396 kB) free.
[1565152405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565152525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565152645] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565153018] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8981636 kB) free!
[1565153138] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7582252 kB) free!
[1565153258] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11316472 kB) free.
[1565154458] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9397632 kB) free!
[1565154578] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8013628 kB) free!
[1565154698] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6338524 kB) free!
[1565154758] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1565154878] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1565154998] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1565155298] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11093180 kB) free.
[1565156107] HOST ALERT: sig5-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1565156366] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 97.3% (34070096 kB) free.
[1565156439] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.10, 0.08
[1565156481] SERVICE ALERT: sig5-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1565156512] SERVICE ALERT: sig5-vm2;CPU_Procs;OK;HARD;1;CPU OK: 200 processes
[1565156533] SERVICE ALERT: sig5-vm2;Total Processes;OK;HARD;1;PROCS OK: 200 processes
[1565156559] SERVICE ALERT: sig5-vm2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1565156559] SERVICE ALERT: sig5-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 241252 MB (96% inode=99%):
[1565157146] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.08, 5.22, 5.03
[1565157171] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 78.8% (19449580 kB) free.
[1565157221] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.8% (22265696 kB) free.
[1565157266] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.93, 5.21, 5.06
[1565157275] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.1% (40447020 kB) free.
[1565157335] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.9% (21797480 kB) free.
[1565157386] SERVICE ALERT: sig4-vm2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig4-vm2' was out of bounds)
[1565157386] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.01, 5.76, 5.27
[1565157395] SERVICE ALERT: sig4-vm2;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4-vm2' was out of bounds)
[1565157402] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.8% (14433096 kB) free.
[1565157424] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'sig4-vm2' was out of bounds)
[1565157476] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4-vm2' was out of bounds)
[1565157506] SERVICE ALERT: sig4-vm2;Current Users;CRITICAL;SOFT;2;(Return code of 255 for service 'Current Users' on host 'sig4-vm2' was out of bounds)
[1565157515] SERVICE ALERT: sig4-vm2;HDD_Root status;CRITICAL;SOFT;2;(Return code of 255 for service 'HDD_Root status' on host 'sig4-vm2' was out of bounds)
[1565157531] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1565157544] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;SOFT;2;(Return code of 255 for service 'Memory' on host 'sig4-vm2' was out of bounds)
[1565157596] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;2;(Return code of 255 for service 'CPU_Procs' on host 'sig4-vm2' was out of bounds)
[1565157626] SERVICE ALERT: sig4-vm2;Current Users;CRITICAL;HARD;3;(Return code of 255 for service 'Current Users' on host 'sig4-vm2' was out of bounds)
[1565157635] SERVICE ALERT: sig4-vm2;HDD_Root status;CRITICAL;HARD;3;(Return code of 255 for service 'HDD_Root status' on host 'sig4-vm2' was out of bounds)
[1565157651] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1565157664] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;3;(Return code of 255 for service 'Memory' on host 'sig4-vm2' was out of bounds)
[1565157716] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;HARD;3;(Return code of 255 for service 'CPU_Procs' on host 'sig4-vm2' was out of bounds)
[1565157771] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1565157771] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191324 kB) free!
[1565157892] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168600 kB) free!
[1565157986] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.33, 4.54, 4.90
[1565158012] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169504 kB) free!
[1565158898] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6815548 kB) free!
[1565159018] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11281512 kB) free.
[1565159211] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1253428 kB) free!
[1565159331] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (1251856 kB) free!
[1565159451] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1239168 kB) free!
[1565159786] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.54, 5.79, 5.18
[1565159906] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.08, 6.24, 5.42
[1565160027] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.68, 6.61, 5.65
[1565160398] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1565160965] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.2% (18687692 kB) free!
[1565160998] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1565161085] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.9% (24966136 kB) free!
[1565161118] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1565161206] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.6% (24482948 kB) free!
[1565161238] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1565161480] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243588 kB) free!
[1565161600] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (245828 kB) free!
[1565161720] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213856 kB) free!
[1565162018] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5870572 kB) free!
[1565162138] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10821856 kB) free.
[1565163446] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565163938] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8284272 kB) free!
[1565164046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565164059] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7241696 kB) free!
[1565164166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565164179] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.7% (5758080 kB) free!
[1565164238] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1565164238] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1565164286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565164532] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 13.7% (6730300 kB) free!
[1565164652] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 13.0% (6427508 kB) free!
[1565164772] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.4% (6123692 kB) free!
[1565164886] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565165372] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4750692 kB) free!
[1565165406] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9528788 kB) free!
[1565165486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565165606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565165726] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565165979] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1565165979] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10842260 kB) free.
[1565166006] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.6% (16553544 kB) free!
[1565166327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565166412] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 58.4% (14417776 kB) free.
[1565166447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565166606] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1565166606] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (8328904 kB) free!
[1565167199] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1565167206] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.3% (14853944 kB) free!
[1565167558] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 250 processes
[1565167647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1565167748] SERVICE ALERT: webserver2;Total Processes;OK;HARD;1;PROCS OK: 170 processes with STATE = RSZDT
[1565167767] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565167826] SERVICE ALERT: sig4-vm2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1565167833] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.07, 0.06
[1565167835] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 116007 MB (90% inode=99%):
[1565167864] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 95.4% (33236592 kB) free.
[1565167916] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 244 processes
[1565167971] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1565168367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565168379] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7381212 kB) free!
[1565168406] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6389136 kB) free!
[1565168487] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565168499] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11216372 kB) free.
[1565169006] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.9% (14433812 kB) free!
[1565169087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565169208] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565169328] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1565169328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565169599] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565169928] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565170206] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12493204 kB) free!
[1565170299] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7079124 kB) free!
[1565170420] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11220092 kB) free.
[1565171019] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7148876 kB) free!
[1565171139] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11210020 kB) free.
[1565171427] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.82, 4.06, 4.73
[1565171728] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565171739] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7180796 kB) free!
[1565171848] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565171859] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11218220 kB) free.
[1565172027] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.76, 6.63, 5.58
[1565172147] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.22, 5.71, 5.37
[1565172267] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.33, 5.29, 5.26
[1565172448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565172459] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7355248 kB) free!
[1565172569] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565172579] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11217116 kB) free.
[1565173012] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (235020 kB) free!
[1565173132] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187064 kB) free!
[1565173168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565173179] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7604388 kB) free!
[1565173238] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1565173252] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194524 kB) free!
[1565173288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565173299] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6350936 kB) free!
[1565173408] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565173419] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11101016 kB) free.
[1565173467] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.21, 3.99, 4.77
[1565174608] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565174620] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8054668 kB) free!
[1565174728] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565174740] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6552800 kB) free!
[1565174848] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565174860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11210456 kB) free.
[1565174999] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1565175449] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565175460] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6543664 kB) free!
[1565175569] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565175580] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11213000 kB) free.
[1565176169] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565176180] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6593316 kB) free!
[1565176289] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565176300] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11026136 kB) free.
[1565176799] SERVICE FLAPPING ALERT: secondary-server;Current Users;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1565176799] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565176852] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 53.1% (13096300 kB) free.
[1565176889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565176900] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9541252 kB) free!
[1565177009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565177020] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8175724 kB) free!
[1565177129] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565177140] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6677764 kB) free!
[1565177143] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4767852 kB) free!
[1565177263] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4706136 kB) free!
[1565177383] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4646288 kB) free!
[1565177626] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201112 kB) free!
[1565177729] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565177740] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11193564 kB) free.
[1565177746] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194896 kB) free!
[1565177866] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180308 kB) free!
[1565178052] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178180 kB) free!
[1565178172] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187768 kB) free!
[1565178292] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1565178292] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188752 kB) free!
[1565178330] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565178450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565178570] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1565178868] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.21, 7.73, 5.56
[1565178941] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9525056 kB) free!
[1565178988] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.67, 8.04, 5.94
[1565179061] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8172208 kB) free!
[1565179108] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1565179108] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.86, 7.65, 6.03
[1565179181] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6607424 kB) free!
[1565179770] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565179781] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11003468 kB) free.
[1565179783] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2105136 kB) free!
[1565180370] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565180381] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9447972 kB) free!
[1565180490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565180501] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8063324 kB) free!
[1565180610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565180621] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.2% (6498048 kB) free!
[1565181006] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565181210] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565181221] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11165904 kB) free.
[1565181392] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4812696 kB) free!
[1565181512] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.5% (4789520 kB) free!
[1565181633] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.5% (4769824 kB) free!
[1565182410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565182421] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6652676 kB) free!
[1565182530] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565182541] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11151780 kB) free.
[1565183130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565183141] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6679792 kB) free!
[1565183250] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565183261] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11160100 kB) free.
[1565183850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565183862] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6658840 kB) free!
[1565183970] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565183982] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11162204 kB) free.
[1565184570] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565184582] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6590588 kB) free!
[1565184606] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.1% (14686844 kB) free!
[1565184691] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565184702] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11158076 kB) free.
[1565185199] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1565185290] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565185302] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6038960 kB) free!
[1565185410] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565185422] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11158160 kB) free.
[1565185708] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1565185806] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (11940720 kB) free!
[1565186010] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565186022] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6291564 kB) free!
[1565186130] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565186142] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11147148 kB) free.
[1565186730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565186742] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6366720 kB) free!
[1565186850] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565186862] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11154776 kB) free.
[1565187450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565187462] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6351384 kB) free!
[1565187571] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565187582] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11154688 kB) free.
[1565188170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565188183] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6476384 kB) free!
[1565188290] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565188303] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11184624 kB) free.
[1565188890] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565188903] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6619996 kB) free!
[1565189010] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565189023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11187892 kB) free.
[1565189095] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565189610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565189623] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8071076 kB) free!
[1565189730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565189743] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7142444 kB) free!
[1565189850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565189863] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11180640 kB) free.
[1565189908] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.11, 3.75, 4.51
[1565189956] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.87, 5.81, 6.13
[1565190076] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.55, 4.39, 5.57
[1565190085] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.93, 5.89, 5.35
[1565190196] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.07, 3.29, 5.02
[1565190205] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.78, 10.53, 7.30
[1565190325] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.68, 9.09, 7.17
[1565190463] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8530880 kB) free!
[1565190583] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7535888 kB) free!
[1565190703] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6045304 kB) free!
[1565190797] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.80, 2.86
[1565191199] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1565191206] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.3% (13587036 kB) free!
[1565191650] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565191806] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1565191806] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (10924308 kB) free!
[1565192125] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.60, 2.90
[1565192504] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11073928 kB) free.
[1565192850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565192970] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565193570] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565193690] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565194206] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.7% (14117300 kB) free!
[1565194291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565194304] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7758256 kB) free!
[1565194411] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565194424] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6107336 kB) free!
[1565194531] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565194544] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11174748 kB) free.
[1565195131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565195144] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6542236 kB) free!
[1565195251] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565195264] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11182400 kB) free.
[1565195851] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565195864] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6559596 kB) free!
[1565195971] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565195984] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11183440 kB) free.
[1565196006] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12430736 kB) free!
[1565196571] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565196584] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6752428 kB) free!
[1565196606] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (14173292 kB) free!
[1565196692] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565196704] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11183144 kB) free.
[1565197291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565197305] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6991052 kB) free!
[1565197411] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565197425] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11182728 kB) free.
[1565197818] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1565197938] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565198011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565198025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7112576 kB) free!
[1565198131] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565198145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11174720 kB) free.
[1565198731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565198745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7115264 kB) free!
[1565198851] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565198865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11183460 kB) free.
[1565199451] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565199465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7254424 kB) free!
[1565199571] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565199585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11182584 kB) free.
[1565199633] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2457512 kB) free!
[1565200172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565200185] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7252640 kB) free!
[1565200292] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565200305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11180112 kB) free.
[1565200892] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565200905] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7270596 kB) free!
[1565201012] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565201025] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11182984 kB) free.
[1565201612] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565201626] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7337520 kB) free!
[1565201732] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565201746] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11177120 kB) free.
[1565202006] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12927096 kB) free!
[1565202332] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565202346] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7472016 kB) free!
[1565202453] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565202466] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11181396 kB) free.