[1586025451] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20640076 kB) free!
[1586026051] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35896800 kB) free.
[1586026651] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16436908 kB) free!
[1586026711] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (612888 kB) free!
[1586026771] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12166408 kB) free!
[1586026893] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (8074848 kB) free!
[1586027111] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181968 kB) free!
[1586027171] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197456 kB) free!
[1586027233] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (189464 kB) free!
[1586027294] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (172112 kB) free!
[1586027493] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23877188 kB) free!
[1586027895] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 76.5% (18750428 kB) free.
[1586028093] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2532620 kB) free!
[1586028304] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1586028511] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.3% (4604096 kB) free.
[1586028693] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14630704 kB) free!
[1586029111] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.2% (1821196 kB) free!
[1586029231] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2803968 kB) free!
[1586029293] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34666180 kB) free.
[1586029351] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.6% (3832332 kB) free.
[1586029573] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1586029893] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9651996 kB) free!
[1586030013] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5077568 kB) free!
[1586030133] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403440 kB) free!
[1586030704] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.50, 7.62, 6.23
[1586030733] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17053748 kB) free!
[1586030824] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.79, 7.72, 6.43
[1586030944] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.66, 7.67, 6.57
[1586031333] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30068804 kB) free.
[1586031751] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2972768 kB) free!
[1586031871] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4037924 kB) free.
[1586031933] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8563120 kB) free!
[1586032054] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3534616 kB) free!
[1586032174] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387568 kB) free!
[1586032461] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (24952296 kB) free!
[1586032472] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3233492 kB) free!
[1586032581] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.9% (30187580 kB) free.
[1586032592] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.9% (3881948 kB) free.
[1586032774] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17172164 kB) free!
[1586033192] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3026568 kB) free!
[1586033312] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3712344 kB) free.
[1586033374] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380648 kB) free!
[1586033912] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.5% (1710752 kB) free!
[1586034032] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.2% (2305832 kB) free!
[1586034152] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2938200 kB) free!
[1586034349] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 64.9% (31992896 kB) free.
[1586034381] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (24986500 kB) free!
[1586034501] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.0% (39522268 kB) free.
[1586034574] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21590860 kB) free!
[1586034643] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.7% (41763852 kB) free.
[1586034752] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (915788 kB) free!
[1586034934] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.5% (9437828 kB) free.
[1586035174] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8888360 kB) free!
[1586035352] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3276252 kB) free.
[1586035774] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32300604 kB) free.
[1586035808] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2774056 kB) free!
[1586035843] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(211) < 103.16.101.72 [05/Apr/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1586035964] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586036106] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2313) < 172.105.59.173 [05/Apr/2020:03:02:02 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586036223] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(268) < 2020/04/05 03:00:22 [error] 22125#0: *81205730 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1586036226] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586036343] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586036374] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4864276 kB) free!
[1586036408] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1510512 kB) free!
[1586036494] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (586628 kB) free!
[1586036614] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411432 kB) free!
[1586037214] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (16014116 kB) free!
[1586037544] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.43, 3.01, 4.29
[1586037814] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34490376 kB) free.
[1586038414] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7996356 kB) free!
[1586038535] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3283112 kB) free!
[1586038655] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381960 kB) free!
[1586038824] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1586038944] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1586039255] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14314680 kB) free!
[1586039855] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30116220 kB) free.
[1586040455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8644964 kB) free!
[1586040575] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4988592 kB) free!
[1586040695] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26149512 kB) free!
[1586041895] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27501244 kB) free.
[1586042495] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7072904 kB) free!
[1586042615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3569784 kB) free!
[1586042735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33836624 kB) free.
[1586043335] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185240 kB) free!
[1586043335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11091616 kB) free!
[1586043455] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170372 kB) free!
[1586043455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7390868 kB) free!
[1586043575] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189784 kB) free!
[1586043575] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17720676 kB) free!
[1586044424] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 60.4% (14793696 kB) free.
[1586044656] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 51.8% (12789864 kB) free.
[1586044808] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 64.6% (15822380 kB) free.
[1586045375] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5980564 kB) free!
[1586045975] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14392504 kB) free!
[1586046752] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586047175] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35982272 kB) free.
[1586047775] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12482628 kB) free!
[1586047895] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8948628 kB) free!
[1586048015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35489236 kB) free.
[1586048351] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1586048471] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1586048615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6679292 kB) free!
[1586048735] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2280484 kB) free!
[1586048855] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36246064 kB) free.
[1586049455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10034328 kB) free!
[1586049575] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6101572 kB) free!
[1586049695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36246564 kB) free.
[1586049795] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4667284 kB) free!
[1586049855] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4665380 kB) free!
[1586049915] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.0% (4666724 kB) free!
[1586049976] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.0% (4660032 kB) free!
[1586050295] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11975628 kB) free!
[1586050415] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7965304 kB) free!
[1586050536] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36197240 kB) free.
[1586051135] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11874504 kB) free!
[1586051255] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6742516 kB) free!
[1586051375] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36080704 kB) free.
[1586051975] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11100552 kB) free!
[1586052095] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7257668 kB) free!
[1586052215] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (971988 kB) free!
[1586052815] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15413004 kB) free!
[1586053056] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186448 kB) free!
[1586053176] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180580 kB) free!
[1586053296] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167716 kB) free!
[1586054015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32618132 kB) free.
[1586054615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6248604 kB) free!
[1586054735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35205976 kB) free.
[1586055335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8745520 kB) free!
[1586055455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4896096 kB) free!
[1586055544] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.63, 5.03, 5.08
[1586055575] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36696224 kB) free.
[1586055664] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.62, 5.27, 5.15
[1586055784] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.92, 5.61, 5.30
[1586056176] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19247964 kB) free!
[1586056296] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15227460 kB) free!
[1586056416] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38382180 kB) free.
[1586056871] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1586056991] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1586057016] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17920148 kB) free!
[1586057136] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (14020296 kB) free!
[1586057256] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37477240 kB) free.
[1586057549] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1586057624] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (1510452 kB) free!
[1586057669] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1586057744] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (1289392 kB) free!
[1586057856] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13532740 kB) free!
[1586057864] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (1076760 kB) free!
[1586057976] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9034192 kB) free!
[1586058096] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33802528 kB) free.
[1586058500] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1586058608] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180664 kB) free!
[1586058620] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1586058696] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11147728 kB) free!
[1586058728] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186880 kB) free!
[1586058741] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1586058816] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7556472 kB) free!
[1586058848] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179032 kB) free!
[1586058936] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25543288 kB) free!
[1586058973] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1586059093] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1586059213] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1586059341] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1586059536] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3214196 kB) free!
[1586059813] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1586060414] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1586060534] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1586060654] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1586060736] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (19024320 kB) free!
[1586060940] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 68.7% (16845304 kB) free.
[1586061336] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27973236 kB) free.
[1586061376] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (336900 kB) free!
[1586061937] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5246176 kB) free!
[1586062057] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1117992 kB) free!
[1586062177] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27616012 kB) free.
[1586062777] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3404740 kB) free!
[1586062897] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (430712 kB) free!
[1586063017] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30708796 kB) free.
[1586063617] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5341196 kB) free!
[1586063737] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1033788 kB) free!
[1586063857] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31158016 kB) free.
[1586064457] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7305304 kB) free!
[1586064577] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2960264 kB) free!
[1586064697] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30560492 kB) free.
[1586064854] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1586064854] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1586065297] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6143720 kB) free!
[1586065417] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1605696 kB) free!
[1586065537] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31968572 kB) free.
[1586066137] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8236664 kB) free!
[1586066258] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3262480 kB) free!
[1586066378] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388752 kB) free!
[1586066503] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4419080 kB) free!
[1586066623] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (3920752 kB) free!
[1586066648] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 31.9% (7821500 kB) free.
[1586066743] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.3% (3445228 kB) free!
[1586067343] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1078420 kB) free!
[1586067578] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16763224 kB) free!
[1586068178] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393820 kB) free!
[1586068544] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.1% (15740716 kB) free.
[1586069378] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36099852 kB) free.
[1586069648] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (4303516 kB) free!
[1586069768] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (4179856 kB) free!
[1586069888] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (4060328 kB) free!
[1586069940] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.4% (4514912 kB) free!
[1586069978] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10466120 kB) free!
[1586070060] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.6% (4303428 kB) free!
[1586070098] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6068972 kB) free!
[1586070180] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.6% (4072292 kB) free!
[1586070218] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1793196 kB) free!
[1586070818] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15721476 kB) free!
[1586071380] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2421828 kB) free!
[1586071418] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31126824 kB) free.
[1586071688] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2445356 kB) free!
[1586072018] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6463072 kB) free!
[1586072138] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2271352 kB) free!
[1586072258] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32515752 kB) free.
[1586072654] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1586072859] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9113144 kB) free!
[1586072979] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4714700 kB) free!
[1586073099] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (531220 kB) free!
[1586073349] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9173924 kB) free!
[1586073469] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (8995968 kB) free!
[1586073589] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (8826700 kB) free!
[1586074299] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24921720 kB) free!
[1586074899] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374072 kB) free!
[1586076099] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18184632 kB) free!
[1586076344] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3602232 kB) free!
[1586076464] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3431564 kB) free!
[1586076584] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3351856 kB) free!
[1586076643] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9500204 kB) free!
[1586076699] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394256 kB) free!
[1586076763] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9396192 kB) free!
[1586076883] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9308708 kB) free!
[1586077299] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17419748 kB) free!
[1586077784] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1586077784] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2182944 kB) free!
[1586077899] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31537364 kB) free.
[1586078226] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(123) < 92.42.104.166 [05/Apr/2020:14:40:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1586078346] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586078389] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4681080 kB) free!
[1586078499] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7679552 kB) free!
[1586078584] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.26, 4.18, 4.79
[1586078619] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3575908 kB) free!
[1586078739] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389352 kB) free!
[1586079546] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(151) < 92.42.104.166 [05/Apr/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1586079584] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.2% (15766680 kB) free.
[1586079666] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586079784] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.96, 7.49, 6.17
[1586079904] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.36, 8.63, 6.74
[1586079939] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23467700 kB) free!
[1586080024] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.35, 8.74, 7.00
[1586080540] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399124 kB) free!
[1586080991] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1586081111] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1586081140] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16990484 kB) free!
[1586081740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32389792 kB) free.
[1586082283] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4505032 kB) free!
[1586082340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6143952 kB) free!
[1586082460] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1714436 kB) free!
[1586082580] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380748 kB) free!
[1586083780] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32375892 kB) free.
[1586084380] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6385364 kB) free!
[1586084500] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1992452 kB) free!
[1586084620] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396896 kB) free!
[1586085220] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13425532 kB) free!
[1586085820] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32390632 kB) free.
[1586086420] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6069676 kB) free!
[1586086540] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1531688 kB) free!
[1586086660] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380128 kB) free!
[1586086784] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2808592 kB) free!
[1586086904] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.6% (2384204 kB) free!
[1586087024] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1928004 kB) free!
[1586087261] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13516452 kB) free!
[1586087861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34798644 kB) free.
[1586088461] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9268824 kB) free!
[1586088581] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4739308 kB) free!
[1586088701] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417724 kB) free!
[1586089301] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19643180 kB) free!
[1586089901] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (434396 kB) free!
[1586090024] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.7% (15875200 kB) free.
[1586090111] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1586090231] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1586090501] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15003468 kB) free!
[1586091101] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377416 kB) free!
[1586091753] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3222684 kB) free!
[1586091873] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3631172 kB) free.
[1586093301] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26282512 kB) free!
[1586093421] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.8% (36640388 kB) free.
[1586093501] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22841620 kB) free!
[1586093600] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1586093673] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2481724 kB) free!
[1586093720] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1586093793] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3106668 kB) free!
[1586093913] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3622340 kB) free.
[1586094101] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410508 kB) free!
[1586094702] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23679912 kB) free!
[1586095302] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391528 kB) free!
[1586095713] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2095892 kB) free!
[1586095833] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2842960 kB) free!
[1586095902] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26505956 kB) free.
[1586095953] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3674412 kB) free.
[1586096421] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26270832 kB) free!
[1586096502] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (586200 kB) free!
[1586096541] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.1% (33085832 kB) free.
[1586096553] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2537140 kB) free!
[1586096622] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (425956 kB) free!
[1586096624] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.3% (3435456 kB) free!
[1586096673] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3415324 kB) free.
[1586096742] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398552 kB) free!
[1586096744] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.1% (2943844 kB) free!
[1586096831] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1586096864] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.8% (2432404 kB) free!
[1586096951] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1586097273] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.1% (2452300 kB) free!
[1586097342] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16962248 kB) free!
[1586097393] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3156536 kB) free!
[1586097464] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (191532 kB) free!
[1586097514] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2818932 kB) free!
[1586097854] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1586097942] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383756 kB) free!
[1586097974] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1586098094] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1586098542] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18128520 kB) free!
[1586098694] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1586098714] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1418816 kB) free!
[1586099142] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397056 kB) free!
[1586099356] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (5796416 kB) free!
[1586099476] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 15.2% (5325560 kB) free!
[1586099596] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 13.9% (4875980 kB) free!
[1586099742] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23013556 kB) free!
[1586100141] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.0% (22437548 kB) free!
[1586100196] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (3165784 kB) free!
[1586100261] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.0% (34221260 kB) free.
[1586100342] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386160 kB) free!
[1586100796] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.2% (31589884 kB) free.
[1586100942] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22826720 kB) free!
[1586101065] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.4% (16026012 kB) free.
[1586101151] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1586101271] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1586101462] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.0% (19787204 kB) free!
[1586101542] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380820 kB) free!
[1586101582] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (30876672 kB) free.
[1586101714] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.8% (5011176 kB) free.
[1586102142] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24276292 kB) free!
[1586102294] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1586102414] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 353 processes
[1586102534] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1586102743] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (502076 kB) free!
[1586102914] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.3% (1670220 kB) free!
[1586103035] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (2685312 kB) free!
[1586103155] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3759936 kB) free.
[1586103343] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25356524 kB) free!
[1586103734] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1586103734] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1586103943] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (437860 kB) free!
[1586104543] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28978416 kB) free.
[1586104582] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (25213560 kB) free!
[1586104702] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.5% (37622656 kB) free.
[1586105143] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2468628 kB) free!
[1586105263] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (447956 kB) free!
[1586105302] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (22783396 kB) free!
[1586105383] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403532 kB) free!
[1586105422] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.4% (33527948 kB) free.
[1586105983] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17747588 kB) free!
[1586106155] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (916340 kB) free!
[1586106275] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.9% (1939128 kB) free!
[1586106395] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1586106395] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2936532 kB) free!
[1586106583] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380952 kB) free!
[1586106995] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3387788 kB) free.
[1586107183] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20540216 kB) free!
[1586107783] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1142284 kB) free!
[1586108383] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19876160 kB) free!
[1586108795] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3251888 kB) free!
[1586108824] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.02, 3.70, 4.71
[1586108916] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (580792 kB) free!
[1586108983] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386512 kB) free!
[1586109036] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1598224 kB) free!
[1586109134] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1586109254] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1586109465] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1586109583] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (21987876 kB) free!
[1586109636] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1695136 kB) free!
[1586110184] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2736984 kB) free!
[1586110454] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1586110574] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1586110694] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1586110784] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19123212 kB) free!
[1586111294] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1586111384] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (741124 kB) free!