[1586889113] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (4196064 kB) free!
[1586889201] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12482080 kB) free!
[1586889321] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16465216 kB) free!
[1586889441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31867684 kB) free.
[1586889713] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189112 kB) free!
[1586890041] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10747004 kB) free!
[1586890161] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34301420 kB) free.
[1586890633] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1586890761] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12982556 kB) free!
[1586890882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34530820 kB) free.
[1586890942] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1586890956] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168068 kB) free!
[1586891021] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193524 kB) free!
[1586891062] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1586891081] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (197440 kB) free!
[1586891143] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (197208 kB) free!
[1586891443] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 72.0% (17638292 kB) free.
[1586891481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14761752 kB) free!
[1586891601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10959164 kB) free!
[1586891721] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33315248 kB) free.
[1586891833] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (878964 kB) free!
[1586891953] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (875416 kB) free!
[1586892073] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (879344 kB) free!
[1586892321] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12411836 kB) free!
[1586892441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34545612 kB) free.
[1586893041] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12973236 kB) free!
[1586893161] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34511480 kB) free.
[1586893761] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13939188 kB) free!
[1586893881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34541280 kB) free.
[1586894145] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1586894157] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3926088 kB) free!
[1586894197] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (418064 kB) free!
[1586894278] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (222640 kB) free!
[1586894317] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (611720 kB) free!
[1586894398] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222272 kB) free!
[1586894437] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2614328 kB) free!
[1586894481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14928048 kB) free!
[1586894601] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.8% (26091048 kB) free!
[1586894721] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31952724 kB) free.
[1586895321] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12270696 kB) free!
[1586895441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32775468 kB) free.
[1586895564] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.13, 5.30, 5.11
[1586895685] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.33, 5.44, 5.20
[1586895806] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.09, 5.02, 5.08
[1586896041] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12940628 kB) free!
[1586896162] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33332944 kB) free.
[1586896761] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12700124 kB) free!
[1586896761] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(124) < 92.42.104.166 [15/Apr/2020:02:00:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1586896798] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.6% (6719960 kB) free!
[1586896881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33452720 kB) free.
[1586896881] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586897481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13918004 kB) free!
[1586897601] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586897601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33748676 kB) free.
[1586897606] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.76, 4.17, 4.92
[1586897634] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6632272 kB) free!
[1586897754] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 74.1% (36531020 kB) free.
[1586898201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14097604 kB) free!
[1586898291] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.7% (43209944 kB) free.
[1586898321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34503916 kB) free.
[1586898921] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15060604 kB) free!
[1586899041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34676532 kB) free.
[1586899313] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.8% (9763628 kB) free.
[1586899641] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14828816 kB) free!
[1586899761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33613648 kB) free.
[1586899873] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (2559596 kB) free!
[1586899881] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2309) < 203.212.70.170 [15/Apr/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586899977] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(275) < 103.16.101.72 [15/Apr/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1586899993] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(344) < 2020/04/15 03:01:32 [error] 22120#0: *84123983 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1586900001] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(96) < 172.105.57.57 [15/Apr/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586900098] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586900113] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586900121] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586900361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14255064 kB) free!
[1586900399] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.1% (44394964 kB) free.
[1586900481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32741560 kB) free.
[1586900721] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [15/Apr/2020:03:13:14 +0530] TCP 502 0 0 63.010 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 3.008, 60.001, 0.000"
[1586900841] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586901681] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16694088 kB) free!
[1586901802] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34825496 kB) free.
[1586902401] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15968272 kB) free!
[1586902521] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33191988 kB) free.
[1586903721] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18054636 kB) free!
[1586903841] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34976244 kB) free.
[1586904441] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17078196 kB) free!
[1586904561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33561628 kB) free.
[1586905761] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17237948 kB) free!
[1586905881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33769840 kB) free.
[1586907081] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14279800 kB) free!
[1586907114] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (277596 kB) free!
[1586907201] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34222412 kB) free.
[1586907234] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170052 kB) free!
[1586907354] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181876 kB) free!
[1586907806] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1586907855] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.3% (19432256 kB) free.
[1586907951] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 78.4% (19213080 kB) free.
[1586908062] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 46.5% (11471668 kB) free.
[1586908401] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15039948 kB) free!
[1586908522] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33442280 kB) free.
[1586910321] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20767468 kB) free!
[1586910441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35743180 kB) free.
[1586911662] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1623536 kB) free!
[1586911782] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166136 kB) free!
[1586911902] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165576 kB) free!
[1586912006] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.18, 5.31, 5.11
[1586912126] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.44, 5.11, 5.06
[1586912241] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19411316 kB) free!
[1586912246] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.82, 4.31, 4.78
[1586912361] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35643464 kB) free.
[1586912437] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 34.0% (44864828 kB) free.
[1586915361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19485840 kB) free!
[1586915481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34703944 kB) free.
[1586921846] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.20, 7.64, 6.10
[1586921966] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.46, 7.11, 6.08
[1586922087] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.89, 6.71, 6.06
[1586923287] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.98, 4.07, 4.68
[1586923881] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13057080 kB) free!
[1586923887] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.34, 5.67, 5.25
[1586924001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30189640 kB) free.
[1586924007] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.33, 6.14, 5.48
[1586924127] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.12, 5.48, 5.33
[1586926401] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (9025380 kB) free!
[1586926521] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27556424 kB) free.
[1586927151] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4338832 kB) free!
[1586927271] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4244592 kB) free!
[1586927391] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (4151224 kB) free!
[1586928321] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5597812 kB) free!
[1586928441] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.3% (25479380 kB) free!
[1586928561] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20506636 kB) free!
[1586929041] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(511) < 103.255.146.164 [15/Apr/2020:11:02:03 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1586929162] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586930332] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 70.0% (34485736 kB) free.
[1586930889] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.3% (3215400 kB) free!
[1586930961] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397172 kB) free!
[1586930991] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1887572 kB) free!
[1586931009] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.8% (2641164 kB) free!
[1586931129] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2096116 kB) free!
[1586931532] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.2% (7971200 kB) free!
[1586931652] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (7891240 kB) free!
[1586931772] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (7813004 kB) free!
[1586932161] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (19999528 kB) free!
[1586932275] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1586932330] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.5% (16268688 kB) free.
[1586932395] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1586932817] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181516 kB) free!
[1586932937] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (154272 kB) free!
[1586933058] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (250332 kB) free!
[1586933114] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/04/15 12:11:51 [error] 22135#0: *84221347 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1586933235] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586933961] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1586933961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26487336 kB) free.
[1586934172] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (985532 kB) free!
[1586935873] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1543480 kB) free!
[1586936361] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10243524 kB) free!
[1586936481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33570600 kB) free.
[1586936595] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1586936715] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1586937081] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8234032 kB) free!
[1586937201] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30792936 kB) free.
[1586937673] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (1814988 kB) free!
[1586938273] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1027996 kB) free!
[1586938401] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8663312 kB) free!
[1586938462] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1586938521] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31461204 kB) free.
[1586938583] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1586938703] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1586938956] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1586939055] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192468 kB) free!
[1586939076] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1586939121] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6188112 kB) free!
[1586939175] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183464 kB) free!
[1586939196] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1586939241] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29768136 kB) free.
[1586939295] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183184 kB) free!
[1586939303] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1586939617] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 42.2% (10355096 kB) free.
[1586939796] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1586939841] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (389464 kB) free!
[1586939903] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1586939962] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38235200 kB) free.
[1586940561] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6828540 kB) free!
[1586940673] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1586940673] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1679444 kB) free!
[1586940681] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18802892 kB) free!
[1586940801] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31460716 kB) free.
[1586941237] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.3% (25489192 kB) free!
[1586941273] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1524336 kB) free!
[1586941330] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3967524 kB) free!
[1586941357] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.0% (29013292 kB) free.
[1586941401] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1301296 kB) free!
[1586941450] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (3712136 kB) free!
[1586941515] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1586941521] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34938152 kB) free.
[1586941571] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3355840 kB) free!
[1586941635] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1586942121] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4267316 kB) free!
[1586942171] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1586942171] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2250892 kB) free!
[1586942241] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36082732 kB) free.
[1586942840] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 72.7% (17816788 kB) free.
[1586942841] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5851096 kB) free!
[1586942903] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1586942961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35492236 kB) free.
[1586943023] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1586943073] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (2141256 kB) free!
[1586943143] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1586943143] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1586943157] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.1% (22580160 kB) free!
[1586943277] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (28899864 kB) free.
[1586943371] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.0% (16166664 kB) free.
[1586943561] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8089500 kB) free!
[1586943561] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(155) < 92.42.104.166 [15/Apr/2020:15:00:22 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1586943681] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34304512 kB) free.
[1586943682] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586943954] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9400208 kB) free!
[1586944074] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (9224516 kB) free!
[1586944194] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (9046612 kB) free!
[1586944281] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8179664 kB) free!
[1586944401] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33233184 kB) free.
[1586944495] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9225112 kB) free!
[1586944615] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8991724 kB) free!
[1586944735] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.9% (8825412 kB) free!
[1586944873] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (305192 kB) free!
[1586945001] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5191248 kB) free!
[1586945077] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.7% (19374252 kB) free!
[1586945122] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34255532 kB) free.
[1586945197] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.0% (22349376 kB) free!
[1586945317] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.6% (29798356 kB) free.
[1586945481] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 172.105.59.173 [15/Apr/2020:15:35:25 +0530] TCP 502 0 0 180.002 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 60.000, 60.001, 60.001"
[1586945601] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586945721] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8085272 kB) free!
[1586945841] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1586945841] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33844500 kB) free.
[1586946073] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3688728 kB) free.
[1586946441] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8070756 kB) free!
[1586946517] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.8% (23519580 kB) free!
[1586946561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33585240 kB) free.
[1586946637] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.8% (28676884 kB) free.
[1586947095] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 36.1% (8855016 kB) free.
[1586947161] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6367848 kB) free!
[1586947237] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.2% (18767380 kB) free!
[1586947273] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.0% (2273292 kB) free!
[1586947281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33480516 kB) free.
[1586947357] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.9% (23539632 kB) free!
[1586947393] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.1% (2290256 kB) free!
[1586947477] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26524100 kB) free.
[1586947513] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2694940 kB) free!
[1586947695] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196188 kB) free!
[1586947735] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4728244 kB) free!
[1586947815] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184568 kB) free!
[1586947881] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6349292 kB) free!
[1586947935] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192620 kB) free!
[1586948001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32659764 kB) free.
[1586948077] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.8% (19539384 kB) free!
[1586948197] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.1% (27769100 kB) free.
[1586948394] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4664832 kB) free!
[1586948601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8049648 kB) free!
[1586948601] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(12645) < 203.212.70.170 [15/Apr/2020:16:33:13 +0530] TCP 502 0 0 78.032 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "15.024, 3.008, 60.000, 0.000"
[1586948697] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(726) < 172.105.57.57 [15/Apr/2020:16:34:48 +0530] TCP 502 0 0 15.023 "10.147.212.37:4000, backend3" "0, 0" "0, 0" "15.023, 0.000"
[1586948721] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32202440 kB) free.
[1586948721] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1780) < 182.18.144.246 [15/Apr/2020:16:35:10 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586948818] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(53) < 103.16.101.72 [15/Apr/2020:16:36:43 +0530] TCP 502 0 0 150.054 "10.147.212.38:4000, 10.147.212.49:4000, 10.147.212.61:4000, 10.147.212.37:4000" "0, 0, 0, 0" "0, 0, 0, 0" "15.030, 60.000, 15.024, 60.000"
[1586948834] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2575) < 2020/04/15 16:37:13 [error] 22131#0: *84276374 no live upstreams while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1586948836] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1586948842] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586948938] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;HARD;3;(310) < 49.50.86.152 [15/Apr/2020:16:37:39 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1586948954] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(408) < 2020/04/15 16:38:38 [error] 22131#0: *84277172 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1586948956] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1586949068] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1586949074] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586949188] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1586949308] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1586949321] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6657284 kB) free!
[1586949397] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.6% (21940172 kB) free!
[1586949441] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 172.105.59.173 [15/Apr/2020:16:46:32 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1586949441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31581512 kB) free.
[1586949517] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (28806360 kB) free.
[1586949539] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586949561] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586949743] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1586950041] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (4999824 kB) free!
[1586950117] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.9% (18293448 kB) free!
[1586950161] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11746204 kB) free!
[1586950237] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.0% (23758792 kB) free!
[1586950281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27129644 kB) free.
[1586950357] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.7% (16700376 kB) free!
[1586950756] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1586950876] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1586950882] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (379508 kB) free!
[1586950935] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1586950935] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (3667012 kB) free!
[1586950943] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1586951002] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30580456 kB) free.
[1586951063] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1586951183] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1586951536] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (205572 kB) free!
[1586951557] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27146408 kB) free.
[1586951602] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3199784 kB) free!
[1586951722] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30765264 kB) free.
[1586952157] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.0% (17157972 kB) free!
[1586952278] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.8% (24842084 kB) free!
[1586952322] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2602900 kB) free!
[1586952371] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 10.8% (2429780 kB) free!
[1586952383] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1586952398] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.5% (16518664 kB) free!
[1586952442] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378904 kB) free!
[1586952491] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (1926760 kB) free!
[1586952562] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32487808 kB) free.
[1586952611] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1451540 kB) free!
[1586952736] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 29.9% (7316664 kB) free.
[1586952908] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1586952983] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1586953103] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1586953162] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3774028 kB) free!
[1586953223] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1586953282] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16224776 kB) free!
[1586953336] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (205004 kB) free!
[1586953402] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29064336 kB) free.
[1586953456] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181096 kB) free!
[1586953513] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (487620 kB) free!
[1586953576] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188984 kB) free!
[1586953811] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.5% (16279588 kB) free.
[1586953823] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1586954002] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1618524 kB) free!
[1586954113] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3048240 kB) free!
[1586954122] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32262856 kB) free.
[1586954176] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (3509620 kB) free!
[1586954423] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1586954543] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1586954663] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1586954722] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2602440 kB) free!
[1586954776] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182084 kB) free!
[1586954840] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.2% (4463528 kB) free!
[1586954842] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2167728 kB) free!
[1586954960] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.6% (4302036 kB) free!
[1586954963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29837184 kB) free.
[1586955080] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.9% (4146028 kB) free!
[1586955263] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1586955562] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2642752 kB) free!
[1586955682] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32515644 kB) free.
[1586955913] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (485340 kB) free!
[1586956280] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2320892 kB) free!
[1586956282] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (4977588 kB) free!
[1586956402] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (424808 kB) free!
[1586956522] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32083280 kB) free.
[1586957113] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3777596 kB) free.
[1586957122] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3638280 kB) free!
[1586957176] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 30.7% (7518464 kB) free.
[1586957242] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403372 kB) free!
[1586957362] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26723512 kB) free.
[1586957798] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1586957798] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (10876148 kB) free!
[1586957962] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (438436 kB) free!
[1586958082] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30943860 kB) free.
[1586958263] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1586958313] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2156752 kB) free!
[1586958376] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1899260 kB) free!
[1586958384] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1586958398] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.7% (21965548 kB) free!
[1586958433] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.2% (2310312 kB) free!
[1586958496] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199152 kB) free!
[1586958504] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1586958554] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (2551176 kB) free!
[1586958616] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169252 kB) free!
[1586958682] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4800636 kB) free!
[1586958802] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32048060 kB) free.
[1586959104] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1586959402] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5912824 kB) free!
[1586959522] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26525044 kB) free.
[1586959754] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (752280 kB) free!
[1586960123] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7825932 kB) free!
[1586960243] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3615492 kB) free!
[1586960354] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (1890492 kB) free!
[1586960363] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30872152 kB) free.
[1586960727] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.41, 3.38, 4.35
[1586960963] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7065924 kB) free!
[1586961083] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2771144 kB) free!
[1586961203] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30415648 kB) free.
[1586961676] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1586961796] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1586961803] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5734780 kB) free!
[1586961923] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1442732 kB) free!
[1586962043] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27890944 kB) free.
[1586962154] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1027288 kB) free!
[1586962216] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 31.3% (7669624 kB) free.
[1586962643] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2574072 kB) free!
[1586962754] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (1890052 kB) free!
[1586962763] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30402664 kB) free.
[1586963363] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3664204 kB) free!
[1586963483] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401976 kB) free!
[1586963603] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30585340 kB) free.
[1586963954] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.9% (3885536 kB) free.
[1586964203] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2648332 kB) free!
[1586964324] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12386468 kB) free!
[1586964398] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27174004 kB) free.
[1586964444] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29721000 kB) free.
[1586964554] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.6% (2052868 kB) free!
[1586964611] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1586964674] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.6% (2214504 kB) free!
[1586964794] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (2373192 kB) free!
[1586964927] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.67, 6.55, 5.56
[1586964998] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.8% (19558312 kB) free!
[1586965044] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4053952 kB) free!
[1586965047] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.85, 6.39, 5.63
[1586965119] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.7% (25998568 kB) free!
[1586965164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32265444 kB) free.
[1586965167] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.03, 6.37, 5.72
[1586965239] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.2% (20026128 kB) free!
[1586965396] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1586965516] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1586965636] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1586965764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9669468 kB) free!
[1586965884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5542540 kB) free!
[1586966004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30845972 kB) free.
[1586966604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7636876 kB) free!
[1586966724] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31783700 kB) free.
[1586967324] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10579004 kB) free!
[1586967444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6610476 kB) free!
[1586967564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31110436 kB) free.
[1586967924] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1620 processes
[1586968164] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8852828 kB) free!
[1586968284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31500640 kB) free.
[1586968394] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (432264 kB) free!
[1586968524] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1819 processes
[1586968884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9775816 kB) free!
[1586969005] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5814420 kB) free!
[1586969125] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31188708 kB) free.
[1586969725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7540252 kB) free!
[1586969725] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1698 processes
[1586969845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31130240 kB) free.
[1586970325] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1586970325] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1821 processes
[1586970445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8919464 kB) free!
[1586970504] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586970565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4983260 kB) free!
[1586970616] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1586970685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29905920 kB) free.
[1586970794] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.4% (2177300 kB) free!
[1586971161] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(178) < 94.237.68.73 [15/Apr/2020:22:44:59 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1586971167] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.17, 4.04, 4.68
[1586971239] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.4% (32151792 kB) free.
[1586971274] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(347) < 2020/04/15 22:50:43 [error] 22134#0: *84333262 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1586971282] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586971285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6869164 kB) free!
[1586971395] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586971405] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30762888 kB) free.
[1586971740] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(278) < 172.105.57.57 [15/Apr/2020:22:49:46 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1586971860] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586972005] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8329568 kB) free!
[1586972125] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4534192 kB) free!
[1586972245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28649272 kB) free.
[1586972845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6115592 kB) free!
[1586972965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30426784 kB) free.
[1586973194] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (796688 kB) free!
[1586973565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7447236 kB) free!
[1586973685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30288796 kB) free.
[1586973794] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.1% (1964784 kB) free!
[1586974286] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8805188 kB) free!
[1586974406] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4934368 kB) free!
[1586974526] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28348780 kB) free.
[1586975126] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6116012 kB) free!
[1586975246] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30010672 kB) free.