[1565462444] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3724132 kB) free!
[1565462510] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 390 processes
[1565462564] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (3700100 kB) free!
[1565462684] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (3676256 kB) free!
[1565463094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.04, 28.27, 19.86
[1565463110] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 5 warn out of 388 processes
[1565463260] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1565463350] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.71, 8.38, 14.60
[1565463380] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1565463559] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 117.20, 38.01, 21.46
[1565463846] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 320 processes
[1565463966] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1565464123] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565464159] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.07, 5.93, 11.67
[1565464294] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.35, 5.54, 11.33
[1565465031] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.9% (10591744 kB) free.
[1565465150] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.34, 26.41, 19.58
[1565465684] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1835244 kB) free!
[1565465775] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6532384 kB) free!
[1565465895] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11676296 kB) free.
[1565466011] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565466350] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.86, 8.33, 13.89
[1565466495] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7087176 kB) free!
[1565466615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11665428 kB) free.
[1565466710] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 395 processes
[1565467211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565467215] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7511092 kB) free!
[1565467311] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1565467311] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 4 warn out of 374 processes
[1565467331] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565467335] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11664588 kB) free.
[1565467580] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565467700] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1565467931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565467935] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7934788 kB) free!
[1565468052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565468055] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6291612 kB) free!
[1565468172] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565468175] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11653312 kB) free.
[1565468772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565468775] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6549292 kB) free!
[1565468892] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565468896] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11661420 kB) free.
[1565468959] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 125.30, 31.23, 18.40
[1565469492] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565469495] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7064112 kB) free!
[1565469559] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.52, 4.77, 9.97
[1565469612] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565469615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11648772 kB) free.
[1565470212] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565470215] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7538572 kB) free!
[1565470332] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565470335] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11659588 kB) free.
[1565470431] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198012 kB) free!
[1565470551] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195128 kB) free!
[1565470671] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191548 kB) free!
[1565470933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565470935] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7694732 kB) free!
[1565471053] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565471055] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11656584 kB) free.
[1565471562] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.1% (41934908 kB) free.
[1565471653] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565471655] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8036836 kB) free!
[1565471773] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565471775] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6323168 kB) free!
[1565471893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565471895] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11001476 kB) free.
[1565472094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 50.01, 31.16, 18.52
[1565472350] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.40, 12.62, 8.24
[1565472493] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565472500] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 399 processes
[1565472620] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1565472950] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.06, 8.30, 10.68
[1565473095] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565473294] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.70, 8.23, 13.32
[1565473484] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 45.5% (11147064 kB) free.
[1565473594] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.8% (10739864 kB) free.
[1565473693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565473696] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8007392 kB) free!
[1565473759] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.46, 25.88, 18.64
[1565473814] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565473816] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6215860 kB) free!
[1565473934] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565473936] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11651216 kB) free.
[1565474959] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 4.87, 10.27
[1565475134] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565475254] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565475350] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.67, 24.70, 20.33
[1565475374] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1565475374] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565475711] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 377 processes
[1565476336] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9244292 kB) free!
[1565476456] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7980660 kB) free!
[1565476576] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6285588 kB) free!
[1565476911] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 371 processes
[1565477150] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.51, 7.94, 14.39
[1565477776] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10064844 kB) free.
[1565478094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 72.24, 31.42, 20.77
[1565478111] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 383 processes
[1565478376] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9193916 kB) free!
[1565478497] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7924988 kB) free!
[1565478559] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 39.26, 30.01, 19.56
[1565478617] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6128504 kB) free!
[1565478694] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 4.72, 11.14
[1565478711] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 370 processes
[1565479159] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.82, 4.81, 10.66
[1565479217] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1565479217] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11289684 kB) free.
[1565480174] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565480417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9356320 kB) free!
[1565480537] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8346348 kB) free!
[1565480657] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7145200 kB) free!
[1565481257] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11624900 kB) free.
[1565481350] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 21.51, 27.21, 17.59
[1565481374] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565481494] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565481615] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565483057] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9822020 kB) free!
[1565483150] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.50, 6.79, 12.80
[1565483177] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8555544 kB) free!
[1565483298] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7127868 kB) free!
[1565483415] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565483511] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 373 processes
[1565483959] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 48.77, 31.93, 21.19
[1565484015] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565484094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 57.29, 25.31, 13.55
[1565484135] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565484255] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565484498] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11354792 kB) free.
[1565484620] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565484694] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 3.84, 7.34
[1565484711] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 370 processes
[1565484741] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565485199] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 610 processes
[1565485259] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 610 processes
[1565485294] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.67, 1.05, 4.11
[1565485594] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4832000 kB) free!
[1565485698] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7349524 kB) free!
[1565485714] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4857756 kB) free!
[1565485759] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 27.75, 24.61, 19.37
[1565485818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11601432 kB) free.
[1565485834] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (4823396 kB) free!
[1565485894] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.51, 5.27, 7.36
[1565486014] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.42, 3.68, 6.52
[1565486134] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 105.03, 35.97, 17.40
[1565487559] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 152.66, 40.89, 21.23
[1565488159] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.92, 6.27, 11.54
[1565488818] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9450936 kB) free!
[1565488834] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185656 kB) free!
[1565488938] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8224988 kB) free!
[1565489059] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6609436 kB) free!
[1565489084] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (252548 kB) free!
[1565489134] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.12, 6.82, 13.42
[1565489150] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.96, 37.46, 26.65
[1565489204] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (236404 kB) free!
[1565489323] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565489324] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (205084 kB) free!
[1565489443] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1565489563] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1565489655] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565489659] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11602716 kB) free.
[1565491159] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.49, 19.62, 23.17
[1565491455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565491459] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7713316 kB) free!
[1565491575] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565491579] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11591300 kB) free.
[1565492150] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.06, 6.93, 14.96
[1565492159] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 617 processes
[1565492175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565492179] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8030932 kB) free!
[1565492219] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 622 processes
[1565492295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565492299] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6367168 kB) free!
[1565492415] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565492419] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11586544 kB) free.
[1565493015] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565493019] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7028960 kB) free!
[1565493135] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565493139] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11593992 kB) free.
[1565493560] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.69, 9.07, 17.66
[1565493736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565493739] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7585060 kB) free!
[1565493856] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565493860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11594140 kB) free.
[1565494456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565494459] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7956952 kB) free!
[1565494534] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 35.25, 30.94, 20.59
[1565494576] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565494579] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6169344 kB) free!
[1565494696] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565494699] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11590316 kB) free.
[1565495134] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 4.70, 11.08
[1565495296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565495299] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6772412 kB) free!
[1565495361] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.92, 26.20, 24.53
[1565495416] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565495419] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11582148 kB) free.
[1565495511] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565496016] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565496019] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7324588 kB) free!
[1565496136] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565496139] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11589840 kB) free.
[1565496737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565496739] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7454292 kB) free!
[1565496857] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565496859] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11592096 kB) free.
[1565497161] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.01, 5.43, 13.02
[1565497457] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565497459] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7743576 kB) free!
[1565497534] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 161.16, 39.47, 20.82
[1565497577] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1565497577] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565497579] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11583972 kB) free.
[1565498134] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.82, 5.89, 11.20
[1565498150] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.44, 32.47, 21.88
[1565498177] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565498179] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8023244 kB) free!
[1565498297] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565498300] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6222200 kB) free!
[1565498417] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565498420] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11579508 kB) free.
[1565498750] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.10, 6.82, 12.79
[1565498917] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.1% (23053296 kB) free.
[1565498973] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39708940 kB) free.
[1565499017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565499020] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6590088 kB) free!
[1565499028] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.7% (22723988 kB) free.
[1565499073] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.4% (23059388 kB) free.
[1565499129] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.7% (14097212 kB) free.
[1565499137] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565499140] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11592932 kB) free.
[1565499584] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216848 kB) free!
[1565499705] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235012 kB) free!
[1565499733] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (174384 kB) free!
[1565499738] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565499740] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6793496 kB) free!
[1565499825] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248672 kB) free!
[1565499854] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166516 kB) free!
[1565499858] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565499860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11590028 kB) free.
[1565499974] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175056 kB) free!
[1565500161] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 59.27, 27.60, 19.39
[1565500311] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 378 processes
[1565500458] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565500460] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7822760 kB) free!
[1565500578] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565500580] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11583668 kB) free.
[1565500761] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.21, 4.66, 10.65
[1565500911] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 369 processes
[1565501134] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 44.15, 31.83, 20.53
[1565501150] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.32, 26.80, 21.16
[1565501178] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565501180] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8848976 kB) free!
[1565501298] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565501300] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7713940 kB) free!
[1565501418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565501420] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6184376 kB) free!
[1565501734] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.81, 4.84, 11.06
[1565502140] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565502261] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565502621] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10820944 kB) free.
[1565502950] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.91, 7.17, 13.62
[1565503218] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565503818] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565503821] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6883852 kB) free!
[1565503938] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565503941] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11586092 kB) free.
[1565504538] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565504541] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7711180 kB) free!
[1565504658] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565504661] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11588772 kB) free.
[1565504734] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 134.21, 38.14, 20.86
[1565505258] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565505261] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8034548 kB) free!
[1565505334] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.94, 5.85, 11.32
[1565505379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565505381] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6273060 kB) free!
[1565505499] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565505501] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11587472 kB) free.
[1565506099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565506101] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7354304 kB) free!
[1565506219] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565506221] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11578280 kB) free.
[1565506550] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 17.60, 28.37, 21.19
[1565506819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565506821] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7748224 kB) free!
[1565506939] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565506942] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11589204 kB) free.
[1565506963] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565507199] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1565507539] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565507541] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7954928 kB) free!
[1565507659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565507661] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6091284 kB) free!
[1565507750] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.42, 7.83, 14.23
[1565507780] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565507781] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11575748 kB) free.
[1565507961] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 29.92, 27.84, 18.92
[1565508036] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 363 processes
[1565508156] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 359 processes
[1565508276] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1565508379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565508381] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6878612 kB) free!
[1565508499] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565508501] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11580460 kB) free.
[1565508561] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.23, 4.60, 10.39
[1565508876] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1565509099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565509101] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7659860 kB) free!
[1565509219] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565509221] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11572332 kB) free.
[1565509550] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 84.18, 34.53, 17.38
[1565509819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565509821] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7987444 kB) free!
[1565509939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565509941] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6219632 kB) free!
[1565510059] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565510061] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11583904 kB) free.
[1565510076] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 2 warn out of 362 processes
[1565510150] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.71, 6.91, 10.35
[1565510197] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 2 warn out of 366 processes
[1565510317] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 363 processes
[1565510660] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565510661] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6941364 kB) free!
[1565510668] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.69, 6.72, 5.59
[1565510780] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565510781] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11574352 kB) free.
[1565510788] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.72, 5.68, 5.35
[1565510908] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.14, 5.55, 5.34
[1565510961] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 17.11, 26.49, 20.33
[1565511380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565511382] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7572056 kB) free!
[1565511500] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565511502] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11583080 kB) free.
[1565511934] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 43.48, 25.21, 16.71
[1565512100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565512102] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7965952 kB) free!
[1565512109] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.08, 3.49, 4.57
[1565512161] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.25, 4.71, 11.69
[1565512220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565512222] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11459028 kB) free.
[1565512284] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199176 kB) free!
[1565512340] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565512404] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193792 kB) free!
[1565512524] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188316 kB) free!
[1565512534] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.14, 5.32, 9.82
[1565512717] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1565512717] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1565512822] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8264948 kB) free!
[1565512940] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565512942] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6676400 kB) free!
[1565513061] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565513062] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11580112 kB) free.
[1565513150] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.25, 32.64, 22.47
[1565513660] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565513662] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7480248 kB) free!
[1565513750] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.59, 6.75, 13.09
[1565513780] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565513782] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11575208 kB) free.
[1565513961] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 52.43, 37.75, 22.02
[1565514334] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 18.66, 25.77, 17.88
[1565514380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565514382] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7876416 kB) free!
[1565514500] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565514502] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11581872 kB) free.
[1565514561] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.03, 5.88, 11.99
[1565514935] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 4.06, 9.71
[1565514950] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.81, 26.47, 20.54
[1565515100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565515102] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8352972 kB) free!
[1565515220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565515223] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6948164 kB) free!
[1565515340] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565515343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11578968 kB) free.
[1565515551] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.43, 6.15, 12.14
[1565515911] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 380 processes
[1565515941] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565515943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7545836 kB) free!
[1565516061] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565516063] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565516063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11579900 kB) free.
[1565516511] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 4 warn out of 387 processes
[1565516661] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565516663] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7983536 kB) free!
[1565516781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565516783] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6124308 kB) free!
[1565516901] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565516903] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11578148 kB) free.
[1565517335] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.49, 29.05, 19.55
[1565517501] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565517503] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6496892 kB) free!
[1565517561] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.00, 32.34, 22.28
[1565517621] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565517623] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11571504 kB) free.
[1565517935] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 17.98, 23.28, 18.63
[1565518161] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.08, 5.31, 12.16
[1565518221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565518223] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6797616 kB) free!
[1565518341] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565518343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11578280 kB) free.
[1565518362] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9562200 kB) free!
[1565518482] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9441788 kB) free!
[1565518602] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9338044 kB) free!
[1565518942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565518943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7316824 kB) free!
[1565519062] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565519064] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11577312 kB) free.
[1565519135] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 48.89, 25.75, 17.96
[1565519151] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.98, 23.26, 21.10
[1565519662] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565519663] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7595112 kB) free!
[1565519735] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.74, 4.10, 9.76
[1565519782] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565519783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11570160 kB) free.
[1565520351] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.90, 6.45, 12.60
[1565520382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565520383] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7932768 kB) free!
[1565520502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565520503] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11450120 kB) free.
[1565520622] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565521103] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8362740 kB) free!
[1565521161] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 27.86, 28.57, 19.96
[1565521222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565521223] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6786320 kB) free!
[1565521342] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565521343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11570584 kB) free.
[1565521460] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.2% (14340924 kB) free.
[1565521504] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 79.1% (19519616 kB) free.
[1565521599] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1565521672] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 63.1% (31102864 kB) free.
[1565521719] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1565521761] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.98, 4.48, 10.82
[1565521943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565521943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7105752 kB) free!
[1565522060] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (559720 kB) free!
[1565522063] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565522063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11573972 kB) free.
[1565522180] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173576 kB) free!
[1565522300] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (153584 kB) free!
[1565522317] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1565522663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565522663] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7544392 kB) free!
[1565522735] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 112.61, 31.99, 18.45
[1565522751] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.97, 25.94, 20.22
[1565522783] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565522783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11565900 kB) free.
[1565522919] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1565523039] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1565523159] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1565523304] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171412 kB) free!
[1565523335] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.78, 5.33, 10.38
[1565523351] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.79, 6.16, 12.01
[1565523383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565523383] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7764732 kB) free!
[1565523424] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178656 kB) free!
[1565523503] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565523504] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11574224 kB) free.
[1565523544] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (373244 kB) free!
[1565524103] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8055724 kB) free!
[1565524103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565524223] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6225704 kB) free!
[1565524223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565524343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11571168 kB) free.
[1565524344] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565524602] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4746912 kB) free!
[1565524685] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (260288 kB) free!
[1565524805] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225516 kB) free!
[1565524911] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 388 processes
[1565524925] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233096 kB) free!
[1565524943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565524943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6549652 kB) free!
[1565525063] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565525063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11569916 kB) free.
[1565525511] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1565525511] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 369 processes
[1565525663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565525663] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6759488 kB) free!
[1565525736] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 42.58, 30.14, 17.27
[1565525783] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565525783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11562652 kB) free.
[1565526336] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 4.65, 9.37
[1565526383] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7558328 kB) free!
[1565526383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565526503] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11564432 kB) free.
[1565526503] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565527103] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8002552 kB) free!
[1565527103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565527223] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6215564 kB) free!
[1565527224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565527343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10844724 kB) free.
[1565527344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565527460] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1565527551] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.17, 23.47, 20.78
[1565527580] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1565528151] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.46, 21.08, 19.59
[1565528751] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.91, 24.58, 21.04
[1565529144] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565529351] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.98, 19.33, 19.37
[1565529561] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 44.74, 12.15, 11.89
[1565529711] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 390 processes
[1565529743] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7037264 kB) free!
[1565529744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565529863] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11555856 kB) free.
[1565529864] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565529951] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.80, 31.54, 23.52
[1565530161] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 4.55, 9.59
[1565530311] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 4 warn out of 372 processes
[1565530463] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7426860 kB) free!
[1565530464] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565530551] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.80, 6.56, 13.57
[1565530583] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11548636 kB) free.
[1565530584] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565531183] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7655484 kB) free!
[1565531185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565531303] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11556072 kB) free.
[1565531305] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565531736] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 72.80, 35.87, 21.79
[1565531903] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7848368 kB) free!
[1565531905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565532023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11549536 kB) free.
[1565532025] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565532286] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.7% (14578716 kB) free.
[1565532327] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.7% (15140352 kB) free.
[1565532359] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 70.0% (34498180 kB) free.
[1565532391] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.8% (22404584 kB) free.
[1565532623] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7444796 kB) free!
[1565532625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565532744] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11552732 kB) free.
[1565532745] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565532936] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 4.40, 10.94
[1565533343] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7810368 kB) free!
[1565533345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565533463] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11557140 kB) free.
[1565533465] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565533591] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183084 kB) free!
[1565533712] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173640 kB) free!
[1565533832] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (402548 kB) free!
[1565534063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8051460 kB) free!
[1565534066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565534180] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1565534183] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6259072 kB) free!
[1565534186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565534300] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1565534303] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11552712 kB) free.
[1565534306] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565534751] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 29.04, 35.43, 24.20
[1565534903] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6579080 kB) free!
[1565534906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565534962] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 105.64, 30.20, 18.43
[1565535023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11542192 kB) free.
[1565535026] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565535294] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (150620 kB) free!
[1565535363] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232876 kB) free!
[1565535414] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (156016 kB) free!
[1565535483] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214840 kB) free!
[1565535534] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167440 kB) free!
[1565535562] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.77, 4.68, 9.99
[1565535603] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243828 kB) free!
[1565535623] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6975264 kB) free!
[1565535626] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565535743] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11550568 kB) free.
[1565535746] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565535936] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 86.70, 21.81, 15.27
[1565535951] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.63, 7.15, 14.47
[1565536343] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7406392 kB) free!
[1565536347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565536463] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11553184 kB) free.
[1565536467] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565536536] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.83, 4.76, 9.88
[1565536700] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565536821] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1565537064] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7659988 kB) free!
[1565537067] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565537184] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11544660 kB) free.
[1565537187] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565537784] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7750324 kB) free!
[1565537787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565537904] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11553812 kB) free.
[1565537907] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565538504] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8150864 kB) free!
[1565538507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565538624] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6548528 kB) free!
[1565538627] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565538744] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11551168 kB) free.
[1565538747] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565539162] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 58.57, 26.95, 18.03
[1565539311] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 380 processes
[1565539344] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6946356 kB) free!
[1565539348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565539464] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11553248 kB) free.
[1565539468] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565539762] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.59, 4.08, 9.68
[1565539911] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 4 warn out of 373 processes
[1565540064] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7270628 kB) free!
[1565540068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565540151] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 29.30, 31.57, 21.85
[1565540184] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11540760 kB) free.
[1565540188] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565540784] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7502720 kB) free!
[1565540788] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565540905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11551676 kB) free.
[1565540908] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565541336] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.90, 4.81
[1565541504] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7727600 kB) free!
[1565541508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565541624] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11542132 kB) free.
[1565541628] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565541936] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.71, 7.81, 8.57
[1565542056] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.40, 5.35, 7.58
[1565542162] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 29.71, 27.93, 18.37
[1565542176] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 46.16, 13.25, 9.80
[1565542224] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7871040 kB) free!
[1565542229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565542344] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11553492 kB) free.
[1565542349] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565542776] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.91, 6.24, 10.25
[1565542912] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 380 processes
[1565542944] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8192080 kB) free!
[1565542949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565543064] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6545980 kB) free!
[1565543069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565543184] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11542124 kB) free.
[1565543189] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565543784] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6828904 kB) free!
[1565543789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565543904] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11550716 kB) free.
[1565543909] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565543962] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.40, 5.99, 12.39
[1565544112] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 4 warn out of 375 processes
[1565544163] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565544283] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1565544403] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1565544504] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7107444 kB) free!
[1565544509] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565544620] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565544624] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11539892 kB) free.
[1565544629] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565544740] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1565545225] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7217568 kB) free!
[1565545229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1565545345] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11548056 kB) free.
[1565545350] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565545945] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7225900 kB) free!
[1565545949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565546065] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11548924 kB) free.
[1565546069] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565546376] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 51.44, 26.89, 18.33
[1565546665] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7306748 kB) free!
[1565546669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565546751] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.72, 4.57, 12.12
[1565546785] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11540872 kB) free.
[1565546789] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565546962] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 146.14, 38.72, 20.02
[1565546976] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.85, 4.17, 9.89
[1565547112] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 380 processes
[1565547385] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7609016 kB) free!
[1565547389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565547403] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565547505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11548956 kB) free.
[1565547509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565547713] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 385 processes
[1565548105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7722836 kB) free!
[1565548109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes