[1595788618] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595788738] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 7328 open files (0% of max 2424156)
[1595788856] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (2450124 kB) free!
[1595788920] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 9.87, 31.36, 41.71
[1595788976] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (2405352 kB) free!
[1595789096] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2392732 kB) free!
[1595789696] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.8% (21178288 kB) free.
[1595790120] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1595790120] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 60.16, 52.50, 44.56
[1595792477] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197188 kB) free!
[1595792597] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191272 kB) free!
[1595792718] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.9% (21310036 kB) free.
[1595793317] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (2126792 kB) free!
[1595793372] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.0% (14493384 kB) free!
[1595793437] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167796 kB) free!
[1595793492] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4901196 kB) free!
[1595793557] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187016 kB) free!
[1595793593] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234604 kB) free!
[1595793612] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (453020 kB) free!
[1595793713] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214856 kB) free!
[1595793833] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233712 kB) free!
[1595794157] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2506332 kB) free!
[1595794757] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.4% (20930344 kB) free.
[1595795521] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 54.32, 44.22, 44.23
[1595795633] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 26.0% (12811580 kB) free.
[1595795822] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4399884 kB) free!
[1595795942] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.5% (21199924 kB) free.
[1595795957] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (216300 kB) free!
[1595796077] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206244 kB) free!
[1595796197] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1595796197] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191068 kB) free!
[1595797305] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 80.5% (39684056 kB) free.
[1595797397] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.4% (21420772 kB) free.
[1595797921] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 68.91, 48.33, 44.33
[1595799721] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 33.58, 41.47, 44.59
[1595799797] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195240 kB) free!
[1595799918] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 82.6% (20249668 kB) free.
[1595800321] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 62.26, 46.25, 45.32
[1595800517] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163020 kB) free!
[1595800637] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188644 kB) free!
[1595800757] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193044 kB) free!
[1595800921] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 52.67, 45.52, 44.60
[1595801357] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.9% (22277128 kB) free.
[1595802121] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 65.22, 50.43, 45.97
[1595802721] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 38.30, 42.38, 44.23
[1595803321] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 52.13, 49.38, 46.51
[1595803921] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 32.59, 38.81, 42.17
[1595805721] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 60.06, 46.88, 43.15
[1595806321] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 43.63, 47.80, 44.68
[1595811721] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 123.58, 72.27, 55.99
[1595812757] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595812812] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 46.4% (61148916 kB) free.
[1595818321] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1595821238] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 319 processes
[1595821358] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1595826558] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16119.13, 13521.91, 7159.09
[1595826672] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15067.21, 6831.11, 2621.64
[1595826678] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16121.23, 14381.97, 8245.74
[1595826792] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15951.37, 9894.43, 4254.60
[1595826798] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16122.45, 14958.05, 9200.86
[1595826851] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16083.50, 13122.09, 6751.03
[1595826912] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16071.00, 11944.18, 5689.55
[1595826971] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16086.80, 14103.21, 7883.05
[1595827091] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16087.47, 14759.74, 8877.82
[1595828558] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 315 processes
[1595828566] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1595828678] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1595828686] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 329 processes
[1595828807] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 333 processes
[1595829798] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.68, 3.23, 412.28
[1595830512] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.46, 2.53, 397.67
[1595830691] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.68, 2.07, 262.98
[1595833012] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.8% (6480136 kB) free!
[1595833132] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.5% (6372924 kB) free!
[1595833252] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.2% (6279528 kB) free!
[1595833515] SERVICE ALERT: TempHP1;zombie_procs;UNKNOWN;SOFT;1;NRPE: Command 'check_zombie_procs' not defined
[1595833607] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1595833635] SERVICE ALERT: TempHP1;zombie_procs;UNKNOWN;SOFT;2;NRPE: Command 'check_zombie_procs' not defined
[1595833711] SERVICE ALERT: TempHP1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1595835548] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595835552] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595835618] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595835621] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595835668] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595835668] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1595835672] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1595835738] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1595835746] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=2.95% system=5.87% iowait=0.00% idle=91.19%
[1595835761] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;(21) < Jul 27 13:11:03 pservercluster2 xinetd[1487]: nrpe: fork failed: Cannot allocate memory (errno = 12)
[1595835788] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 330 processes
[1595835788] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 7488 open files (0% of max 2424156)
[1595835792] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 43.5% (10659688 kB) free.
[1595835881] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595836114] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595836115] SERVICE ALERT: sservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595836167] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595836234] SERVICE ALERT: sservercluster1;Total Processes;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595836235] SERVICE ALERT: sservercluster1;zombie_procs;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595836287] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 78.7% (19290704 kB) free.
[1595836354] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 332 processes
[1595836355] SERVICE ALERT: sservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1595836555] SERVICE ALERT: sservercluster1;Messages Log Status;CRITICAL;SOFT;1;(19) < Jul 27 13:20:17 sservercluster1 ksmtuned: /usr/sbin/ksmtuned: fork: Cannot allocate memory
[1595836675] SERVICE ALERT: sservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595836905] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9481084 kB) free!
[1595837025] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9207580 kB) free!
[1595837145] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8936412 kB) free!
[1595837598] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15584.58, 7772.96, 3072.42
[1595837718] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16124.30, 10563.94, 4665.11
[1595837838] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16196.05, 12431.18, 6064.59
[1595838052] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (3143764 kB) free!
[1595838652] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 12.4% (4066700 kB) free!
[1595838798] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595838919] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 7872 open files (0% of max 2424156)
[1595839545] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3801876 kB) free!
[1595839852] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1595839852] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2807684 kB) free!
[1595840425] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1595840838] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.92, 4.93, 676.69
[1595841313] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4199448 kB) free!
[1595841433] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.1% (3829520 kB) free!
[1595841438] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1595841438] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.69, 2.64, 357.49
[1595841553] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (3453864 kB) free!
[1595842153] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1842872 kB) free!
[1595842753] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 95.2% (21366344 kB) free.
[1595845512] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15780.77, 8770.80, 3602.25
[1595845632] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16054.65, 11194.96, 5117.27
[1595845752] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16091.95, 12817.25, 6448.68
[1595845892] SERVICE ALERT: sservercluster1;Memory;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595845971] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595845971] SERVICE ALERT: sservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595845971] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595845985] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595845992] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595846012] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595846022] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595846045] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595846091] SERVICE ALERT: sservercluster1;Total Processes;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595846091] SERVICE ALERT: sservercluster1;zombie_procs;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595846091] SERVICE ALERT: sservercluster1;HDD_Home status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595846105] SERVICE ALERT: sservercluster1;CPU_Procs;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595846112] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595846132] SERVICE ALERT: sservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38684 MB (75% inode=99%):
[1595846142] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 62.0% (15195112 kB) free.
[1595846170] SERVICE ALERT: sservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.56% system=12.90% iowait=0.04% idle=85.49%
[1595846211] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 166901 MB (75% inode=99%):
[1595846211] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 393 processes
[1595846211] SERVICE ALERT: sservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1595846225] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 392 processes
[1595846276] SERVICE ALERT: sservercluster1;Messages Log Status;CRITICAL;SOFT;1;(40) < Jul 27 16:04:27 sservercluster1 ksmtuned: /usr/sbin/ksmtuned: fork: Cannot allocate memory
[1595846396] SERVICE ALERT: sservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595848929] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9634796 kB) free!
[1595849049] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9524568 kB) free!
[1595849169] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (9404412 kB) free!
[1595849352] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.09, 2.68, 343.08
[1595849835] SERVICE ALERT: pservercluster1;Messages Log Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595849842] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595849864] SERVICE ALERT: pservercluster1;Current Users;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595849888] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595849901] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595849918] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595849931] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595849935] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595849965] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595849972] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595849987] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595849994] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850018] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850021] SERVICE ALERT: pservercluster1;CPU Load;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595850048] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850061] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850065] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850090] SERVICE ALERT: pservercluster1;Messages Log Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595850092] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1595850107] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 334 processes
[1595850114] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1595850138] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1595850139] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;UNKNOWN;SOFT;1;Could not open pipe: /usr/bin/ps
[1595850141] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1595850168] SERVICE ALERT: pservercluster1;Open-Files;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595850172] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850181] SERVICE ALERT: pservercluster1;Sensors Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595850195] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850269] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850287] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850296] SERVICE ALERT: pservercluster1;Memory;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595850299] SERVICE ALERT: pservercluster1;HDD_Home status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595850394] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595850417] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850426] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850429] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850430] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 324 processes
[1595850547] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850550] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 324 processes
[1595850556] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850652] SERVICE FLAPPING ALERT: mrq4;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1595850670] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 322 processes
[1595850692] SERVICE ALERT: pservercluster1;HDD_Root status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595850724] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850778] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850791] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850854] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595850984] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851004] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851292] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1595851300] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851327] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851339] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851350] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851371] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851398] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851402] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851404] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851408] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851457] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851469] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851480] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851491] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851501] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851518] SERVICE ALERT: pservercluster2;Total Processes;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595851518] SERVICE ALERT: pservercluster2;Open-Files;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595851522] SERVICE ALERT: pservercluster2;Memory;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595851534] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851538] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851587] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851599] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851610] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851621] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851626] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851631] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851643] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595851648] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851648] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851652] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851654] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595851663] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851668] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851751] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851756] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851773] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851778] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851793] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851886] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851900] SERVICE ALERT: pservercluster1;Messages Log Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595851903] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851917] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595851923] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595852047] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595852177] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595852184] SERVICE ALERT: pservercluster1;Current Users;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595852497] SERVICE FLAPPING ALERT: pservercluster1;HDD_Root status;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1595852497] SERVICE ALERT: pservercluster1;HDD_Root status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595852500] SERVICE FLAPPING ALERT: pservercluster1;Messages Log Status;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1595852500] SERVICE ALERT: pservercluster1;Messages Log Status;WARNING;HARD;3;NRPE: Unable to read output
[1595852548] SERVICE ALERT: pservercluster1;CPU Load;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595852794] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595852804] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595853102] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595853110] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595853158] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595853414] SERVICE FLAPPING ALERT: pservercluster1;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1595853414] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595854943] SERVICE ALERT: pservercluster1;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595855553] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595855769] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4900216 kB) free!
[1595857068] HOST ALERT: pservercluster2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.38)
[1595857073] HOST ALERT: pservercluster1;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1595857258] HOST ALERT: pservercluster2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1595857262] HOST ALERT: pservercluster1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1595857287] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1595857302] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39507 MB (77% inode=99%):
[1595857303] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.75, 0.45, 0.18
[1595857323] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595857358] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1595857358] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.08, 0.10, 0.05
[1595857358] SERVICE FLAPPING ALERT: pservercluster1;CPU-STATISTICS;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1595857358] SERVICE ALERT: pservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=0.01% system=0.03% iowait=0.03% idle=99.94%
[1595857379] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 4224 open files (0% of max 2440681)
[1595857391] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1595857395] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1595857577] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 318 processes
[1595857587] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1595857594] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1595857594] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1595857599] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595857610] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1595857614] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595857629] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 172668 MB (78% inode=99%):
[1595857631] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 147534 MB (65% inode=99%):
[1595857648] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 327 processes
[1595857652] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.0% (23280088 kB) free.
[1595857659] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=2.07% system=4.77% iowait=0.51% idle=92.64%
[1595857668] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38293 MB (74% inode=98%):
[1595857747] SERVICE ALERT: pservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1595857756] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.4% (23541016 kB) free.
[1595857778] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 6432 open files (0% of max 2424156)
[1595857911] SERVICE ALERT: pservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595858329] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595860270] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1595860961] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 89.3% (44020684 kB) free.
[1595861122] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861142] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861152] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861165] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861175] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861176] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861221] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861221] SERVICE ALERT: sservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861221] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861235] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861240] SERVICE ALERT: sservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861252] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861272] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861282] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861295] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861305] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861306] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861351] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861351] SERVICE ALERT: sservercluster1;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861351] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861362] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861366] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861370] SERVICE ALERT: sservercluster1;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861382] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861402] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861406] SERVICE ALERT: sservercluster1;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861412] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861415] SERVICE ALERT: sservercluster1;Current Users;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595861435] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861436] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861481] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861481] SERVICE ALERT: sservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861481] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861490] SERVICE ALERT: sservercluster1;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595861492] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861496] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861536] SERVICE ALERT: sservercluster1;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861622] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595861666] SERVICE ALERT: sservercluster1;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595862025] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595863782] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595863802] SERVICE ALERT: sservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 40508 MB (79% inode=99%):
[1595863812] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.4% (20680532 kB) free.
[1595863825] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1595863836] SERVICE ALERT: sservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1595863841] SERVICE ALERT: sservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=2.98% system=11.37% iowait=0.04% idle=85.61%
[1595863852] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 25.5% (8349256 kB) free.
[1595863881] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 149760 MB (67% inode=99%):
[1595863881] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 346 processes
[1595863881] SERVICE ALERT: sservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1595863890] SERVICE ALERT: sservercluster1;Open-Files;OK;HARD;1;OK: 12480 open files (0% of max 2424092)
[1595863896] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1595863904] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1595864022] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.75, 2.26, 2.46
[1595864666] SERVICE ALERT: sservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595864795] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1595865102] SERVICE FLAPPING ALERT: pservercluster1;HDD_Root status;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1595865111] SERVICE FLAPPING ALERT: pservercluster1;Messages Log Status;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1595865158] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1595865414] SERVICE FLAPPING ALERT: pservercluster1;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1595867564] SERVICE FLAPPING ALERT: pservercluster1;CPU-STATISTICS;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1595867638] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1520 processes
[1595867758] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1522 processes
[1595867879] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1519 processes
[1595872871] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1595874129] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 32.37, 33.66, 38.76