[1583346792] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20815036 kB) free!
[1583346912] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41572084 kB) free.
[1583346987] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.6% (29378880 kB) free.
[1583347003] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201668 kB) free!
[1583347123] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200848 kB) free!
[1583347243] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170344 kB) free!
[1583347512] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20316076 kB) free!
[1583347632] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40874624 kB) free.
[1583348232] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19829756 kB) free!
[1583348352] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39556860 kB) free.
[1583348782] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184332 kB) free!
[1583348842] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195652 kB) free!
[1583348902] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.5% (17030604 kB) free.
[1583348944] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4160492 kB) free!
[1583348952] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19688424 kB) free!
[1583349072] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38597088 kB) free.
[1583349544] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (205308 kB) free!
[1583350144] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.2% (9916508 kB) free.
[1583350273] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22117168 kB) free!
[1583350393] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41335936 kB) free.
[1583350744] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1688268 kB) free!
[1583350864] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173024 kB) free!
[1583350984] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583350984] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196988 kB) free!
[1583350993] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20943184 kB) free!
[1583351113] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38899392 kB) free.
[1583352313] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (23000596 kB) free!
[1583352433] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38695944 kB) free.
[1583354833] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23475312 kB) free!
[1583354953] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41491164 kB) free.
[1583355011] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.0% (10061708 kB) free.
[1583355083] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7587648 kB) free!
[1583355203] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7409196 kB) free!
[1583355323] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (7389896 kB) free!
[1583355923] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 78.7% (38808040 kB) free.
[1583356091] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.8% (41810540 kB) free.
[1583356753] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23187896 kB) free!
[1583356874] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38599572 kB) free.
[1583357593] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(179) < 49.50.86.152 [05/Mar/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583357692] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(490) < 172.105.59.173 [05/Mar/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583357715] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583357772] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(234) < 2020/03/05 03:00:22 [error] 12247#0: *65093649 no live upstreams while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583357812] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583357892] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583362384] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583362900] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 403 processes
[1583363021] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1583365211] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188180 kB) free!
[1583365331] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182860 kB) free!
[1583365451] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181528 kB) free!
[1583365847] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.6% (20499236 kB) free.
[1583365984] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.6% (12981972 kB) free.
[1583368971] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.2% (31925400 kB) free.
[1583375344] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39283576 kB) free.
[1583375422] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.6% (21720504 kB) free.
[1583375448] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 73.5% (18017364 kB) free.
[1583375892] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(33) < 2020/03/05 08:01:14 [error] 12254#0: *65183986 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583376012] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583376184] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (276356 kB) free!
[1583376304] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (281460 kB) free!
[1583376424] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (281012 kB) free!
[1583378620] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1583378741] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1583380273] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26172588 kB) free!
[1583380393] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40041368 kB) free.
[1583380987] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21927080 kB) free.
[1583383332] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1583383452] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 705 processes
[1583383573] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 705 processes
[1583383994] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26215780 kB) free!
[1583384114] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21489708 kB) free!
[1583384234] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17053600 kB) free!
[1583384773] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1583386034] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4868428 kB) free!
[1583386634] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24846188 kB) free!
[1583387051] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 362 processes
[1583387169] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3319584 kB) free!
[1583387171] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 363 processes
[1583387234] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1583387234] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29485740 kB) free.
[1583387289] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (3011852 kB) free!
[1583387291] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 361 processes
[1583387409] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.0% (2696532 kB) free!
[1583388009] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (961408 kB) free!
[1583388191] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4397244 kB) free!
[1583388311] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4255824 kB) free!
[1583388431] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (4086344 kB) free!
[1583388973] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1583389093] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1583389172] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.1% (15523416 kB) free.
[1583389213] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1583389254] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.7% (4082724 kB) free!
[1583389374] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.6% (3813952 kB) free!
[1583389494] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.8% (3624212 kB) free!
[1583389587] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7593244 kB) free!
[1583389707] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7211752 kB) free!
[1583389828] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.4% (30280216 kB) free.
[1583389847] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3554636 kB) free!
[1583389886] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.1% (3492700 kB) free!
[1583389967] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4234716 kB) free!
[1583390006] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 12.8% (3150800 kB) free!
[1583390087] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (4102312 kB) free!
[1583390126] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 11.2% (2771248 kB) free!
[1583390213] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(39) < 172.105.59.173 [05/Mar/2020:12:03:40 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1583390231] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (2000368 kB) free!
[1583390252] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1583390291] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1583390333] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583390372] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1583390412] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(122) < 2020/03/05 12:01:40 [error] 12247#0: *65291629 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583390432] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (4486960 kB) free!
[1583390492] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1583390532] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583390552] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (4302588 kB) free!
[1583390672] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (3593692 kB) free!
[1583390694] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1499968 kB) free!
[1583390726] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (1042028 kB) free!
[1583391092] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1583391213] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1583391275] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21647624 kB) free.
[1583391326] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.1% (22731164 kB) free.
[1583391333] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583391382] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1583391382] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.2% (18544872 kB) free.
[1583391453] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.7% (21507808 kB) free.
[1583391533] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 49.50.64.122 [05/Mar/2020:12:25:38 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583391653] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583391732] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/05 12:27:06 [error] 12257#0: *65306308 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:17/0, bytes from/to upstream:0/17
[1583391777] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1583391852] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583391872] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2175456 kB) free!
[1583391897] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1583391992] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4594728 kB) free!
[1583392035] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14943816 kB) free!
[1583392064] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (4315744 kB) free!
[1583392112] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4469448 kB) free!
[1583392145] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9477568 kB) free!
[1583392155] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36723180 kB) free.
[1583392184] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (4218736 kB) free!
[1583392232] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4377768 kB) free!
[1583392265] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9478040 kB) free!
[1583392304] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (4094100 kB) free!
[1583392385] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (9477348 kB) free!
[1583392854] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2576) < 103.255.146.164 [05/Mar/2020:12:45:35 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583392974] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583393355] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13530932 kB) free!
[1583393475] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34302212 kB) free.
[1583393504] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1583393504] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2058220 kB) free!
[1583393967] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.7% (4835932 kB) free!
[1583394032] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (454260 kB) free!
[1583394075] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9886644 kB) free!
[1583394087] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.6% (4815736 kB) free!
[1583394195] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43694784 kB) free.
[1583394207] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.6% (4793024 kB) free!
[1583394491] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9165224 kB) free!
[1583394611] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8887348 kB) free!
[1583394731] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8604412 kB) free!
[1583395385] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4630360 kB) free!
[1583395395] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22124456 kB) free!
[1583395515] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44115796 kB) free.
[1583395994] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46885 MB (20% inode=99%):
[1583396114] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46885 MB (20% inode=99%):
[1583396115] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20374016 kB) free!
[1583396234] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46885 MB (20% inode=99%):
[1583396235] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42898012 kB) free.
[1583396531] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4621464 kB) free!
[1583396574] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1137) < 103.255.146.164 [05/Mar/2020:13:49:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583396694] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583396835] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19150612 kB) free!
[1583396955] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41867732 kB) free.
[1583397551] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1583397555] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17844856 kB) free!
[1583397585] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22074724 kB) free.
[1583397642] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.3% (37597232 kB) free.
[1583397671] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583397674] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 45.9% (22629148 kB) free.
[1583397676] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40773584 kB) free.
[1583397712] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.8% (21514232 kB) free.
[1583397791] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583397883] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194292 kB) free!
[1583398003] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186628 kB) free!
[1583398052] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.9% (21051380 kB) free.
[1583398104] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 89.4% (21915276 kB) free.
[1583398123] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 36.1% (8845716 kB) free.
[1583398184] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4497220 kB) free!
[1583398275] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17213160 kB) free!
[1583398304] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4294288 kB) free!
[1583398395] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39713352 kB) free.
[1583398424] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.1% (4055884 kB) free!
[1583398494] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1338) < 103.255.146.164 [05/Mar/2020:14:18:08 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583398615] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583398634] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51709 MB (23% inode=99%):
[1583398792] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.0% (3212536 kB) free!
[1583398912] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (1523476 kB) free!
[1583398992] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1583398995] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16552696 kB) free!
[1583399032] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (831456 kB) free!
[1583399115] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12480976 kB) free!
[1583399235] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35877912 kB) free.
[1583399624] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1679160 kB) free!
[1583399835] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12047088 kB) free!
[1583399955] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583399955] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36204788 kB) free.
[1583400192] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1583400312] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1583400539] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 11.2% (2756388 kB) free!
[1583400555] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12246324 kB) free!
[1583400659] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 10.5% (2577516 kB) free!
[1583400675] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36285360 kB) free.
[1583400779] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2325124 kB) free!
[1583400825] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.2% (15535204 kB) free.
[1583400923] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9303592 kB) free!
[1583401014] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(179) < 92.42.104.166 [05/Mar/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583401043] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (9112396 kB) free!
[1583401134] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583401163] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8941128 kB) free!
[1583401275] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12049280 kB) free!
[1583401328] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 389 processes
[1583401379] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 67.5% (16549396 kB) free.
[1583401395] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (7998836 kB) free!
[1583401448] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 387 processes
[1583401512] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1583401515] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32441088 kB) free.
[1583401632] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583401753] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583402115] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8438740 kB) free!
[1583402235] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32984380 kB) free.
[1583402579] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179716 kB) free!
[1583402649] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 389 processes
[1583402699] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179932 kB) free!
[1583402769] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 387 processes
[1583402819] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181864 kB) free!
[1583402835] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (916888 kB) free!
[1583402956] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391344 kB) free!
[1583403076] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1338352 kB) free!
[1583403369] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 390 processes
[1583403471] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1583403489] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 390 processes
[1583403512] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.4% (4744564 kB) free!
[1583403609] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 390 processes
[1583403632] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 10.5% (2583244 kB) free!
[1583403752] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 1.5% (366576 kB) free!
[1583404071] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4503640 kB) free!
[1583404134] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1299) < 103.255.146.164 [05/Mar/2020:15:51:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583404163] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4246136 kB) free!
[1583404191] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (4334108 kB) free!
[1583404209] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 390 processes
[1583404213] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1583404254] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583404276] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15160432 kB) free!
[1583404311] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (823960 kB) free!
[1583404333] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1583404352] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 88.3% (21629960 kB) free.
[1583404453] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1583404753] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1583404753] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1583404876] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27398200 kB) free.
[1583404971] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25018176 kB) free!
[1583405053] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1583405091] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.1% (33092836 kB) free.
[1583405219] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1583405219] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 63.2% (15501344 kB) free.
[1583405303] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3182684 kB) free!
[1583405353] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583405423] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3305228 kB) free.
[1583405473] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583405476] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2140888 kB) free!
[1583405594] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583405596] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39885728 kB) free.
[1583405819] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.5% (3552092 kB) free!
[1583405939] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (1123776 kB) free!
[1583406059] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184912 kB) free!
[1583406196] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10150204 kB) free!
[1583406291] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.5% (25720204 kB) free!
[1583406316] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6323956 kB) free!
[1583406411] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.6% (19186284 kB) free!
[1583406436] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38984200 kB) free.
[1583406531] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27138548 kB) free.
[1583406623] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3093124 kB) free!
[1583406652] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9433532 kB) free!
[1583406743] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.0% (5356060 kB) free.
[1583406773] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9434168 kB) free!
[1583406893] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (9433348 kB) free!
[1583407036] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14416896 kB) free!
[1583407131] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25781976 kB) free!
[1583407156] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10539888 kB) free!
[1583407251] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.1% (34392976 kB) free.
[1583407276] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37820316 kB) free.
[1583407487] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.8% (22157344 kB) free.
[1583407532] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.5% (44603628 kB) free.
[1583407565] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39249864 kB) free.
[1583407599] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.0% (21564328 kB) free.
[1583407634] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.3% (21647964 kB) free.
[1583407809] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 391 processes
[1583407876] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12682548 kB) free!
[1583407929] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 392 processes
[1583407994] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1583407996] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8611676 kB) free!
[1583408049] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 388 processes
[1583408117] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35558952 kB) free.
[1583408543] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2154576 kB) free!
[1583408663] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2540132 kB) free!
[1583408694] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (285008 kB) free!
[1583408716] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10886268 kB) free!
[1583408783] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (2854820 kB) free!
[1583408814] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194124 kB) free!
[1583408836] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6984564 kB) free!
[1583408843] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (1747364 kB) free!
[1583408934] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202092 kB) free!
[1583408956] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34625228 kB) free.
[1583408963] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (1510228 kB) free!
[1583409052] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (24844128 kB) free!
[1583409083] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1287320 kB) free!
[1583409172] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.9% (18295040 kB) free!
[1583409225] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4114260 kB) free!
[1583409292] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (28246312 kB) free.
[1583409345] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.1% (3830032 kB) free!
[1583409383] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3916400 kB) free.
[1583409465] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.7% (3530564 kB) free!
[1583409556] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9641784 kB) free!
[1583409634] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1583409676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5686312 kB) free!
[1583409723] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.0% (2701656 kB) free!
[1583409754] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1583409796] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33558128 kB) free.
[1583409843] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.5% (2579804 kB) free!
[1583409874] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1583409964] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (2468104 kB) free!
[1583409983] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.7% (2068208 kB) free!
[1583410065] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1583410065] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2011064 kB) free!
[1583410103] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2457708 kB) free!
[1583410224] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (1826560 kB) free!
[1583410283] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1583410283] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (4003020 kB) free!
[1583410396] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8568128 kB) free!
[1583410492] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25635952 kB) free!
[1583410516] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4618952 kB) free!
[1583410564] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2044972 kB) free!
[1583410612] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.6% (19310888 kB) free!
[1583410636] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31879096 kB) free.
[1583410732] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.3% (29386736 kB) free.
[1583410824] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583410824] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3602068 kB) free.
[1583411236] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7828392 kB) free!
[1583411356] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3801120 kB) free!
[1583411476] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28465748 kB) free.
[1583411483] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1831560 kB) free!
[1583411866] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.4% (15582124 kB) free.
[1583412076] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7068008 kB) free!
[1583412197] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2942784 kB) free!
[1583412274] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1583412317] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1583412317] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402904 kB) free!
[1583412532] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.1% (22548424 kB) free!
[1583412624] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.1% (1633700 kB) free!
[1583412652] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.5% (15133748 kB) free!
[1583412744] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.3% (2319332 kB) free!
[1583412772] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.5% (32239276 kB) free.
[1583412864] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (3041364 kB) free!
[1583413254] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(64) < 103.255.146.164 [05/Mar/2020:18:25:09 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583413375] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583413464] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (621988 kB) free!
[1583413517] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20251640 kB) free!
[1583413602] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1583414064] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (2392228 kB) free!
[1583414574] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(64) < 103.255.146.164 [05/Mar/2020:18:47:46 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583414664] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.5% (4800396 kB) free.
[1583414694] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583414717] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (8296412 kB) free!
[1583415264] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.1% (2289944 kB) free!
[1583415294] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1112) < 103.255.146.164 [05/Mar/2020:18:57:57 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583415298] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1583415317] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17761884 kB) free!
[1583415384] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2886352 kB) free!
[1583415415] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583415418] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1583415504] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3466256 kB) free.
[1583415683] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9096900 kB) free!
[1583415804] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8905580 kB) free!
[1583415917] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30659764 kB) free.
[1583415924] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.7% (8709020 kB) free!
[1583415969] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9212400 kB) free!
[1583416089] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9210596 kB) free!
[1583416209] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (9211064 kB) free!
[1583416517] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5532920 kB) free!
[1583416603] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.7% (4579936 kB) free!
[1583416637] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1489680 kB) free!
[1583416704] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3038512 kB) free!
[1583416723] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.7% (4341884 kB) free!
[1583416757] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32175624 kB) free.
[1583416825] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3491144 kB) free.
[1583416843] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.7% (4103708 kB) free!
[1583417357] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7980264 kB) free!
[1583417424] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (720248 kB) free!
[1583417477] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3861148 kB) free!
[1583417544] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (1188268 kB) free!
[1583417594] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583417597] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400016 kB) free!
[1583417664] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1616124 kB) free!
[1583418043] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1803316 kB) free!
[1583418264] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (2956160 kB) free!
[1583418324] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4361696 kB) free!
[1583418534] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583418798] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25662572 kB) free!
[1583418852] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(260) < 2020/03/05 20:02:09 [error] 12250#0: *65553653 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583418864] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.7% (4819820 kB) free.
[1583418933] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 392 processes
[1583418972] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583419372] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.7% (20733124 kB) free!
[1583419398] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2231892 kB) free!
[1583419464] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2493228 kB) free!
[1583419492] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.1% (13334744 kB) free!
[1583419533] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1583419584] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2904148 kB) free!
[1583419612] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.8% (28802816 kB) free.
[1583419653] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1583419704] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3377504 kB) free.
[1583419773] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1583420409] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3587036 kB) free!
[1583420598] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26279244 kB) free!
[1583420866] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (3797144 kB) free!
[1583420904] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2861252 kB) free!
[1583420986] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (3706996 kB) free!
[1583421024] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3304320 kB) free.
[1583421106] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.1% (3619208 kB) free!
[1583421198] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3972224 kB) free!
[1583421573] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1583421798] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16187080 kB) free!
[1583422032] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (2872564 kB) free!
[1583422153] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (2866212 kB) free!
[1583422224] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2718688 kB) free!
[1583422273] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (2872696 kB) free!
[1583422283] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583422306] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.4% (15574236 kB) free.
[1583422344] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3142036 kB) free!
[1583422398] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28455348 kB) free.
[1583422465] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3590756 kB) free.
[1583422998] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4532028 kB) free!
[1583423118] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (710488 kB) free!
[1583423212] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (24967076 kB) free!
[1583423238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32597732 kB) free.
[1583423332] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.4% (32118684 kB) free.
[1583423838] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8449880 kB) free!
[1583423958] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4541168 kB) free!
[1583424078] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (683236 kB) free!
[1583424532] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24434228 kB) free!
[1583424652] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (30936812 kB) free.
[1583425173] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583425252] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26192576 kB) free!
[1583425278] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23955720 kB) free!
[1583425293] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 384 processes
[1583425372] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.8% (31431872 kB) free.
[1583425879] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2553564 kB) free!
[1583426479] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13319096 kB) free!
[1583427679] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3656984 kB) free!
[1583428279] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14721036 kB) free!
[1583428879] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27885548 kB) free.
[1583429064] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2841768 kB) free!
[1583429184] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.1% (4573864 kB) free.
[1583429479] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7622004 kB) free!
[1583429599] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4062056 kB) free!
[1583429719] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31403124 kB) free.
[1583430319] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9747880 kB) free!
[1583430439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6179680 kB) free!
[1583430559] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32779408 kB) free.
[1583430984] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3009460 kB) free!
[1583431104] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583431104] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.6% (4486076 kB) free.
[1583431159] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11579068 kB) free!
[1583431279] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7940172 kB) free!
[1583431399] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4341624 kB) free!
[1583431827] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234976 kB) free!
[1583431947] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231536 kB) free!
[1583431972] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (416292 kB) free!
[1583432067] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235208 kB) free!
[1583432092] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (879040 kB) free!
[1583432212] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3865952 kB) free!
[1583432600] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22792104 kB) free!