[1550688831] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1550689359] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1550689431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550689551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550689671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550690271] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550690872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550690992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550691112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550691712] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550692312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550692432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550692552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550695366] SERVICE FLAPPING ALERT: primary-server;Current Users;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1550697059] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4800956 kB) free!
[1550697179] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4771060 kB) free!
[1550697299] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (4740140 kB) free!
[1550703952] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550705752] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550706299] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2439576 kB) free!
[1550706352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550706472] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550706592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550707192] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550707792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1550707913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550708033] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1550708033] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550708633] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1550709233] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550709353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550709473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550710348] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 137.72, 34.03, 11.75
[1550710468] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 19.30, 23.03, 10.42
[1550710588] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.43, 15.69, 9.26
[1550711788] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.78, 2.93
[1550712473] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1550713073] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550713193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550713313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550716913] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1550717513] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550717633] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1550718233] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550718353] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550718953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550719074] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550719673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550719793] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550720393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550720513] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550720633] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550721833] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1550722433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550722554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550722674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550724372] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.1% (22721208 kB) free.
[1550724445] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.0% (14786168 kB) free.
[1550725059] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 53 users currently logged in
[1550725074] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550725119] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 53 users currently logged in
[1550725179] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 53 users currently logged in
[1550725192] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1550725239] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 53 users currently logged in
[1550725312] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1550725674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550725794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550725915] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550726525] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2033648 kB) free!
[1550726645] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.7% (1740576 kB) free!
[1550726765] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1697300 kB) free!
[1550727137] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 64.9% (10542260 kB) free.
[1550727343] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (31.6% change >= 20.0% threshold)
[1550727384] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 54 users currently logged in
[1550727705] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1550728791] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3182064 kB) free!
[1550728911] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3133572 kB) free!
[1550728915] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550729031] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.0% (3085764 kB) free!
[1550729515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550729635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550729756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550729825] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (14944188 kB) free.
[1550730356] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550730956] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550731076] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550731196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550732761] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24662 MB (10% inode=99%):
[1550733554] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (4079172 kB) free!
[1550733674] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4870764 kB) free!
[1550733794] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (4804592 kB) free!
[1550734189] SERVICE ALERT: sig6;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'sig6' was out of bounds)
[1550734192] SERVICE ALERT: sig6;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig6' was out of bounds)
[1550734192] SERVICE ALERT: sig6;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig6' was out of bounds)
[1550734192] HOST ALERT: sig6;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1550734195] SERVICE ALERT: sig6;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig6' was out of bounds)
[1550734214] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6' was out of bounds)
[1550734216] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 337 processes
[1550734217] SERVICE ALERT: sig6;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6' was out of bounds)
[1550734309] SERVICE ALERT: sig6;Total Processes;OK;HARD;1;PROCS OK: 337 processes
[1550734312] SERVICE ALERT: sig6;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1550734312] SERVICE ALERT: sig6;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.16, 0.06
[1550734313] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1550734315] SERVICE ALERT: sig6;Sensors Status;OK;HARD;1;SENSORS OK
[1550734334] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550734337] SERVICE ALERT: sig6;HDD_Root status;OK;HARD;1;DISK OK - free space: / 27104 MB (52% inode=98%):
[1550734994] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 30.3% (7421316 kB) free.
[1550735161] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27033 MB (12% inode=99%):
[1550736144] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550736147] SERVICE ALERT: sig6;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550736174] HOST ALERT: sig6;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1550736183] SERVICE ALERT: sig6;HDD_Home status;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550736187] SERVICE ALERT: sig6;Memory;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550736264] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550736267] SERVICE ALERT: sig6;HDD_Root status;OK;HARD;1;DISK OK - free space: / 27104 MB (52% inode=98%):
[1550736268] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1550736303] SERVICE ALERT: sig6;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 212561 MB (94% inode=99%):
[1550736307] SERVICE ALERT: sig6;Memory;OK;HARD;1;OK - 94.7% (15396224 kB) free.
[1550736361] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24625 MB (10% inode=99%):
[1550736394] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1550736454] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1550736514] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1550736574] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1550736719] SERVICE ALERT: sig6;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550736721] SERVICE ALERT: sig6;Sensors Status;CRITICAL;SOFT;1;(Return code of 255 for service 'Sensors Status' on host 'sig6' was out of bounds)
[1550736721] HOST ALERT: sig6;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1550736721] SERVICE ALERT: sig6;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig6' was out of bounds)
[1550736721] SERVICE ALERT: sig6;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig6' was out of bounds)
[1550736839] SERVICE ALERT: sig6;Total Processes;OK;HARD;1;PROCS OK: 315 processes
[1550736841] SERVICE ALERT: sig6;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1550736841] SERVICE ALERT: sig6;CPU Load;OK;HARD;1;OK - load average: 0.11, 0.10, 0.05
[1550736841] SERVICE ALERT: sig6;Sensors Status;OK;HARD;1;SENSORS OK
[1550736843] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1550737196] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550737796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550737916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550738037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550738374] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 52 users currently logged in
[1550738501] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.37, 5.85, 6.75
[1550738621] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.49, 4.40, 6.11
[1550738639] SERVICE ALERT: sig6;Total Processes;WARNING;SOFT;1;PROCS WARNING: 576 processes
[1550738741] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.65, 3.44, 5.55
[1550738759] SERVICE ALERT: sig6;Total Processes;WARNING;SOFT;2;PROCS WARNING: 594 processes
[1550738879] SERVICE ALERT: sig6;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 612 processes
[1550739341] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.30, 1.59, 3.53
[1550740651] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3217608 kB) free!
[1550740772] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3201228 kB) free!
[1550740892] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.5% (3170492 kB) free!
[1550740994] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4804060 kB) free!
[1550741114] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 32.8% (8042972 kB) free.
[1550741161] SERVICE FLAPPING ALERT: pservercluster2;HDD_Home status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1550741161] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27925 MB (12% inode=99%):
[1550741637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550741895] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550742015] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550742016] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 340 processes
[1550742135] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550742237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550742357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550742477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550742735] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550742874] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1550742934] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1550742994] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1550743054] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1550743163] SERVICE ALERT: sig5;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig5' was out of bounds)
[1550743163] HOST ALERT: sig5;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.57)
[1550743230] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;HARD;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1550743234] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1550743283] SERVICE ALERT: sig5;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.18, 0.07
[1550743350] SERVICE ALERT: sig5;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1550745927] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550745957] HOST ALERT: sig6;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1550746047] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;HARD;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1550746077] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550746079] SERVICE ALERT: sig6;Total Processes;OK;HARD;1;PROCS OK: 310 processes
[1550746083] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1550746648] SERVICE ALERT: sig6;CPU_Procs;OK;HARD;1;CPU OK: 387 processes
[1550746677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550746797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550746918] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550747247] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 452 processes
[1550747367] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 451 processes
[1550747487] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 445 processes
[1550747717] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 322 processes
[1550747837] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1550748361] SERVICE FLAPPING ALERT: pservercluster2;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1550748437] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 323 processes
[1550748558] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 324 processes
[1550748678] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 323 processes
[1550749892] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1580392 kB) free!
[1550750163] SERVICE ALERT: sig5;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550750193] HOST ALERT: sig5;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1550750283] SERVICE ALERT: sig5;Total Processes;OK;HARD;1;PROCS OK: 285 processes
[1550750287] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1550750518] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550750550] HOST ALERT: sig6;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1550750614] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1550751087] SERVICE FLAPPING ALERT: sig6;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1550751087] SERVICE ALERT: sig6;CPU_Procs;OK;HARD;1;CPU OK: 392 processes
[1550751118] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550751239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550751359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550752055] HOST ALERT: webservercluster2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1550752119] HOST ALERT: webservercluster2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1550753254] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 67 users currently logged in
[1550753561] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550753681] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550753801] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550754616] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 338 processes
[1550754754] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1550754959] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1550755001] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550755560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550755680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550755800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550757000] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550757600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550757720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550757840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550759041] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550759488] SERVICE FLAPPING ALERT: sig6;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1550759641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550759761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550759854] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 52 users currently logged in
[1550759881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550760481] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1550761081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550761201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550761321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550765521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550766121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550766241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550766362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550769962] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550771762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550771882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550772482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550772602] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550773202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550773322] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes