[1568832037] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1568832158] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193036 kB) free!
[1568832233] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187424 kB) free!
[1568832278] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198180 kB) free!
[1568832353] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176092 kB) free!
[1568832474] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188192 kB) free!
[1568835127] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 464 processes
[1568835163] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 12.4% (3037144 kB) free!
[1568835187] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 460 processes
[1568835223] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 42.7% (10470976 kB) free.
[1568835348] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 257 processes
[1568835773] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (392584 kB) free!
[1568835809] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (236180 kB) free!
[1568835893] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407176 kB) free!
[1568835929] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (225168 kB) free!
[1568835949] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 253 processes
[1568836013] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1087428 kB) free!
[1568836049] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (245360 kB) free!
[1568836074] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.5% (3324208 kB) free!
[1568836674] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164088 kB) free!
[1568836867] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9567448 kB) free!
[1568836987] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9523840 kB) free!
[1568837107] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (9520444 kB) free!
[1568839995] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (247120 kB) free!
[1568841368] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 141387 MB (62% inode=99%):
[1568841795] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.7% (46184784 kB) free.
[1568843078] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 52.3% (12820508 kB) free.
[1568843332] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 46.1% (11288940 kB) free.
[1568846512] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 68.59, 50.85, 25.23
[1568846632] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 68.55, 56.74, 30.50
[1568846752] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.40, 47.32, 30.51
[1568847032] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 57.28, 21.96, 8.60
[1568847152] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 65.53, 36.83, 15.67
[1568847272] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 66.19, 46.63, 21.83
[1568847352] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.94, 6.96, 16.28
[1568847872] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 10.14, 15.52
[1568848552] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.64, 4.88
[1568848677] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 177.14, 58.24, 21.01
[1568848797] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 225.73, 116.15, 46.74
[1568848917] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 232.21, 154.94, 69.37
[1568849072] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.84, 1.08, 4.96
[1568849274] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2888580 kB) free!
[1568850717] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 0.93, 14.15
[1568851917] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.15, 0.21, 4.06
[1568853450] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.0% (33283260 kB) free.
[1568853475] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176748 kB) free!
[1568854013] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.3% (30748016 kB) free.
[1568854613] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.4% (24214472 kB) free!
[1568854733] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.2% (18669796 kB) free!
[1568854853] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 41.5% (54655400 kB) free.
[1568855081] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184868 kB) free!
[1568855206] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188240 kB) free!
[1568855329] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194284 kB) free!
[1568855333] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171404 kB) free!
[1568855454] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185292 kB) free!
[1568855574] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198240 kB) free!
[1568855887] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 506 processes
[1568855947] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 508 processes
[1568856007] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 506 processes
[1568856607] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 499 processes
[1568856667] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 491 processes
[1568868432] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.5% (23410760 kB) free.
[1568868494] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.1% (21492784 kB) free.
[1568868577] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.0% (23031296 kB) free.
[1568868792] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.9% (40359592 kB) free.
[1568869097] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178152 kB) free!
[1568869217] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164544 kB) free!
[1568869306] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.8% (23393116 kB) free.
[1568870747] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.30, 7.39, 6.10
[1568870867] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.38, 6.35, 5.88
[1568870987] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.34, 5.74, 5.71
[1568871114] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175748 kB) free!
[1568871234] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174508 kB) free!
[1568871354] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165288 kB) free!
[1568871818] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 307 processes
[1568871939] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1568872124] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.7% (23369236 kB) free.
[1568872992] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (7472864 kB) free!
[1568873112] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 22.1% (10895080 kB) free.
[1568873260] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1568873596] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7894920 kB) free!
[1568873626] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1568873712] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7717364 kB) free!
[1568873716] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.9% (7343532 kB) free!
[1568873747] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1568873833] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.7% (7718188 kB) free!
[1568873836] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.8% (6797240 kB) free!
[1568873953] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.7% (7719352 kB) free!
[1568874436] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 65.9% (32461220 kB) free.
[1568875530] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40321516 kB) free.
[1568875733] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1678456 kB) free!
[1568875853] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (1531424 kB) free!
[1568875973] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1568875973] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1368508 kB) free!
[1568876954] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.2% (23256460 kB) free.
[1568877584] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4511144 kB) free!
[1568877704] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4362368 kB) free!
[1568877817] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.09, 6.26, 5.28
[1568877824] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4229064 kB) free!
[1568877940] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.28, 7.19, 5.77
[1568878063] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.83, 5.46, 5.31
[1568878298] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.9% (22537100 kB) free.
[1568878663] SERVICE ALERT: sig3-vm2;CPU Load;OK;HARD;1;OK - load average: 1.46, 3.81, 4.83
[1568879260] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 578 processes
[1568879263] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.97, 5.94, 5.49
[1568879383] SERVICE ALERT: sig3-vm2;CPU Load;OK;HARD;1;OK - load average: 2.31, 4.42, 4.98
[1568879964] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.4% (3557124 kB) free!
[1568880084] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167840 kB) free!
[1568880204] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173508 kB) free!
[1568880653] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.3% (24098328 kB) free!
[1568880773] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.2% (29199680 kB) free.
[1568881302] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4856648 kB) free!
[1568881422] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4753252 kB) free!
[1568881542] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (4627316 kB) free!
[1568881660] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 622 processes
[1568882703] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (2778424 kB) free!
[1568882748] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.6% (2495440 kB) free!
[1568882823] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (2781580 kB) free!
[1568882856] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9781672 kB) free!
[1568882868] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (2032928 kB) free!
[1568882943] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (2777572 kB) free!
[1568882976] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9324536 kB) free!
[1568882988] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1550924 kB) free!
[1568883096] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (8992288 kB) free!
[1568883173] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25555452 kB) free!
[1568883293] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (31027368 kB) free.
[1568883342] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1568883342] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2205940 kB) free!
[1568883825] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 557 processes
[1568883894] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (22870452 kB) free!
[1568883946] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 554 processes
[1568884014] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.3% (25496328 kB) free!
[1568884066] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 559 processes
[1568884134] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.6% (35048064 kB) free.
[1568884296] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1568884296] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4815768 kB) free!
[1568884944] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.5% (23078020 kB) free.
[1568884989] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.1% (22572256 kB) free.
[1568885024] SERVICE ALERT: sig3-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1568885025] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.7% (23038124 kB) free.
[1568885067] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 42.5% (20940524 kB) free.
[1568885144] SERVICE ALERT: sig3-vm2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1568885835] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4779076 kB) free!
[1568885866] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 477 processes
[1568885955] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4670040 kB) free!
[1568886075] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4555432 kB) free!
[1568886586] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.0% (14468748 kB) free.
[1568887134] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25777852 kB) free!
[1568887254] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (30911332 kB) free.
[1568888475] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2055732 kB) free!
[1568888987] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.35, 4.75, 4.80
[1568889587] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.33, 5.25, 5.07
[1568889707] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.99, 5.76, 5.31
[1568889827] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.29, 5.23, 5.17
[1568890356] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1568890427] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.36, 4.35, 4.67
[1568894136] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9719748 kB) free!
[1568894256] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9231208 kB) free!
[1568894376] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (9233928 kB) free!
[1568894556] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.8% (3661244 kB) free!
[1568894594] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1568894676] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568894677] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.4% (3555072 kB) free!
[1568894797] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.1% (3469896 kB) free!
[1568895997] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180904 kB) free!
[1568896393] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189316 kB) free!
[1568896473] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 67.98, 18.37, 7.00
[1568896513] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197380 kB) free!
[1568896593] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.78, 12.98, 6.41
[1568896633] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (210936 kB) free!
[1568896713] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.83, 9.15, 5.80
[1568897076] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8698784 kB) free!
[1568897196] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8479324 kB) free!
[1568897316] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8287388 kB) free!
[1568899713] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.80, 1.69, 4.94
[1568899917] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 161.23, 77.13, 30.47
[1568900024] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 650 processes
[1568900037] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 168.46, 107.74, 47.35
[1568900144] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 664 processes
[1568900157] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 169.51, 128.27, 62.19
[1568900197] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (2598416 kB) free!
[1568900264] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 420 processes
[1568900313] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.04, 7.62, 7.65
[1568900316] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4637196 kB) free!
[1568900433] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.64, 5.32, 6.81
[1568900553] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 3.77, 6.06
[1568900653] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1568900783] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1568900798] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1568900798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1444168 kB) free!
[1568900903] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1568901153] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1568901153] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.47, 1.02, 3.48
[1568901357] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 3.16, 18.55
[1568902171] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.9% (23425084 kB) free.
[1568902213] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.8% (22987380 kB) free.
[1568902246] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 51.4% (25321612 kB) free.
[1568902291] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.3% (23357332 kB) free.
[1568903157] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 0.55, 3.13
[1568904576] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4881224 kB) free!
[1568906454] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.5% (17825008 kB) free!
[1568906574] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8497896 kB) free!
[1568906694] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2272004 kB) free!
[1568908677] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568908797] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1568909750] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.7% (47098776 kB) free.
[1568909789] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.7% (40279996 kB) free.
[1568911777] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4591252 kB) free!
[1568911897] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4561564 kB) free!
[1568912018] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4532552 kB) free!
[1568912553] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)