[1592332899] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592336100] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (180576 kB) free!
[1592337003] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 80.9% (19834308 kB) free.
[1592337300] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1592337420] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1592337492] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398192 kB) free!
[1592342225] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 80.6% (39704508 kB) free.
[1592342915] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.7% (42225000 kB) free.
[1592343256] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.1% (5424300 kB) free.
[1592343415] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(244) < 2020/06/17 03:01:22 [error] 12474#0: *94654508 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.62.98, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1592343515] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1296) < 172.105.59.173 [17/Jun/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592343535] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592343537] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(189) < 172.105.62.98 [17/Jun/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1592343635] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592343657] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592343857] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (307368 kB) free!
[1592343977] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (297728 kB) free!
[1592344097] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (296272 kB) free!
[1592345820] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1592345941] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1592346671] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 57.29, 47.88, 41.95
[1592347271] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 32.33, 35.79, 37.06
[1592348897] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.3% (10365600 kB) free.
[1592348940] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1592349060] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1592349671] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 66.30, 44.54, 37.78
[1592350697] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (210540 kB) free!
[1592350818] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174492 kB) free!
[1592350828] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 16.6% (4066264 kB) free!
[1592350871] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1592350871] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 26.54, 30.97, 34.99
[1592350938] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.7% (5568308 kB) free.
[1592350948] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 16.6% (4057220 kB) free!
[1592351068] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 16.6% (4059448 kB) free!
[1592351430] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 51.2% (12633528 kB) free.
[1592351471] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.54, 39.91, 36.35
[1592351538] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (226956 kB) free!
[1592351624] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.3% (19437332 kB) free.
[1592351658] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (218880 kB) free!
[1592351778] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1592351778] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (212704 kB) free!
[1592352071] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 17.74, 31.00, 34.79
[1592354598] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 73.3% (17955776 kB) free.
[1592355060] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1592355180] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1592355671] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 75.50, 47.29, 42.03
[1592357292] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 61.7% (81354880 kB) free.
[1592358317] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1592358437] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1592359785] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1592359830] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173184 kB) free!
[1592359905] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1592359950] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183640 kB) free!
[1592360070] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202508 kB) free!
[1592361378] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1592364071] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1592364180] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1592364300] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 330 processes
[1592364420] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1592364722] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.74, 6.16, 5.23
[1592364842] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.21, 6.07, 5.29
[1592364962] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.83, 6.29, 5.47
[1592369599] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4398328 kB) free!
[1592369719] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (4223568 kB) free!
[1592369839] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (4050336 kB) free!
[1592371039] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2002620 kB) free!
[1592371572] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3099.83, 1532.42, 593.89
[1592371604] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3624.91, 1474.09, 549.79
[1592371621] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 340 processes
[1592371692] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 420.05, 1026.08, 522.14
[1592371724] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 538.37, 1032.73, 502.57
[1592371741] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1592371812] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 59.32, 687.51, 459.17
[1592371844] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 74.62, 691.79, 441.92
[1592371872] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2225.48, 590.81, 201.23
[1592371992] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3943.49, 1788.80, 687.47
[1592372112] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 695.21, 1342.62, 664.60
[1592372239] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.8% (3381400 kB) free!
[1592372436] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [17/Jun/2020:11:04:36 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1592372501] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.3% (3218416 kB) free!
[1592372556] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592372621] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.2% (2747608 kB) free!
[1592372712] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.88, 182.11, 349.41
[1592372741] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.5% (2367684 kB) free!
[1592372839] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1592372839] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2073168 kB) free!
[1592373341] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (176240 kB) free!
[1592375624] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4588476 kB) free!
[1592375742] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.6% (15398892 kB) free.
[1592375744] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4412784 kB) free!
[1592375864] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4236768 kB) free!
[1592375915] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9492784 kB) free!
[1592376035] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9211348 kB) free!
[1592376155] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8913092 kB) free!
[1592376378] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2634432 kB) free!
[1592376978] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1580164 kB) free!
[1592377664] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1912092 kB) free!
[1592377741] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1592377861] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1592377955] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4793516 kB) free!
[1592378605] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 777.51, 1449.81, 717.02
[1592378725] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 106.47, 970.75, 630.32
[1592379103] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592379105] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1592379156] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(17) < 203.212.70.170 [17/Jun/2020:12:53:10 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592379223] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1592379225] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 351 processes
[1592379276] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592379343] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 386 processes
[1592379345] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1592379747] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 821 processes
[1592379867] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 343 processes
[1592379901] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 60.5% (14836612 kB) free.
[1592379943] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1592379946] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1592380477] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(453) < 103.255.146.164 [17/Jun/2020:13:14:50 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1592380543] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1592380597] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592380664] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1592380784] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1592381384] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1592381384] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1592381701] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 17.6% (5757528 kB) free!
[1592381821] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 17.0% (5569176 kB) free!
[1592381941] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 16.5% (5400800 kB) free!
[1592382342] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (2246556 kB) free!
[1592382425] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9589536 kB) free!
[1592382462] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (1822732 kB) free!
[1592382545] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9412708 kB) free!
[1592382583] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1443856 kB) free!
[1592382665] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (9218708 kB) free!
[1592384239] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592384702] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.4% (4021508 kB) free!
[1592384822] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.3% (3754836 kB) free!
[1592384942] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.4% (3520012 kB) free!
[1592385535] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/06/17 14:46:06 [error] 12489#0: *94749088 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1592385541] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2936328 kB) free!
[1592385655] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592385665] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4593028 kB) free!
[1592385998] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 172.105.62.98 [17/Jun/2020:14:52:20 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1592386118] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592386142] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1863052 kB) free!
[1592386185] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1592386718] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(261) < 103.13.97.241 [17/Jun/2020:15:01:23 +0530] TCP 502 0 0 60.000 "10.147.212.155:13001, backend_13001" "0, 0" "0, 0" "60.000, 0.000"
[1592386741] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 11.5% (3761784 kB) free!
[1592386783] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.7% (15653140 kB) free.
[1592386838] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592387341] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1592387341] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (3233444 kB) free!
[1592388061] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 339 processes
[1592388181] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1592389238] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(27) < 203.212.70.170 [17/Jun/2020:15:43:23 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1592389358] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592393385] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1592394158] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(31) < 172.105.62.98 [17/Jun/2020:17:10:52 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1592394278] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592395783] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3299564 kB) free!
[1592395903] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3074128 kB) free!
[1592396023] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (2851124 kB) free!
[1592396623] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1396360 kB) free!
[1592397223] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1592397223] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.2% (15530312 kB) free.
[1592398741] SERVICE FLAPPING ALERT: mrq4;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592403878] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(412) < 103.255.146.164 [17/Jun/2020:19:54:38 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1592403998] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(604) < 103.255.146.164 [17/Jun/2020:19:54:44 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1592404118] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592404255] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(484) < 2020/06/17 19:59:41 [error] 12482#0: *94801167 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.62.98, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1592404258] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(390) < 49.50.86.152 [17/Jun/2020:19:58:35 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1592404375] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592404378] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592404718] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(381) < 172.105.59.173 [17/Jun/2020:20:03:49 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1592404839] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592405438] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3149) < 103.255.146.164 [17/Jun/2020:20:19:41 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1592405558] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(3096) < 103.255.146.164 [17/Jun/2020:20:22:11 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1592405678] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592405762] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.28, 4.66, 4.93
[1592408086] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592408212] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.1% (40451268 kB) free.
[1592408341] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 65.8% (21546028 kB) free.
[1592408623] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592411913] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592412038] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=9.36% system=27.64% iowait=0.70% idle=62.30%
[1592415362] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.04, 5.75, 5.05
[1592415482] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.85, 5.56, 5.09
[1592415602] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.50, 5.51, 5.12
[1592416202] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.05, 4.75, 4.90
[1592418071] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 21.68, 38.28, 44.22