[1567276255] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1567276304] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 407 processes
[1567276424] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1567276645] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.5% (3311792 kB) free!
[1567276766] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166648 kB) free!
[1567276855] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184324 kB) free!
[1567276886] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175804 kB) free!
[1567276976] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167640 kB) free!
[1567277096] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176256 kB) free!
[1567277177] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180060 kB) free!
[1567277297] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (211772 kB) free!
[1567277417] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196952 kB) free!
[1567280097] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.5% (9497232 kB) free.
[1567280148] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 293 processes
[1567280254] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1120748 kB) free!
[1567280374] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1761628 kB) free!
[1567280494] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385416 kB) free!
[1567280697] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (4451284 kB) free!
[1567280735] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (1036184 kB) free!
[1567280748] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 288 processes
[1567280817] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.2% (3743980 kB) free!
[1567280855] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (1021596 kB) free!
[1567280937] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.6% (3105052 kB) free!
[1567280975] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (983296 kB) free!
[1567281537] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1567281537] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195148 kB) free!
[1567282424] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567282544] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567285286] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231864 kB) free!
[1567285406] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (218072 kB) free!
[1567285414] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 493 processes
[1567285527] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.2% (44934212 kB) free.
[1567285534] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 487 processes
[1567287617] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 51.0% (12492048 kB) free.
[1567287687] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 47.6% (11669284 kB) free.
[1567287944] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567288064] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567288777] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 40.85, 16.46, 7.11
[1567288897] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 46.98, 26.72, 11.99
[1567289017] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.63, 34.27, 16.54
[1567289337] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (2556328 kB) free!
[1567289617] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 8.88, 12.98
[1567290537] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2320016 kB) free!
[1567291180] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.27, 5.28, 5.92
[1567291300] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.37, 3.66, 5.25
[1567291420] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 2.55, 4.65
[1567292017] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.53, 0.80, 3.45
[1567298375] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 96.1% (33673112 kB) free.
[1567298494] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.4% (20239344 kB) free!
[1567299094] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12812336 kB) free!
[1567300064] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567300184] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567300294] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 39.4% (51903928 kB) free.
[1567300737] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.0% (12839280 kB) free.
[1567301261] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 363 processes
[1567301381] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1567303137] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180392 kB) free!
[1567303222] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201904 kB) free!
[1567303257] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190408 kB) free!
[1567303290] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183652 kB) free!
[1567303346] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197004 kB) free!
[1567303377] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179656 kB) free!
[1567303410] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181484 kB) free!
[1567303466] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189896 kB) free!
[1567303530] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173268 kB) free!
[1567303781] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1567303901] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1567309184] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567309304] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1567312979] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1567316504] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567316624] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567317835] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 78.6% (19272624 kB) free.
[1567317900] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40731424 kB) free.
[1567318012] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 67.5% (16546416 kB) free.
[1567318069] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.7% (18444628 kB) free.
[1567318427] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1567318443] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1567318447] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175392 kB) free!
[1567318547] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.78, 3.86, 3.94
[1567318566] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 346 processes
[1567318567] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173532 kB) free!
[1567318608] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (205704 kB) free!
[1567318677] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165004 kB) free!
[1567318687] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177884 kB) free!
[1567318730] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177216 kB) free!
[1567318797] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 21.08, 9.23, 5.72
[1567318797] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167040 kB) free!
[1567318850] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195608 kB) free!
[1567318917] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183072 kB) free!
[1567318917] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 28.14, 15.89, 8.58
[1567319047] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1567320249] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.63, 12.55, 17.47
[1567320347] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 142.95, 62.70, 25.38
[1567320467] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 154.79, 93.81, 41.30
[1567320587] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 158.70, 115.53, 55.63
[1567320820] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 162.88, 103.24, 45.43
[1567320941] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 163.52, 123.18, 59.75
[1567321061] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 60.82, 107.01, 62.18
[1567321087] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2664828 kB) free!
[1567321454] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 108.11, 30.65, 16.41
[1567322261] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.09, 2.79, 17.64
[1567322287] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1567322287] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2290356 kB) free!
[1567322387] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.89, 2.43, 12.51
[1567322706] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.9% (7828952 kB) free!
[1567322826] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.5% (7632564 kB) free!
[1567322946] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.5% (7632812 kB) free!
[1567324061] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.70, 0.67, 3.17
[1567324187] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.50, 2.18, 3.99
[1567325346] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (1998384 kB) free!
[1567326255] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1567326255] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.08, 4.28, 19.43
[1567328055] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.16, 1.53, 4.73
[1567332265] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22652292 kB) free.
[1567332313] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.0% (23281800 kB) free.
[1567332350] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40776252 kB) free.
[1567332377] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567332424] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.2% (23330048 kB) free.
[1567334624] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1567334744] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 402 processes
[1567334864] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567335276] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168832 kB) free!
[1567335396] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198632 kB) free!
[1567335516] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194932 kB) free!
[1567335814] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4911516 kB) free!
[1567335934] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 20.0% (4904332 kB) free!
[1567336054] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 20.0% (4895776 kB) free!
[1567337160] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8567556 kB) free!
[1567337280] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8567824 kB) free!
[1567337401] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (8565276 kB) free!
[1567338255] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1567339801] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (638828 kB) free!
[1567339901] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1567340022] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1567340254] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187924 kB) free!
[1567340589] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 93.8% (23001416 kB) free.
[1567340609] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.8% (23152616 kB) free.
[1567340641] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1567340641] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.8% (40784368 kB) free.
[1567341821] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1567341941] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1567342061] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1567342412] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (4257548 kB) free!
[1567342532] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.6% (3861592 kB) free!
[1567342652] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (3850636 kB) free!
[1567342661] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1567342664] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 403 processes
[1567342784] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1567343252] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1567343252] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (399164 kB) free!
[1567343862] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1567343982] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1567345526] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9829452 kB) free!
[1567345646] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9749180 kB) free!
[1567345766] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (9665928 kB) free!
[1567347525] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.3% (23024264 kB) free.
[1567347845] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9098440 kB) free!
[1567347965] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9097064 kB) free!
[1567348085] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (9094944 kB) free!
[1567348782] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1567348902] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1567349022] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1567349622] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1567349622] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1567350222] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1567350342] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1567350463] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1567350485] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4555544 kB) free!
[1567352262] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1567352382] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1567352502] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1567353255] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.92, 8.40, 8.80
[1567353375] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.14, 5.93, 7.84
[1567353495] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 4.20, 6.98
[1567353533] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1567354096] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.47, 1.39, 4.14
[1567354133] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4558680 kB) free!
[1567354253] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.1% (3721312 kB) free!
[1567354373] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (3634544 kB) free!
[1567356174] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1520380 kB) free!
[1567360085] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1567361366] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4929192 kB) free!
[1567361502] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)