[1559413869] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (289204 kB) free!
[1559414469] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.4% (2668544 kB) free!
[1559414649] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 392 processes
[1559414753] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1559415069] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (945168 kB) free!
[1559415249] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 386 processes
[1559415390] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.43, 1.56, 4.07
[1559415669] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 20.0% (3238080 kB) free!
[1559416269] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (205484 kB) free!
[1559416674] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559416869] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.4% (1849316 kB) free!
[1559418069] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (883388 kB) free!
[1559419270] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.0% (3075192 kB) free!
[1559419673] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559419686] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10068552 kB) free.
[1559419863] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 93.77, 24.37, 8.41
[1559419870] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165260 kB) free!
[1559419983] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 178.93, 79.89, 30.69
[1559420050] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 395 processes
[1559420103] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 204.07, 121.02, 51.63
[1559420273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559420286] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3435156 kB) free!
[1559420393] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559420406] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10064188 kB) free.
[1559420470] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.1% (3090964 kB) free!
[1559420805] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2860796 kB) free!
[1559420994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559421007] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4299824 kB) free!
[1559421070] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162640 kB) free!
[1559421114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559421127] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (2007424 kB) free!
[1559421234] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559421247] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10071512 kB) free.
[1559421250] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 392 processes
[1559421670] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.3% (2314780 kB) free!
[1559421834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559421847] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4514884 kB) free!
[1559421903] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 0.80, 11.14
[1559421954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559421967] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (2135908 kB) free!
[1559422074] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559422087] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10068832 kB) free.
[1559422674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559422687] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4436492 kB) free!
[1559422794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559422807] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (2039520 kB) free!
[1559422870] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.7% (3839828 kB) free.
[1559422915] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559422927] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10068812 kB) free.
[1559423103] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.33, 3.30
[1559423514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559423527] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4246116 kB) free!
[1559423634] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559423647] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (1849648 kB) free!
[1559423754] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1559423767] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10057908 kB) free.
[1559424354] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559424368] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4108172 kB) free!
[1559424474] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559424488] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (1700776 kB) free!
[1559424594] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559424608] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10053108 kB) free.
[1559424670] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.0% (1781528 kB) free!
[1559424790] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.8% (4674864 kB) free.
[1559425194] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559425208] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3970604 kB) free!
[1559425315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559425328] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (1595896 kB) free!
[1559425435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559425448] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10067328 kB) free.
[1559425450] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 388 processes
[1559426035] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559426048] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3911004 kB) free!
[1559426050] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 373 processes
[1559426155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559426168] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (1509036 kB) free!
[1559426190] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 466.39, 197.99, 75.51
[1559426275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559426288] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10062916 kB) free.
[1559426310] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 505.33, 301.70, 128.37
[1559426430] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 511.11, 371.24, 174.87
[1559426875] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559426888] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3692564 kB) free!
[1559426995] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559427008] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10054296 kB) free.
[1559427054] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2556664 kB) free!
[1559427116] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559427608] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5506932 kB) free!
[1559427715] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559427729] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (3786280 kB) free!
[1559427835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559427849] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9917584 kB) free.
[1559427850] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 363 processes
[1559427955] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1559428254] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452636 kB) free!
[1559428308] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 137.15, 60.15, 23.03
[1559428428] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 146.76, 89.30, 38.21
[1559428449] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5613288 kB) free!
[1559428548] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 148.22, 108.85, 51.57
[1559428555] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559428569] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (3755704 kB) free!
[1559428675] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559428689] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10058084 kB) free.
[1559428795] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559429138] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (2898108 kB) free!
[1559429289] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5498084 kB) free!
[1559429395] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559429409] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3544736 kB) free!
[1559429430] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.61, 1.74, 13.37
[1559429515] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559429529] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10049936 kB) free.
[1559429650] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1559429748] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.10, 2.17, 14.52
[1559430116] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559430129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5406428 kB) free!
[1559430236] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559430249] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3422572 kB) free!
[1559430356] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1559430369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10044848 kB) free.
[1559430630] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.52, 1.66, 4.89
[1559430850] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1559430948] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.28, 0.30, 4.16
[1559430956] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559430969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5359636 kB) free!
[1559431076] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559431090] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3385184 kB) free!
[1559431196] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559431210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10035856 kB) free.
[1559431450] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 364 processes
[1559431796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559431810] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5218048 kB) free!
[1559431917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559431930] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (3131416 kB) free!
[1559432037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559432050] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 354 processes
[1559432050] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10051448 kB) free.
[1559432637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559432650] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 364 processes
[1559432650] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5034360 kB) free!
[1559432757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559432770] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (2887024 kB) free!
[1559432877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559432890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10041604 kB) free.
[1559433477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559433490] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4861920 kB) free!
[1559433597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559433610] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (2585064 kB) free!
[1559433717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559433730] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9875140 kB) free!
[1559433790] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559434006] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2438900 kB) free!
[1559434330] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10027876 kB) free.
[1559434450] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 389 processes
[1559435050] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 393 processes
[1559435139] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2414328 kB) free!
[1559435590] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (443776 kB) free!
[1559435711] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (1120288 kB) free!
[1559435831] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (271748 kB) free!
[1559435890] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4868720 kB) free!
[1559436010] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4772564 kB) free!
[1559436130] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9845668 kB) free!
[1559436130] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4760760 kB) free!
[1559436250] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7825132 kB) free!
[1559436250] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 397 processes
[1559436370] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6174220 kB) free!
[1559436431] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 34.0% (5516632 kB) free.
[1559436850] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 377 processes
[1559437230] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 206.33, 147.80, 69.70
[1559437350] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 207.36, 167.50, 86.38
[1559437470] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 207.51, 180.78, 101.09
[1559437474] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5102 MB (9% inode=97%):
[1559437854] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (2690520 kB) free!
[1559438074] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6995 MB (13% inode=98%):
[1559438517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559438638] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559438758] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559439271] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1559439271] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.57, 2.14, 16.51
[1559439850] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 386 processes
[1559440254] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2423540 kB) free!
[1559440450] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 372 processes
[1559440930] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.7% (5066248 kB) free.
[1559441071] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.91, 1.69, 3.79
[1559441650] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1559442250] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1559442370] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1559442490] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 373 processes
[1559443880] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 31.7% (7760084 kB) free.
[1559444739] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2808272 kB) free!
[1559444748] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.88, 243.09, 158.03
[1559444868] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.42, 162.72, 138.89
[1559444988] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.14, 108.92, 122.07
[1559446040] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 605 processes
[1559446100] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 598 processes
[1559446471] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 512.78, 410.39, 208.04
[1559446474] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5574 MB (10% inode=97%):
[1559446591] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 512.52, 444.18, 244.96
[1559446711] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 71.03, 297.88, 215.50
[1559446788] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.05, 0.50, 17.89
[1559446987] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.5% (23577964 kB) free.
[1559447036] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.7% (22957396 kB) free.
[1559447074] SERVICE FLAPPING ALERT: pservercluster2;HDD_Root status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1559447074] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6987 MB (13% inode=98%):
[1559447103] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.14, 12.63, 8.68
[1559447155] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.7% (14748316 kB) free.
[1559447202] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.4% (14862504 kB) free.
[1559447223] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.31, 8.49, 7.64
[1559447344] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.43, 5.83, 6.77
[1559447530] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4544388 kB) free!
[1559447650] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4396248 kB) free!
[1559447739] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2434840 kB) free!
[1559447770] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (4340072 kB) free!
[1559447944] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 1.11, 3.75
[1559448588] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.38, 2.85
[1559449111] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.73, 1.93, 18.00
[1559449509] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24504 MB (10% inode=99%):
[1559450911] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.49, 2.08, 4.26
[1559451490] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1559451638] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.8% (22258680 kB) free.
[1559452031] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3184900 kB) free!
[1559452090] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 399 processes
[1559452151] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.7% (4173884 kB) free.
[1559452210] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 395 processes
[1559452330] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 391 processes
[1559452571] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2380524 kB) free!
[1559454233] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 348 processes
[1559454353] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 348 processes
[1559454366] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175424 kB) free!
[1559454411] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (174060 kB) free!
[1559454473] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 348 processes
[1559454486] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176716 kB) free!
[1559454531] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155944 kB) free!
[1559454606] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163268 kB) free!
[1559454652] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173500 kB) free!
[1559454731] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1559455632] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.5% (14868272 kB) free.
[1559455673] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559455681] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 69.3% (16986368 kB) free.
[1559455719] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.6% (15057536 kB) free.
[1559455751] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (3048332 kB) free!
[1559455871] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.4% (2981048 kB) free!
[1559455931] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1559455991] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.2% (2468588 kB) free!
[1559456051] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559456171] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1559457191] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (3306472 kB) free.
[1559457791] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (2814660 kB) free!
[1559457911] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.0% (2762556 kB) free!
[1559458031] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.8% (2718220 kB) free!
[1559458474] SERVICE FLAPPING ALERT: pservercluster2;HDD_Root status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559458631] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1559458631] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1342252 kB) free!
[1559458659] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (1644944 kB) free!
[1559459259] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159124 kB) free!
[1559459832] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (3550480 kB) free.
[1559459924] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178448 kB) free!
[1559460044] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176664 kB) free!
[1559460164] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (177096 kB) free!
[1559460432] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.1% (2935616 kB) free!
[1559460552] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.5% (2673668 kB) free!
[1559460672] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.1% (2288176 kB) free!
[1559461111] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559461272] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1220124 kB) free!
[1559461872] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.2% (2627204 kB) free!
[1559462259] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (1664008 kB) free!
[1559462544] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1559462753] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1559462819] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;SOFT;1;USERS CRITICAL - 13 users currently logged in
[1559462860] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1559462860] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162864 kB) free!
[1559462939] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;SOFT;2;USERS CRITICAL - 13 users currently logged in
[1559463059] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1559463971] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1559464571] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559464691] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 376 processes
[1559464811] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 386 processes
[1559465472] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1209048 kB) free!
[1559465873] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 360 processes
[1559465994] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 361 processes
[1559466072] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.0% (1783112 kB) free!
[1559466114] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1559466611] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 385 processes
[1559466714] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 361 processes
[1559466834] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 360 processes
[1559466909] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 124604 MB (55% inode=99%):
[1559466954] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 361 processes
[1559467042] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (186340 kB) free!
[1559467162] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (195228 kB) free!
[1559467212] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1559467282] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (192816 kB) free!
[1559467332] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1559467873] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1248580 kB) free!
[1559467932] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1559468052] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1559468154] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1559468170] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4914540 kB) free!
[1559468172] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 388 processes
[1559468344] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 153.74, 108.84, 50.33
[1559468464] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 61.77, 98.01, 54.09
[1559468473] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.5% (2669720 kB) free!
[1559468584] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.40, 65.62, 47.54
[1559468770] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5062312 kB) free!
[1559469073] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (477524 kB) free!
[1559469673] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.8% (2717988 kB) free!
[1559469761] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.1% (14974976 kB) free.
[1559469784] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 1.76, 13.41
[1559469812] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.5% (15035184 kB) free.
[1559470273] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (752536 kB) free!
[1559470362] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171092 kB) free!
[1559470482] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169196 kB) free!
[1559470555] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 360 processes
[1559470572] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 380 processes
[1559470602] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159608 kB) free!
[1559470675] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 360 processes
[1559470795] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1559470859] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1559470873] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.4% (2819288 kB) free!
[1559470985] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.29, 3.89
[1559471172] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 390 processes
[1559471292] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 385 processes
[1559471683] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559471770] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4696248 kB) free!
[1559471892] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1559471995] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 360 processes
[1559472012] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1559472073] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1039088 kB) free!
[1559472116] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559472132] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 381 processes
[1559472236] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 359 processes
[1559472673] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (3062932 kB) free!
[1559472732] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1559472836] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1559472836] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1559473332] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 383 processes
[1559473453] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 381 processes
[1559473573] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 381 processes
[1559473859] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1559474473] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (650092 kB) free!
[1559474636] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 368 processes
[1559474756] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1559474773] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 386 processes
[1559475073] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.4% (3150700 kB) free!
[1559475373] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 369 processes
[1559475493] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1559475613] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1559475673] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1235848 kB) free!
[1559476213] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1559476273] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (3359280 kB) free.
[1559477388] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 105.08, 291.13, 179.13
[1559477413] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 386 processes
[1559477508] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.48, 194.91, 157.45
[1559477534] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1559477637] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.03, 126.25, 136.92
[1559477654] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1559478254] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 385 processes
[1559478356] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1559478374] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 379 processes
[1559478476] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1559478494] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 375 processes
[1559479111] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 139.81, 66.40, 26.75
[1559479231] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 146.62, 93.28, 41.41
[1559479273] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (2610384 kB) free!
[1559479351] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 147.41, 111.25, 54.29
[1559479393] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155968 kB) free!
[1559479474] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5276 MB (10% inode=97%):
[1559479513] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (663660 kB) free!
[1559480074] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6999 MB (13% inode=98%):
[1559481151] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.33, 2.14, 13.77
[1559481237] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.02, 0.52, 11.89
[1559482095] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 392 processes
[1559482351] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.70, 1.69, 4.98
[1559482437] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.20, 0.28, 3.46
[1559482695] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 379 processes
[1559482815] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 372 processes
[1559482935] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1559483113] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.1% (1635172 kB) free!
[1559484476] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559485513] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.8% (3377640 kB) free.
[1559486551] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 485.14, 228.14, 89.48
[1559486671] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 508.86, 322.24, 140.78
[1559486791] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1559486791] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 512.13, 385.23, 185.86
[1559487135] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1559487735] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559487855] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559487975] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1559489791] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.78, 1.91, 13.89
[1559491591] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.75, 1.97, 3.61
[1559492770] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.8% (9770340 kB) free!
[1559493313] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (2614204 kB) free!
[1559493358] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559493434] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (3433188 kB) free.
[1559494033] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.5% (2522096 kB) free!
[1559494153] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (3305392 kB) free.
[1559494386] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 96.34, 269.19, 160.35
[1559494506] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.12, 180.17, 140.93
[1559494626] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.22, 120.72, 123.90
[1559494753] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.7% (1730100 kB) free!
[1559494873] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.0% (2438536 kB) free!
[1559494993] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.4% (3145988 kB) free!
[1559495158] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559495278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559495398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559495593] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (971264 kB) free!
[1559496193] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (3523344 kB) free.
[1559496426] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 0.86, 18.26
[1559496793] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1033036 kB) free!
[1559496913] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (1305776 kB) free!
[1559496975] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 392 processes
[1559497033] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1227740 kB) free!
[1559497576] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1559497696] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1559498226] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.28, 0.35, 2.94
[1559498233] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.9% (2251284 kB) free!
[1559498296] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 393 processes
[1559498416] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 390 processes
[1559498833] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.8% (3378716 kB) free.
[1559499016] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 393 processes
[1559499136] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1559499736] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 398 processes
[1559499856] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 393 processes
[1559499976] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 392 processes