[1585508543] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1585513348] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25708808 kB) free!
[1585515148] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16503.50, 12050.45, 5670.87
[1585515268] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16524.64, 13531.98, 6987.30
[1585515288] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 613 processes
[1585515297] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(438) < 203.212.70.170 [30/Mar/2020:02:24:56 +0530] TCP 502 0 0 120.003 "10.147.212.142:3001, 10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.154:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 60.001, 60.001"
[1585515325] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.6% (9470380 kB) free.
[1585515388] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16528.34, 14523.50, 8144.18
[1585515408] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 632 processes
[1585515417] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(15) < 203.212.70.170 [30/Mar/2020:02:26:16 +0530] TCP 502 0 0 120.002 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 0.000, 0.001"
[1585515528] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1332 processes
[1585515538] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(2) < 172.105.59.173 [30/Mar/2020:02:27:01 +0530] TCP 502 0 0 120.002 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1585516128] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 356 processes
[1585516738] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585517337] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16535.81, 12289.47, 5853.98
[1585517339] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(278) < 203.212.70.170 [30/Mar/2020:02:58:58 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1585517428] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(269) < 103.16.101.72 [30/Mar/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1585517457] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7024.47, 11084.90, 6260.26
[1585517459] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1938) < 203.212.70.170 [30/Mar/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1585517548] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(11) < 103.16.101.72 [30/Mar/2020:03:01:31 +0530] TCP 502 0 0 60.000 "10.147.212.37:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1585517577] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 948.86, 7417.70, 5501.39
[1585517579] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(114) < 172.105.59.173 [30/Mar/2020:03:01:52 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1585517669] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585517725] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 21.9% (5400704 kB) free.
[1585517924] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(369) < 2020/03/30 03:01:31 [error] 22129#0: *78712056 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585518044] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585518179] SERVICE FLAPPING ALERT: primary-server;Nginx Access Log Status;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1585518179] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585518388] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.05, 3.55, 516.40
[1585518748] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12671888 kB) free!
[1585519348] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26051632 kB) free!
[1585519977] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.76, 4.12, 416.32
[1585521325] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4297608 kB) free!
[1585521445] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (4289684 kB) free!
[1585521565] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (4286288 kB) free!
[1585522948] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1585522948] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (13140872 kB) free!
[1585524148] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26111448 kB) free!
[1585524926] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163676 kB) free!
[1585525046] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (141348 kB) free!
[1585525165] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 31.8% (7850612 kB) free.
[1585525166] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2623408 kB) free!
[1585525183] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 48.8% (11951180 kB) free.
[1585525766] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1368168 kB) free!
[1585527149] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27272520 kB) free.
[1585527749] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24198476 kB) free!
[1585527869] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20616012 kB) free!
[1585527989] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16988500 kB) free!
[1585529579] SERVICE FLAPPING ALERT: primary-server;Nginx Access Log Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585531589] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27105844 kB) free.
[1585532189] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22573024 kB) free!
[1585532309] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.9% (18389088 kB) free!
[1585532429] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14504320 kB) free!
[1585535829] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1585535949] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1585540149] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1585540269] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1585540596] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4700212 kB) free!
[1585541213] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.85, 7.60, 5.56
[1585541333] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.41, 7.79, 5.88
[1585541365] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4487816 kB) free!
[1585541429] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9055148 kB) free!
[1585541453] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.26, 7.36, 5.96
[1585541485] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (4424256 kB) free!
[1585541605] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (4334064 kB) free!
[1585542029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39309532 kB) free.
[1585542257] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4718776 kB) free!
[1585542857] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (5830896 kB) free!
[1585543405] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2265580 kB) free!
[1585543783] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4865392 kB) free!
[1585543903] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4754488 kB) free!
[1585544023] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4652536 kB) free!
[1585544057] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3141488 kB) free!
[1585544430] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26360172 kB) free!
[1585544506] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 37.1% (18276708 kB) free.
[1585544550] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (21776064 kB) free!
[1585544670] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36372588 kB) free.
[1585545069] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1585545189] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1585546470] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24817964 kB) free!
[1585546590] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.2% (20094080 kB) free!
[1585546596] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 62.8% (30970288 kB) free.
[1585546710] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35449504 kB) free.
[1585546929] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1585547046] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1585547049] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585547166] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1585547170] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1585547271] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165240 kB) free!
[1585547391] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179904 kB) free!
[1585547511] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170456 kB) free!
[1585547623] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1293516 kB) free!
[1585547721] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4466604 kB) free!
[1585547841] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4085520 kB) free!
[1585547961] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.6% (3733508 kB) free!
[1585548561] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1904824 kB) free!
[1585549110] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25230196 kB) free!
[1585549230] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19778428 kB) free!
[1585549348] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4452284 kB) free!
[1585549350] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14580676 kB) free!
[1585549468] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (4011196 kB) free!
[1585549566] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1585549588] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.7% (3636632 kB) free!
[1585549686] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585549806] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1585549854] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 48.2% (11802232 kB) free.
[1585550188] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2168472 kB) free!
[1585550362] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.9% (15255624 kB) free.
[1585551388] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.2% (22769716 kB) free.
[1585551606] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1585552389] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 327 processes
[1585552509] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1585552950] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4510468 kB) free!
[1585553506] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9565552 kB) free!
[1585553511] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 37.4% (9159576 kB) free.
[1585553626] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9403860 kB) free!
[1585553746] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (9245176 kB) free!
[1585555254] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.9% (3885160 kB) free!
[1585555374] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.7% (3590928 kB) free!
[1585555494] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.7% (3360360 kB) free!
[1585555950] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13780728 kB) free!
[1585556694] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1823352 kB) free!
[1585557150] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12799708 kB) free!
[1585557346] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4806088 kB) free!
[1585558162] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (3783760 kB) free!
[1585558282] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.7% (3310580 kB) free!
[1585558402] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (2840916 kB) free!
[1585558509] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1585558629] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1585558911] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4367652 kB) free!
[1585559002] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1585559002] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (385940 kB) free!
[1585559031] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4168396 kB) free!
[1585559151] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (4003604 kB) free!
[1585559550] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21365296 kB) free!
[1585559788] SERVICE ALERT: db1-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1890380 kB) free!
[1585559796] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9525096 kB) free!
[1585559908] SERVICE ALERT: db1-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (1230008 kB) free!
[1585559916] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9176420 kB) free!
[1585560028] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (595800 kB) free!
[1585560037] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (8851924 kB) free!
[1585560351] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2316912 kB) free!
[1585560750] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28015020 kB) free.
[1585560779] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(144) < 92.42.104.166 [30/Mar/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585560803] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.5% (15374076 kB) free.
[1585560899] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585561350] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1731576 kB) free!
[1585561470] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27563496 kB) free.
[1585561828] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.4% (22797756 kB) free.
[1585561837] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3873084 kB) free!
[1585562070] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (813952 kB) free!
[1585562190] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27200368 kB) free.
[1585562790] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1313860 kB) free!
[1585562910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27353416 kB) free.
[1585563510] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2352364 kB) free!
[1585563630] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27367560 kB) free.
[1585564230] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3186148 kB) free!
[1585564350] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (392480 kB) free!
[1585564470] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24788444 kB) free!
[1585564937] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (2653412 kB) free!
[1585565057] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 36.0% (5849812 kB) free.
[1585565070] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390900 kB) free!
[1585566053] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.71, 25.42, 19.49
[1585566653] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.32, 16.80, 17.96
[1585566803] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4489492 kB) free!
[1585566857] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3157484 kB) free!
[1585566871] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24835728 kB) free!
[1585566923] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4155636 kB) free!
[1585566977] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3543664 kB) free.
[1585567043] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.2% (3854776 kB) free!
[1585567577] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (1998036 kB) free!
[1585567697] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.3% (2485132 kB) free!
[1585567817] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (2953784 kB) free!
[1585568071] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3687468 kB) free!
[1585568243] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2090204 kB) free!
[1585568418] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3342076 kB) free.
[1585569271] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15377640 kB) free!
[1585570471] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26824424 kB) free.
[1585570818] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (2418220 kB) free!
[1585570938] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2997316 kB) free!
[1585571058] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3529540 kB) free.
[1585571071] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2510692 kB) free!
[1585571191] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27236520 kB) free.
[1585571429] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4650912 kB) free!
[1585571549] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4382332 kB) free!
[1585571658] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2403964 kB) free!
[1585571669] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.7% (4134392 kB) free!
[1585571778] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2899380 kB) free!
[1585571791] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4208896 kB) free!
[1585571843] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.9% (15466992 kB) free.
[1585571898] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3592412 kB) free.
[1585571911] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394360 kB) free!
[1585572031] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.7% (25903400 kB) free!
[1585572269] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2477440 kB) free!
[1585572498] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (2701664 kB) free!
[1585572618] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3377044 kB) free.
[1585572631] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1557576 kB) free!
[1585572869] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1585572869] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22874528 kB) free.
[1585573218] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2481464 kB) free!
[1585573338] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3251840 kB) free!
[1585573458] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3590344 kB) free.
[1585573831] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15200752 kB) free!
[1585574658] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2889804 kB) free!
[1585574779] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3145676 kB) free!
[1585574899] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3122712 kB) free!
[1585575631] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5184740 kB) free!
[1585576831] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21874644 kB) free!
[1585577431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28356932 kB) free.
[1585577899] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1585577899] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1571364 kB) free!
[1585578031] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4924116 kB) free!
[1585578151] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (673088 kB) free!
[1585578271] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27135336 kB) free.
[1585578653] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.01, 4.19, 5.00
[1585578871] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2558004 kB) free!
[1585578991] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17664848 kB) free!
[1585579111] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24680740 kB) free!
[1585579711] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (844560 kB) free!
[1585579853] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.77, 5.42, 5.19
[1585579973] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.00, 5.38, 5.18
[1585580093] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.82, 5.16, 5.12
[1585580243] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1585580911] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18235196 kB) free!
[1585582112] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3339904 kB) free!
[1585582493] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.59, 3.62, 4.35
[1585582699] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3117872 kB) free!
[1585583312] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17485016 kB) free!
[1585584269] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585584512] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2714944 kB) free!
[1585585099] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1496832 kB) free!
[1585585699] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 32.0% (5195224 kB) free.
[1585585712] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16225348 kB) free!
[1585586899] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3049104 kB) free!
[1585586912] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (2975580 kB) free!
[1585587019] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3505752 kB) free.
[1585587857] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (5382192 kB) free!
[1585588112] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15499276 kB) free!
[1585588854] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 5300.17, 1365.18, 460.68
[1585588974] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 10125.30, 4512.65, 1723.74
[1585589094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10775.37, 6618.38, 2833.55
[1585589279] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 6030.96, 7561.86, 4072.59
[1585589312] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1044904 kB) free!
[1585589399] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 814.65, 5060.27, 3578.95
[1585589519] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 111.23, 3386.50, 3145.20
[1585590719] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.58, 62.89, 865.08
[1585591113] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18264544 kB) free!
[1585591319] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.04, 10.30, 454.32
[1585591494] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.70, 6.81, 390.88
[1585592094] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.85, 2.39, 205.60
[1585592313] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3327176 kB) free!