[1590518175] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16097512 kB) free!
[1590518775] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31938352 kB) free.
[1590519051] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.9% (3724804 kB) free.
[1590519375] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7949268 kB) free!
[1590519495] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4057664 kB) free!
[1590519615] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (558908 kB) free!
[1590520251] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (183392 kB) free!
[1590520371] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2889280 kB) free!
[1590520491] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3285920 kB) free.
[1590520816] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23437256 kB) free!
[1590521091] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (788648 kB) free!
[1590521212] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (220264 kB) free!
[1590521332] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3474112 kB) free.
[1590521416] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2615252 kB) free!
[1590521511] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (372656 kB) free!
[1590521571] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173756 kB) free!
[1590521631] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (184128 kB) free!
[1590521697] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (184944 kB) free!
[1590522597] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 76.7% (18787564 kB) free.
[1590522616] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24399484 kB) free!
[1590523216] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3397372 kB) free!
[1590524416] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21573468 kB) free!
[1590525016] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32285432 kB) free.
[1590525616] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10360976 kB) free!
[1590525736] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6692408 kB) free!
[1590525856] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3361576 kB) free!
[1590526342] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5608276 kB) free!
[1590526434] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238824 kB) free!
[1590526462] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (260880 kB) free!
[1590526583] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226624 kB) free!
[1590527056] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20390112 kB) free!
[1590527444] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (3181316 kB) free!
[1590527656] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31579504 kB) free.
[1590527675] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2766820 kB) free!
[1590527783] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.1% (44918916 kB) free.
[1590528044] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2298864 kB) free!
[1590528234] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (45253408 kB) free.
[1590528256] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8762652 kB) free!
[1590528275] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2189552 kB) free!
[1590528376] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5140896 kB) free!
[1590528497] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31967816 kB) free.
[1590528533] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1590528602] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2715764 kB) free!
[1590528644] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (3505652 kB) free!
[1590528704] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1590528927] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2853) < 172.105.59.173 [27/May/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590529047] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590529096] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2325324 kB) free!
[1590529202] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1407764 kB) free!
[1590529216] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393848 kB) free!
[1590529336] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35181560 kB) free.
[1590529936] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7552228 kB) free!
[1590530056] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3554568 kB) free!
[1590530176] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23615668 kB) free!
[1590530776] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (7112500 kB) free!
[1590531376] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18439852 kB) free!
[1590531875] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.8% (10004848 kB) free.
[1590531976] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30230652 kB) free.
[1590532244] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1590532244] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2251472 kB) free!
[1590532475] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (1503944 kB) free!
[1590532576] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6289744 kB) free!
[1590532595] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (211024 kB) free!
[1590532696] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3382808 kB) free!
[1590532715] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1590532715] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (220204 kB) free!
[1590532816] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32723856 kB) free.
[1590533416] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12307832 kB) free!
[1590533536] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34975620 kB) free.
[1590534137] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19263372 kB) free!
[1590534257] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15804720 kB) free!
[1590534377] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37066156 kB) free.
[1590534516] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.8% (9499220 kB) free.
[1590534977] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18124644 kB) free!
[1590535097] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14665900 kB) free!
[1590535217] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35249060 kB) free.
[1590535817] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16425660 kB) free!
[1590535937] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (36984256 kB) free.
[1590536316] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (231528 kB) free!
[1590536436] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180088 kB) free!
[1590536537] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18790560 kB) free!
[1590536556] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184188 kB) free!
[1590536657] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15257220 kB) free!
[1590536777] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35661188 kB) free.
[1590537284] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.2% (18188484 kB) free.
[1590537377] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16757060 kB) free!
[1590537497] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36873792 kB) free.
[1590537603] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.6% (12482244 kB) free.
[1590537644] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 57.9% (14200708 kB) free.
[1590538097] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18813692 kB) free!
[1590538217] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15191456 kB) free!
[1590538338] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35805800 kB) free.
[1590538937] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15891904 kB) free!
[1590539057] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36635720 kB) free.
[1590539657] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16268552 kB) free!
[1590539777] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12402732 kB) free!
[1590539897] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34617392 kB) free.
[1590540497] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13925672 kB) free!
[1590540617] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36114968 kB) free.
[1590541203] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (4015632 kB) free!
[1590541217] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13927124 kB) free!
[1590541306] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.2% (18706588 kB) free!
[1590541323] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (748492 kB) free!
[1590541337] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35686164 kB) free.
[1590541426] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.0% (14522996 kB) free!
[1590541443] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188212 kB) free!
[1590541546] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (10392736 kB) free!
[1590541844] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3918580 kB) free!
[1590541888] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1590541904] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590541937] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16256280 kB) free!
[1590541964] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3812968 kB) free!
[1590542008] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1590542024] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1590542057] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12213652 kB) free!
[1590542084] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (3724648 kB) free!
[1590542084] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192140 kB) free!
[1590542128] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1590542144] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590542177] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34601460 kB) free.
[1590542204] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182812 kB) free!
[1590542324] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (205224 kB) free!
[1590542728] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1590542745] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1590542777] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13321700 kB) free!
[1590542887] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/27 06:52:24 [error] 22135#0: *91174271 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590542897] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35026312 kB) free.
[1590543007] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590543497] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13238556 kB) free!
[1590543616] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 74.3% (18209616 kB) free.
[1590543618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34694224 kB) free.
[1590544217] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14794064 kB) free!
[1590544337] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10928780 kB) free!
[1590544457] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32545976 kB) free.
[1590544546] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 73.8% (97251564 kB) free.
[1590544983] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (183504 kB) free!
[1590545057] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12803444 kB) free!
[1590545084] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461500 kB) free!
[1590545177] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34191508 kB) free.
[1590545777] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14270552 kB) free!
[1590545897] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33983548 kB) free.
[1590546497] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15675120 kB) free!
[1590546617] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11861860 kB) free!
[1590546737] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590546737] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32407368 kB) free.
[1590546756] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1590547337] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11483432 kB) free!
[1590547457] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33376844 kB) free.
[1590548057] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11992556 kB) free!
[1590548177] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20862184 kB) free!
[1590548298] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30461980 kB) free.
[1590548897] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9101224 kB) free!
[1590549017] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32092696 kB) free.
[1590549298] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.22757.747> (10.147.212.147:59224 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590549418] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590549617] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10274544 kB) free!
[1590549737] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31813540 kB) free.
[1590550337] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11057448 kB) free!
[1590550457] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6210952 kB) free!
[1590550577] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30892196 kB) free.
[1590551177] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7125200 kB) free!
[1590551297] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2916508 kB) free!
[1590551417] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27669452 kB) free.
[1590551439] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.13, 9.38, 7.00
[1590551559] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.19, 10.05, 7.53
[1590551679] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.97, 10.49, 8.01
[1590552017] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4563412 kB) free!
[1590552137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28977284 kB) free.
[1590552737] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6233372 kB) free!
[1590552857] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1758824 kB) free!
[1590552978] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25336568 kB) free!
[1590553578] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2277236 kB) free!
[1590554084] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1590554679] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.04, 3.62, 4.44
[1590555378] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24164456 kB) free!
[1590555879] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.53, 8.85, 6.73
[1590555978] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1590555978] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394860 kB) free!
[1590555999] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.62, 8.07, 6.72
[1590556119] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.46, 7.16, 6.55
[1590556217] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.1% (3949972 kB) free!
[1590556337] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.4% (3783252 kB) free!
[1590556457] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.8% (3618776 kB) free!
[1590557178] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20943208 kB) free!
[1590557407] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/27 10:58:02 [error] 22135#0: *91220274 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590557527] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590557657] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1272100 kB) free!
[1590557778] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401980 kB) free!
[1590558263] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.1% (2707188 kB) free!
[1590558383] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.4% (2345252 kB) free!
[1590558503] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2009436 kB) free!
[1590558857] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.9% (4626788 kB) free!
[1590560178] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14303992 kB) free!
[1590560640] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1590560657] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1590560657] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1611512 kB) free!
[1590560760] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1590560920] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1590560920] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.80, 4.21, 4.71
[1590561378] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403100 kB) free!
[1590561504] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.4% (15815260 kB) free.
[1590561520] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.02, 5.86, 5.51
[1590561532] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3212924 kB) free!
[1590561640] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.31, 6.15, 5.64
[1590561652] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4200068 kB) free.
[1590561760] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.22, 7.00, 6.02
[1590561979] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16324720 kB) free!
[1590562579] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32243712 kB) free.
[1590563179] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5290260 kB) free!
[1590563248] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(366) < 103.255.146.164 [27/May/2020:12:32:37 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590563290] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.7% (2386140 kB) free!
[1590563299] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (843568 kB) free!
[1590563368] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590563419] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409696 kB) free!
[1590564619] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22885404 kB) free!
[1590564834] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9225740 kB) free!
[1590564954] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8891164 kB) free!
[1590565074] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.4% (8595432 kB) free!
[1590565219] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399376 kB) free!
[1590565768] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(817) < 103.255.146.164 [27/May/2020:13:14:12 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590565852] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2719108 kB) free!
[1590565889] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590565972] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.2% (2473896 kB) free!
[1590566092] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.2% (2315312 kB) free!
[1590566160] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 8 warn out of 344 processes
[1590566280] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1590566419] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26037096 kB) free!
[1590566874] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3820212 kB) free!
[1590567019] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (430956 kB) free!
[1590567088] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(428) < 103.255.146.164 [27/May/2020:13:34:34 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590567208] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590567383] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8643184 kB) free!
[1590567503] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (8413792 kB) free!
[1590567504] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (3517816 kB) free!
[1590567619] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20443196 kB) free!
[1590567623] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (8129276 kB) free!
[1590567624] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.5% (3024732 kB) free!
[1590567744] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.2% (2510592 kB) free!
[1590568080] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 343 processes
[1590568200] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1590568219] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395436 kB) free!
[1590568328] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.9489.6174> (10.147.212.53:44282 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1590568344] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1590568344] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (322308 kB) free!
[1590568448] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590568492] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (789524 kB) free!
[1590568819] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13432756 kB) free!
[1590569092] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.8% (2401696 kB) free!
[1590569420] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398736 kB) free!
[1590569423] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4429556 kB) free!
[1590569728] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590569745] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590569848] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590569866] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590569969] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1590569986] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590570569] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1590570586] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1590570620] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29252908 kB) free.
[1590570727] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/27 14:33:33 [error] 22135#0: *91294490 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590570847] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590570857] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 39.0% (9546964 kB) free.
[1590570892] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1590570892] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (907092 kB) free!
[1590570968] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (7012732 kB) free!
[1590571088] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.6% (6508976 kB) free!
[1590571186] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1590571208] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 17.3% (6062396 kB) free!
[1590571220] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1703996 kB) free!
[1590571306] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1590571340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403532 kB) free!
[1590571408] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(231) < 103.255.146.164 [27/May/2020:14:48:18 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590571426] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1590571460] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (459024 kB) free!
[1590571492] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (2284920 kB) free!
[1590571528] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590571809] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 86.6% (30323184 kB) free.
[1590571944] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.3% (16001860 kB) free.
[1590572129] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(171) < 185.255.8.58 [27/May/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590572249] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590572660] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31089152 kB) free.
[1590573160] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590573257] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.9% (3646624 kB) free!
[1590573260] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4353992 kB) free!
[1590573297] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4912736 kB) free!
[1590573357] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4916672 kB) free.
[1590573378] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 13.7% (3346536 kB) free!
[1590573380] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (460616 kB) free!
[1590573498] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 12.5% (3051584 kB) free!
[1590573500] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412400 kB) free!
[1590573600] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1590573720] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 334 processes
[1590573840] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1590573892] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (251304 kB) free!
[1590574098] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1631968 kB) free!
[1590574100] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14042124 kB) free!
[1590574492] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2628300 kB) free!
[1590574700] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (450944 kB) free!
[1590575692] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1290836 kB) free!
[1590575901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30083016 kB) free.
[1590576226] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1590576226] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590576240] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1590576292] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (2414856 kB) free!
[1590576360] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1590576501] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4082028 kB) free!
[1590576621] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396556 kB) free!
[1590576741] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (473776 kB) free!
[1590576760] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.17, 3.71, 4.88
[1590577360] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.94, 6.95, 5.89
[1590577480] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.75, 7.27, 6.14
[1590577492] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (786272 kB) free!
[1590577600] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.21, 9.27, 6.97
[1590577649] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(487) < 103.255.146.164 [27/May/2020:16:31:56 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590577769] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590578760] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1590578881] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1590579001] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590579144] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2774972 kB) free!
[1590579265] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (2228440 kB) free!
[1590579385] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1854484 kB) free!
[1590579569] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(269) < 103.255.146.164 [27/May/2020:17:01:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590579689] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590579893] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3425116 kB) free.
[1590580447] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/27 17:18:37 [error] 22135#0: *91345768 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590580567] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590581093] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3033212 kB) free!
[1590581213] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (437828 kB) free!
[1590581333] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (1790508 kB) free!
[1590581541] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29170756 kB) free.
[1590582141] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3268516 kB) free!
[1590582262] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (442636 kB) free!
[1590582382] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433948 kB) free!
[1590582982] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17185552 kB) free!
[1590582985] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16076544 kB) free.
[1590583133] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1155656 kB) free!
[1590583426] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1590583582] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (458276 kB) free!
[1590583733] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2784240 kB) free!
[1590584200] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.02, 4.51, 4.95
[1590584333] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3329836 kB) free.
[1590584933] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (1339952 kB) free!
[1590585053] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3018380 kB) free!
[1590585173] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3475212 kB) free.
[1590585400] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.12, 5.63, 5.31
[1590585498] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590585520] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.56, 4.84, 5.06
[1590585640] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.12, 4.12, 4.76
[1590585773] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1211196 kB) free!
[1590585893] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (792044 kB) free!
[1590586013] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2168572 kB) free!
[1590586614] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3597508 kB) free.
[1590587182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29938560 kB) free.
[1590587214] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.5% (1875332 kB) free!
[1590587334] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3364248 kB) free.
[1590587782] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2211528 kB) free!
[1590587902] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386728 kB) free!
[1590587934] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (640288 kB) free!
[1590588022] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (444780 kB) free!
[1590588054] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.8% (1921604 kB) free!
[1590588174] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.2% (3941132 kB) free.
[1590588622] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15961216 kB) free!
[1590588774] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (916940 kB) free!
[1590588894] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2455220 kB) free!
[1590589014] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.0% (3900052 kB) free.
[1590589222] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392140 kB) free!
[1590589614] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (671240 kB) free!
[1590589734] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.7% (2551992 kB) free!
[1590589823] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13907476 kB) free!
[1590589854] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3793448 kB) free.
[1590590423] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (438320 kB) free!
[1590590454] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (851528 kB) free!
[1590590575] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (2947684 kB) free!
[1590590695] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4029488 kB) free.
[1590591183] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (1652340 kB) free!
[1590591295] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1171708 kB) free!
[1590591385] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1590591415] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (2677548 kB) free!
[1590591535] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3853752 kB) free.
[1590591784] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1626476 kB) free!
[1590592135] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (654096 kB) free!
[1590592255] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.2% (2139048 kB) free!
[1590592375] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3668984 kB) free.
[1590593575] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (1438076 kB) free!
[1590593695] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2983632 kB) free!
[1590593815] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3638412 kB) free.
[1590594023] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30735592 kB) free.
[1590594623] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3382908 kB) free!
[1590594743] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (377428 kB) free!
[1590594863] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395240 kB) free!
[1590595015] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1217312 kB) free!
[1590595135] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2429076 kB) free!
[1590595256] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3533544 kB) free.
[1590595463] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16771004 kB) free!
[1590595855] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2022212 kB) free!
[1590595975] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3152524 kB) free!
[1590596063] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448348 kB) free!
[1590596095] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4037504 kB) free.
[1590598495] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (2932004 kB) free!
[1590598615] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3783552 kB) free.
[1590599215] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (2683080 kB) free!
[1590599335] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3376584 kB) free.
[1590599663] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29370960 kB) free.
[1590599936] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.8% (1917560 kB) free!
[1590600056] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2528312 kB) free!
[1590600176] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3135064 kB) free!
[1590600263] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4621996 kB) free!
[1590600383] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (479452 kB) free!
[1590600503] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (427152 kB) free!
[1590600776] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1592424 kB) free!
[1590601103] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14115796 kB) free!
[1590601376] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3589828 kB) free.
[1590601703] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425860 kB) free!
[1590601976] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.7% (1740344 kB) free!
[1590602096] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (2200100 kB) free!
[1590602216] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (2731492 kB) free!
[1590602816] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3483564 kB) free.
[1590602903] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29229124 kB) free.
[1590603503] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5218192 kB) free!
[1590603623] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1169280 kB) free!
[1590603743] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400984 kB) free!
[1590604016] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (1512424 kB) free!
[1590604136] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (1434364 kB) free!