[1585679661] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2832120 kB) free!
[1585679781] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3625320 kB) free.
[1585680381] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2836308 kB) free!
[1585680501] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3582048 kB) free.
[1585680518] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (6783072 kB) free!
[1585681101] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2856280 kB) free!
[1585681118] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16560620 kB) free!
[1585681221] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3558892 kB) free.
[1585681387] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201260 kB) free!
[1585681447] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187592 kB) free!
[1585681507] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (178676 kB) free!
[1585681567] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (190892 kB) free!
[1585681718] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26805576 kB) free.
[1585681821] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (2419372 kB) free!
[1585681941] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2970520 kB) free!
[1585682061] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3535756 kB) free.
[1585682167] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 77.9% (19098476 kB) free.
[1585682318] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7521000 kB) free!
[1585682439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4127172 kB) free!
[1585682559] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28956240 kB) free.
[1585682661] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2482624 kB) free!
[1585682781] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3098940 kB) free!
[1585682863] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (3732520 kB) free!
[1585682901] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585682901] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3508976 kB) free.
[1585683159] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9692284 kB) free!
[1585683279] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6331292 kB) free!
[1585683399] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2947204 kB) free!
[1585683463] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2051072 kB) free!
[1585683501] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (2497056 kB) free!
[1585683621] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2983688 kB) free!
[1585683741] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3483028 kB) free.
[1585684341] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.0% (1793980 kB) free!
[1585684461] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.0% (2276948 kB) free!
[1585684578] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10824600 kB) free!
[1585684581] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2698528 kB) free!
[1585684599] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20997448 kB) free!
[1585684698] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10828456 kB) free!
[1585684818] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (7203088 kB) free!
[1585685026] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 10.7% (5251784 kB) free!
[1585685146] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5249872 kB) free!
[1585685181] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3431792 kB) free.
[1585685199] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2946392 kB) free!
[1585685266] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.6% (5248180 kB) free!
[1585685781] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (2683320 kB) free!
[1585685902] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2972884 kB) free!
[1585686022] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3396696 kB) free.
[1585686399] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19337012 kB) free!
[1585686999] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28583372 kB) free.
[1585687066] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 29.6% (14575424 kB) free.
[1585687222] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (1513612 kB) free!
[1585687342] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.2% (1666076 kB) free!
[1585687462] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (1801144 kB) free!
[1585687599] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8884448 kB) free!
[1585687719] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5542972 kB) free!
[1585687839] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2247152 kB) free!
[1585688386] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 71.0% (34989104 kB) free.
[1585688575] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585688581] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(155) < 92.42.104.166 [01/Apr/2020:02:25:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585688701] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585689039] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19090488 kB) free!
[1585689279] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.1% (45894892 kB) free.
[1585689640] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1264488 kB) free!
[1585689812] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 34.6% (8490496 kB) free.
[1585689862] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1585689862] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 38.4% (6237216 kB) free.
[1585690372] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(260) < 2020/04/01 03:01:22 [error] 22134#0: *79600307 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585690439] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (2855588 kB) free!
[1585690492] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585690501] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2157) < 49.50.86.152 [01/Apr/2020:03:01:48 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1585690591] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(174) < 103.16.101.72 [01/Apr/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1585690621] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585690663] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (2943604 kB) free!
[1585690712] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585690840] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14966020 kB) free!
[1585691039] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2411136 kB) free!
[1585691221] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(28) < 103.13.97.241 [01/Apr/2020:03:14:10 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1585691263] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1585691263] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2435792 kB) free!
[1585691342] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585692018] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.4% (17647988 kB) free!
[1585692040] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5155796 kB) free!
[1585693240] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19744596 kB) free!
[1585693840] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1309440 kB) free!
[1585695040] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16518232 kB) free!
[1585696240] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3801484 kB) free!
[1585697440] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14726508 kB) free!
[1585697612] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (158668 kB) free!
[1585697733] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166664 kB) free!
[1585697854] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187032 kB) free!
[1585698641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27791948 kB) free.
[1585698839] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.2% (19424032 kB) free.
[1585699048] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 57.0% (14059380 kB) free.
[1585699063] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 56.5% (13850864 kB) free.
[1585699241] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6906036 kB) free!
[1585699361] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3683256 kB) free!
[1585699462] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1585699481] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23159100 kB) free!
[1585700081] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385052 kB) free!
[1585701281] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17049824 kB) free!
[1585702218] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.8% (47232584 kB) free.
[1585703081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30160084 kB) free.
[1585704881] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13078764 kB) free!
[1585705001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31287972 kB) free.
[1585705063] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1585706201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15139720 kB) free!
[1585706322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31244528 kB) free.
[1585706575] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.44, 6.09, 5.04
[1585706695] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.83, 7.24, 5.58
[1585706815] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.86, 7.38, 5.84
[1585708015] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.89, 2.98, 4.09
[1585708121] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15310764 kB) free!
[1585708241] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31501284 kB) free.
[1585708841] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23423172 kB) free!
[1585708961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28971648 kB) free.
[1585709815] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.21, 5.55, 5.05
[1585709935] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.10, 5.02, 4.92
[1585710161] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10970088 kB) free!
[1585710281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28577476 kB) free.
[1585710535] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.37, 6.17, 5.50
[1585710655] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.87, 6.34, 5.64
[1585710775] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.07, 6.70, 5.86
[1585711481] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12602380 kB) free!
[1585711601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28977852 kB) free.
[1585712201] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4905360 kB) free!
[1585712321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26795756 kB) free.
[1585712922] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3378800 kB) free!
[1585713042] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26831104 kB) free.
[1585713448] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164108 kB) free!
[1585713568] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167940 kB) free!
[1585713642] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7365412 kB) free!
[1585713688] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188972 kB) free!
[1585713762] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.8% (24761384 kB) free!
[1585713882] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20600528 kB) free!
[1585715263] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4773940 kB) free!
[1585715383] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4719376 kB) free!
[1585715503] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4504104 kB) free!
[1585715575] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1585715575] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.34, 3.80, 4.73
[1585715682] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (6899368 kB) free!
[1585716175] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.95, 8.04, 6.25
[1585716296] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.18, 8.59, 6.69
[1585716416] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.92, 8.88, 7.02
[1585717303] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1457180 kB) free!
[1585717482] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13675288 kB) free!
[1585718612] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4412112 kB) free!
[1585718682] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397216 kB) free!
[1585718732] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (3859580 kB) free!
[1585718817] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.54, 20.76, 20.03
[1585718852] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (3331532 kB) free!
[1585719008] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1585719128] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1585719248] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1585719282] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16641368 kB) free!
[1585719417] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.23, 11.41, 15.78
[1585719452] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (1205504 kB) free!
[1585720590] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (4110516 kB) free!
[1585720710] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.5% (3823324 kB) free!
[1585720830] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.3% (3518592 kB) free!
[1585721082] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2225044 kB) free!
[1585721430] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2102084 kB) free!
[1585722282] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17652172 kB) free!
[1585722670] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1585722790] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1585722882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28830500 kB) free.
[1585723052] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.6% (15395384 kB) free.
[1585723390] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1585723482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (381056 kB) free!
[1585723510] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1585723602] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384276 kB) free!
[1585723722] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32519832 kB) free.
[1585723830] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.5% (22837184 kB) free.
[1585724217] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.72, 19.04, 14.04
[1585724322] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9380004 kB) free!
[1585724442] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5283244 kB) free!
[1585724562] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1590728 kB) free!
[1585724817] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.77, 10.98, 12.52
[1585725762] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16542752 kB) free!
[1585725839] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4557020 kB) free!
[1585725959] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4434256 kB) free!
[1585726079] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4306128 kB) free!
[1585726962] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2619388 kB) free!
[1585727648] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1585727879] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1950076 kB) free!
[1585727986] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4924772 kB) free!
[1585728106] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (919028 kB) free!
[1585728163] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24766284 kB) free!
[1585728226] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (266304 kB) free!
[1585728763] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1844888 kB) free!
[1585728826] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 30.2% (14877780 kB) free.
[1585728880] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (496948 kB) free!
[1585728911] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1585729000] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.6% (8165928 kB) free!
[1585729031] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1585729120] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.0% (7901644 kB) free!
[1585729218] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.9% (22240432 kB) free!
[1585729338] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (27678640 kB) free.
[1585729426] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234344 kB) free!
[1585729547] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (2358216 kB) free!
[1585729652] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3421516 kB) free!
[1585729667] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (5935652 kB) free!
[1585729772] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.7% (2848068 kB) free!
[1585729893] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (2264992 kB) free!
[1585729938] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.4% (17606252 kB) free!
[1585729963] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25620740 kB) free!
[1585730058] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.7% (23340964 kB) free!
[1585730178] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.6% (16611628 kB) free!
[1585730267] SERVICE FLAPPING ALERT: qservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1585730267] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 79.5% (39171824 kB) free.
[1585730493] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1585730493] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (174316 kB) free!
[1585730563] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398196 kB) free!
[1585730817] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 46.98, 36.81, 26.64
[1585730920] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4384544 kB) free!
[1585731030] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3589152 kB) free!
[1585731150] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 11.5% (2843892 kB) free!
[1585731270] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2129448 kB) free!
[1585731378] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12384244 kB) free!
[1585731541] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(166) < 92.42.104.166 [01/Apr/2020:14:20:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585731661] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585731978] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.2% (18656916 kB) free!
[1585732874] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 13.8% (4827352 kB) free!
[1585732963] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15171852 kB) free!
[1585732994] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 12.1% (4243492 kB) free!
[1585733114] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 10.6% (3729284 kB) free!
[1585733461] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(119) < 92.42.104.166 [01/Apr/2020:15:00:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585733493] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.4% (15579784 kB) free.
[1585733563] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (459324 kB) free!
[1585733581] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585733714] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.2% (31575424 kB) free.
[1585734163] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15085544 kB) free!
[1585734378] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1585734378] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (10311372 kB) free!
[1585734417] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.45, 13.68, 17.55
[1585734763] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (426064 kB) free!
[1585734870] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.9% (22933996 kB) free.
[1585734978] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.4% (16362524 kB) free!
[1585735336] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585735456] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1585735576] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1585735963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31030544 kB) free.
[1585736178] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12955968 kB) free!
[1585736563] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6499272 kB) free!
[1585736663] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1857508 kB) free!
[1585736684] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2096068 kB) free!
[1585736776] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1585736783] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.1% (2284824 kB) free!
[1585736804] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (456264 kB) free!
[1585736817] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.44, 24.45, 21.33
[1585736903] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2847868 kB) free!
[1585737176] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 51.8% (12702792 kB) free.
[1585737378] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.6% (15353024 kB) free!
[1585737404] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15977444 kB) free!
[1585737417] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.12, 11.80, 16.16
[1585738004] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (470988 kB) free!
[1585738103] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3395784 kB) free.
[1585738703] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2634420 kB) free!
[1585738823] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3325532 kB) free.
[1585739204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30093284 kB) free.
[1585739423] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (2783084 kB) free!
[1585739543] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3513080 kB) free.
[1585739804] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4847196 kB) free!
[1585739924] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (563264 kB) free!
[1585740044] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (460680 kB) free!
[1585740143] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3094852 kB) free!
[1585740263] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3976728 kB) free.
[1585740644] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13585764 kB) free!
[1585741244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33316712 kB) free.
[1585741667] SERVICE FLAPPING ALERT: qservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585741844] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9349148 kB) free!
[1585741964] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5114596 kB) free!
[1585742084] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (693360 kB) free!
[1585742493] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1585742663] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3131900 kB) free!
[1585742684] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17020348 kB) free!
[1585742783] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3073144 kB) free!
[1585742879] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 29.2% (7163392 kB) free.
[1585742904] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (859360 kB) free!
[1585743093] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (3517072 kB) free!
[1585743176] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4883868 kB) free!
[1585743213] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3332988 kB) free!
[1585743284] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383860 kB) free!
[1585743297] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.5% (4775656 kB) free!
[1585743333] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3140300 kB) free!
[1585743417] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.0% (4661184 kB) free!
[1585743504] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1585743504] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.4% (4293000 kB) free.
[1585744104] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (422824 kB) free!
[1585744224] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.9% (1934748 kB) free!
[1585744344] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.6% (4969328 kB) free.
[1585744484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29187580 kB) free.
[1585744533] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.2% (15531732 kB) free.
[1585744944] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (1494480 kB) free!
[1585745064] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.8% (4847252 kB) free.
[1585745084] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3705608 kB) free!
[1585745178] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5892548 kB) free!
[1585745204] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (377988 kB) free!
[1585745217] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1243216 kB) free!
[1585745324] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400708 kB) free!
[1585745664] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (1631488 kB) free!
[1585745778] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (24875568 kB) free!
[1585745784] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2889896 kB) free!
[1585745879] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4804780 kB) free!
[1585745904] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4090200 kB) free.
[1585745924] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15036568 kB) free!
[1585745999] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4738012 kB) free!
[1585746119] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4667124 kB) free!
[1585746378] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9091208 kB) free!
[1585746524] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401268 kB) free!
[1585747105] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1101140 kB) free!
[1585747225] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.1% (2295744 kB) free!
[1585747296] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 22.4% (5502396 kB) free.
[1585747319] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175592 kB) free!
[1585747345] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3516544 kB) free.
[1585747578] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (14179772 kB) free!
[1585747724] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15917324 kB) free!
[1585748324] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398168 kB) free!
[1585748545] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (705188 kB) free!
[1585748665] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.8% (1909904 kB) free!
[1585748785] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3207984 kB) free!
[1585748817] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585748924] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18682464 kB) free!
[1585749119] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1585749119] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (2576032 kB) free!
[1585749385] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3533064 kB) free.
[1585749524] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392988 kB) free!
[1585749848] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1585749985] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (901704 kB) free!
[1585750105] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2457400 kB) free!
[1585750124] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15289352 kB) free!
[1585750225] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3593988 kB) free.
[1585750667] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9825056 kB) free!
[1585750724] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (413544 kB) free!
[1585750779] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 13 users currently logged in
[1585750787] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (9738836 kB) free!
[1585750899] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 13 users currently logged in
[1585750907] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (9642136 kB) free!
[1585751019] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 13 users currently logged in
[1585751324] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14426304 kB) free!
[1585751425] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1137204 kB) free!
[1585751545] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (2740048 kB) free!
[1585751619] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1585751665] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3548500 kB) free.
[1585751778] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (27692776 kB) free.
[1585751925] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398080 kB) free!
[1585752378] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26331336 kB) free!
[1585752498] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.1% (18575160 kB) free!
[1585752618] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (27021856 kB) free.
[1585753125] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34009004 kB) free.
[1585753218] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.5% (21697608 kB) free!
[1585753338] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.6% (15234660 kB) free!
[1585753458] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.3% (22861776 kB) free!
[1585753465] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (1313928 kB) free!
[1585753585] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.6% (2693212 kB) free!
[1585753705] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.6% (3835716 kB) free.
[1585753725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7603232 kB) free!
[1585753845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3438408 kB) free!
[1585753965] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381596 kB) free!
[1585754021] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 13 users currently logged in
[1585754141] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 13 users currently logged in
[1585754261] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 13 users currently logged in
[1585754519] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2388876 kB) free!
[1585754565] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19292640 kB) free!
[1585754905] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (936932 kB) free!
[1585755025] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2094032 kB) free!
[1585755145] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3322984 kB) free.
[1585755165] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388332 kB) free!
[1585755765] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14853468 kB) free!
[1585755858] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (26761716 kB) free.
[1585756365] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34755604 kB) free.
[1585756661] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1585756661] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1585756965] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11134964 kB) free!
[1585757085] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6754136 kB) free!
[1585757205] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2609208 kB) free!
[1585757805] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19263692 kB) free!
[1585758405] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382764 kB) free!
[1585759006] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13330596 kB) free!
[1585759346] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.2% (1815092 kB) free!
[1585759466] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.2% (2476792 kB) free!
[1585759586] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (3092784 kB) free!
[1585759606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31632604 kB) free.
[1585759897] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.9% (4871988 kB) free!
[1585760017] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.8% (4850228 kB) free!
[1585760137] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.7% (4827748 kB) free!
[1585760186] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (682560 kB) free!
[1585760206] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2641876 kB) free!
[1585760326] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406344 kB) free!
[1585760446] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403612 kB) free!
[1585760786] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3302708 kB) free.
[1585761119] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1585761386] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3188932 kB) free!
[1585761506] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (571856 kB) free!
[1585761626] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (798800 kB) free!
[1585762226] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 35.2% (5715916 kB) free.
[1585762846] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32845732 kB) free.
[1585763446] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8415504 kB) free!
[1585763566] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4074888 kB) free!
[1585763686] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396920 kB) free!
[1585764286] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16444376 kB) free!
[1585764417] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.31, 4.06, 4.54
[1585764886] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34460820 kB) free.
[1585765448] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1585765458] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1585765487] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9510540 kB) free!
[1585765606] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5375216 kB) free!
[1585765726] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1398224 kB) free!