[1592678279] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1592678991] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1592679111] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592679231] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1592679831] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592679951] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1592680071] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1592680671] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592680791] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1592680911] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592681714] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (186152 kB) free!
[1592682614] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 79.3% (19428596 kB) free.
[1592682680] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 15.75, 19.62, 25.38
[1592682711] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1592683257] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232200 kB) free!
[1592683377] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236924 kB) free!
[1592683498] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242072 kB) free!
[1592683881] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 29.39, 32.27, 30.44
[1592683911] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592684001] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 48.13, 35.92, 31.82
[1592684031] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1592684121] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1592684121] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 28.79, 31.67, 30.70
[1592684721] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 22.51, 22.26, 25.96
[1592685298] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 52.5% (25872688 kB) free.
[1592685321] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 43.39, 30.70, 29.35
[1592685441] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 34.06, 30.78, 29.52
[1592686041] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 27.95, 32.64, 31.65
[1592686161] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 12.65, 25.29, 29.08
[1592687832] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 84.1% (41445412 kB) free.
[1592688876] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 80.8% (39811784 kB) free.
[1592689120] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(425) < 172.105.62.98 [21/Jun/2020:03:01:22 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592689240] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592689431] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592689551] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1592690151] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592690271] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592690347] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 338 processes
[1592690391] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1592690467] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1592692161] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 34.58, 34.05, 30.06
[1592692281] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 43.87, 37.47, 31.84
[1592692401] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 21.15, 31.66, 30.45
[1592692702] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (2577464 kB) free!
[1592693001] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 28.08, 26.07, 28.08
[1592693302] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (244488 kB) free!
[1592693391] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592693511] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1592694202] SERVICE ALERT: mrq4;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 55.83, 34.99, 30.12
[1592694322] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 26.33, 32.50, 29.96
[1592694922] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 42.14, 34.23, 30.79
[1592695042] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 27.12, 32.56, 30.69
[1592695102] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.5% (10166012 kB) free.
[1592695162] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 22.16, 29.27, 29.74
[1592695312] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592695433] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592695553] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1592696302] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (208724 kB) free!
[1592696362] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 50.34, 36.01, 31.13
[1592696422] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195536 kB) free!
[1592696482] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 28.58, 34.07, 31.10
[1592696542] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1592696542] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (315124 kB) free!
[1592696602] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 20.55, 30.06, 30.04
[1592696872] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 47.7% (11684988 kB) free.
[1592697397] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 54.9% (13543632 kB) free.
[1592698402] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 17.47, 25.32, 30.00
[1592699002] SERVICE ALERT: mrq4;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 69.72, 37.98, 31.68
[1592699122] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 27.68, 33.60, 30.91
[1592699242] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 28.70, 32.31, 30.74
[1592699842] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 59.78, 42.44, 35.85
[1592700014] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4418508 kB) free!
[1592700075] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4416088 kB) free!
[1592700135] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 18.0% (4415280 kB) free!
[1592700195] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 18.0% (4416028 kB) free!
[1592700277] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 18.0% (4399956 kB) free!
[1592700397] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 18.0% (4400668 kB) free!
[1592700442] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 34.46, 34.75, 34.67
[1592700517] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 18.0% (4399972 kB) free!
[1592700572] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 76.3% (18709212 kB) free.
[1592700667] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 335 processes
[1592700787] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 333 processes
[1592700907] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 334 processes
[1592701214] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.7% (24669472 kB) free!
[1592701507] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1592701642] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 69.09, 69.25, 54.64
[1592701814] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.3% (32014680 kB) free.
[1592705197] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4389076 kB) free!
[1592705317] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179368 kB) free!
[1592705387] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1592705437] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174320 kB) free!
[1592705507] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1592707342] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1592712307] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1592712427] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1592712442] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1592715849] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.02, 5.25, 5.04
[1592715969] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.96, 4.91, 4.95
[1592717372] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4578944 kB) free!
[1592717492] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 29.1% (7133932 kB) free.
[1592718472] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4757332 kB) free!
[1592718592] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4708856 kB) free!
[1592718712] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4645300 kB) free!
[1592719312] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187240 kB) free!
[1592719553] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1592719673] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1592719793] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1592719892] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4444208 kB) free!
[1592719907] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1592720012] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (4300680 kB) free!
[1592720027] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1592720132] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (4119660 kB) free!
[1592720279] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.9% (3572712 kB) free!
[1592720399] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.4% (3227660 kB) free!
[1592720519] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3038712 kB) free!
[1592720993] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1592721120] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.2% (15530608 kB) free.
[1592721545] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 65.6% (16067760 kB) free.
[1592721932] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1845796 kB) free!
[1592723393] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1592723514] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1592723634] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1592724969] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.97, 5.60, 5.08
[1592725089] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.59, 4.91, 4.89
[1592725434] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1592725434] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1592727834] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1592727954] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1592729676] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9853096 kB) free!
[1592729796] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9741872 kB) free!
[1592729917] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.5% (9624536 kB) free!
[1592731840] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(273) < 172.105.62.98 [21/Jun/2020:15:00:13 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1592731960] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(34) < 185.255.8.58 [21/Jun/2020:15:01:30 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1592732080] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592735317] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4581128 kB) free!
[1592735754] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1592738945] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.4% (4262280 kB) free!
[1592739065] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.8% (4126540 kB) free!
[1592739185] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.2% (3974188 kB) free!
[1592739354] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592739474] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1592740075] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1592740195] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592740315] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1592740985] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2389336 kB) free!
[1592741515] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592741635] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1592742835] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592742955] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1592743027] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1592743147] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1592743555] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1592743675] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1592743689] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.29, 5.57, 5.22
[1592743717] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 10.0% (2445488 kB) free!
[1592743809] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.21, 5.37, 5.18
[1592743929] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.57, 5.40, 5.21
[1592744529] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.70, 4.37, 4.78
[1592745475] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592745595] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1592745729] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.65, 5.04, 5.05
[1592745783] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 18.9% (6191616 kB) free!
[1592745849] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.57, 5.84, 5.35
[1592745903] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 18.6% (6078736 kB) free!
[1592745970] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.83, 5.31, 5.23
[1592746023] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.3% (5976280 kB) free!
[1592746195] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1592746315] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1592746570] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1592746570] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.92, 4.79, 4.91
[1592746623] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 20.6% (6752164 kB) free.
[1592746915] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592747035] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1592747223] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.1% (6250064 kB) free!
[1592747343] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 18.8% (6155176 kB) free!
[1592747464] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.5% (6052704 kB) free!
[1592747635] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592747755] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1592747875] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1592748475] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1592748970] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.49, 5.75, 5.33
[1592749032] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9763876 kB) free!
[1592749075] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592749090] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.69, 5.59, 5.32
[1592749152] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9641616 kB) free!
[1592749196] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1592749210] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.77, 4.90, 5.10
[1592749272] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (9516688 kB) free!
[1592749795] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1592749915] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1592751121] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.1% (3162148 kB) free!
[1592751241] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.6% (2820444 kB) free!
[1592751361] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.0% (2464316 kB) free!
[1592751664] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1592751664] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (2413364 kB) free!
[1592751961] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (731268 kB) free!
[1592752214] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5067720 kB) free!
[1592752334] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1992104 kB) free!
[1592752454] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (1996520 kB) free!
[1592753347] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1592753464] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 36.5% (11938004 kB) free.
[1592753467] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 327 processes
[1592753472] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4568412 kB) free!
[1592753515] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1592753588] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1592753635] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1592753761] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.4% (15583256 kB) free.
[1592754835] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1592754955] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1592755210] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 5.51, 4.75, 4.81
[1592757955] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1592758042] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 51.79, 44.01, 44.72
[1592758075] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592758196] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592758642] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 59.81, 38.91, 40.67
[1592758796] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1592758796] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1592759242] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 30.76, 34.53, 36.96
[1592760442] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1592760442] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 61.32, 48.06, 41.86
[1592761210] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1592761642] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 39.83, 43.22, 43.92
[1592763664] SERVICE FLAPPING ALERT: mrq4;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)