[1555525827] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5944636 kB) free!
[1555525947] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.6% (13116016 kB) free.
[1555526547] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8816832 kB) free!
[1555526667] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.8% (5831932 kB) free!
[1555526787] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13138524 kB) free.
[1555527388] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.8% (7766976 kB) free!
[1555527508] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4906072 kB) free!
[1555527628] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.4% (12001664 kB) free.
[1555528228] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7181940 kB) free!
[1555528348] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.8% (13202764 kB) free.
[1555528948] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8461396 kB) free!
[1555529068] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.9% (5347416 kB) free!
[1555529188] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.8% (12240784 kB) free.
[1555529788] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.4% (7104888 kB) free!
[1555529908] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.9% (13269984 kB) free.
[1555530508] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.8% (7788128 kB) free!
[1555530628] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.0% (13313832 kB) free.
[1555531228] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8725672 kB) free!
[1555531349] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.0% (5937952 kB) free!
[1555531469] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.0% (13316980 kB) free.
[1555532069] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.5% (7121804 kB) free!
[1555532189] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.2% (13384960 kB) free.
[1555532789] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8690300 kB) free!
[1555532909] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5970236 kB) free!
[1555533029] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.8% (13712560 kB) free.
[1555533629] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8903704 kB) free!
[1555533749] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.4% (14006360 kB) free.
[1555534349] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9201352 kB) free!
[1555534469] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.5% (14060900 kB) free.
[1555535069] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9262000 kB) free!
[1555535189] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.6% (14116144 kB) free.
[1555535790] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9796952 kB) free!
[1555535910] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 15.7% (7715488 kB) free!
[1555536030] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.7% (13631372 kB) free.
[1555536630] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9702232 kB) free!
[1555536750] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555536750] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.9% (14251160 kB) free.
[1555537536] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 285 processes
[1555537656] SERVICE ALERT: sig2-vm1;CPU_Procs;OK;HARD;1;CPU OK: 280 processes
[1555539150] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9241808 kB) free!
[1555539270] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.0% (14273172 kB) free.
[1555539837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1555539870] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8857156 kB) free!
[1555539990] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.1% (13866468 kB) free.
[1555540437] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555540557] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555540677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555541190] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9601968 kB) free!
[1555541311] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.3% (14434800 kB) free.
[1555542510] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9585096 kB) free!
[1555542630] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.1% (14333012 kB) free.
[1555545477] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1555546077] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555546198] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555546318] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1555546318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1555547220] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24628 MB (10% inode=99%):
[1555548718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1555549319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555549439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555549559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555551630] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9831464 kB) free!
[1555551750] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.0% (14308764 kB) free.
[1555553220] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47161 MB (20% inode=99%):
[1555553235] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.1% (13673352 kB) free.
[1555554150] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9021400 kB) free!
[1555554270] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.2% (13395204 kB) free.
[1555554927] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.2% (8988444 kB) free!
[1555554937] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11121712 kB) free.
[1555554993] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.2% (20281416 kB) free.
[1555555081] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.9% (22036772 kB) free.
[1555555144] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.6% (22444448 kB) free.
[1555555192] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.1% (21839168 kB) free.
[1555555542] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2783768 kB) free!
[1555555662] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (1689836 kB) free!
[1555555782] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (526860 kB) free!
[1555556070] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8215700 kB) free!
[1555556190] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.7% (12659676 kB) free.
[1555557390] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8208648 kB) free!
[1555557510] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.1% (12859016 kB) free.
[1555560959] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555561191] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1555562356] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1555562356] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10905480 kB) free.
[1555562961] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3142016 kB) free!
[1555563081] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (1894572 kB) free!
[1555563201] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (528224 kB) free!
[1555564081] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555564201] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555564321] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555564885] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1555565005] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1555565125] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1555566040] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10542372 kB) free.
[1555566646] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4654684 kB) free!
[1555566766] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2956772 kB) free!
[1555566886] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (956360 kB) free!
[1555569494] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10686884 kB) free.
[1555570101] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5400040 kB) free!
[1555570221] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (3635576 kB) free!
[1555570342] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (2612160 kB) free!
[1555570711] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8328524 kB) free!
[1555570831] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.9% (14262316 kB) free.
[1555571431] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7390252 kB) free!
[1555571551] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.9% (12255116 kB) free.
[1555572435] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (2983656 kB) free!
[1555572555] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.7% (2883748 kB) free!
[1555572675] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.0% (2760576 kB) free!
[1555573321] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555573351] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9740612 kB) free!
[1555573471] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 32.5% (16013464 kB) free.
[1555573499] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555573624] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10498432 kB) free.
[1555573689] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.8% (15080064 kB) free.
[1555574099] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555574219] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555574237] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4091624 kB) free!
[1555574339] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555574357] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (4122952 kB) free!
[1555574477] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (2297396 kB) free!
[1555576096] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (810832 kB) free!
[1555576216] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (762764 kB) free!
[1555576336] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (707884 kB) free!
[1555576603] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4942500 kB) free!
[1555576723] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 20.4% (5025312 kB) free.
[1555577323] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4738780 kB) free!
[1555577339] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555577349] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 557 processes
[1555577410] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 567 processes
[1555577443] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4670832 kB) free!
[1555577521] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555577563] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4611664 kB) free!
[1555577641] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555577761] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555577939] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555578059] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555578179] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1555578179] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555578271] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.2% (6999760 kB) free!
[1555578359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1555578392] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.5% (14039108 kB) free.
[1555578959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 322 processes
[1555578991] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9786824 kB) free!
[1555579079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555579111] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.1% (6930644 kB) free!
[1555579199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555579231] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.5% (13559816 kB) free.
[1555579802] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4781488 kB) free!
[1555579809] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 574 processes
[1555579831] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9707152 kB) free!
[1555579869] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 574 processes
[1555579923] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4715744 kB) free!
[1555579929] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 570 processes
[1555579951] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6818640 kB) free!
[1555579963] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163780 kB) free!
[1555579989] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 0 warn out of 571 processes
[1555580043] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4682416 kB) free!
[1555580071] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.6% (13090284 kB) free.
[1555580290] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 563 processes
[1555580294] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (164748 kB) free!
[1555580414] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (172552 kB) free!
[1555580534] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155784 kB) free!
[1555580643] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (289144 kB) free!
[1555580671] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9524168 kB) free!
[1555580702] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10314720 kB) free.
[1555580791] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.4% (6608940 kB) free!
[1555580814] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 60.0% (29563380 kB) free.
[1555580878] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.6% (20888880 kB) free.
[1555581134] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1555581134] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 63.4% (10303796 kB) free.
[1555581316] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2875764 kB) free!
[1555581420] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57160 MB (25% inode=99%):
[1555581436] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (1463332 kB) free!
[1555581556] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229868 kB) free!
[1555581919] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1555581922] HOST ALERT: pservercluster1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.37)
[1555581983] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'pservercluster1' was out of bounds)
[1555582001] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'pservercluster1' was out of bounds)
[1555582040] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'pservercluster1' was out of bounds)
[1555582049] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'pservercluster1' was out of bounds)
[1555582085] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1555582085] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'pservercluster1' was out of bounds)
[1555582133] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'pservercluster1' was out of bounds)
[1555582299] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'pservercluster1' was out of bounds)
[1555582371] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'pservercluster1' was out of bounds)
[1555582464] HOST ALERT: pservercluster1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.37)
[1555583161] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555583584] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2809644 kB) free!
[1555583704] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2809564 kB) free!
[1555583824] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2809228 kB) free!
[1555584421] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.8% (5336556 kB) free!
[1555584541] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.0% (13787320 kB) free.
[1555584605] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 96.7% (23873836 kB) free.
[1555584609] HOST ALERT: pservercluster1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1555584631] SERVICE ALERT: webservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1555584634] HOST ALERT: webservercluster2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.52)
[1555584653] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.48, 0.31, 0.13
[1555584680] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'webservercluster2' was out of bounds)
[1555584723] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10104604 kB) free.
[1555584754] SERVICE ALERT: webservercluster2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'webservercluster2' was out of bounds)
[1555584819] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 306 processes
[1555584845] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 299 processes
[1555584876] SERVICE ALERT: webservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'webservercluster2' was out of bounds)
[1555584891] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 153065 MB (69% inode=99%):
[1555584928] SERVICE ALERT: webservercluster2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'webservercluster2' was out of bounds)
[1555584961] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555585012] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'webservercluster2' was out of bounds)
[1555585027] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;(Return code of 255 for service 'Memory' on host 'webservercluster2' was out of bounds)
[1555585039] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1555585071] SERVICE ALERT: webservercluster2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'webservercluster2' was out of bounds)
[1555585081] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555585092] SERVICE ALERT: webservercluster2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'webservercluster2' was out of bounds)
[1555585103] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1555585121] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1555585125] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'webservercluster2' was out of bounds)
[1555585169] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 40324 MB (78% inode=99%):
[1555585197] HOST ALERT: webservercluster2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.52)
[1555585201] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1555585201] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555585333] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2889736 kB) free!
[1555585453] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (938616 kB) free!
[1555585561] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10343448 kB) free.
[1555585626] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.0% (22783864 kB) free.
[1555585741] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4827116 kB) free!
[1555585861] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (1940352 kB) free!
[1555585982] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.4% (12020768 kB) free.
[1555586109] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 81.4% (19955308 kB) free.
[1555586161] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1589560 kB) free!
[1555586282] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (339424 kB) free!
[1555586355] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3543632 kB) free!
[1555586402] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229448 kB) free!
[1555586475] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (575204 kB) free!
[1555586581] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.3% (8023716 kB) free!
[1555586595] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186872 kB) free!
[1555586701] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4911284 kB) free!
[1555586715] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172892 kB) free!
[1555586821] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (1919636 kB) free!
[1555586835] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174964 kB) free!
[1555586901] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4873020 kB) free!
[1555586955] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1555586955] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175080 kB) free!
[1555587021] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.8% (4855712 kB) free!
[1555587141] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.8% (4842984 kB) free!
[1555587179] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555587421] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11177704 kB) free.
[1555587734] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.7% (2880044 kB) free!
[1555587854] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (2796976 kB) free!
[1555587974] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.6% (2696332 kB) free!
[1555588021] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8093604 kB) free!
[1555588141] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (4886416 kB) free!
[1555588261] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (1838232 kB) free!
[1555588633] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4688568 kB) free!
[1555588753] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4560564 kB) free!
[1555588861] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1555588861] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.2% (10953476 kB) free.
[1555588873] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4496036 kB) free!
[1555589140] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47011 MB (20% inode=99%):
[1555589260] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46883 MB (20% inode=99%):
[1555589380] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46886 MB (20% inode=99%):
[1555589461] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9094372 kB) free!
[1555589482] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3064568 kB) free!
[1555589581] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.9% (5859076 kB) free!
[1555589595] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 67.0% (16424420 kB) free.
[1555589602] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3011472 kB) free!
[1555589645] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 298 processes
[1555589657] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.9% (22514272 kB) free.
[1555589701] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2745716 kB) free!
[1555589723] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.2% (2961600 kB) free!
[1555589774] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1342656 kB) free!
[1555589999] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1555590199] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.5% (21929992 kB) free.
[1555590270] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.8% (22751372 kB) free.
[1555590333] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.0% (15120188 kB) free.
[1555590404] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 85.4% (13881544 kB) free.
[1555590599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555590719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555590839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555590901] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.1% (10879488 kB) free.
[1555591399] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1555591439] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1555591502] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8101204 kB) free!
[1555591622] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4702296 kB) free!
[1555591742] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1556544 kB) free!
[1555592040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1555592063] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188012 kB) free!
[1555592160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555592183] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (143700 kB) free!
[1555592280] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1555592280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1555592303] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187828 kB) free!
[1555592342] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11919692 kB) free.
[1555592820] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47081 MB (20% inode=99%):
[1555592903] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.5% (10422660 kB) free.
[1555592940] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47009 MB (20% inode=99%):
[1555593060] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46977 MB (20% inode=99%):
[1555593179] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1555593245] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 297 processes
[1555593542] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7903236 kB) free!
[1555593605] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1555593662] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4602044 kB) free!
[1555593782] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (1498232 kB) free!
[1555593876] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194440 kB) free!
[1555593996] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 40.1% (9819316 kB) free.
[1555594103] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176540 kB) free!
[1555594181] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 58595 MB (26% inode=99%):
[1555594223] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170968 kB) free!
[1555594343] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181764 kB) free!
[1555594382] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.7% (12149636 kB) free.
[1555594668] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.8% (22506632 kB) free.
[1555595045] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1555595045] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 292 processes
[1555595391] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1555595402] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555595582] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8058380 kB) free!
[1555595702] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4755916 kB) free!
[1555595822] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (1603796 kB) free!
[1555596423] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.6% (12599652 kB) free.
[1555597623] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8775600 kB) free!
[1555597743] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.7% (5255888 kB) free!
[1555597863] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2020816 kB) free!
[1555597995] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4862236 kB) free!
[1555598116] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4821248 kB) free!
[1555598236] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (4785864 kB) free!
[1555599063] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13144188 kB) free.
[1555600263] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7874156 kB) free!
[1555600383] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (4375712 kB) free!
[1555600503] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1174428 kB) free!
[1555601103] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.5% (13077892 kB) free.
[1555601140] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555601392] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1555601402] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555601522] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555601642] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555602304] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9264688 kB) free!
[1555602396] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4886452 kB) free!
[1555602424] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5949780 kB) free!
[1555602517] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4864892 kB) free!
[1555602544] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2826696 kB) free!
[1555602637] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4844844 kB) free!
[1555603681] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555603744] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.5% (12092568 kB) free.
[1555604344] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9390420 kB) free!
[1555604464] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.4% (6099996 kB) free!
[1555604584] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (3014416 kB) free!
[1555604880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1555605245] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555605472] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555605480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1555605600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1555605720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555605779] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555605784] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11322880 kB) free.
[1555605899] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555606019] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555606384] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7190040 kB) free!
[1555606504] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4182636 kB) free!
[1555606624] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1271208 kB) free!
[1555607225] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.9% (9303740 kB) free!
[1555607643] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1555607819] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555608419] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555608425] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4332628 kB) free!
[1555608539] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555608659] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555609259] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1555609259] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555609625] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.9% (10782864 kB) free.
[1555610225] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7922892 kB) free!
[1555610345] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.2% (5043420 kB) free!
[1555610459] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555610465] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (2297080 kB) free!
[1555610579] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555610699] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555610872] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196764 kB) free!
[1555610992] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178524 kB) free!
[1555611065] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.1% (11362092 kB) free.
[1555611112] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184296 kB) free!
[1555611665] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8909492 kB) free!
[1555611785] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5971540 kB) free!
[1555611905] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3050216 kB) free!