[1589308212] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1589308272] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (2123824 kB) free!
[1589308301] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418436 kB) free!
[1589308901] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17212660 kB) free!
[1589309077] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194156 kB) free!
[1589309413] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1589309501] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38661188 kB) free.
[1589310013] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1589310102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15108624 kB) free!
[1589310133] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1589310222] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11048404 kB) free!
[1589310253] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589310342] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (7053160 kB) free!
[1589310942] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25117116 kB) free!
[1589311396] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 75.4% (18470228 kB) free.
[1589311542] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (4985216 kB) free!
[1589311872] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1357388 kB) free!
[1589312142] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21345736 kB) free!
[1589312472] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (1894708 kB) free!
[1589312742] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (1985528 kB) free!
[1589312826] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1589313253] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589313342] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17279080 kB) free!
[1589313373] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1589313456] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7634392 kB) free!
[1589313493] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1589313576] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (474416 kB) free!
[1589313696] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (447264 kB) free!
[1589313746] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (226564 kB) free!
[1589313866] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217800 kB) free!
[1589313942] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34415052 kB) free.
[1589313986] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242968 kB) free!
[1589314093] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589314273] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1356916 kB) free!
[1589314542] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13059080 kB) free!
[1589314662] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9264664 kB) free!
[1589314693] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589314782] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5706148 kB) free!
[1589314813] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1589314873] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (1881544 kB) free!
[1589314933] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1589315382] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20038844 kB) free!
[1589315534] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1589315654] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1589315774] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1589315786] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.0% (8369240 kB) free!
[1589315982] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36130452 kB) free.
[1589316582] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14769476 kB) free!
[1589316619] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (609412 kB) free!
[1589316703] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11074584 kB) free!
[1589316739] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (283772 kB) free!
[1589316823] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7426156 kB) free!
[1589316859] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (285368 kB) free!
[1589317423] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20535320 kB) free!
[1589318023] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35694876 kB) free.
[1589318059] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.7% (43211296 kB) free.
[1589318494] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.3% (43994068 kB) free.
[1589318623] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14422180 kB) free!
[1589318743] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10587260 kB) free!
[1589318863] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (6947084 kB) free!
[1589319101] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(457) < 2020/05/13 03:01:32 [error] 22123#0: *88010126 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589319221] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589319375] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4280860 kB) free!
[1589319386] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.0% (44348524 kB) free.
[1589319404] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3406) < 203.212.70.170 [13/May/2020:03:02:32 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1589319426] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1589319458] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(350) < 182.18.144.246 [13/May/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589319463] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19087248 kB) free!
[1589319525] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589319579] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589319975] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1230192 kB) free!
[1589320063] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32602456 kB) free.
[1589320663] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12711428 kB) free!
[1589320783] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9157580 kB) free!
[1589320903] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37434228 kB) free.
[1589321473] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1023728 kB) free!
[1589321503] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17234264 kB) free!
[1589321623] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13600312 kB) free!
[1589321743] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (9864616 kB) free!
[1589321774] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589322074] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3698692 kB) free.
[1589322137] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.6% (10189944 kB) free.
[1589322343] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21159796 kB) free!
[1589322944] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33244508 kB) free.
[1589323337] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (233368 kB) free!
[1589323457] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (2166728 kB) free!
[1589323544] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13917500 kB) free!
[1589323577] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1818852 kB) free!
[1589323664] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10404444 kB) free!
[1589323784] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37282272 kB) free.
[1589324384] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17528548 kB) free!
[1589324504] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (14021376 kB) free!
[1589324624] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (10609796 kB) free!
[1589325224] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20684468 kB) free!
[1589325824] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31131080 kB) free.
[1589325977] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.8% (9756116 kB) free.
[1589326424] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12239060 kB) free!
[1589326544] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8924600 kB) free!
[1589326574] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589326578] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167060 kB) free!
[1589326664] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34664216 kB) free.
[1589326694] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589326698] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202028 kB) free!
[1589326819] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1589326819] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183372 kB) free!
[1589327264] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15049228 kB) free!
[1589327384] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (11509804 kB) free!
[1589327504] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37053984 kB) free.
[1589327895] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1589328015] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 707 processes
[1589328104] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17519780 kB) free!
[1589328135] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1589328225] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (14005320 kB) free!
[1589328277] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.2% (17441480 kB) free.
[1589328345] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (10573212 kB) free!
[1589328357] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.1% (11858096 kB) free.
[1589328375] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 44.6% (10942952 kB) free.
[1589328735] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589328945] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19804688 kB) free!
[1589329335] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1589329455] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1589329545] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29023672 kB) free.
[1589329575] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589330145] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8084220 kB) free!
[1589330265] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36342560 kB) free.
[1589330865] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7978896 kB) free!
[1589330985] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3554132 kB) free!
[1589331105] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395028 kB) free!
[1589331696] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 41.0% (54082172 kB) free.
[1589332305] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26844156 kB) free.
[1589332696] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186476 kB) free!
[1589332756] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 28.7% (7031128 kB) free.
[1589332905] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4716808 kB) free!
[1589333025] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35672112 kB) free.
[1589333056] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 14.1% (3466240 kB) free!
[1589333116] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 13.4% (3277008 kB) free!
[1589333176] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 13.3% (3271888 kB) free!
[1589333236] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 13.4% (3274572 kB) free!
[1589333474] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589333625] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11635240 kB) free!
[1589333745] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7824692 kB) free!
[1589333757] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192892 kB) free!
[1589333865] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4231892 kB) free!
[1589333877] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178072 kB) free!
[1589333997] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188244 kB) free!
[1589334466] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20222828 kB) free!
[1589334736] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.8% (2400332 kB) free!
[1589335066] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29288904 kB) free.
[1589336227] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1589336266] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18666580 kB) free!
[1589336347] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 342 processes
[1589336386] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14854420 kB) free!
[1589336467] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1589336506] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (11062428 kB) free!
[1589336775] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589336895] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1589337015] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1589337106] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19196836 kB) free!
[1589337615] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1589337706] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28580952 kB) free.
[1589338219] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589338306] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7874860 kB) free!
[1589338426] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34036680 kB) free.
[1589338815] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1589338935] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1589339026] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12976772 kB) free!
[1589339056] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1589339146] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8712184 kB) free!
[1589339266] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34548576 kB) free.
[1589339866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13247004 kB) free!
[1589339986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8952908 kB) free!
[1589340106] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4896380 kB) free!
[1589340255] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589340375] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1589340375] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202752 kB) free!
[1589340495] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1589340495] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (228896 kB) free!
[1589340615] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (217548 kB) free!
[1589340707] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14329964 kB) free!
[1589340848] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.34, 8.37, 5.76
[1589340968] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.21, 8.25, 6.06
[1589341088] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.38, 10.96, 7.36
[1589341095] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589341907] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389068 kB) free!
[1589342507] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18672104 kB) free!
[1589342933] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (1723928 kB) free!
[1589342986] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 18.8% (4600032 kB) free!
[1589343106] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4904348 kB) free!
[1589343107] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35257768 kB) free.
[1589343221] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/13 09:35:05 [error] 22130#0: *88070835 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589343226] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4896196 kB) free!
[1589343341] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589343707] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10931300 kB) free!
[1589343827] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6607696 kB) free!
[1589343947] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2358024 kB) free!
[1589344267] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1589344734] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1589344854] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1589344874] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (2615996 kB) free!
[1589344975] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1589344995] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.9% (4371484 kB) free.
[1589345077] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181380 kB) free!
[1589345147] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22201068 kB) free!
[1589345197] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191368 kB) free!
[1589345317] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189000 kB) free!
[1589345575] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1589345747] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (435268 kB) free!
[1589346341] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/13 10:33:17 [error] 22133#0: *88080040 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589346461] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589346533] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(480) < 103.255.146.164 [13/May/2020:10:32:17 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589346653] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589347061] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(621) < 2020/05/13 10:45:28 [error] 22126#0: *88083305 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589347178] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(505) < 172.105.57.57 [13/May/2020:10:44:18 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589347181] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589347253] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1126) < 203.212.70.170 [13/May/2020:10:43:59 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589347298] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589347373] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589347538] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 10.9% (2438996 kB) free!
[1589347658] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (1953932 kB) free!
[1589347778] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1457796 kB) free!
[1589347782] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/13 10:53:53 [error] 22135#0: *88084929 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589347815] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.1% (2959648 kB) free!
[1589347902] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589348147] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30278284 kB) free.
[1589348150] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 51.9% (12715828 kB) free.
[1589348415] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (788040 kB) free!
[1589348573] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [13/May/2020:11:11:31 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.000, 0.000"
[1589348693] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589348748] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5503848 kB) free!
[1589348868] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1187788 kB) free!
[1589348988] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386384 kB) free!
[1589349588] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15555100 kB) free!
[1589349656] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 363 processes
[1589349776] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 356 processes
[1589349794] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2953532 kB) free!
[1589349896] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 355 processes
[1589349914] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2458972 kB) free!
[1589350034] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3410560 kB) free.
[1589350188] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439080 kB) free!
[1589350663] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (5860564 kB) free!
[1589350783] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 15.5% (5436660 kB) free!
[1589350788] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13493520 kB) free!
[1589350856] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (4151556 kB) free!
[1589350896] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26295644 kB) free!
[1589350903] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (4929660 kB) free!
[1589350976] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3774248 kB) free!
[1589351016] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.9% (32782572 kB) free.
[1589351096] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3464992 kB) free!
[1589351150] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.7% (4100724 kB) free!
[1589351270] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.6% (3815676 kB) free!
[1589351333] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (635364 kB) free!
[1589351388] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401060 kB) free!
[1589351390] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.1% (3461016 kB) free!
[1589351503] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.1% (31902768 kB) free.
[1589351616] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.6% (21945456 kB) free!
[1589351694] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(674) < 103.255.146.164 [13/May/2020:11:57:49 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589351696] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1487064 kB) free!
[1589351737] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.0% (33001224 kB) free.
[1589351814] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589351979] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.7% (16090644 kB) free.
[1589351990] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1821472 kB) free!
[1589352094] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9675212 kB) free!
[1589352214] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9344244 kB) free!
[1589352334] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9021828 kB) free!
[1589352434] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2118360 kB) free!
[1589352495] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589352555] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.2% (1824244 kB) free!
[1589352588] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33303660 kB) free.
[1589352675] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1579264 kB) free!
[1589352702] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/13 12:13:51 [error] 22135#0: *88101554 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589352822] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589352896] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22891852 kB) free.
[1589353188] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7483420 kB) free!
[1589353275] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (2751272 kB) free!
[1589353308] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3087968 kB) free!
[1589353428] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380528 kB) free!
[1589353536] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.9% (19591956 kB) free!
[1589353614] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(547) < 103.255.146.164 [13/May/2020:12:31:08 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589353656] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.4% (32126356 kB) free.
[1589353734] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(683) < 203.212.70.170 [13/May/2020:12:38:50 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1589353855] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589353875] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3297148 kB) free.
[1589353898] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(405) < 49.50.86.152 [13/May/2020:12:41:38 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589354018] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(88) < 172.105.57.57 [13/May/2020:12:41:44 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589354022] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(685) < 2020/05/13 12:42:44 [error] 22135#0: *88109904 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589354026] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.6% (2355128 kB) free!
[1589354028] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20622092 kB) free!
[1589354096] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1589354134] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3635740 kB) free!
[1589354139] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589354142] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589354256] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25807596 kB) free!
[1589354376] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.7% (33881784 kB) free.
[1589354475] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2825912 kB) free!
[1589354595] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (910688 kB) free!
[1589354628] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446176 kB) free!
[1589354715] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1589354715] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (895600 kB) free!
[1589354933] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (1633816 kB) free!
[1589355054] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(679) < 103.255.146.164 [13/May/2020:12:59:45 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589355174] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589355228] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18210636 kB) free!
[1589355576] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.1% (19949268 kB) free!
[1589355696] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.9% (35401096 kB) free.
[1589355829] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432044 kB) free!
[1589355860] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8663624 kB) free!
[1589355980] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8356808 kB) free!
[1589356100] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (8094204 kB) free!
[1589356117] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.1% (4923084 kB) free.
[1589356515] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (2444792 kB) free!
[1589356542] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/13 13:17:08 [error] 22122#0: *88118354 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589356662] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589356717] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (4079220 kB) free!
[1589356837] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3929188 kB) free!
[1589356957] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (3746672 kB) free!
[1589356975] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(687) < 103.255.146.164 [13/May/2020:13:31:15 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589357095] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589357115] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (310672 kB) free!
[1589357496] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25888008 kB) free!
[1589357616] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.2% (18673868 kB) free!
[1589357715] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (2388364 kB) free!
[1589357736] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.3% (25425720 kB) free!
[1589357900] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3989408 kB) free!
[1589357980] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (1933772 kB) free!
[1589358100] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (1365332 kB) free!
[1589358158] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2241404 kB) free!
[1589358220] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (811736 kB) free!
[1589358336] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (30625004 kB) free.
[1589358915] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (961032 kB) free!
[1589359516] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (2517792 kB) free!
[1589360029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30472348 kB) free.
[1589360136] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.7% (15484276 kB) free!
[1589360256] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.9% (23621480 kB) free!
[1589360376] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.1% (17288296 kB) free!
[1589360629] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3227580 kB) free!
[1589360749] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384232 kB) free!
[1589360869] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396264 kB) free!
[1589361316] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1588108 kB) free!
[1589361469] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15111060 kB) free!
[1589361703] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 15.8% (5542256 kB) free!
[1589361823] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 15.1% (5292636 kB) free!
[1589361896] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.1% (4213120 kB) free!
[1589361916] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3504532 kB) free.
[1589361943] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 14.3% (5023520 kB) free!
[1589362016] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3981940 kB) free!
[1589362069] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (428184 kB) free!
[1589362136] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3607656 kB) free!
[1589362420] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.2% (16221128 kB) free.
[1589362495] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(312) < 103.255.146.164 [13/May/2020:15:01:51 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589362516] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (620756 kB) free!
[1589362543] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1589362543] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.2% (31955604 kB) free.
[1589362615] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589362636] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (1465588 kB) free!
[1589362662] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/13 14:58:50 [error] 22135#0: *88141252 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589362669] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22846472 kB) free!
[1589362736] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1589362736] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2212924 kB) free!
[1589362756] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (2291740 kB) free!
[1589362783] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589363215] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(154) < 103.255.146.164 [13/May/2020:15:15:10 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589363270] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402540 kB) free!
[1589363336] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589363336] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.9% (22929660 kB) free.
[1589363870] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22882684 kB) free!
[1589364470] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416388 kB) free!
[1589364576] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1589364576] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (30999164 kB) free.
[1589365070] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30664728 kB) free.
[1589365135] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 183.82.105.122 [13/May/2020:15:48:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589365255] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589365670] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2162988 kB) free!
[1589365774] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 733 processes
[1589365776] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.2% (18769964 kB) free!
[1589365790] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396984 kB) free!
[1589365791] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 21.5% (5280532 kB) free.
[1589365894] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 744 processes
[1589365896] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.9% (30234012 kB) free.
[1589365910] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (440692 kB) free!
[1589366014] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 742 processes
[1589366289] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.46, 4.39, 4.80
[1589366356] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1441716 kB) free!
[1589366391] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.7% (4574676 kB) free!
[1589366456] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [13/May/2020:16:07:47 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1589366467] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1589366510] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19384400 kB) free!
[1589366511] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.0% (4403764 kB) free!
[1589366576] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 203.212.70.170 [13/May/2020:16:12:37 +0530] TCP 502 0 0 60.003 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.003, 0.000"
[1589366614] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 364 processes
[1589366631] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.1% (4180664 kB) free!
[1589366696] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(17) < 203.212.70.170 [13/May/2020:16:14:52 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1589366889] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.70, 10.78, 7.89
[1589367009] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.24, 10.45, 8.14
[1589367096] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.0% (15783328 kB) free!
[1589367110] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383668 kB) free!
[1589367129] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.31, 9.35, 8.02
[1589367216] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28129424 kB) free.
[1589367296] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589367710] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26569700 kB) free.
[1589367816] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11864128 kB) free!
[1589367831] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2213636 kB) free!
[1589367896] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(617) < 103.255.146.164 [13/May/2020:16:31:25 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589367936] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.3% (18906872 kB) free!
[1589368017] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589368056] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11889684 kB) free!
[1589368156] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2787372 kB) free!
[1589368310] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (442152 kB) free!
[1589368430] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386760 kB) free!
[1589368550] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417008 kB) free!
[1589368656] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.0% (23686836 kB) free!
[1589368756] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3697200 kB) free.
[1589369150] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16932680 kB) free!
[1589369356] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.2% (1655972 kB) free!
[1589369467] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1589369476] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3029152 kB) free!
[1589369587] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1589369596] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3467024 kB) free.
[1589369751] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (447624 kB) free!
[1589369856] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (30937572 kB) free.
[1589370196] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.2% (1658324 kB) free!
[1589370316] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3204200 kB) free!
[1589370351] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23418688 kB) free!
[1589370436] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3435536 kB) free.
[1589370787] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1589370907] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 339 processes
[1589370951] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392556 kB) free!
[1589371027] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1589371036] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.5% (2192136 kB) free!
[1589371056] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.5% (15220028 kB) free!
[1589371156] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3590688 kB) free.
[1589371176] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26468220 kB) free.
[1589371420] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3408212 kB) free!
[1589371540] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.2% (2961912 kB) free!
[1589371551] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28391076 kB) free.
[1589371660] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.2% (2519304 kB) free!
[1589371757] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (1612908 kB) free!
[1589371776] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.7% (24623652 kB) free!
[1589371877] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3465640 kB) free.
[1589371897] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.7% (15374328 kB) free!
[1589372017] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (8684920 kB) free!
[1589372151] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (477924 kB) free!
[1589372261] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (318628 kB) free!
[1589372271] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (419040 kB) free!
[1589372333] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (243268 kB) free!
[1589372391] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420392 kB) free!
[1589372477] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (625156 kB) free!
[1589372597] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (148656 kB) free!
[1589372617] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.8% (18142872 kB) free!
[1589372717] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1241092 kB) free!
[1589372991] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26758560 kB) free.
[1589373462] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1589373462] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.3% (16230856 kB) free.
[1589373591] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (393032 kB) free!
[1589373711] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (414684 kB) free!
[1589373831] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (479772 kB) free!
[1589373917] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.6% (4324376 kB) free.
[1589373943] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589374417] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7665768 kB) free!
[1589374431] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23643212 kB) free!
[1589374517] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (2802304 kB) free!
[1589374637] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.7% (4340184 kB) free.
[1589374733] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (2602196 kB) free!
[1589374736] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589375017] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.8% (15526484 kB) free!
[1589375031] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408320 kB) free!
[1589375838] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (1997908 kB) free!
[1589375958] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.5% (4637912 kB) free.
[1589376217] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (28818940 kB) free.
[1589376558] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3082904 kB) free!
[1589376678] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.4% (4610536 kB) free.
[1589376817] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11971228 kB) free!
[1589376937] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (27057968 kB) free.
[1589377278] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (2419664 kB) free!
[1589377398] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (150252 kB) free!
[1589377519] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.3% (4440476 kB) free.
[1589377537] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9537272 kB) free!
[1589377657] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.7% (22032548 kB) free!
[1589377777] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (14192424 kB) free!
[1589378631] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23114372 kB) free!
[1589378718] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2124012 kB) free!
[1589378838] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.3% (4768324 kB) free.
[1589379231] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384696 kB) free!
[1589379339] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1589379438] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (2920520 kB) free!
[1589379459] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1589379558] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.3% (4595688 kB) free.
[1589379577] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (28168496 kB) free.
[1589379729] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.91, 3.62, 4.69
[1589379831] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29962772 kB) free.
[1589380431] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (461016 kB) free!
[1589380551] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400812 kB) free!
[1589380616] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 203.212.70.170 [13/May/2020:20:01:09 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.000, 0.000"
[1589380627] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1589380671] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414800 kB) free!
[1589380736] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589380758] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (2256720 kB) free!
[1589380878] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.8% (4676664 kB) free.
[1589381228] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1589381260] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1589381271] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29341532 kB) free.
[1589381380] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1589381478] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3162700 kB) free!
[1589381598] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.2% (4580728 kB) free.
[1589381871] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (982024 kB) free!
[1589381991] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (431596 kB) free!
[1589382111] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (422016 kB) free!
[1589382578] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.9% (23627920 kB) free!
[1589382698] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.9% (32882516 kB) free.
[1589382712] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16691984 kB) free!
[1589382730] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.65, 6.08, 5.69
[1589382798] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (1283236 kB) free!
[1589382850] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.85, 5.93, 5.68
[1589382918] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3418756 kB) free.
[1589382970] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.27, 5.72, 5.63
[1589383312] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (457336 kB) free!
[1589383519] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (834064 kB) free!
[1589383570] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.99, 2.72, 4.14
[1589383639] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2906436 kB) free!
[1589383759] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.8% (4519828 kB) free.
[1589383898] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26340188 kB) free!
[1589383912] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27992188 kB) free.
[1589384018] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.9% (19683812 kB) free!
[1589384138] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (26987176 kB) free.
[1589384262] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589384359] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2033724 kB) free!
[1589384479] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.2% (4419336 kB) free.
[1589384512] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (390780 kB) free!
[1589384632] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402288 kB) free!
[1589384752] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410832 kB) free!
[1589385079] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1171248 kB) free!
[1589385199] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3030356 kB) free!
[1589385319] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.6% (4317808 kB) free.
[1589385352] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21219324 kB) free!
[1589385952] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (479896 kB) free!
[1589386519] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589386552] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33268444 kB) free.
[1589387119] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2038688 kB) free!
[1589387152] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4291272 kB) free!
[1589387239] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3535096 kB) free.
[1589387272] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402236 kB) free!
[1589387392] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (434424 kB) free!
[1589387992] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30332836 kB) free.
[1589388439] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (1059120 kB) free!
[1589388559] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.2% (2476816 kB) free!
[1589388592] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3510276 kB) free!
[1589388680] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3700132 kB) free.
[1589388712] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399864 kB) free!
[1589388832] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432812 kB) free!
[1589389433] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23070908 kB) free!
[1589389879] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (724708 kB) free!
[1589389999] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.9% (1937888 kB) free!
[1589390033] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415156 kB) free!
[1589390119] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3171224 kB) free!
[1589390138] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1589390633] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33265960 kB) free.
[1589390719] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3618156 kB) free.
[1589391233] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5444584 kB) free!
[1589391353] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1174332 kB) free!
[1589391473] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383500 kB) free!
[1589391919] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3071788 kB) free!
[1589392028] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1589392039] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3432416 kB) free.
[1589392073] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23022068 kB) free!
[1589392673] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417084 kB) free!
[1589393239] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2756604 kB) free!
[1589393273] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29794868 kB) free.
[1589393359] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3365624 kB) free.
[1589393873] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3992232 kB) free!
[1589393960] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2852984 kB) free!
[1589393993] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382908 kB) free!
[1589394080] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3309436 kB) free.
[1589394113] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415132 kB) free!