[1589567614] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.8% (1759440 kB) free!
[1589567734] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.1% (1803060 kB) free!
[1589567776] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402992 kB) free!
[1589567854] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (2082840 kB) free!
[1589568376] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30615424 kB) free.
[1589568976] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3687432 kB) free!
[1589569096] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (380388 kB) free!
[1589569216] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409752 kB) free!
[1589569816] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20950180 kB) free!
[1589570416] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397564 kB) free!
[1589570961] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 85.3% (20917824 kB) free.
[1589571016] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23253716 kB) free!
[1589571616] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385216 kB) free!
[1589572216] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26274500 kB) free!
[1589572654] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1589572654] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (512852 kB) free!
[1589572816] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2513772 kB) free!
[1589573068] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232524 kB) free!
[1589573138] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2110344 kB) free!
[1589573188] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225416 kB) free!
[1589573254] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (1979552 kB) free!
[1589573258] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3375072 kB) free!
[1589573308] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237700 kB) free!
[1589573378] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3500652 kB) free!
[1589573417] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27245988 kB) free.
[1589574017] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (2970128 kB) free!
[1589574137] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (387432 kB) free!
[1589574257] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433044 kB) free!
[1589574857] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15539240 kB) free!
[1589575108] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.7% (8699764 kB) free!
[1589575457] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410724 kB) free!
[1589575654] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.3% (4104812 kB) free.
[1589575654] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3528948 kB) free!
[1589575820] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (410092 kB) free!
[1589576057] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17243000 kB) free!
[1589576254] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (969904 kB) free!
[1589576374] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (560752 kB) free!
[1589576494] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4034468 kB) free.
[1589576657] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403088 kB) free!
[1589577095] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2997104 kB) free!
[1589577215] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2099396 kB) free!
[1589577257] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17245956 kB) free!
[1589577335] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1693388 kB) free!
[1589577454] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.4% (44554960 kB) free.
[1589577620] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.4% (43066040 kB) free.
[1589577858] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408244 kB) free!
[1589578261] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2606) < 103.13.97.241 [16/May/2020:03:00:35 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589578304] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(415) < 172.105.57.57 [16/May/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589578381] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(204) < 49.50.86.152 [16/May/2020:03:01:42 +0530] TCP 502 0 0 60.001 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1589578425] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589578458] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15853404 kB) free!
[1589578502] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(1) < 203.212.70.170 [16/May/2020:03:03:32 +0530] TCP 502 0 0 180.000 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1589578535] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 35.3% (5741344 kB) free.
[1589578543] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(541) < 2020/05/16 03:01:42 [error] 22135#0: *88782474 no live upstreams while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589578584] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (3266424 kB) free!
[1589578663] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589578708] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.0% (44373736 kB) free.
[1589578761] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193624 kB) free!
[1589578821] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198772 kB) free!
[1589578881] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 1.0% (237720 kB) free!
[1589578942] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.9% (228296 kB) free!
[1589579058] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423180 kB) free!
[1589579102] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589579184] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179320 kB) free!
[1589580259] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32999056 kB) free.
[1589580863] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11142108 kB) free!
[1589580983] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7239864 kB) free!
[1589581103] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3541964 kB) free!
[1589581703] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20833576 kB) free!
[1589582303] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1840212 kB) free!
[1589582903] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18097364 kB) free!
[1589583503] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401804 kB) free!
[1589583578] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.2% (13418144 kB) free!
[1589584103] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14326256 kB) free!
[1589584703] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33757540 kB) free.
[1589584778] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (13243964 kB) free!
[1589584940] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.1% (9576080 kB) free.
[1589585303] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2561856 kB) free!
[1589585423] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407272 kB) free!
[1589585540] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (142432 kB) free!
[1589585543] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408588 kB) free!
[1589585660] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178204 kB) free!
[1589585780] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187676 kB) free!
[1589586143] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17494076 kB) free!
[1589586381] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 27.1% (6652572 kB) free.
[1589586578] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (13654132 kB) free!
[1589586743] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418844 kB) free!
[1589586981] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171204 kB) free!
[1589587101] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194596 kB) free!
[1589587221] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1589587221] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189032 kB) free!
[1589587342] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 28.6% (7008192 kB) free.
[1589587343] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13886232 kB) free!
[1589587360] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.8% (20296520 kB) free.
[1589587437] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.7% (20162460 kB) free.
[1589587584] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 80.8% (19794088 kB) free.
[1589587642] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3619676 kB) free!
[1589587702] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3620524 kB) free!
[1589587762] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 14.8% (3618568 kB) free!
[1589587822] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 14.8% (3616076 kB) free!
[1589588543] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (9863860 kB) free!
[1589588978] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1589588978] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12874236 kB) free!
[1589589143] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32021664 kB) free.
[1589589340] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589589351] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.4% (4743356 kB) free!
[1589589471] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.1% (4688684 kB) free!
[1589589591] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.9% (4628824 kB) free!
[1589589743] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8665152 kB) free!
[1589589863] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4480624 kB) free!
[1589589983] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439476 kB) free!
[1589590389] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1589590509] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1589590583] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18227372 kB) free!
[1589590629] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1589591183] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34391536 kB) free.
[1589591783] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12580972 kB) free!
[1589591903] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8836272 kB) free!
[1589591978] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 59.6% (78590496 kB) free.
[1589592023] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5182452 kB) free!
[1589592623] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21564856 kB) free!
[1589592837] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (2060224 kB) free!
[1589592957] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191840 kB) free!
[1589593078] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168112 kB) free!
[1589593223] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36685840 kB) free.
[1589593791] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2127184 kB) free!
[1589593823] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15531324 kB) free!
[1589593943] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11372248 kB) free!
[1589594063] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (7177424 kB) free!
[1589594664] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20531584 kB) free!
[1589595264] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35599564 kB) free.
[1589595864] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14547820 kB) free!
[1589595984] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10622644 kB) free!
[1589596104] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (6778416 kB) free!
[1589596704] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (18950868 kB) free!
[1589597172] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.97, 6.55, 5.18
[1589597292] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.39, 6.52, 5.34
[1589597304] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34693080 kB) free.
[1589597412] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.50, 5.97, 5.30
[1589597904] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12342312 kB) free!
[1589598012] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.20, 3.65, 4.48
[1589598024] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8187828 kB) free!
[1589598144] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (4008320 kB) free!
[1589598621] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589598744] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17891296 kB) free!
[1589599344] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33602104 kB) free.
[1589599944] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10546480 kB) free!
[1589600064] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5731024 kB) free!
[1589600184] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1128180 kB) free!
[1589600412] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.60, 9.27, 6.72
[1589600532] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.80, 9.22, 7.02
[1589600652] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.22, 8.32, 6.95
[1589600785] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19256476 kB) free!
[1589600829] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1589600978] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1589601385] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446968 kB) free!
[1589601852] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1589601852] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 37.71, 19.14, 11.91
[1589602452] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.23, 12.21, 11.77
[1589603184] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.4% (4029744 kB) free!
[1589603304] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3687024 kB) free!
[1589603424] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3330472 kB) free!
[1589603702] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(494) < 103.255.146.164 [16/May/2020:10:01:58 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589603785] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23658604 kB) free!
[1589603822] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589604024] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1239768 kB) free!
[1589604385] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (461324 kB) free!
[1589604985] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14356856 kB) free!
[1589605079] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.8% (2412776 kB) free!
[1589605453] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 41.17, 29.23, 20.90
[1589605585] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416172 kB) free!
[1589605664] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/16 10:35:08 [error] 22122#0: *88860217 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589605697] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 761 processes
[1589605784] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589605786] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4424264 kB) free!
[1589605817] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 762 processes
[1589605906] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.0% (3811020 kB) free!
[1589605937] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 762 processes
[1589606026] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3152040 kB) free!
[1589606185] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41825864 kB) free.
[1589606627] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167068 kB) free!
[1589606785] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14398784 kB) free!
[1589606905] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9787668 kB) free!
[1589607025] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.0% (3439132 kB) free!
[1589607025] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5397972 kB) free!
[1589607054] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.21247.367> (10.147.212.158:49524 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1589607137] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 690 processes
[1589607173] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (1640176 kB) free!
[1589607174] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589607625] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1589607625] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (644516 kB) free!
[1589607625] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24822552 kB) free!
[1589607737] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 762 processes
[1589607760] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4210204 kB) free!
[1589607773] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1630984 kB) free!
[1589607857] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 758 processes
[1589607880] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (4034784 kB) free!
[1589607977] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 757 processes
[1589608000] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (3847120 kB) free!
[1589608225] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374720 kB) free!
[1589608664] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (7011864 kB) free!
[1589608784] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.7% (6562296 kB) free!
[1589608826] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36971332 kB) free.
[1589608859] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4848684 kB) free!
[1589608904] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 17.4% (6081296 kB) free!
[1589608979] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4452560 kB) free!
[1589609053] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 15.53, 14.17, 16.96
[1589609099] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.2% (3991224 kB) free!
[1589609140] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (428100 kB) free!
[1589609200] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2234924 kB) free!
[1589609260] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (607648 kB) free!
[1589609380] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (839884 kB) free!
[1589609426] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5447716 kB) free!
[1589609505] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (3320012 kB) free!
[1589609546] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (419036 kB) free!
[1589609666] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412136 kB) free!
[1589609700] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1525420 kB) free!
[1589609722] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (205376 kB) free!
[1589609980] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3699728 kB) free.
[1589610266] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26571036 kB) free.
[1589610454] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8768064 kB) free!
[1589610574] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8415400 kB) free!
[1589610694] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.5% (8112644 kB) free!
[1589610705] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.8% (31785252 kB) free.
[1589610828] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.2% (15997124 kB) free.
[1589610866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (463944 kB) free!
[1589610986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386152 kB) free!
[1589611106] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425560 kB) free!
[1589611180] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3199676 kB) free!
[1589611300] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3098472 kB) free!
[1589611420] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3056292 kB) free!
[1589611498] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1589611618] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589611706] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19509784 kB) free!
[1589611738] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1589612100] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.1% (22982028 kB) free.
[1589612306] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379136 kB) free!
[1589612494] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3852852 kB) free!
[1589612822] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(589) < 103.255.146.164 [16/May/2020:12:30:44 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589612906] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21785804 kB) free!
[1589612938] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1589612943] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589613506] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (422352 kB) free!
[1589613620] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9331284 kB) free!
[1589613740] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (9025040 kB) free!
[1589613861] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (9181556 kB) free!
[1589613977] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1589613977] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 694 processes
[1589614106] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33051644 kB) free.
[1589614191] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 57.2% (14016644 kB) free.
[1589614420] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1589614420] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (643412 kB) free!
[1589614453] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.77, 22.52, 20.30
[1589614706] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (418972 kB) free!
[1589614827] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (411852 kB) free!
[1589614947] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (451972 kB) free!
[1589615061] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4794172 kB) free!
[1589615342] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(522) < 103.255.146.164 [16/May/2020:13:12:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589615462] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589615547] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29823972 kB) free.
[1589615620] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (2916228 kB) free!
[1589616147] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (374228 kB) free!
[1589616220] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (757940 kB) free!
[1589616267] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (424992 kB) free!
[1589616387] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (442212 kB) free!
[1589616584] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/16 13:32:49 [error] 22121#0: *88894855 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net"
[1589616704] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589616820] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (2892368 kB) free!
[1589616829] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1886656 kB) free!
[1589616949] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (1280820 kB) free!
[1589616977] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 759 processes
[1589616987] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38459376 kB) free.
[1589617069] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (634032 kB) free!
[1589617097] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 759 processes
[1589617191] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.5% (3544176 kB) free!
[1589617217] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 761 processes
[1589617311] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 12.7% (3121160 kB) free!
[1589617420] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1381108 kB) free!
[1589617432] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 11.1% (2724724 kB) free!
[1589617587] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5297668 kB) free!
[1589617707] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408432 kB) free!
[1589617725] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1589617827] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (373732 kB) free!
[1589617845] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1589617862] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(336) < 103.255.146.164 [16/May/2020:13:55:23 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589617982] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589618020] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3311124 kB) free.
[1589618032] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (654244 kB) free!
[1589618279] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 37.7% (9249396 kB) free.
[1589618427] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34293220 kB) free.
[1589618620] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (2907592 kB) free!
[1589618740] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2848100 kB) free!
[1589618860] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (836792 kB) free!
[1589619025] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589619027] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (1099180 kB) free!
[1589619147] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (416492 kB) free!
[1589619267] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448452 kB) free!
[1589619579] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.6% (24561460 kB) free!
[1589619699] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.9% (47259836 kB) free.
[1589619705] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 12.2% (4282344 kB) free!
[1589619825] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 10.7% (3755476 kB) free!
[1589619900] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.0% (3702328 kB) free!
[1589619945] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (3362248 kB) free!
[1589620020] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3356556 kB) free!
[1589620060] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (1888532 kB) free!
[1589620140] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.2% (3005528 kB) free!
[1589620373] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (2471284 kB) free!
[1589620453] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.27, 10.93, 17.26
[1589620661] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1034156 kB) free!
[1589620740] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1589620740] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (1144732 kB) free!
[1589620982] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(375) < 103.255.146.164 [16/May/2020:14:45:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589621102] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589621667] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18843060 kB) free!
[1589621702] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(144) < 185.255.8.58 [16/May/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1589621745] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.1% (31903948 kB) free.
[1589621822] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589621861] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (1730572 kB) free!
[1589621870] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.2% (16220744 kB) free.
[1589622267] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401172 kB) free!
[1589622540] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.0% (22959756 kB) free.
[1589622867] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33712244 kB) free.
[1589623061] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163448 kB) free!
[1589623467] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2402012 kB) free!
[1589623587] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393432 kB) free!
[1589623622] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(61) < 103.255.146.164 [16/May/2020:15:31:24 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589623661] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (1947916 kB) free!
[1589623707] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (459840 kB) free!
[1589623742] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589624307] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33675612 kB) free.
[1589624861] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (987680 kB) free!
[1589624907] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2428512 kB) free!
[1589625027] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407756 kB) free!
[1589625147] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423052 kB) free!
[1589625461] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (3239852 kB) free!
[1589625747] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30451836 kB) free.
[1589626061] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (882088 kB) free!
[1589626218] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1589626305] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1589626347] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (465496 kB) free!
[1589626425] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1589626454] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1589626467] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406660 kB) free!
[1589626588] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416380 kB) free!
[1589626661] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3335644 kB) free.
[1589627188] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29373696 kB) free.
[1589627230] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1589627261] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (2800872 kB) free!
[1589627381] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3391168 kB) free.
[1589627788] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (387936 kB) free!
[1589627908] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (465528 kB) free!
[1589627982] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.8% (1921056 kB) free!
[1589628028] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (462692 kB) free!
[1589628102] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3325512 kB) free.
[1589628628] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24629636 kB) free!
[1589628702] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1248184 kB) free!
[1589628822] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3275384 kB) free.
[1589629228] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412064 kB) free!
[1589629422] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (1197596 kB) free!
[1589629542] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3017784 kB) free!
[1589629662] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3689740 kB) free.
[1589629670] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (3913504 kB) free!
[1589629790] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3562804 kB) free!
[1589629910] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (3193280 kB) free!
[1589630262] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2936376 kB) free!
[1589630382] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3637500 kB) free.
[1589630510] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1679444 kB) free!
[1589630982] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3117500 kB) free!
[1589631102] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3742608 kB) free.
[1589631628] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21510460 kB) free!
[1589632228] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439116 kB) free!
[1589632303] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.4% (2335204 kB) free!
[1589632310] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1589632310] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16179092 kB) free.
[1589632423] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2607648 kB) free!
[1589632543] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.7% (2228412 kB) free!
[1589632740] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1589633743] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (838932 kB) free!
[1589634343] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.3% (4431020 kB) free.
[1589636743] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2411988 kB) free!
[1589636864] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.2% (4745576 kB) free.
[1589638063] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2817004 kB) free!
[1589638183] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.2% (4741168 kB) free.
[1589639653] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.58, 2.87, 4.50
[1589639983] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3210560 kB) free!
[1589640103] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176204 kB) free!
[1589640173] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176852 kB) free!
[1589640223] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3028468 kB) free!
[1589641229] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29589456 kB) free.
[1589641424] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1271500 kB) free!
[1589641829] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (402388 kB) free!
[1589641949] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399028 kB) free!
[1589642024] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.5% (4789300 kB) free.
[1589642053] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.05, 6.19, 5.53
[1589642069] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389192 kB) free!
[1589642173] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.94, 5.81, 5.47
[1589642293] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.18, 5.33, 5.35
[1589642510] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4271516 kB) free!
[1589642630] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4178904 kB) free!
[1589642750] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.3% (4101204 kB) free!
[1589642893] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.31, 2.96, 4.25
[1589643224] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2126196 kB) free!
[1589643344] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.2% (4253468 kB) free.
[1589643350] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.5% (16272540 kB) free.
[1589643944] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1243740 kB) free!
[1589644064] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3017612 kB) free!
[1589644184] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.2% (4413732 kB) free.
[1589644784] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1367684 kB) free!
[1589644904] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3058360 kB) free!
[1589645024] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.5% (4308340 kB) free.
[1589645624] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (917504 kB) free!
[1589645669] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26725800 kB) free.
[1589645744] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.8% (2249172 kB) free!
[1589645864] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3866016 kB) free.
[1589646269] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (443552 kB) free!
[1589646389] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397828 kB) free!
[1589646509] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432324 kB) free!
[1589647109] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24526916 kB) free!
[1589647664] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (786716 kB) free!
[1589647710] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439548 kB) free!
[1589647784] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.1% (1811280 kB) free!
[1589647904] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (2915400 kB) free!
[1589648396] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1695 processes
[1589649322] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.6% (4804600 kB) free!
[1589649704] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1498676 kB) free!
[1589650304] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3632080 kB) free.
[1589651310] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14132476 kB) free!
[1589651504] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (1715460 kB) free!
[1589651624] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (2192680 kB) free!
[1589651744] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2618576 kB) free!
[1589651910] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424272 kB) free!
[1589652493] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.37, 6.79, 5.75
[1589652510] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19713088 kB) free!
[1589652613] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.32, 7.35, 6.08
[1589652733] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.18, 6.85, 6.06
[1589652944] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.9% (5184488 kB) free.
[1589653110] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409552 kB) free!
[1589653710] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28517044 kB) free.