[1562005826] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12412212 kB) free.
[1562005875] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562006268] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192952 kB) free!
[1562006388] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (242220 kB) free!
[1562006426] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8681400 kB) free!
[1562006475] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562006508] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (226372 kB) free!
[1562006546] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12424620 kB) free.
[1562006595] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562006787] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (1742928 kB) free!
[1562006907] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1735916 kB) free!
[1562007027] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1725772 kB) free!
[1562007147] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8380304 kB) free!
[1562007267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12425324 kB) free.
[1562007281] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 78.6% (12771844 kB) free.
[1562007295] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.8% (21769540 kB) free.
[1562007307] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.1% (22581424 kB) free.
[1562007327] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 68.9% (17010776 kB) free.
[1562007795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562007885] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (150680 kB) free!
[1562007915] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562007937] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180160 kB) free!
[1562008005] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176132 kB) free!
[1562008036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562008057] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177848 kB) free!
[1562008125] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1562008125] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (149184 kB) free!
[1562008177] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1562008177] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194120 kB) free!
[1562008635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562008755] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562008841] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 216.51, 67.57, 23.89
[1562008961] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 291.46, 145.51, 57.75
[1562009081] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562009081] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 301.60, 197.71, 87.52
[1562009355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562009475] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562009906] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.59, 4.60, 4.89
[1562010075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562010195] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562010795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562010881] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 1.10, 17.59
[1562010916] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562011467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9207580 kB) free!
[1562011515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562011587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12768248 kB) free.
[1562011635] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562012188] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9270040 kB) free!
[1562012235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562012308] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12777888 kB) free.
[1562012355] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562012681] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.09, 0.21, 2.74
[1562012908] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9324836 kB) free!
[1562012955] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562013028] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12769076 kB) free.
[1562013075] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562013075] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562013339] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (185904 kB) free!
[1562013459] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (212100 kB) free!
[1562013579] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (218060 kB) free!
[1562013628] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9339712 kB) free!
[1562013675] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562013748] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12772824 kB) free.
[1562013796] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562014055] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562014348] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9398468 kB) free!
[1562014395] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562014468] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12774728 kB) free.
[1562014515] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562015068] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9301184 kB) free!
[1562015115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562015188] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12765292 kB) free.
[1562015235] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562015788] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9317160 kB) free!
[1562015835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562015908] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12774648 kB) free.
[1562015955] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562016509] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9199892 kB) free!
[1562016555] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562016629] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12763700 kB) free.
[1562016676] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562017106] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1562017229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9025824 kB) free!
[1562017275] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562017349] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12774028 kB) free.
[1562017395] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562017949] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8789920 kB) free!
[1562017995] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562018069] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12772992 kB) free.
[1562018115] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562018377] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562018669] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8813596 kB) free!
[1562018715] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562018789] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12754448 kB) free.
[1562018835] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562018925] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562019389] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8879936 kB) free!
[1562019435] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562019509] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12764860 kB) free.
[1562019555] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562019881] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 34.77, 8.60, 3.05
[1562020001] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 80.98, 34.94, 13.37
[1562020109] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8973636 kB) free!
[1562020121] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 87.62, 52.65, 22.48
[1562020156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562020230] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12754316 kB) free.
[1562020276] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1562020721] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.16, 14.34, 19.42
[1562020829] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9027416 kB) free!
[1562020876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562020949] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12761152 kB) free.
[1562020996] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562021549] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9021556 kB) free!
[1562021596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562021669] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12763488 kB) free.
[1562021716] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562022269] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8985448 kB) free!
[1562022316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562022389] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12756088 kB) free.
[1562022436] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562022521] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.27, 3.01
[1562022711] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 374 processes
[1562022832] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1562022989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8975792 kB) free!
[1562023037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562023055] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 73.37, 56.14, 28.57
[1562023109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12765092 kB) free.
[1562023157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562023175] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 73.80, 61.96, 34.04
[1562023295] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.69, 51.93, 34.04
[1562023709] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9129388 kB) free!
[1562023757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562023829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12754760 kB) free.
[1562023877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562023895] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.20, 9.46, 19.19
[1562024429] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9168308 kB) free!
[1562024477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562024550] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12767108 kB) free.
[1562024597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562025149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8875648 kB) free!
[1562025197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562025269] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12763216 kB) free.
[1562025317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562025438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562025695] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 233.95, 129.81, 57.08
[1562026638] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562027069] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8457348 kB) free!
[1562027189] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12503944 kB) free.
[1562027238] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562027358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562027478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562027980] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 48.2% (7827984 kB) free.
[1562028095] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1562028095] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.37, 3.70, 15.91
[1562028580] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (2619272 kB) free!
[1562028670] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562028700] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (184928 kB) free!
[1562028820] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 58.0% (9434320 kB) free.
[1562028989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8492780 kB) free!
[1562029109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12479344 kB) free.
[1562029427] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 531 processes
[1562029487] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 527 processes
[1562029706] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 60.77, 46.68, 24.49
[1562029803] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1209 processes
[1562029826] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.32, 37.66, 24.09
[1562029895] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.28, 2.88, 4.67
[1562029924] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1205 processes
[1562029946] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.18, 25.53, 21.29
[1562030044] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 933 processes
[1562030478] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1562030478] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562030546] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 3.88, 11.42
[1562030644] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 327 processes
[1562030910] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8519168 kB) free!
[1562031030] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12534604 kB) free.
[1562031630] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8537288 kB) free!
[1562031678] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562031746] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.73, 0.70, 3.57
[1562031750] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12541740 kB) free.
[1562031798] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562032350] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8225276 kB) free!
[1562032470] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12543112 kB) free.
[1562032721] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562032998] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562033070] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8122012 kB) free!
[1562033118] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562033190] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12024788 kB) free.
[1562034318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562034438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562034558] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562036358] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562036790] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7869912 kB) free!
[1562036910] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12538752 kB) free.
[1562038110] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9095160 kB) free!
[1562038158] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562038230] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12531276 kB) free.
[1562038278] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562038830] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9052228 kB) free!
[1562038878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562038950] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12532744 kB) free.
[1562038998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562039550] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8965356 kB) free!
[1562039598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562039670] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12534648 kB) free.
[1562039718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562039872] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.36, 6.27, 5.01
[1562039992] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.69, 5.11, 4.75
[1562040095] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562040270] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8785016 kB) free!
[1562040318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562040391] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12524332 kB) free.
[1562040439] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562040990] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8537676 kB) free!
[1562041038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562041110] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12530832 kB) free.
[1562041158] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562041710] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8091064 kB) free!
[1562041805] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.9% (17995540 kB) free.
[1562041830] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12522264 kB) free.
[1562041870] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1562041909] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.8% (14428992 kB) free.
[1562041990] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1562042110] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1562042336] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.1% (2614376 kB) free!
[1562042358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562042405] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187308 kB) free!
[1562042430] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7667596 kB) free!
[1562042456] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.7% (2393328 kB) free!
[1562042478] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562042515] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (156900 kB) free!
[1562042525] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180404 kB) free!
[1562042550] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12525596 kB) free.
[1562042576] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.5% (2200656 kB) free!
[1562042635] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (150216 kB) free!
[1562042646] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191312 kB) free!
[1562042755] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169060 kB) free!
[1562043078] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562043177] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179868 kB) free!
[1562043198] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562043220] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151140 kB) free!
[1562043340] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171824 kB) free!
[1562043460] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (285788 kB) free!
[1562043750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8658752 kB) free!
[1562043799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562043870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12526284 kB) free.
[1562043919] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562044470] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8127884 kB) free!
[1562044590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12515200 kB) free.
[1562044790] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.2% (18071796 kB) free.
[1562045119] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562045191] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7843352 kB) free!
[1562045239] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562045311] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12526072 kB) free.
[1562045321] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 22.28, 140.24, 98.10
[1562045397] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175784 kB) free!
[1562045441] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.17, 93.88, 86.22
[1562045448] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 585 processes
[1562045517] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196808 kB) free!
[1562045561] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.61, 62.91, 75.80
[1562045568] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 327 processes
[1562045637] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1562045637] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184376 kB) free!
[1562045839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562045959] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562045992] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.26, 6.63, 5.96
[1562046112] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.15, 8.58, 6.74
[1562046177] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 37.1% (6033600 kB) free.
[1562046195] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 181749 MB (80% inode=99%):
[1562046232] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.57, 7.53, 6.60
[1562046511] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9138844 kB) free!
[1562046559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562046631] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12517664 kB) free.
[1562046680] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562046695] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 118.62, 98.38, 52.83
[1562046815] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 62.16, 90.44, 55.87
[1562046935] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.90, 61.51, 49.46
[1562047231] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8839168 kB) free!
[1562047279] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562047351] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12520336 kB) free.
[1562047361] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 0.70, 11.43
[1562047377] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (3083652 kB) free!
[1562047399] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562047497] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.0% (2755072 kB) free!
[1562047618] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1562047618] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.1% (2452580 kB) free!
[1562047787] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 575 processes
[1562047847] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 568 processes
[1562047999] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562048135] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.97, 3.92, 15.74
[1562048218] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1269704 kB) free!
[1562048348] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (1369272 kB) free!
[1562048468] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.2% (1192464 kB) free!
[1562048551] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8928508 kB) free!
[1562048561] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.34, 3.42
[1562048588] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (1034844 kB) free!
[1562048599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562048671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12519980 kB) free.
[1562048719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562048885] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2731232 kB) free!
[1562049005] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (2822792 kB) free!
[1562049125] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2803324 kB) free!
[1562049189] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.1% (6776384 kB) free.
[1562049271] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8490664 kB) free!
[1562049319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562049392] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12510692 kB) free.
[1562049439] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562049918] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4361628 kB) free!
[1562049935] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.34, 3.15, 4.94
[1562049991] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8058064 kB) free!
[1562050038] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4237028 kB) free!
[1562050111] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12515668 kB) free.
[1562050158] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (4119416 kB) free!
[1562050639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562050660] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3396576 kB) free.
[1562050711] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7719936 kB) free!
[1562050759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562050831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12505168 kB) free.
[1562051099] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5373 MB (10% inode=97%):
[1562051135] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 291.26, 250.75, 138.52
[1562051218] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.1% (1805552 kB) free!
[1562051255] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 52.72, 178.63, 126.18
[1562051260] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (170344 kB) free!
[1562051359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562051375] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1562051375] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 9.76, 120.57, 111.27
[1562051380] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (157260 kB) free!
[1562051479] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562051500] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (414008 kB) free!
[1562051561] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 49.95, 17.43, 6.67
[1562051681] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 60.94, 32.44, 13.48
[1562051699] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6931 MB (13% inode=98%):
[1562051801] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1562051801] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 62.49, 42.47, 19.45
[1562051818] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (364244 kB) free!
[1562051959] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2446320 kB) free!
[1562052031] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8317312 kB) free!
[1562052151] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12243572 kB) free.
[1562052401] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 15.45, 19.32
[1562052679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562052800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562052920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562053175] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.58, 4.48, 19.21
[1562053848] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 593 processes
[1562053908] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 589 processes
[1562054201] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.56, 0.90, 3.67
[1562054403] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.61, 6.61, 5.52
[1562054501] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1562054501] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 40.9% (6642932 kB) free.
[1562054523] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.32, 6.95, 5.78
[1562054589] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (1335940 kB) free!
[1562054643] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.96, 6.95, 5.92
[1562054709] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.8% (1015880 kB) free!
[1562054829] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (713444 kB) free!
[1562055101] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176080 kB) free!
[1562055146] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 71.18, 161.66, 95.13
[1562055221] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (170856 kB) free!
[1562055266] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.56, 109.28, 84.00
[1562055341] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160260 kB) free!
[1562055386] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.39, 74.23, 74.23
[1562055575] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.40, 2.93, 4.43
[1562055843] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.39, 4.27, 4.85
[1562056437] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562056443] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.25, 6.75, 5.72
[1562056563] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.81, 5.51, 5.42
[1562056683] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.95, 3.87, 4.82
[1562056909] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 600 processes
[1562056969] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 592 processes
[1562057141] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (2417324 kB) free!
[1562057186] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.86, 2.20, 12.56
[1562057741] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1549968 kB) free!
[1562057883] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.64, 4.85, 5.13
[1562058003] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.54, 3.38, 4.56
[1562058151] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9405020 kB) free!
[1562058271] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7655492 kB) free!
[1562058320] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562058391] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12482232 kB) free.
[1562058920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562058986] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.38, 1.99, 3.83
[1562058991] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8901720 kB) free!
[1562059040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562059111] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12483904 kB) free.
[1562059160] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562059711] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8779280 kB) free!
[1562059760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562059803] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.18, 6.13, 5.10
[1562059831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12475456 kB) free.
[1562059880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562059923] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.36, 6.03, 5.18
[1562060043] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.36, 6.02, 5.27
[1562060186] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 46.79, 26.94, 12.45
[1562060229] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6840548 kB) free.
[1562060306] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 47.44, 33.79, 16.72
[1562060426] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1562060426] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.60, 38.40, 20.48
[1562060431] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8449296 kB) free!
[1562060480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562060552] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12484196 kB) free.
[1562060601] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562061026] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.59, 7.58, 13.21
[1562061151] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8143296 kB) free!
[1562061271] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12486076 kB) free.
[1562061295] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4536144 kB) free!
[1562061415] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4438296 kB) free!
[1562061535] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (4317040 kB) free!
[1562061800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562061920] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1562062226] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.57, 1.40, 4.33
[1562062465] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 13.5% (2190164 kB) free!
[1562062471] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8904744 kB) free!
[1562062520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562062585] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 12.4% (2008984 kB) free!
[1562062591] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12432992 kB) free.
[1562062640] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562062705] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 11.5% (1870948 kB) free!
[1562063191] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6647756 kB) free!
[1562063219] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562063305] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1108400 kB) free!
[1562063311] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12211344 kB) free.
[1562063335] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2242344 kB) free!
[1562063710] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 357 processes
[1562064401] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562065040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562065160] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562065175] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562065629] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.6% (778420 kB) free!
[1562065711] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8954656 kB) free!
[1562065749] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (540356 kB) free!
[1562065760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562065832] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12390460 kB) free.
[1562065870] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (279964 kB) free!
[1562065880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562066431] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8921948 kB) free!
[1562066480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562066555] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12393816 kB) free.
[1562066600] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562067155] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8734532 kB) free!
[1562067200] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562067275] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12394844 kB) free.
[1562067321] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562067875] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8729080 kB) free!
[1562067920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562067995] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12384860 kB) free.
[1562068040] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562068595] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8503172 kB) free!
[1562068640] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562068715] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12393468 kB) free.
[1562068760] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562069141] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562069315] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8257780 kB) free!
[1562069360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562069435] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12383492 kB) free.
[1562069480] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562069632] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.98, 3.80, 4.50
[1562069741] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (1784788 kB) free!
[1562069801] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 136.03, 35.39, 12.30
[1562069921] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 259.29, 115.84, 44.58
[1562069975] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 152.20, 82.05, 35.09
[1562070035] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7885956 kB) free!
[1562070041] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 276.15, 169.73, 72.97
[1562070095] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 156.08, 106.75, 49.84
[1562070156] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12389132 kB) free.
[1562070215] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.89, 80.07, 47.59
[1562070535] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (2565080 kB) free!
[1562070671] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.6% (6812404 kB) free.
[1562070680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562070800] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562070941] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1356032 kB) free!
[1562071106] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.7% (13767932 kB) free.
[1562071135] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1562071135] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2128104 kB) free!
[1562071355] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8267888 kB) free!
[1562071400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562071415] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.22, 4.38, 15.25
[1562071475] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12387608 kB) free.
[1562071520] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562071841] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 1.09, 16.55
[1562072032] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.74, 6.47, 5.52
[1562072063] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.3% (14518364 kB) free.
[1562072075] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7785140 kB) free!
[1562072136] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.4% (22306612 kB) free.
[1562072152] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.27, 6.25, 5.54
[1562072187] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.8% (22254428 kB) free.
[1562072195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12378452 kB) free.
[1562072230] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 68.6% (33815168 kB) free.
[1562072272] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.5% (22175000 kB) free.
[1562072272] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.62, 5.62, 5.40
[1562072315] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.4% (14201672 kB) free.
[1562072344] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 67.8% (11016756 kB) free.
[1562072720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562072840] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562072916] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (161124 kB) free!
[1562073026] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562073036] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169064 kB) free!
[1562073041] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.42, 4.85
[1562073156] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156812 kB) free!
[1562073215] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.96, 3.24, 4.97
[1562073346] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191092 kB) free!
[1562073395] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9083836 kB) free!
[1562073440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562073466] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185960 kB) free!
[1562073473] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.93, 4.91, 4.87
[1562073515] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12386268 kB) free.
[1562073561] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562073587] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192768 kB) free!
[1562073815] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.93, 10.07, 9.62
[1562073935] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.97, 7.70, 8.81
[1562074055] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1562074055] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.11, 6.16, 8.11
[1562074073] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.72, 5.53, 5.30
[1562074115] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9181736 kB) free!
[1562074160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562074193] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.35, 5.06, 5.16
[1562074235] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11553228 kB) free.
[1562074280] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562074313] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1562074313] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.06, 4.72, 5.02
[1562074443] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.97, 4.12, 4.82
[1562074835] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9095996 kB) free!
[1562074880] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562074913] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.30, 3.55, 4.17
[1562074956] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12270964 kB) free.
[1562075000] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562075043] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.81, 5.68, 5.06
[1562075163] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.41, 7.27, 5.71
[1562075239] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8939824 kB) free!
[1562075255] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.62, 3.01, 4.39
[1562075283] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.67, 8.11, 6.21
[1562075359] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (7800544 kB) free!
[1562075479] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (6844336 kB) free!
[1562075555] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8891720 kB) free!
[1562075600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562075675] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12385044 kB) free.
[1562075720] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562076026] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 226.39, 81.89, 30.50
[1562076079] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1296172 kB) free!
[1562076113] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.53, 6.08, 5.48
[1562076146] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 267.18, 145.35, 59.98
[1562076233] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.38, 5.63, 5.40
[1562076266] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 272.68, 187.75, 85.87
[1562076275] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8513988 kB) free!
[1562076321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562076353] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.00, 5.53, 5.40
[1562076395] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12373520 kB) free.
[1562076441] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562076995] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8106720 kB) free!
[1562077055] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 294.77, 237.21, 122.01
[1562077115] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12385076 kB) free.
[1562077175] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 249.56, 246.51, 139.78
[1562077295] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.06, 165.84, 123.16
[1562077553] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.69, 4.12, 4.69
[1562077641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562077715] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7538624 kB) free!
[1562077761] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562077835] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12373584 kB) free.
[1562078066] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.84, 2.24, 19.96
[1562078215] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562078335] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562078361] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562078481] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562079035] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8798284 kB) free!
[1562079081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562079155] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12380020 kB) free.
[1562079202] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562079695] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.30, 2.46, 11.71
[1562079756] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8321300 kB) free!
[1562079801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562079866] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.49, 0.57, 3.39
[1562079876] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12379880 kB) free.
[1562079921] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562080476] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7914740 kB) free!
[1562080596] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12368932 kB) free.
[1562081121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562081241] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562081495] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 70.51, 46.99, 21.62
[1562081796] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8868252 kB) free!
[1562081841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562081916] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12380136 kB) free.
[1562081961] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562082095] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.18, 8.31, 12.46
[1562082478] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562082516] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8554904 kB) free!
[1562082562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562082636] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12369116 kB) free.
[1562082682] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562083236] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8009584 kB) free!
[1562083356] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12376084 kB) free.
[1562083882] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562083895] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.29, 2.40, 3.86
[1562083957] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7730084 kB) free!
[1562084002] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562084077] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12378240 kB) free.
[1562084602] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562084722] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562085277] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8559020 kB) free!
[1562085322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562085397] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12369992 kB) free.
[1562085443] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562085483] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.26, 4.62, 4.98
[1562085997] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8286328 kB) free!
[1562086042] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562086083] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.86, 6.32, 5.73
[1562086117] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12377804 kB) free.
[1562086162] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562086203] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.21, 5.51, 5.50
[1562086323] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1562086323] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.11, 5.37, 5.45
[1562086717] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8176928 kB) free!
[1562086837] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12366528 kB) free.
[1562086923] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.09, 4.47, 4.95
[1562087362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562087437] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8054156 kB) free!
[1562087441] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 150.79, 43.54, 15.14
[1562087482] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562087543] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 35.5% (17483772 kB) free.
[1562087561] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 226.48, 107.98, 42.20
[1562087592] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22655148 kB) free.
[1562087652] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.3% (39554780 kB) free.
[1562087681] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 236.56, 151.05, 65.96
[1562087707] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.5% (14226756 kB) free.
[1562088082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562088198] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4754620 kB) free!
[1562088202] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562088318] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (1701600 kB) free!
[1562088318] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (153100 kB) free!
[1562088353] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562088438] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168820 kB) free!
[1562088438] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (156916 kB) free!
[1562088559] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (150568 kB) free!
[1562088749] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9700872 kB) free!
[1562088803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562088869] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8151056 kB) free!
[1562088910] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 363 processes
[1562088923] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562088989] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12741760 kB) free.
[1562089323] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.73, 6.93, 5.24
[1562089443] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.86, 7.32, 5.60
[1562089481] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 0.91, 14.35
[1562089510] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 365 processes
[1562089523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562089563] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.80, 7.60, 5.91
[1562089589] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8057400 kB) free!
[1562089643] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562089709] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12745588 kB) free.
[1562090081] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 100.01, 52.26, 28.44
[1562090243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562090309] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7962352 kB) free!
[1562090363] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562090429] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12740516 kB) free.
[1562090763] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.70, 4.13, 4.94
[1562090963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562091083] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562091258] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.1% (6474616 kB) free!
[1562091281] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1562091281] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.49, 3.18, 15.21
[1562091363] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.10, 6.40, 5.75
[1562091379] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.2% (5507204 kB) free!
[1562091483] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.38, 6.05, 5.70
[1562091499] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4573688 kB) free!
[1562091603] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.81, 6.04, 5.74
[1562091629] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9452928 kB) free!
[1562091683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562091750] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12747844 kB) free.
[1562091804] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes