[1584729383] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584729528] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171260 kB) free!
[1584729621] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584729648] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189868 kB) free!
[1584729720] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202176 kB) free!
[1584729768] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194612 kB) free!
[1584729782] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1584729807] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2433924 kB) free!
[1584729840] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189744 kB) free!
[1584729902] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1584729960] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173664 kB) free!
[1584730735] SERVICE FLAPPING ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584730787] SERVICE FLAPPING ALERT: smscapp2;Sensors Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584730980] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (2435468 kB) free!
[1584731040] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (1445476 kB) free!
[1584731089] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584731101] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (191428 kB) free!
[1584731153] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (2723028 kB) free!
[1584731167] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15825060 kB) free!
[1584731167] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (190092 kB) free!
[1584731191] SERVICE FLAPPING ALERT: smscapp2;HDD_Root status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584731273] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3002304 kB) free!
[1584731393] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3280024 kB) free.
[1584731767] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388504 kB) free!
[1584731828] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (4204604 kB) free!
[1584732067] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 75.1% (18402328 kB) free.
[1584732367] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21533816 kB) free!
[1584732428] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190704 kB) free!
[1584732593] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2529716 kB) free!
[1584732713] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (2721664 kB) free!
[1584732833] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (2971604 kB) free!
[1584732967] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394656 kB) free!
[1584733068] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584733433] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 36.9% (5992036 kB) free.
[1584733502] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1584733567] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26138636 kB) free!
[1584733622] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1584733628] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (4188384 kB) free!
[1584734167] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2851632 kB) free!
[1584734228] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180292 kB) free!
[1584734444] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584734689] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236076 kB) free!
[1584734691] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1564528 kB) free!
[1584734809] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236488 kB) free!
[1584734811] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8560104 kB) free!
[1584734929] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224872 kB) free!
[1584734931] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382168 kB) free!
[1584736729] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 53.4% (26331504 kB) free.
[1584736795] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1584736844] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.32, 5.09, 5.04
[1584736964] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.82, 5.02, 5.02
[1584737084] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.90, 5.03, 5.03
[1584737489] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2437152 kB) free!
[1584737609] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234004 kB) free!
[1584737684] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.16, 4.40, 4.73
[1584737729] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240604 kB) free!
[1584738284] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.61, 6.23, 5.49
[1584738404] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.20, 5.92, 5.47
[1584738524] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.98, 5.56, 5.39
[1584739124] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1584739124] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.54, 3.66, 4.60
[1584739131] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.7% (16705884 kB) free!
[1584739620] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 83.8% (41296612 kB) free.
[1584740129] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.1% (44394208 kB) free.
[1584740408] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584740529] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584740767] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.2% (9353500 kB) free.
[1584740822] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 334 processes
[1584740942] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1584741367] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23178820 kB) free!
[1584741967] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389328 kB) free!
[1584742567] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27932316 kB) free.
[1584743167] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2737292 kB) free!
[1584743287] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402656 kB) free!
[1584743407] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401960 kB) free!
[1584744007] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19655828 kB) free!
[1584744595] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1584744607] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377056 kB) free!
[1584745207] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17138068 kB) free!
[1584745629] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584745807] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400328 kB) free!
[1584746407] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17951124 kB) free!
[1584747007] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1023388 kB) free!
[1584747369] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (218648 kB) free!
[1584747489] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177444 kB) free!
[1584747608] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19232548 kB) free!
[1584747609] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177700 kB) free!
[1584747892] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2271.21, 1485.01, 605.08
[1584748012] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 308.09, 994.29, 531.94
[1584748132] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 43.53, 666.04, 467.72
[1584748208] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32935088 kB) free.
[1584748732] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.44, 91.47, 246.25
[1584748808] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15461628 kB) free!
[1584748928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12070848 kB) free!
[1584748968] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.3% (17478348 kB) free.
[1584749048] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (8684884 kB) free!
[1584749228] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.2% (10665784 kB) free.
[1584749648] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18092532 kB) free!
[1584749932] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 29.8% (39323260 kB) free.
[1584749995] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(136) < 92.42.104.166 [21/Mar/2020:05:45:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584750115] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584750248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36387068 kB) free.
[1584750524] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584750848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12905036 kB) free!
[1584750968] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8837364 kB) free!
[1584751088] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4787732 kB) free!
[1584751132] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.2% (16060760 kB) free!
[1584751252] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11838328 kB) free!
[1584751372] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6047216 kB) free!
[1584751688] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22711888 kB) free!
[1584752288] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4026308 kB) free!
[1584752572] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.3% (35969532 kB) free.
[1584752888] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22738912 kB) free!
[1584753172] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4562572 kB) free!
[1584753292] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (906996 kB) free!
[1584753412] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1584753412] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432212 kB) free!
[1584753488] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36461924 kB) free.
[1584754012] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 37.2% (49025220 kB) free.
[1584754088] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18478692 kB) free!
[1584754208] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (14994784 kB) free!
[1584754329] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (11294772 kB) free!
[1584754929] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24595368 kB) free!
[1584755529] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37386080 kB) free.
[1584756129] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20592116 kB) free!
[1584756249] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17112968 kB) free!
[1584756369] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13617400 kB) free!
[1584757569] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1127172 kB) free!
[1584758169] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14664668 kB) free!
[1584758767] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.7% (3353376 kB) free!
[1584758769] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30040316 kB) free.
[1584758827] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3354004 kB) free!
[1584758874] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 5118.96, 1514.05, 525.06
[1584758887] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 13.6% (3344824 kB) free!
[1584758905] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4539.20, 1556.81, 557.49
[1584758947] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 13.6% (3345396 kB) free!
[1584758994] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7377.41, 3570.59, 1398.59
[1584758995] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1584759007] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.32751.349> (10.147.212.158:47942 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1584759025] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 5561.34, 2975.81, 1206.06
[1584759056] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3101.17, 2611.61, 1142.73
[1584759115] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7681.87, 4946.63, 2166.20
[1584759115] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1584759127] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584759145] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1150.44, 2325.12, 1196.89
[1584759176] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 420.95, 1748.48, 1004.54
[1584759235] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 595 processes
[1584759296] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 58.86, 1170.67, 883.00
[1584759369] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (685448 kB) free!
[1584759428] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (4084000 kB) free!
[1584759489] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396188 kB) free!
[1584759548] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.5% (3829372 kB) free!
[1584759609] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395092 kB) free!
[1584759668] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2724836 kB) free!
[1584759715] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 182.18.144.246 [21/Mar/2020:08:22:39 +0530] TCP 502 0 0 121.002 "10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 1.002, 0.000"
[1584759745] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.22, 320.32, 632.33
[1584759835] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584759896] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.10, 159.50, 463.99
[1584760209] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23031232 kB) free!
[1584760268] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1363636 kB) free!
[1584760315] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.17, 113.65, 701.64
[1584760435] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1584760556] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 336 processes
[1584760676] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1584760724] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.28, 7.35, 5.96
[1584760809] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394872 kB) free!
[1584760844] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.68, 7.37, 6.13
[1584760915] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.69, 17.35, 368.80
[1584760964] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.67, 8.45, 6.69
[1584761276] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1584761410] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20167876 kB) free!
[1584762010] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420292 kB) free!
[1584762574] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < {writer,send_failed,{error,timeout}}
[1584762610] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20077808 kB) free!
[1584762695] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584763170] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584763210] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403208 kB) free!
[1584763290] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584764118] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45897 MB (20% inode=99%):
[1584764238] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45897 MB (20% inode=99%):
[1584764358] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45897 MB (20% inode=99%):
[1584765010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31455536 kB) free.
[1584765412] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584765610] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4275368 kB) free!
[1584765730] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389468 kB) free!
[1584765850] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376904 kB) free!
[1584766132] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 576.03, 227.74, 84.83
[1584766252] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 80.55, 153.36, 74.89
[1584766296] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584766368] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4227660 kB) free!
[1584766417] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584766435] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 103.13.97.241 [21/Mar/2020:10:19:40 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584766450] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24857856 kB) free!
[1584766488] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (4050660 kB) free!
[1584766539] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584766555] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584766608] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (3867832 kB) free!
[1584767050] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384476 kB) free!
[1584767155] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(50) < 103.255.146.164 [21/Mar/2020:10:32:10 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584767212] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.3% (25422660 kB) free!
[1584767275] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584767332] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.5% (17737352 kB) free!
[1584767452] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.9% (20982552 kB) free!
[1584767808] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1830976 kB) free!
[1584767875] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(21) < 103.255.146.164 [21/Mar/2020:10:44:40 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584767996] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584768042] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4238868 kB) free!
[1584768052] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.4% (32119344 kB) free.
[1584768162] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.1% (3836960 kB) free!
[1584768282] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (3449352 kB) free!
[1584768339] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584768460] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584768882] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1587944 kB) free!
[1584769252] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.5% (17853992 kB) free!
[1584769372] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10227728 kB) free!
[1584769451] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17697576 kB) free!
[1584769492] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.8% (24717324 kB) free!
[1584769529] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9288636 kB) free!
[1584769649] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8920676 kB) free!
[1584769769] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.4% (8564368 kB) free!
[1584769795] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(46) < 103.255.146.164 [21/Mar/2020:11:11:45 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584769915] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584770051] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409364 kB) free!
[1584770183] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.6% (3592876 kB) free!
[1584770303] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 12.6% (3111380 kB) free!
[1584770423] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.9% (2694088 kB) free!
[1584770651] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25767428 kB) free!
[1584770969] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4888660 kB) free!
[1584771021] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (5840136 kB) free!
[1584771023] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (283384 kB) free!
[1584771141] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 15.4% (5381144 kB) free!
[1584771251] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (555520 kB) free!
[1584771261] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (4937996 kB) free!
[1584771715] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(73) < 103.255.146.164 [21/Mar/2020:11:51:32 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584771834] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21707932 kB) free.
[1584771836] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584771861] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2852368 kB) free!
[1584771865] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.6% (21377316 kB) free.
[1584771892] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1584771892] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.3% (30753416 kB) free.
[1584771911] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34178812 kB) free.
[1584771955] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.0% (21324728 kB) free.
[1584771987] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.2% (20881432 kB) free.
[1584772019] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 35.3% (17408912 kB) free.
[1584772020] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9535520 kB) free!
[1584772140] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9342288 kB) free!
[1584772146] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.4% (6782712 kB) free!
[1584772260] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9539568 kB) free!
[1584772266] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.0% (31519648 kB) free.
[1584772461] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.5% (32218632 kB) free.
[1584772483] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.5% (15372956 kB) free.
[1584773035] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(36) < 103.255.146.164 [21/Mar/2020:12:13:44 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584773092] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.4% (22931296 kB) free!
[1584773119] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11752952 kB) free!
[1584773155] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584773212] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.7% (16793916 kB) free!
[1584773239] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7828256 kB) free!
[1584773332] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.3% (24121604 kB) free!
[1584773359] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4025596 kB) free!
[1584773423] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22868280 kB) free.
[1584773668] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169544 kB) free!
[1584773788] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178204 kB) free!
[1584773860] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584773908] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167340 kB) free!
[1584773981] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584774355] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(61) < 103.255.146.164 [21/Mar/2020:12:31:00 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584774476] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584774582] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584774660] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4622488 kB) free!
[1584774704] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584774834] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3248640 kB) free!
[1584774954] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3567728 kB) free.
[1584775075] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(31) < 103.255.146.164 [21/Mar/2020:12:41:49 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584775158] SERVICE ALERT: primary-server;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23808 MB (10% inode=99%):
[1584775159] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25362336 kB) free!
[1584775195] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584775369] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (213856 kB) free!
[1584775489] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (206668 kB) free!
[1584775609] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (203072 kB) free!
[1584775732] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12804764 kB) free!
[1584775759] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (859092 kB) free!
[1584775795] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(39) < 103.255.146.164 [21/Mar/2020:12:56:53 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584775915] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584776159] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179948 kB) free!
[1584776191] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172788 kB) free!
[1584776279] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194740 kB) free!
[1584776311] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201464 kB) free!
[1584776359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28696160 kB) free.
[1584776400] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170400 kB) free!
[1584776431] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191216 kB) free!
[1584776516] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(29) < 103.255.146.164 [21/Mar/2020:13:11:30 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584776636] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584776932] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (28609524 kB) free.
[1584776959] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3159352 kB) free!
[1584777079] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (390104 kB) free!
[1584777199] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378304 kB) free!
[1584777430] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (221660 kB) free!
[1584777532] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9484452 kB) free!
[1584777550] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219936 kB) free!
[1584777652] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.4% (22949016 kB) free!
[1584777670] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235108 kB) free!
[1584777772] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.9% (15642916 kB) free!
[1584778436] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(47) < 103.255.146.164 [21/Mar/2020:13:42:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584778556] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584779154] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.1% (1638920 kB) free!
[1584779194] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/21 13:49:25 [error] 20713#0: *74872007 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1584779274] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.7% (2393316 kB) free!
[1584779315] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584779394] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3060228 kB) free!
[1584779599] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16916476 kB) free!
[1584780199] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400544 kB) free!
[1584780239] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.0% (3680460 kB) free!
[1584780283] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4204276 kB) free!
[1584780356] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(64) < 103.255.146.164 [21/Mar/2020:14:11:59 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584780360] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (1537228 kB) free!
[1584780403] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.6% (3948812 kB) free!
[1584780476] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584780480] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (398136 kB) free!
[1584780523] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.4% (3678612 kB) free!
[1584780558] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24912 MB (11% inode=99%):
[1584780572] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1584780692] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1584780704] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584780812] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1584780825] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584781080] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.9% (4391032 kB) free!
[1584781158] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 62353 MB (27% inode=99%):
[1584781372] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (10133776 kB) free!
[1584781680] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1584781680] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 64.9% (15909788 kB) free.
[1584781723] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (683996 kB) free!
[1584781794] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.7% (5475240 kB) free.
[1584781863] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.2% (17826480 kB) free.
[1584781897] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 28.1% (13858596 kB) free.
[1584781927] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 32.1% (42315076 kB) free.
[1584781959] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22918464 kB) free!
[1584781962] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33126432 kB) free.
[1584781989] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (8878256 kB) free!
[1584781992] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 43.5% (21429272 kB) free.
[1584782021] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1584782021] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.4% (15359516 kB) free.
[1584782053] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 27.2% (13417020 kB) free.
[1584782078] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 65.8% (16132620 kB) free.
[1584782096] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.1% (18641748 kB) free.
[1584782128] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.2% (14172688 kB) free.
[1584782280] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194932 kB) free!
[1584782394] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2763312 kB) free!
[1584782400] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183264 kB) free!
[1584782423] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 12.9% (3185808 kB) free!
[1584782470] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192372 kB) free!
[1584782514] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 34.1% (5545928 kB) free.
[1584782520] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 70.2% (17198872 kB) free.
[1584782544] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 11.6% (2874104 kB) free!
[1584782590] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194232 kB) free!
[1584782599] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.6% (5731460 kB) free!
[1584782612] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1584782664] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2448952 kB) free!
[1584782678] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183104 kB) free!
[1584782710] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (162888 kB) free!
[1584782719] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.7% (5293600 kB) free!
[1584782799] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189208 kB) free!
[1584782839] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (4995196 kB) free!
[1584782876] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(46) < 103.255.146.164 [21/Mar/2020:14:54:23 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584782919] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1584782919] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177788 kB) free!
[1584782996] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584783115] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (2004212 kB) free!
[1584783169] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13359800 kB) free!
[1584783235] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.4% (5425252 kB) free.
[1584783289] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8856584 kB) free!
[1584783409] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4637180 kB) free!
[1584783439] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1584783439] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4385740 kB) free!
[1584783835] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2109924 kB) free!
[1584783955] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3452736 kB) free.
[1584783974] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 21 users currently logged in
[1584784009] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25888484 kB) free!
[1584784094] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 21 users currently logged in
[1584784197] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(206) < 103.255.146.164 [21/Mar/2020:15:14:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584784215] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 21 users currently logged in
[1584784317] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584784464] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22876584 kB) free.
[1584784609] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383304 kB) free!
[1584785155] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (746816 kB) free!
[1584785275] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.0% (2281404 kB) free!
[1584785395] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.2% (4257532 kB) free.
[1584785626] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584785747] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584785809] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23129012 kB) free!
[1584785942] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1584785995] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (518328 kB) free!
[1584786015] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 27 users currently logged in
[1584786062] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584786115] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1894748 kB) free!
[1584786235] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.9% (4368404 kB) free.
[1584786300] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178552 kB) free!
[1584786409] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392128 kB) free!
[1584786420] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193644 kB) free!
[1584786514] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/21 15:49:04 [error] 20706#0: *74923500 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/DlrLoad/msisdn"
[1584786540] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164488 kB) free!
[1584786634] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584786663] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1584786783] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1584786864] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 372 processes
[1584786984] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1584787320] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.2% (3961152 kB) free!
[1584787435] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (163364 kB) free!
[1584787440] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.7% (3611356 kB) free!
[1584787555] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.5% (2029664 kB) free!
[1584787560] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.4% (3292632 kB) free!
[1584787676] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.6% (4331324 kB) free.
[1584788160] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1681904 kB) free!
[1584788875] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2480484 kB) free!
[1584788995] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.8% (4523636 kB) free.
[1584789739] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 58.4% (14418320 kB) free.
[1584789772] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 63.0% (15432132 kB) free.
[1584789803] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16952804 kB) free!
[1584789831] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 49.5% (24385136 kB) free.
[1584789857] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 60.8% (14899408 kB) free.
[1584789878] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.5% (20217236 kB) free.
[1584790195] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2095264 kB) free!
[1584790315] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.1% (4249104 kB) free.
[1584790350] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168756 kB) free!
[1584790409] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406384 kB) free!
[1584790440] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4795748 kB) free!
[1584790470] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183032 kB) free!
[1584790560] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (4654324 kB) free!
[1584790590] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180888 kB) free!
[1584790680] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4529536 kB) free!
[1584790915] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2555400 kB) free!
[1584791035] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.3% (4270400 kB) free.
[1584791630] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3540096 kB) free!
[1584791636] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (150252 kB) free!
[1584791750] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3215840 kB) free!
[1584791756] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.0% (1957684 kB) free!
[1584791870] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.2% (2953360 kB) free!
[1584791876] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.9% (4201932 kB) free.
[1584792470] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1669952 kB) free!
[1584792730] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584792809] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29784744 kB) free.
[1584792847] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (184892 kB) free!
[1584793076] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (346668 kB) free!
[1584793196] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.6% (2047864 kB) free!
[1584793316] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3912916 kB) free.
[1584793409] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (387608 kB) free!
[1584793464] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4816184 kB) free!
[1584793529] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400928 kB) free!
[1584793548] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584793584] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4407236 kB) free!
[1584793649] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398836 kB) free!
[1584793669] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584793704] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.4% (4047508 kB) free!
[1584793917] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(10) < 203.212.70.170 [21/Mar/2020:18:01:09 +0530] TCP 502 0 0 60.001 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1584793945] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4942.06, 2089.13, 787.61
[1584794037] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(61) < 172.105.59.173 [21/Mar/2020:18:02:09 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584794065] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3819.64, 2840.91, 1234.03
[1584794157] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(6) < 103.13.97.241 [21/Mar/2020:18:05:26 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584794185] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 517.86, 1901.80, 1084.71
[1584794209] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 37.4% (9222060 kB) free.
[1584794271] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.6% (15407116 kB) free.
[1584794276] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1584794297] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 36.8% (18151400 kB) free.
[1584794304] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584794304] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2300116 kB) free!
[1584794457] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584794659] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (4028516 kB) free!
[1584794779] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.1% (3935624 kB) free!
[1584794785] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.30, 259.05, 571.00
[1584794809] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193076 kB) free!
[1584794899] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (3863896 kB) free!
[1584794900] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216412 kB) free!
[1584794904] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22874660 kB) free.
[1584794929] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190788 kB) free!
[1584795020] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235856 kB) free!
[1584795050] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1584795050] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172968 kB) free!
[1584795061] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17418464 kB) free!
[1584795116] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2412496 kB) free!
[1584795140] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216300 kB) free!
[1584795236] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3857952 kB) free.
[1584795476] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1584795596] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1584795615] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 24 users currently logged in
[1584795661] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417092 kB) free!
[1584795716] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1584795957] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584796557] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(570) < 49.50.86.152 [21/Mar/2020:18:45:47 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1584796678] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(10) < 172.105.57.57 [21/Mar/2020:18:46:43 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584796739] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584796798] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584796834] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(40) < 2020/03/21 18:44:45 [error] 20706#0: *75010954 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584796860] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 70.7% (34847356 kB) free.
[1584796861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33407520 kB) free.
[1584796954] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584797299] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1882668 kB) free!
[1584797416] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 27 users currently logged in
[1584797461] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (1089908 kB) free!
[1584797582] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389684 kB) free!
[1584797637] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (1095668 kB) free!
[1584797659] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9867276 kB) free!
[1584797702] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1584797702] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433904 kB) free!
[1584797757] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2576456 kB) free!
[1584797779] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9561360 kB) free!
[1584797877] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.1% (4080548 kB) free.
[1584797899] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (9183916 kB) free!
[1584798179] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.2% (3962628 kB) free!
[1584798299] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.8% (3625168 kB) free!
[1584798419] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.4% (3286328 kB) free!
[1584799019] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1652280 kB) free!
[1584799077] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2134396 kB) free!
[1584799197] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3592784 kB) free.
[1584799316] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584799699] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3720832 kB) free!
[1584799798] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2056) < 182.18.144.246 [21/Mar/2020:19:34:00 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584799916] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1584799918] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584800036] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 339 processes
[1584800156] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1584800156] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1584800397] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (148800 kB) free!
[1584800517] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.8% (1749404 kB) free!
[1584800637] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3325692 kB) free.
[1584800680] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4331416 kB) free!
[1584800702] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29676608 kB) free.
[1584800756] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1584800800] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4420416 kB) free!
[1584800920] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4368228 kB) free!
[1584801302] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (385084 kB) free!
[1584801331] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (177436 kB) free!
[1584801422] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389996 kB) free!
[1584801451] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180588 kB) free!
[1584801479] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.4% (3464344 kB) free!
[1584801542] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416084 kB) free!
[1584801571] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180312 kB) free!
[1584801599] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.1% (3175864 kB) free!
[1584801719] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.9% (2888244 kB) free!
[1584802318] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.57.57 [21/Mar/2020:20:19:24 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 60.000, 0.000"
[1584802319] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1595900 kB) free!
[1584802438] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (1461208 kB) free!
[1584802438] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584802558] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3646988 kB) free.
[1584802720] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2320680 kB) free!
[1584802742] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16815320 kB) free!
[1584802816] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 24 users currently logged in
[1584802900] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9393700 kB) free!
[1584803020] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (9133972 kB) free!
[1584803141] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (8975092 kB) free!
[1584803304] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (4104700 kB) free!
[1584803342] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399572 kB) free!
[1584803424] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.4% (3563716 kB) free!
[1584803544] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.4% (3049128 kB) free!
[1584803638] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(27) < 203.212.70.170 [21/Mar/2020:20:38:05 +0530] TCP 502 0 0 180.001 "10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 60.000, 0.000"
[1584803756] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1584803758] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584803758] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.2% (1654688 kB) free!
[1584803876] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1584803878] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.6% (3992752 kB) free.
[1584803965] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21665612 kB) free.
[1584803989] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.2% (11883552 kB) free.
[1584803996] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1584804058] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 40.6% (9118812 kB) free.
[1584804117] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 28.9% (14248432 kB) free.
[1584804139] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 24.1% (11861368 kB) free.
[1584804144] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (383708 kB) free!
[1584804161] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 66.8% (16371580 kB) free.
[1584804187] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 67.6% (16570684 kB) free.
[1584804242] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.2% (14174284 kB) free.
[1584804261] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (6441348 kB) free!
[1584804381] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 17.6% (6127016 kB) free!
[1584804501] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (5677796 kB) free!
[1584804592] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187808 kB) free!
[1584804616] SERVICE FLAPPING ALERT: secondary-server;Current Users;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1584804616] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 26 users currently logged in
[1584804637] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32966552 kB) free.
[1584804712] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187532 kB) free!
[1584804720] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229752 kB) free!
[1584804755] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/03/21 21:02:18 [error] 20699#0: *75071124 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584804832] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177196 kB) free!
[1584804840] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214552 kB) free!
[1584804875] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584804959] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3258) < 49.50.86.152 [21/Mar/2020:21:05:23 +0530] TCP 502 0 0 180.001 "10.147.212.142:3001, 10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.001, 0.000"
[1584804960] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246868 kB) free!
[1584805079] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584805101] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.5% (32237500 kB) free.
[1584805196] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1584805237] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3671952 kB) free!
[1584805357] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399772 kB) free!
[1584805478] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398160 kB) free!
[1584805541] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4778728 kB) free!
[1584805678] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.4% (1698292 kB) free!
[1584805799] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3332780 kB) free.
[1584805944] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.8% (22902808 kB) free.
[1584806137] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 15.6% (3812700 kB) free!
[1584806396] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584806678] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20865432 kB) free!
[1584806998] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (1397136 kB) free!
[1584807118] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2865592 kB) free!
[1584807238] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.0% (4222432 kB) free.
[1584807278] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374112 kB) free!
[1584807341] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (9323216 kB) free!
[1584808079] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(137) < 172.105.59.173 [21/Mar/2020:21:57:54 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584808196] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584808199] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(9) < 172.105.59.173 [21/Mar/2020:21:58:44 +0530] TCP 502 0 0 195.022 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "15.022, 60.000, 60.000, 60.000"
[1584808319] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584808364] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.34, 3.44, 4.37
[1584808370] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584808438] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (381920 kB) free!
[1584808478] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24708384 kB) free!
[1584808558] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.5% (1714292 kB) free!
[1584808625] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.1% (18790004 kB) free.
[1584808652] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 54.1% (26670184 kB) free.
[1584808678] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3061144 kB) free!
[1584809078] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394288 kB) free!
[1584809231] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165292 kB) free!
[1584809267] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229288 kB) free!
[1584809278] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.7% (4170780 kB) free.
[1584809351] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163628 kB) free!
[1584809371] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1584809387] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (644196 kB) free!
[1584809471] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184328 kB) free!
[1584809507] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (576800 kB) free!
[1584809520] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4732) < 203.212.70.170 [21/Mar/2020:22:20:42 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584809640] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584810278] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27556316 kB) free.
[1584810840] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(771) < 172.105.59.173 [21/Mar/2020:22:43:50 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584810878] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (396832 kB) free!
[1584810909] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36617 MB (16% inode=99%):
[1584810961] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(166) < 172.105.59.173 [21/Mar/2020:22:45:52 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584810998] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (377892 kB) free!
[1584811029] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 36601 MB (16% inode=99%):
[1584811081] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584811118] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398952 kB) free!
[1584811149] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 37957 MB (16% inode=99%):
[1584811681] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1167) < 203.212.70.170 [21/Mar/2020:22:50:49 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584811796] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1584811801] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584811916] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1584811964] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.62, 5.00, 5.07
[1584812036] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1584812084] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.16, 4.41, 4.85
[1584812741] SERVICE FLAPPING ALERT: qservercluster2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1584812741] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4908284 kB) free!
[1584812949] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 126271 MB (56% inode=99%):
[1584813001] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(69) < 103.255.146.164 [21/Mar/2020:23:20:01 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584813016] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 23 users currently logged in
[1584813121] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1676) < 203.212.70.170 [21/Mar/2020:23:21:24 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584813236] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1584813241] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(5) < 203.212.70.170 [21/Mar/2020:23:23:24 +0530] TCP 502 0 0 180.000 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1584813660] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584814118] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21424388 kB) free!
[1584814441] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584814718] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411564 kB) free!
[1584815041] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(772) < 182.18.144.246 [21/Mar/2020:23:53:28 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584815161] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584815278] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2559564 kB) free!
[1584815398] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3294300 kB) free.