[1584642602] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.21711.278> (10.147.212.158:34380 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1584642704] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 715 processes
[1584642722] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584642824] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 713 processes
[1584642858] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 28.9% (7084860 kB) free.
[1584642944] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 712 processes
[1584643114] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3258384 kB) free!
[1584643234] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (412272 kB) free!
[1584643283] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170692 kB) free!
[1584643355] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (413800 kB) free!
[1584643403] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192184 kB) free!
[1584643459] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186540 kB) free!
[1584643523] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195684 kB) free!
[1584643579] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197676 kB) free!
[1584643700] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204616 kB) free!
[1584643955] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24488676 kB) free!
[1584644555] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401044 kB) free!
[1584644698] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (200636 kB) free!
[1584645155] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27987236 kB) free.
[1584645599] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 73.1% (17912928 kB) free.
[1584645755] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4220856 kB) free!
[1584645875] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386116 kB) free!
[1584645995] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433928 kB) free!
[1584646579] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584646595] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22238908 kB) free!
[1584647195] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2872228 kB) free!
[1584647779] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11362772 kB) free!
[1584647795] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24245060 kB) free!
[1584647899] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9669080 kB) free!
[1584647929] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6088784 kB) free!
[1584647993] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1584648019] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (9604432 kB) free!
[1584648049] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (2485684 kB) free!
[1584648172] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237348 kB) free!
[1584648395] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4182572 kB) free!
[1584648995] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24596024 kB) free!
[1584649595] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4157640 kB) free!
[1584649852] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 967.94, 210.47, 70.64
[1584649972] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 5544.65, 2211.60, 820.91
[1584649973] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 27.3% (13436772 kB) free.
[1584649975] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 5580.07, 2224.79, 825.63
[1584650092] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6160.55, 3550.06, 1480.01
[1584650095] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6202.23, 3572.85, 1489.27
[1584650195] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25236196 kB) free!
[1584650216] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 839.17, 2391.58, 1309.02
[1584650385] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(570) < 172.105.59.173 [20/Mar/2020:02:06:06 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.000, 0.000"
[1584650393] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1584650505] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584650795] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5401480 kB) free!
[1584651396] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24973508 kB) free!
[1584651416] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.03, 45.85, 361.44
[1584651662] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (228876 kB) free!
[1584651762] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (961660 kB) free!
[1584651782] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238004 kB) free!
[1584651883] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228836 kB) free!
[1584651892] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.76, 21.21, 285.04
[1584651902] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246292 kB) free!
[1584651996] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5792440 kB) free!
[1584652003] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 88.6% (43664152 kB) free.
[1584652596] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24357928 kB) free!
[1584653103] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.7% (43201916 kB) free.
[1584653196] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5862284 kB) free!
[1584653796] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22492668 kB) free!
[1584653900] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.4% (9899780 kB) free.
[1584654348] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584654396] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4256584 kB) free!
[1584654996] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20203884 kB) free!
[1584655596] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39213860 kB) free.
[1584656196] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17391744 kB) free!
[1584656316] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13697544 kB) free!
[1584656436] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (10054756 kB) free!
[1584657036] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25641936 kB) free!
[1584657636] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (7045372 kB) free!
[1584658236] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21682368 kB) free!
[1584658836] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39333148 kB) free.
[1584659436] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18021396 kB) free!
[1584659557] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14393124 kB) free!
[1584659677] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (10801472 kB) free!
[1584660277] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24739064 kB) free!
[1584660877] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39283336 kB) free.
[1584661100] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177976 kB) free!
[1584661220] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196304 kB) free!
[1584661340] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187204 kB) free!
[1584661477] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19280724 kB) free!
[1584661597] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15568240 kB) free!
[1584661717] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (11988024 kB) free!
[1584662317] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25571836 kB) free!
[1584662539] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.4% (10471688 kB) free.
[1584662723] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 70.7% (17333936 kB) free.
[1584662917] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39184308 kB) free.
[1584663517] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19605232 kB) free!
[1584663637] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (16010160 kB) free!
[1584663757] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12455428 kB) free!
[1584664357] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24110164 kB) free!
[1584664670] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4870920 kB) free!
[1584664790] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.8% (4854472 kB) free!
[1584664911] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.7% (4838496 kB) free!
[1584664957] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37892280 kB) free.
[1584665557] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14820440 kB) free!
[1584665677] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10907884 kB) free!
[1584665798] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6561736 kB) free!
[1584666019] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.1% (27769800 kB) free.
[1584666398] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18039024 kB) free!
[1584666619] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1420284 kB) free!
[1584666739] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (415680 kB) free!
[1584666859] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (753404 kB) free!
[1584666998] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32442512 kB) free.
[1584667459] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.0% (46113908 kB) free.
[1584667526] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39488432 kB) free.
[1584667600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9759752 kB) free!
[1584667601] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.4% (20697520 kB) free.
[1584667690] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.4% (50611736 kB) free.
[1584667911] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2228980 kB) free!
[1584668099] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 12.7% (3107612 kB) free!
[1584668159] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 12.7% (3109708 kB) free!
[1584668220] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 12.7% (3109632 kB) free!
[1584668280] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 12.7% (3108600 kB) free!
[1584668539] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (185016 kB) free!
[1584668659] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200748 kB) free!
[1584668779] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174164 kB) free!
[1584669499] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16136864 kB) free!
[1584669619] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11862616 kB) free!
[1584669739] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (7808096 kB) free!
[1584670339] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21410660 kB) free!
[1584670939] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36954080 kB) free.
[1584671539] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11711084 kB) free!
[1584671659] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7470320 kB) free!
[1584671779] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42453308 kB) free.
[1584672234] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1939.10, 515.25, 175.80
[1584672264] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 5666.32, 1853.30, 656.48
[1584672354] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1812.03, 1152.82, 463.52
[1584672379] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17232780 kB) free!
[1584672384] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7229.75, 3713.02, 1483.10
[1584672474] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 245.73, 771.82, 407.48
[1584672500] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12946348 kB) free!
[1584672504] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7440.53, 4957.20, 2209.40
[1584672620] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8413672 kB) free!
[1584672705] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2052) < 103.255.146.164 [20/Mar/2020:08:13:21 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584672826] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584673016] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2856.34, 913.05, 322.62
[1584673080] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 27.7% (6785376 kB) free.
[1584673136] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1287.80, 1188.40, 511.03
[1584673193] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1584673220] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20967420 kB) free!
[1584673256] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 176.24, 796.14, 449.42
[1584673313] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1584673433] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1584673609] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (1233888 kB) free!
[1584673704] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.54, 112.74, 700.93
[1584673729] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189720 kB) free!
[1584673820] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36624024 kB) free.
[1584673849] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178816 kB) free!
[1584674305] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.46, 17.63, 368.70
[1584674420] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (9027588 kB) free!
[1584674540] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4441944 kB) free!
[1584674660] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393684 kB) free!
[1584675260] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16913300 kB) free!
[1584675306] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584675430] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 16.87, 17.53, 14.68
[1584675860] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30821808 kB) free.
[1584676383] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676460] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5020244 kB) free!
[1584676503] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 63.2% (31165032 kB) free.
[1584676580] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (492496 kB) free!
[1584676582] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676680] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676700] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380128 kB) free!
[1584676712] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676729] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676799] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676799] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676802] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676810] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676842] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676857] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676859] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584676920] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 193434 MB (86% inode=99%):
[1584676920] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584676930] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2529 processes
[1584676930] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2529 processes
[1584676978] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584676980] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584677113] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677236] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677243] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677300] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13652536 kB) free!
[1584677366] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677373] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677496] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677529] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677529] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677532] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677540] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677587] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677589] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677651] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584677652] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 189325 MB (84% inode=99%):
[1584677657] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2535 processes
[1584677661] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2536 processes
[1584677717] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677719] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677723] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (3157548 kB) free!
[1584677843] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.5% (2820176 kB) free!
[1584677847] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677849] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584677900] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400136 kB) free!
[1584677963] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2462264 kB) free!
[1584678050] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38771 MB (75% inode=98%):
[1584678259] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584678259] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584678262] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584678270] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584678389] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584678389] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584678390] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2548 processes
[1584678391] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2551 processes
[1584678449] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584678450] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584678514] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 184791 MB (82% inode=99%):
[1584679066] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(22) < {writer,send_failed,{error,timeout}}
[1584679101] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26229840 kB) free!
[1584679174] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.7% (27939132 kB) free.
[1584679187] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(38) < {writer,send_failed,{error,timeout}}
[1584679307] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584679530] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8269064 kB) free!
[1584679650] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9236980 kB) free!
[1584679701] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6611468 kB) free!
[1584679770] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (9213540 kB) free!
[1584679907] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(6) < {writer,send_failed,{error,timeout}}
[1584679956] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21666212 kB) free.
[1584679960] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 157.09, 97.39, 45.54
[1584679985] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1043 processes
[1584680010] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.6% (20632232 kB) free.
[1584680026] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(188) < 103.255.146.164 [20/Mar/2020:10:21:17 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584680028] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584680074] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.6% (20493332 kB) free.
[1584680080] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 160.97, 117.37, 59.06
[1584680106] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1045 processes
[1584680118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38478184 kB) free.
[1584680146] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584680200] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 152.42, 129.00, 70.45
[1584680203] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.6% (20983284 kB) free.
[1584680228] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1046 processes
[1584680828] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 371 processes
[1584680970] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3534140 kB) free!
[1584681228] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(5) < {writer,send_failed,{error,timeout}}
[1584681328] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13562628 kB) free!
[1584681348] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(10) < {writer,send_failed,{error,timeout}}
[1584681448] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9866424 kB) free!
[1584681468] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(16) < {writer,send_failed,{error,timeout}}
[1584681492] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3419832 kB) free!
[1584681568] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30913348 kB) free.
[1584681612] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.0% (2915628 kB) free!
[1584681668] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39461788 kB) free.
[1584681732] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.6% (2383132 kB) free!
[1584681819] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195000 kB) free!
[1584681939] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174144 kB) free!
[1584681992] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584682010] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 79.4% (19591104 kB) free.
[1584682033] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 661 processes
[1584682116] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2539 processes
[1584682159] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 686 processes
[1584682168] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (401320 kB) free!
[1584682283] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 688 processes
[1584682288] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389112 kB) free!
[1584682333] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (191060 kB) free!
[1584682349] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (193408 kB) free!
[1584682409] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384240 kB) free!
[1584682469] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (295988 kB) free!
[1584682546] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(153) < 103.255.146.164 [20/Mar/2020:11:04:09 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584682589] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (309516 kB) free!
[1584682609] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179020 kB) free!
[1584682622] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172544 kB) free!
[1584682666] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584682730] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165184 kB) free!
[1584682742] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195828 kB) free!
[1584682852] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189300 kB) free!
[1584682862] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163096 kB) free!
[1584682938] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.8% (20678852 kB) free.
[1584682986] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.7% (57652968 kB) free.
[1584683019] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.3% (15112172 kB) free.
[1584683033] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584683078] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4669156 kB) free!
[1584683121] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.2% (21619792 kB) free.
[1584683178] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.5% (14058060 kB) free.
[1584683190] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 10.4% (2571580 kB) free!
[1584683198] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4423932 kB) free!
[1584683311] SERVICE ALERT: db1-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (1882036 kB) free!
[1584683314] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/20 11:17:31 [error] 12245#0: *73981001 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1584683318] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4126824 kB) free!
[1584683431] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1204180 kB) free!
[1584683434] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584683545] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584683547] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (1006624 kB) free!
[1584683667] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (628868 kB) free!
[1584683670] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 795 processes
[1584683703] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9790964 kB) free!
[1584683787] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1584683787] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185108 kB) free!
[1584683823] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9432792 kB) free!
[1584683899] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (5744428 kB) free!
[1584683943] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9061296 kB) free!
[1584683974] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8677448 kB) free!
[1584684019] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 14.7% (5119076 kB) free!
[1584684095] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8095724 kB) free!
[1584684139] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (4531736 kB) free!
[1584684216] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.1% (7460844 kB) free!
[1584684518] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1584684518] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1336432 kB) free!
[1584684663] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 13.4% (4677176 kB) free!
[1584684677] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.1% (6967624 kB) free!
[1584684739] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (1329572 kB) free!
[1584684759] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.3% (4252144 kB) free!
[1584684783] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 11.8% (4138700 kB) free!
[1584684786] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4216248 kB) free!
[1584684797] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.9% (6872320 kB) free!
[1584684816] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1584684816] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2452648 kB) free!
[1584684879] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.8% (3881428 kB) free!
[1584684903] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (3523368 kB) free!
[1584684906] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (379064 kB) free!
[1584684918] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1584684918] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.8% (6819524 kB) free!
[1584684936] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22847848 kB) free.
[1584684972] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21948624 kB) free.
[1584684999] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.5% (3547604 kB) free!
[1584685015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39603132 kB) free.
[1584685068] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584685091] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 91.7% (31935632 kB) free.
[1584685217] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 43.7% (21518104 kB) free.
[1584685322] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584685344] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.1% (20360344 kB) free.
[1584685429] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (3784100 kB) free!
[1584685442] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(13) < {writer,send_failed,{error,timeout}}
[1584685503] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (649076 kB) free!
[1584685531] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178376 kB) free!
[1584685549] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3206152 kB) free!
[1584685562] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;3;(27) < {writer,send_failed,{error,timeout}}
[1584685599] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1733480 kB) free!
[1584685603] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8152456 kB) free!
[1584685651] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199132 kB) free!
[1584685669] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(195) < {writer,send_failed,{error,timeout}}
[1584685669] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1584685669] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.8% (2648392 kB) free!
[1584685723] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7813248 kB) free!
[1584685723] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(152) < {writer,send_failed,{error,timeout}}
[1584685743] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4007416 kB) free!
[1584685771] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189892 kB) free!
[1584685789] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(44) < {writer,send_failed,{error,timeout}}
[1584685843] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (7626968 kB) free!
[1584685844] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(9) < {writer,send_failed,{error,timeout}}
[1584685912] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584685965] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584686104] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.1% (31560400 kB) free.
[1584686213] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21626764 kB) free.
[1584686219] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16769100 kB) free!
[1584686269] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.7% (15192052 kB) free.
[1584686289] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (45227532 kB) free.
[1584686339] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11866356 kB) free!
[1584686345] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1584686345] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.1% (21356668 kB) free.
[1584686383] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39163512 kB) free.
[1584686434] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/03/20 12:09:35 [error] 20586#0: *74025238 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: rq-1-1gbps.server, request: "GET /favicon.ico HTTP/1.1", host: "rq-1-1gbps.server:8091", referrer: "http://rq-1-1gbps.server:8091/"
[1584686443] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 89.4% (44065580 kB) free.
[1584686461] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584686512] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(116) < {writer,send_failed,{error,timeout}}
[1584686554] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/03/20 12:11:35 [error] 20710#0: *74027037 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: mrq4-1gbps.server, request: "GET /favicon.ico HTTP/2.0", host: "mrq4-1gbps.server:9443", referrer: "https://mrq4-1gbps.server:9443/"
[1584686565] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < {writer,send_failed,{error,timeout}}
[1584686585] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38771 MB (75% inode=98%):
[1584686642] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584686675] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584686686] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(3) < {writer,send_failed,{error,timeout}}
[1584686718] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4711952 kB) free!
[1584686765] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584686777] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4906604 kB) free!
[1584686789] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584686800] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584686808] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584686859] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584686860] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584686897] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (4502844 kB) free!
[1584686904] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39449380 kB) free.
[1584686915] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1584686925] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2559 processes
[1584686980] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584686985] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584686991] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22026560 kB) free!
[1584687017] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4135516 kB) free!
[1584687027] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584687111] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42237964 kB) free.
[1584687149] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 64.6% (31845592 kB) free.
[1584687202] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584687286] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 64.7% (31900084 kB) free.
[1584687364] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(60) < {writer,send_failed,{error,timeout}}
[1584687407] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(3) < {writer,send_failed,{error,timeout}}
[1584687488] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(9) < {writer,send_failed,{error,timeout}}
[1584687528] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584687611] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584687617] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190440 kB) free!
[1584687649] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;HARD;3;(2) < {writer,send_failed,{error,timeout}}
[1584687711] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6215900 kB) free!
[1584687720] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.1% (21001108 kB) free.
[1584687758] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33081252 kB) free.
[1584687900] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584688022] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.4% (30763188 kB) free.
[1584688110] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9172492 kB) free!
[1584688209] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(43) < {writer,send_failed,{error,timeout}}
[1584688230] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8912360 kB) free!
[1584688324] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170872 kB) free!
[1584688331] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(2) < {writer,send_failed,{error,timeout}}
[1584688351] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (8709024 kB) free!
[1584688368] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12292624 kB) free!
[1584688444] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168824 kB) free!
[1584688452] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(4) < {writer,send_failed,{error,timeout}}
[1584688488] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (47572252 kB) free.
[1584688493] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 78.7% (19418460 kB) free.
[1584688637] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584688765] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.4% (30750380 kB) free.
[1584688950] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186288 kB) free!
[1584689054] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 79.5% (19487640 kB) free.
[1584689088] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5702600 kB) free!
[1584689098] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170500 kB) free!
[1584689148] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 641 processes
[1584689197] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21210004 kB) free.
[1584689208] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397080 kB) free!
[1584689223] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374388 kB) free!
[1584689254] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29490288 kB) free.
[1584689352] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185508 kB) free!
[1584689472] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179908 kB) free!
[1584689483] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 535 processes
[1584689509] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.4% (21671280 kB) free.
[1584689659] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (214336 kB) free!
[1584689732] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.5% (20719608 kB) free.
[1584689859] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9277584 kB) free!
[1584689867] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(300) < 103.255.146.164 [20/Mar/2020:13:00:59 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584689964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21511356 kB) free!
[1584689971] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21301516 kB) free!
[1584689987] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584690050] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584690151] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4580168 kB) free!
[1584690362] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584690396] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194628 kB) free!
[1584690516] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172224 kB) free!
[1584690636] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179276 kB) free!
[1584690938] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187760 kB) free!
[1584691058] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179764 kB) free!
[1584691178] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177176 kB) free!
[1584691179] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (458408 kB) free!
[1584691214] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.3% (21533196 kB) free.
[1584691251] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(4) < {writer,send_failed,{error,timeout}}
[1584691292] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27579388 kB) free.
[1584691320] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192416 kB) free!
[1584691350] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.5% (38674532 kB) free.
[1584691373] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584691393] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.7% (20750744 kB) free.
[1584691433] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191624 kB) free!
[1584691437] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187672 kB) free!
[1584691467] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.6% (21721388 kB) free.
[1584691494] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584691819] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (452072 kB) free!
[1584691883] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 606 processes
[1584691939] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190616 kB) free!
[1584692004] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 620 processes
[1584692059] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169308 kB) free!
[1584692094] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584692125] SERVICE FLAPPING ALERT: sig4-vm2;Total Processes;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1584692125] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 621 processes
[1584692215] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584692269] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.9% (3560544 kB) free!
[1584692388] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(36) < 92.42.104.166 [20/Mar/2020:13:40:20 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584692389] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3073824 kB) free!
[1584692497] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5948284 kB) free!
[1584692508] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584692509] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.5% (2593012 kB) free!
[1584692617] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1865764 kB) free!
[1584692726] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;HARD;3;PROCS WARNING: 563 processes
[1584692737] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392536 kB) free!
[1584692762] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.1983.306> (10.147.212.143:34086 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1584692816] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(12) < {writer,send_failed,{error,timeout}}
[1584692882] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584692937] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(5) < {writer,send_failed,{error,timeout}}
[1584692957] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584693003] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584693058] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;HARD;3;(2) < {writer,send_failed,{error,timeout}}
[1584693108] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(122) < 103.255.146.164 [20/Mar/2020:13:54:50 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584693109] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (191560 kB) free!
[1584693228] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584693327] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 684 processes
[1584693337] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20923976 kB) free!
[1584693602] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(3) < {writer,send_failed,{error,timeout}}
[1584693722] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584693930] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 478 processes
[1584693937] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389396 kB) free!
[1584693941] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (4180136 kB) free!
[1584694061] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3531816 kB) free!
[1584694181] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 11.9% (2929640 kB) free!
[1584694262] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1584694322] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584694383] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1584694443] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584694532] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 616 processes
[1584694656] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 615 processes
[1584694693] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 932 processes
[1584694779] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 621 processes
[1584694781] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202028 kB) free!
[1584694813] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 936 processes
[1584694933] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 932 processes
[1584694944] SERVICE ALERT: qservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7749152 kB) free!
[1584694957] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4779792 kB) free!
[1584694984] SERVICE ALERT: smscapp1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1584695005] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21932768 kB) free.
[1584695029] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(114) < 103.255.146.164 [20/Mar/2020:14:26:50 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584695042] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(7) < {writer,send_failed,{error,timeout}}
[1584695064] SERVICE ALERT: qservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9361392 kB) free!
[1584695072] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28036084 kB) free.
[1584695077] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (4630976 kB) free!
[1584695104] SERVICE ALERT: smscapp1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1584695142] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 66.6% (14956864 kB) free.
[1584695149] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584695162] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584695184] SERVICE ALERT: qservercluster1;Memory;OK;HARD;1;OK - 20.7% (10212408 kB) free.
[1584695197] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4387836 kB) free!
[1584695210] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 63.7% (31388480 kB) free.
[1584695224] SERVICE ALERT: smscapp1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1584695300] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 13.3% (4623084 kB) free!
[1584695379] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 485 processes
[1584695420] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 11.7% (4059724 kB) free!
[1584695540] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (3455468 kB) free!
[1584695610] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180432 kB) free!
[1584695679] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13799104 kB) free!
[1584695705] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (5633872 kB) free!
[1584695730] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170444 kB) free!
[1584695749] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(48) < 103.255.146.164 [20/Mar/2020:14:43:34 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584695799] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9370800 kB) free!
[1584695825] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 14.7% (5139208 kB) free!
[1584695850] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166084 kB) free!
[1584695869] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584695919] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (5322632 kB) free!
[1584695922] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.36, 6.22, 5.39
[1584695945] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 13.4% (4704040 kB) free!
[1584696042] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.39, 6.24, 5.53
[1584696162] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.44, 5.96, 5.54
[1584696197] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (1178272 kB) free!
[1584696233] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1584696267] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166560 kB) free!
[1584696317] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (1039260 kB) free!
[1584696349] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 734 processes
[1584696354] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1584696384] SERVICE ALERT: qservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7188692 kB) free!
[1584696388] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179060 kB) free!
[1584696424] SERVICE ALERT: smscapp1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1584696437] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (927036 kB) free!
[1584696469] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 730 processes
[1584696469] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 172.105.59.173 [20/Mar/2020:14:52:06 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1584696474] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1584696504] SERVICE ALERT: qservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.9% (6864104 kB) free!
[1584696508] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186000 kB) free!
[1584696519] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22993656 kB) free!
[1584696545] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1584696545] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (2279868 kB) free!
[1584696557] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6552668 kB) free!
[1584696589] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 727 processes
[1584696590] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584696625] SERVICE ALERT: qservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4498292 kB) free!
[1584696659] SERVICE FLAPPING ALERT: rq-1;rq1-rabbitmqerror;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1584696659] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584696678] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.8% (30968256 kB) free.
[1584696740] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.4% (32196236 kB) free.
[1584696913] SERVICE ALERT: sigqscluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697034] SERVICE ALERT: sigqscluster2;Total Processes;OK;HARD;1;PROCS OK: 302 processes
[1584697119] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697119] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697119] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395284 kB) free!
[1584697122] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697145] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.0% (31513148 kB) free.
[1584697189] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 103.13.97.241 [20/Mar/2020:15:02:01 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1584697189] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697225] SERVICE ALERT: qservercluster1;Memory;WARNING;HARD;3;WARNING - 12.6% (6201848 kB) free!
[1584697249] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697249] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697252] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697287] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697309] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584697315] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584697379] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697379] SERVICE ALERT: smscapp2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697382] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697417] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697547] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697600] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.30, 5.10, 12.71
[1584697618] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7909368 kB) free!
[1584697674] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584697702] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22865168 kB) free.
[1584697730] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697738] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.7% (7755728 kB) free!
[1584697790] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697858] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.4% (7607312 kB) free!
[1584697860] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697860] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584697909] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(58) < 103.255.146.164 [20/Mar/2020:15:19:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584697919] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697920] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697962] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.70, 4.32, 4.72
[1584697981] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584697990] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584697991] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584698029] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584698049] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584698050] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584698121] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584698179] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584698251] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584698425] SERVICE ALERT: qservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4301864 kB) free!
[1584698458] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3725632 kB) free!
[1584698598] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2670 processes
[1584698801] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 75.24, 59.46, 37.38
[1584698981] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 559 processes
[1584698982] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.4% (18606240 kB) free.
[1584699039] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23464864 kB) free!
[1584699102] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 564 processes
[1584699153] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 58.1% (28644060 kB) free.
[1584699162] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.95, 5.45, 5.05
[1584699208] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584699223] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;HARD;3;PROCS WARNING: 559 processes
[1584699228] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.5% (21436556 kB) free.
[1584699230] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(83) < 103.255.146.164 [20/Mar/2020:15:41:54 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584699281] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.2% (21854300 kB) free.
[1584699282] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.79, 6.76, 5.62
[1584699338] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584699350] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584699403] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.08, 6.98, 5.83
[1584699468] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584699589] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173076 kB) free!
[1584699645] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9473168 kB) free!
[1584699651] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38800460 kB) free.
[1584699689] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1584699710] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180292 kB) free!
[1584699735] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 24 users currently logged in
[1584699765] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (9250088 kB) free!
[1584699810] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1584699823] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 501 processes
[1584699830] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191156 kB) free!
[1584699855] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 24 users currently logged in
[1584699885] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (9050248 kB) free!
[1584699887] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180896 kB) free!
[1584699930] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1584699975] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 24 users currently logged in
[1584700007] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195060 kB) free!
[1584700022] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178236 kB) free!
[1584700074] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584700127] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171664 kB) free!
[1584700143] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192048 kB) free!
[1584700194] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1584700225] SERVICE FLAPPING ALERT: qservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1584700225] SERVICE ALERT: qservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (5263420 kB) free!
[1584700251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (405036 kB) free!
[1584700263] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186128 kB) free!
[1584700314] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1584700371] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407636 kB) free!
[1584700381] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < {writer,send_failed,{error,timeout}}
[1584700491] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394408 kB) free!
[1584700503] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584700530] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1584700562] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(5) < {writer,send_failed,{error,timeout}}
[1584700603] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1584700603] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.02, 4.48, 4.84
[1584700683] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584700825] SERVICE ALERT: qservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2484468 kB) free!
[1584700863] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 79.3% (19441076 kB) free.
[1584700914] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584700962] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8931396 kB) free!
[1584701034] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584701052] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584701082] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.5% (7632584 kB) free!
[1584701155] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584701202] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (8066304 kB) free!
[1584701363] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 9.9% (2426244 kB) free!
[1584701483] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 9.8% (2410104 kB) free!
[1584701603] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.8% (2391172 kB) free!
[1584701623] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 590 processes
[1584701685] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4675372 kB) free!
[1584701743] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 593 processes
[1584701751] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(110) < 103.255.146.164 [20/Mar/2020:16:18:25 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584701755] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1584701755] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1584701864] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;HARD;3;PROCS WARNING: 593 processes
[1584701871] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584702162] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1584702203] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 37.4% (9164236 kB) free.
[1584702282] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1584702330] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584702402] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1584702402] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3345444 kB) free!
[1584702450] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584702482] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(10) < Channel error on connection <0.25803.325> (10.147.212.158:43944 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1584702570] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1584702603] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584702625] SERVICE ALERT: qservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (5012488 kB) free!
[1584702663] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179096 kB) free!
[1584702783] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171992 kB) free!
[1584702839] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167572 kB) free!
[1584702903] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (209072 kB) free!
[1584702959] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192724 kB) free!
[1584702975] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 26 users currently logged in
[1584702979] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584703051] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38769 MB (75% inode=98%):
[1584703069] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1584703069] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2216 processes
[1584703071] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(141) < 103.255.146.164 [20/Mar/2020:16:44:28 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584703079] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188920 kB) free!
[1584703171] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1584703171] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1584703189] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 675 processes
[1584703191] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584703225] SERVICE ALERT: qservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3598504 kB) free!
[1584703297] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 45.51, 35.99, 36.17
[1584703333] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 353 processes
[1584703379] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 109803 MB (48% inode=99%):
[1584703379] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1584703382] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2354 processes
[1584703410] SERVICE ALERT: qservercluster1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10689 MB (20% inode=98%):
[1584703450] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584703452] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584703491] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15876320 kB) free!
[1584703502] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(4) < {writer,send_failed,{error,timeout}}
[1584703530] SERVICE ALERT: qservercluster1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 10570 MB (20% inode=98%):
[1584703547] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 50.0% (24641016 kB) free.
[1584703549] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584703573] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584703623] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584703650] SERVICE ALERT: qservercluster1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 10498 MB (20% inode=98%):
[1584703693] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584703792] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [20/Mar/2020:16:53:15 +0530] TCP 502 0 0 60.001 "10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 60.000, 0.000"
[1584703803] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584703912] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584703923] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584704074] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/03/20 17:00:36 [error] 20710#0: *74292387 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: 148qservercluster1.com, request: "DELETE /rabbitmq/api/queues/%2F/PDC/contents HTTP/1.1", upstream: "http://10.147.212.41:15672/api/queues/%2F/PDC/contents", host: "148qservercluster1.com:8091", referrer: "http://148qservercluster1.com:8091/rabbitmq/"
[1584704091] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407376 kB) free!
[1584704103] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1584704103] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 46.3% (11351956 kB) free.
[1584704149] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (3991024 kB) free!
[1584704175] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 24 users currently logged in
[1584704194] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584704265] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 681 processes
[1584704269] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3593816 kB) free!
[1584704389] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3175368 kB) free!
[1584704523] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584704644] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584704764] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;3;(2) < {writer,send_failed,{error,timeout}}
[1584704850] SERVICE ALERT: qservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 10765 MB (21% inode=98%):
[1584704899] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9787268 kB) free!
[1584704968] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21958388 kB) free.
[1584704989] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (976092 kB) free!
[1584704990] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1584705013] SERVICE ALERT: qservercluster1;Memory;OK;HARD;1;OK - 25.3% (12455268 kB) free.
[1584705019] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9549136 kB) free!
[1584705053] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 22.7% (11204668 kB) free.
[1584705098] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.42, 43.44, 41.85
[1584705110] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1584705112] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42237188 kB) free.
[1584705139] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9267196 kB) free!
[1584705153] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1584705153] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.9% (15239200 kB) free.
[1584705187] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 47.9% (23611508 kB) free.
[1584705214] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.8% (21765880 kB) free.
[1584705230] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1584705248] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.9% (17624748 kB) free.
[1584705375] SERVICE FLAPPING ALERT: secondary-server;Current Users;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1584705375] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 28 users currently logged in
[1584705394] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/03/20 17:24:50 [error] 20714#0: *74317628 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: 148qservercluster1.com, request: "DELETE /rabbitmq/api/queues/%2F/PDC/contents HTTP/1.1", upstream: "http://10.147.212.41:15672/api/queues/%2F/PDC/contents", host: "148qservercluster1.com:8091", referrer: "http://148qservercluster1.com:8091/rabbitmq/"
[1584705424] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584705494] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(64) < {writer,send_failed,{error,timeout}}
[1584705514] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584705545] SERVICE FLAPPING ALERT: rq-1;rq1-rabbitmqerror;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1584705545] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584705614] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(25) < {writer,send_failed,{error,timeout}}
[1584705718] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (427876 kB) free!
[1584705735] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(1) < {writer,send_failed,{error,timeout}}
[1584705830] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1584705830] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 689 processes
[1584705838] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (376576 kB) free!
[1584705958] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416984 kB) free!
[1584706103] SERVICE ALERT: db1-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (2093884 kB) free!
[1584706114] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/03/20 17:31:23 [error] 20714#0: *74321868 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: 148qservercluster1.com, request: "POST /rabbitmq/api/queues/%2F/PDC/get HTTP/1.1", upstream: "http://10.147.212.41:15672/api/queues/%2F/PDC/get", host: "148qservercluster1.com:8091", referrer: "http://148qservercluster1.com:8091/rabbitmq/"
[1584706132] SERVICE ALERT: qservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584706145] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(4) < {writer,send_failed,{error,timeout}}
[1584706176] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (695104 kB) free!
[1584706220] SERVICE ALERT: qservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (609916 kB) free!
[1584706223] SERVICE ALERT: db1-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (1513684 kB) free!
[1584706234] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(142) < 2020/03/20 17:40:32 [error] 20714#0: *74333305 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1584706252] SERVICE ALERT: qservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1584706267] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584706296] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (2207260 kB) free!
[1584706313] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4886) < 92.42.104.166 [20/Mar/2020:17:41:52 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584706340] SERVICE ALERT: qservercluster1;Memory;OK;HARD;1;OK - 97.8% (48176944 kB) free.
[1584706343] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (929408 kB) free!
[1584706354] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(2373) < 2020/03/20 17:42:34 [error] 20707#0: *74337578 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584706388] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;HARD;3;(1) < {writer,send_failed,{error,timeout}}
[1584706390] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.8% (6314236 kB) free!
[1584706416] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2153172 kB) free!
[1584706433] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(8942) < 172.105.59.173 [20/Mar/2020:17:43:53 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584706510] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.4% (6133864 kB) free!
[1584706553] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(7514) < 172.105.59.173 [20/Mar/2020:17:45:53 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584706556] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5872) < 103.16.101.72 [20/Mar/2020:17:45:55 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1584706630] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (5005432 kB) free!
[1584706678] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(528) < 103.16.101.72 [20/Mar/2020:17:46:25 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1584706800] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584706940] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (5456400 kB) free!
[1584707061] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (4987828 kB) free!
[1584707158] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22400244 kB) free!
[1584707181] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 12.9% (4494904 kB) free!
[1584707230] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4136500 kB) free!
[1584707539] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4821848 kB) free!
[1584707554] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584707660] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584707754] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584707757] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584707758] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397520 kB) free!
[1584707765] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584707781] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.6% (32245852 kB) free.
[1584707789] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584707790] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584707861] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584707865] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 404 processes
[1584707879] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584707887] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584707915] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584707919] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584707991] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584708009] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584708017] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584708049] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584708121] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584708136] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2660 processes
[1584708141] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584708189] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584708303] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.7% (4820744 kB) free!
[1584708423] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.7% (4592844 kB) free!
[1584708520] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584708543] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.0% (4411276 kB) free!
[1584708546] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584708608] SERVICE ALERT: sigqscluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584708650] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584708721] SERVICE FLAPPING ALERT: smscapp2;HDD_Root status;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1584708721] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38769 MB (75% inode=98%):
[1584708728] SERVICE ALERT: sigqscluster2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1584708743] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.9% (22929700 kB) free.
[1584708771] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584708958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27772728 kB) free.
[1584709218] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1584709218] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 77.2% (38038824 kB) free.
[1584709249] SERVICE FLAPPING ALERT: smscapp2;Sensors Status;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1584709249] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584709346] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(3) < {writer,send_failed,{error,timeout}}
[1584709354] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(19) < 2020/03/20 18:27:48 [error] 20713#0: *74383503 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1584709391] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(3) < {writer,send_failed,{error,timeout}}
[1584709471] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584709474] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584709512] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584709535] SERVICE ALERT: sigqscluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584709545] SERVICE ALERT: sigqscluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584709554] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(29) < 172.105.59.173 [20/Mar/2020:18:29:19 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 60.000, 60.000"
[1584709555] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584709558] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (421412 kB) free!
[1584709601] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.16, 5.19, 13.92
[1584709655] SERVICE ALERT: sigqscluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 42473 MB (82% inode=99%):
[1584709666] SERVICE ALERT: sigqscluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584709674] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584709675] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584709678] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (438124 kB) free!
[1584709743] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2284212 kB) free!
[1584709795] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584709798] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392904 kB) free!
[1584710062] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1584710068] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(2) < {writer,send_failed,{error,timeout}}
[1584710165] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.12511.348> (10.147.212.158:44578 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1584710193] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584710275] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(219) < 103.255.146.164 [20/Mar/2020:18:46:16 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584710286] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(9) < Channel error on connection <0.13897.348> (10.147.212.158:45702 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1584710395] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584710406] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;3;(1) < Channel error on connection <0.30561.318> (10.147.212.158:45718 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1584710538] SERVICE ALERT: sigqscluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584710582] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 342 processes
[1584710658] SERVICE ALERT: sigqscluster2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1584710702] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 346 processes
[1584710712] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584710822] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1584710833] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584710954] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584711006] SERVICE FLAPPING ALERT: rq-2;rq2-rabbitmqerror;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1584711006] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584711052] SERVICE ALERT: qservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1584711172] SERVICE ALERT: qservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 331 processes
[1584711219] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198668 kB) free!
[1584711263] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170548 kB) free!
[1584711292] SERVICE ALERT: qservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 330 processes
[1584711340] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170976 kB) free!
[1584711383] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182584 kB) free!
[1584711390] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(92) < {writer,send_failed,{error,timeout}}
[1584711401] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 64.27, 28.83, 15.36
[1584711403] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584711460] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184888 kB) free!
[1584711466] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 655 processes
[1584711503] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198848 kB) free!
[1584711511] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584711555] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(5) < {writer,send_failed,{error,timeout}}
[1584711586] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 659 processes
[1584711676] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584711707] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 654 processes
[1584711730] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1584712022] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 3 warn out of 337 processes
[1584712110] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(20) < {writer,send_failed,{error,timeout}}
[1584712142] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1584712231] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(13) < {writer,send_failed,{error,timeout}}
[1584712277] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584712352] SERVICE FLAPPING ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1584712352] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(34) < {writer,send_failed,{error,timeout}}
[1584712398] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584712575] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 25 users currently logged in
[1584712602] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1584712798] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (47649748 kB) free.
[1584712952] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584712999] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(3) < {writer,send_failed,{error,timeout}}
[1584713120] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584713398] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (421324 kB) free!
[1584713518] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397072 kB) free!
[1584713638] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375364 kB) free!
[1584713720] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(4) < {writer,send_failed,{error,timeout}}
[1584713841] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584714113] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2986732 kB) free!
[1584714140] SERVICE FLAPPING ALERT: qservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584714152] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(47) < {writer,send_failed,{error,timeout}}
[1584714233] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3398508 kB) free.
[1584714273] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584714390] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181136 kB) free!
[1584714394] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(79) < {writer,send_failed,{error,timeout}}
[1584714442] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584714510] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (157704 kB) free!
[1584714563] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584714571] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584714630] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (580960 kB) free!
[1584714759] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.5% (3037852 kB) free!
[1584714879] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.7% (2623592 kB) free!
[1584714999] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2245480 kB) free!
[1584715164] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(75) < {writer,send_failed,{error,timeout}}
[1584715285] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584715416] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584715742] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1584715795] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584715859] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584715862] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1584715886] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584715980] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584716007] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584716195] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584716543] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.9% (3684740 kB) free!
[1584716608] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < {writer,send_failed,{error,timeout}}
[1584716663] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 12.9% (3177040 kB) free!
[1584716729] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584716783] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.8% (2674688 kB) free!
[1584717230] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584717238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37157160 kB) free.
[1584717381] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (6671888 kB) free!
[1584717383] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1584717383] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (266292 kB) free!
[1584717396] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(2) < {writer,send_failed,{error,timeout}}
[1584717460] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584717501] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 17.9% (6246028 kB) free!
[1584717516] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584717524] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 55.0% (13477576 kB) free.
[1584717558] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 76.1% (18663664 kB) free.
[1584717621] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 16.6% (5798620 kB) free!
[1584717633] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 63.2% (14187916 kB) free.
[1584717637] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584717659] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 43.5% (21455944 kB) free.
[1584717684] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 23.1% (11378716 kB) free.
[1584717710] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.3% (21161428 kB) free.
[1584717767] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.7% (14085124 kB) free.
[1584717795] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 61.0% (15044508 kB) free.
[1584717838] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (383300 kB) free!
[1584717958] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (431728 kB) free!
[1584718078] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398548 kB) free!
[1584718221] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584718221] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.4% (32187836 kB) free.
[1584718237] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(39) < {writer,send_failed,{error,timeout}}
[1584718339] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718357] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1584718380] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718389] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718389] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718389] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718392] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718469] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718479] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(2) < {writer,send_failed,{error,timeout}}
[1584718505] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584718511] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1584718512] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 69797 MB (31% inode=99%):
[1584718518] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2756 processes
[1584718520] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718530] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584718575] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584718599] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718650] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718651] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584718771] SERVICE ALERT: sigqscluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718828] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718891] SERVICE ALERT: sigqscluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1584718931] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584718958] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719061] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719088] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719093] SERVICE ALERT: qservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1584719110] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719122] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719183] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.0% (22962312 kB) free.
[1584719191] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719240] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584719252] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719252] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584719256] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584719278] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23357520 kB) free!
[1584719374] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584719382] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719467] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6195400 kB) free!
[1584719533] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 625 processes
[1584719587] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.3% (5549168 kB) free!
[1584719598] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176148 kB) free!
[1584719633] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (2737440 kB) free!
[1584719653] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1018 processes
[1584719685] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584719707] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3831736 kB) free!
[1584719718] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168588 kB) free!
[1584719753] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (3161740 kB) free!
[1584719773] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1584719803] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.73, 13.86, 8.51
[1584719838] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191512 kB) free!
[1584719840] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719866] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584719873] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3490148 kB) free.
[1584719878] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412812 kB) free!
[1584719923] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.76, 14.89, 9.68
[1584719960] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584719982] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 1739 processes
[1584719991] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584720044] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.17, 11.61, 9.08
[1584720107] SERVICE FLAPPING ALERT: sig4-vm2;Total Processes;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1584720289] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 88.8% (43749936 kB) free.
[1584720292] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(4) < {writer,send_failed,{error,timeout}}
[1584720391] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38769 MB (75% inode=98%):
[1584720401] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1725 processes
[1584720417] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584720644] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 46.48, 33.90, 20.84
[1584720660] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4465784 kB) free!
[1584720781] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 24.0% (11844032 kB) free.
[1584721078] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30079488 kB) free.
[1584721380] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (2140648 kB) free!
[1584721500] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (1826504 kB) free!
[1584721620] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (7223724 kB) free!
[1584721643] SERVICE ALERT: sigqscluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584721673] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2713936 kB) free!
[1584721678] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (392852 kB) free!
[1584721763] SERVICE ALERT: sigqscluster2;Total Processes;OK;HARD;1;PROCS OK: 303 processes
[1584721767] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584721793] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.7% (5158284 kB) free.
[1584721798] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397400 kB) free!
[1584721898] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.94, 10.21, 33.60
[1584721918] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381972 kB) free!
[1584722407] SERVICE FLAPPING ALERT: rq-2;rq2-rabbitmqerror;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584722444] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.81, 12.76, 17.44
[1584722498] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 12.51, 11.65, 22.90
[1584722507] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 438 processes
[1584722518] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39947788 kB) free.
[1584722586] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.3% (18346144 kB) free.
[1584722620] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 67.0% (33009308 kB) free.
[1584722674] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/03/20 22:11:13 [error] 20714#0: *74522105 upstream prematurely closed connection while reading response header from upstream, client: 183.82.105.122, server: 148qservercluster1.com, request: "DELETE /rabbitmq/api/queues/%2F/PDC/contents HTTP/1.1", upstream: "http://10.147.212.41:15672/api/queues/%2F/PDC/contents", host: "148qservercluster1.com:8091", referrer: "http://148qservercluster1.com:8091/rabbitmq/"
[1584722795] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584722801] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.91, 2.31, 11.98
[1584722813] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584722820] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4029908 kB) free!
[1584722934] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584723062] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1584723118] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3517064 kB) free!
[1584723182] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1584723188] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168060 kB) free!
[1584723238] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393204 kB) free!
[1584723308] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165716 kB) free!
[1584723358] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398520 kB) free!
[1584723428] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1584723428] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168584 kB) free!
[1584723593] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3145392 kB) free!
[1584723713] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3476392 kB) free.
[1584723828] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230364 kB) free!
[1584723948] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244660 kB) free!
[1584724001] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.83, 0.93, 4.06
[1584724068] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231504 kB) free!
[1584725439] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584725444] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1584725444] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.28, 2.39, 4.25
[1584725758] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13250696 kB) free!
[1584725775] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 19 users currently logged in
[1584726113] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3181252 kB) free!
[1584726233] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3772744 kB) free.
[1584726358] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (373588 kB) free!
[1584726958] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23787224 kB) free!
[1584727433] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3188396 kB) free!
[1584727553] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3704264 kB) free.
[1584727558] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402004 kB) free!
[1584727767] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.8% (4599884 kB) free!
[1584727887] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.2% (4457992 kB) free!
[1584728007] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.5% (4280516 kB) free!