[1558981984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558982104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558982173] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2460392 kB) free!
[1558982224] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558982825] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558982945] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558983545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558983586] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558983665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558983706] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558983785] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558983826] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558983973] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.3% (3025236 kB) free!
[1558984385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558984505] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558985105] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558985226] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558985825] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558985945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558986065] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558986665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558986785] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558987385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558987505] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558987534] SERVICE ALERT: redis-1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1558987537] HOST ALERT: redis-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.201)
[1558987547] SERVICE ALERT: redis-1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-1' was out of bounds)
[1558987612] SERVICE ALERT: redis-1;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'redis-1' was out of bounds)
[1558987612] HOST ALERT: redis-2;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1558987647] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'redis-2' was out of bounds)
[1558987657] SERVICE ALERT: redis-1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-1' was out of bounds)
[1558987727] SERVICE ALERT: redis-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'redis-2' was out of bounds)
[1558987727] HOST ALERT: redis-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1558987732] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 963130 MB (90% inode=99%):
[1558987767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 94.2% (15285476 kB) free.
[1558987771] HOST ALERT: redis-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1558987847] SERVICE ALERT: redis-2;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.42, 0.19
[1558988026] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558988106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558988226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558988254] SERVICE ALERT: redis-1;Sensors Status;OK;HARD;1;SENSORS OK
[1558988259] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558988267] SERVICE ALERT: redis-1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 44702 MB (87% inode=99%):
[1558988346] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1558988946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558988974] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3368520 kB) free!
[1558989066] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558989574] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5897784 kB) free!
[1558989666] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558989748] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1558989786] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558989826] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558989869] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1558989946] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558989989] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1558990066] SERVICE FLAPPING ALERT: secondary-server;Current Users;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1558990066] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558990218] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (1019664 kB) free!
[1558990338] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (1019880 kB) free!
[1558990386] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1558990458] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1020208 kB) free!
[1558993386] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558993506] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558993626] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558994102] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1558994222] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558994226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558994259] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.52, 9.33, 5.56
[1558994347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558994379] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.47, 6.35, 4.93
[1558994467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558995667] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558996067] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558996867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558996987] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558997587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558997708] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558998307] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558998427] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558999027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558999147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558999267] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558999867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558999987] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559000588] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559000708] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559000974] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559001308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559001428] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559002067] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1559002582] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.07, 6.58, 5.58
[1559002628] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559002702] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.56, 4.89, 5.09
[1559002748] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559002822] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.59, 3.80, 4.66
[1559003348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559003374] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4555316 kB) free!
[1559003469] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559003658] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.1% (4070860 kB) free.
[1559003789] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559003909] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559004068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559004188] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559004308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559004574] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.7% (9698956 kB) free!
[1559004973] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2439316 kB) free!
[1559005573] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2488412 kB) free!
[1559006108] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1559006108] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559006173] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452176 kB) free!
[1559006709] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559006829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559006949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559007575] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3917516 kB) free!
[1559008749] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559008775] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1559008775] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.4% (9554968 kB) free!
[1559009349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559009375] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4116448 kB) free!
[1559009469] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559010575] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6150648 kB) free!
[1559010670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559010790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559010910] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559011134] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 258.99, 109.61, 41.42
[1559011222] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 271.19, 196.02, 92.42
[1559011254] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 280.94, 167.42, 70.87
[1559011342] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 271.88, 221.12, 114.17
[1559011374] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 283.97, 206.12, 96.76
[1559011462] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 179.68, 216.19, 126.11
[1559011510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559011630] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559011631] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2873624 kB) free!
[1559012230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559012279] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.8% (2881228 kB) free!
[1559012350] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559012909] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1559012950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559012975] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4349800 kB) free!
[1559012979] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.72, 7.38, 5.22
[1559013029] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1559013071] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559013099] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 5.02, 4.61
[1559013575] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6402296 kB) free!
[1559013670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559013699] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.17, 9.24, 7.68
[1559013774] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.36, 2.16, 12.82
[1559013790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559013820] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.31, 6.25, 6.77
[1559013862] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.18, 3.01, 14.92
[1559013910] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559013940] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 4.30, 5.99
[1559014510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559014540] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.78, 3.24
[1559014630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559014750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559014974] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 0.47, 3.84
[1559015150] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1559015157] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.6% (15048708 kB) free.
[1559015204] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.7% (21983240 kB) free.
[1559015247] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.1% (14814560 kB) free.
[1559015270] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1559015303] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.4% (23541440 kB) free.
[1559015351] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 45.5% (22412936 kB) free.
[1559015389] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1559015427] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.9% (22762236 kB) free.
[1559015509] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1559015662] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.50, 1.57, 3.48
[1559016029] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559016149] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559016655] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1559017150] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559017551] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559017671] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559017713] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (437848 kB) free!
[1559017750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559017791] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559017833] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (434420 kB) free!
[1559017870] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559017953] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (425444 kB) free!
[1559017990] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559018140] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.04, 8.24, 6.20
[1559018260] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.45, 5.65, 5.49
[1559018380] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.27, 3.87, 4.86
[1559019055] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1559019190] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559019790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559019910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559020030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559021831] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559022431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559022551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559022671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559023614] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181792 kB) free!
[1559023734] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178256 kB) free!
[1559023833] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189932 kB) free!
[1559023854] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166720 kB) free!
[1559023871] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559023953] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184320 kB) free!
[1559024073] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187848 kB) free!
[1559024471] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559024591] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559024658] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 628 processes
[1559024712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559024718] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 626 processes
[1559025109] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559025149] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559025229] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559025251] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.6% (21478760 kB) free.
[1559025269] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559025302] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.3% (22130356 kB) free.
[1559025369] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.5% (21843784 kB) free.
[1559025551] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559025869] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1559025912] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1559025970] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197460 kB) free!
[1559025990] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 336 processes
[1559026090] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179124 kB) free!
[1559026110] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1559026211] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188848 kB) free!
[1559026512] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559027270] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 349 processes
[1559027310] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 337 processes
[1559027390] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559027762] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169308 kB) free!
[1559027867] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559027882] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178472 kB) free!
[1559027910] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1559027951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9416360 kB) free!
[1559027987] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559028002] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176628 kB) free!
[1559028071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.9% (16192772 kB) free.
[1559028107] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559028590] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559028710] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559028830] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559029110] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1559029110] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 336 processes
[1559029431] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559029710] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1559030662] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 257.91, 104.00, 39.75
[1559030782] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 286.69, 165.96, 70.24
[1559030852] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.6% (2529456 kB) free!
[1559030902] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 290.42, 207.32, 97.01
[1559030910] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 336 processes
[1559030972] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.7% (2387108 kB) free!
[1559031072] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8642396 kB) free!
[1559031092] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.9% (2251636 kB) free!
[1559031192] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.1% (15822784 kB) free.
[1559031258] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (2800484 kB) free!
[1559031378] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.3% (2800436 kB) free!
[1559031498] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.6% (2852940 kB) free!
[1559031692] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1514948 kB) free!
[1559032098] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.8% (3702676 kB) free.
[1559032264] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.54, 5.49, 5.04
[1559032384] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.88, 5.62, 5.14
[1559032504] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.86, 6.10, 5.38
[1559032699] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (707756 kB) free!
[1559032819] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.2% (4579188 kB) free.
[1559032937] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.6% (998996 kB) free!
[1559033030] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559033057] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 11.4% (838044 kB) free!
[1559033104] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.22, 3.93, 4.68
[1559033150] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559033177] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (690864 kB) free!
[1559033303] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.55, 2.04, 14.96
[1559033419] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.8% (2720220 kB) free!
[1559033539] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.6% (3668028 kB) free.
[1559033592] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9721420 kB) free!
[1559033712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.9% (16192264 kB) free.
[1559033712] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559033750] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559033870] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559033910] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 339 processes
[1559034139] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.6% (2524676 kB) free!
[1559034259] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.4% (3641076 kB) free.
[1559034313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559034433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559034553] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559034655] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1559034775] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1559034832] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 361 processes
[1559034896] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1559034952] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559035103] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.85, 1.95, 3.84
[1559035710] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1559036051] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176732 kB) free!
[1559036059] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (1603916 kB) free!
[1559036120] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177428 kB) free!
[1559036152] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 358 processes
[1559036171] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184084 kB) free!
[1559036177] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.4% (6870512 kB) free.
[1559036179] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (3317352 kB) free.
[1559036240] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185176 kB) free!
[1559036272] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 3 warn out of 364 processes
[1559036291] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199208 kB) free!
[1559036310] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1559036353] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559036361] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189896 kB) free!
[1559036392] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 364 processes
[1559036696] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1559036916] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 586 processes
[1559036953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559036980] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.52, 8.55, 5.28
[1559037036] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 336 processes
[1559037073] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559037100] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.58, 5.84, 4.69
[1559037194] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559037312] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8474696 kB) free!
[1559037370] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.9% (2748108 kB) free!
[1559037379] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (493964 kB) free!
[1559037433] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.8% (15675132 kB) free.
[1559037490] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2597552 kB) free!
[1559037499] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.6% (2199896 kB) free!
[1559037510] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 340 processes
[1559037592] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1559037592] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1559037610] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.1% (2455908 kB) free!
[1559037619] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.6% (3993420 kB) free.
[1559038192] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1559038210] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1497136 kB) free!
[1559038313] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1559038394] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559038820] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (903216 kB) free!
[1559038912] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 364 processes
[1559038940] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.2% (2784308 kB) free!
[1559038994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559039032] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 367 processes
[1559039060] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (3760088 kB) free.
[1559039114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559039118] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 637 processes
[1559039152] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 364 processes
[1559039178] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 613 processes
[1559039234] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559039391] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1559039696] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1559039832] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8784124 kB) free!
[1559039952] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.8% (15671352 kB) free.
[1559040174] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 238.45, 123.21, 50.04
[1559040260] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (1341416 kB) free!
[1559040294] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 245.46, 164.02, 73.87
[1559040296] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1559040380] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.7% (3849128 kB) free.
[1559040414] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 227.02, 187.35, 93.52
[1559040416] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1559040536] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1559040536] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1559040611] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2604424 kB) free!
[1559040980] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (798764 kB) free!
[1559041034] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559041070] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559041100] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3139208 kB) free!
[1559041190] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559041212] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1739564 kB) free!
[1559041220] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.5% (4457196 kB) free.
[1559041310] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559041504] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.21, 5.65, 5.23
[1559041552] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 358 processes
[1559041624] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.36, 5.78, 5.35
[1559041635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559041744] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.41, 5.75, 5.37
[1559041755] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559041820] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.9% (2416684 kB) free!
[1559041875] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559041910] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559041940] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.0% (4378744 kB) free.
[1559042030] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559042152] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 361 processes
[1559042215] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 1.03, 13.95
[1559042344] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1559042344] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.75, 4.22, 4.75
[1559042540] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.0% (2440944 kB) free!
[1559042630] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559042660] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.3% (4435396 kB) free.
[1559042750] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559042952] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9420100 kB) free!
[1559043072] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7034884 kB) free!
[1559043192] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.6% (14599712 kB) free.
[1559043260] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.6% (2699196 kB) free!
[1559043352] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 359 processes
[1559043380] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (3549060 kB) free.
[1559043415] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.20, 2.03, 4.99
[1559043504] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 39.77, 14.50, 6.33
[1559043624] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 48.35, 26.16, 11.60
[1559043744] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.92, 34.06, 16.26
[1559043873] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5547 MB (10% inode=97%):
[1559043950] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559043952] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559043977] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 12.2% (900656 kB) free!
[1559043981] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171344 kB) free!
[1559044070] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559044097] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 11.0% (806392 kB) free!
[1559044101] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.1% (1957740 kB) free!
[1559044217] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (689956 kB) free!
[1559044221] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.4% (2495316 kB) free!
[1559044344] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1559044344] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.63, 15.68, 17.85
[1559044473] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6872 MB (13% inode=98%):
[1559044736] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 7 users currently logged in
[1559044821] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1559044821] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (3305300 kB) free.
[1559044992] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9403296 kB) free!
[1559045112] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7013676 kB) free!
[1559045232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.0% (14795072 kB) free.
[1559045270] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559045344] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.75, 5.99, 5.28
[1559045390] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559045421] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.3% (2971876 kB) free!
[1559045464] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.94, 5.89, 5.33
[1559045541] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174684 kB) free!
[1559045584] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.89, 5.16, 5.14
[1559045661] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.7% (1730048 kB) free!
[1559045752] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1559045872] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1559045992] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1559046144] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.73, 1.94, 4.31
[1559046261] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1517136 kB) free!
[1559046536] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1559046592] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559046610] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 86.6% (14082352 kB) free.
[1559046656] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1559046776] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1559046862] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (1753036 kB) free!
[1559047218] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.8% (6900720 kB) free.
[1559047275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559047462] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1119980 kB) free!
[1559047632] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9686184 kB) free!
[1559047752] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7410524 kB) free!
[1559047872] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.7% (14637364 kB) free.
[1559047875] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1559047984] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.86, 3.25, 4.11
[1559047995] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559048115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559048310] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559048392] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1559048512] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1559048632] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559048662] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.1% (3424048 kB) free.
[1559049262] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (3057772 kB) free!
[1559049382] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161148 kB) free!
[1559049502] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.8% (1912436 kB) free!
[1559050102] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (554068 kB) free!
[1559050190] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559050310] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559050515] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1559050872] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9580348 kB) free!
[1559050992] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7357192 kB) free!
[1559051112] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.8% (14194092 kB) free.
[1559051115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559051235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559051355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559052144] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 298.76, 156.42, 63.90
[1559052264] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 306.62, 206.49, 93.46
[1559052384] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 307.71, 240.08, 119.48
[1559053155] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559053307] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1559053310] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559053430] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559053512] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9525416 kB) free!
[1559053632] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7302208 kB) free!
[1559053752] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.6% (14089860 kB) free.
[1559053756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559053876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559053996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559054630] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559054750] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559054784] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.34, 1.71, 14.77
[1559055350] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559055470] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559055796] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559056070] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1559056152] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9033184 kB) free!
[1559056190] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559056272] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.3% (14450300 kB) free.
[1559056396] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559056516] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1559056584] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.91, 1.83, 3.68
[1559056636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559056976] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559057302] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.9% (1762252 kB) free!
[1559057390] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559057510] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559057584] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1559057815] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.52, 15.00, 13.25
[1559057836] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559057935] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.46, 10.01, 11.63
[1559058055] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 6.78, 10.25
[1559058100] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 247.92, 132.67, 53.87
[1559058110] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559058220] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 253.40, 172.90, 78.17
[1559058230] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559058233] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559058340] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 101.60, 157.12, 84.92
[1559058436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559058491] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2776704 kB) free!
[1559058556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559058676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559059255] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 0.45, 3.11
[1559059272] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8318868 kB) free!
[1559059392] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.0% (14275820 kB) free.
[1559059430] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559059550] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559059876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559060140] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.21, 0.70, 12.57
[1559060150] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559060270] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559060477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559060597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559060717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559060902] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1329892 kB) free!
[1559061340] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.28, 3.62
[1559061470] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559061590] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559061793] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9264552 kB) free!
[1559061913] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.2% (14381824 kB) free.
[1559061917] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559062103] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.7% (4168004 kB) free.
[1559062190] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559062310] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559062517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559062637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559062757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559062910] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559063030] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559063303] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.3% (2484684 kB) free!
[1559063423] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.5% (4299924 kB) free.
[1559063631] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559063751] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1559063957] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559064023] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3120580 kB) free!
[1559064143] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.0% (4376880 kB) free.
[1559064313] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9815332 kB) free!
[1559064351] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559064433] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7688660 kB) free!
[1559064471] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559064553] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.7% (13672968 kB) free.
[1559064557] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559064677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559064743] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (2790408 kB) free!
[1559064797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559064863] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.5% (4468576 kB) free.
[1559065671] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1559065791] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559065997] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559066063] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (998172 kB) free!
[1559066183] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.2% (2792488 kB) free!
[1559066303] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.0% (4535248 kB) free.
[1559066391] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559066511] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559066598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559066718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559066784] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559066838] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559066903] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1063680 kB) free!
[1559066954] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8516520 kB) free!
[1559067024] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.5% (2521844 kB) free!
[1559067074] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.8% (14195236 kB) free.
[1559067111] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559067144] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.2% (4089964 kB) free.
[1559067231] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559067351] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559067491] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452924 kB) free!
[1559068038] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559068091] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (2538364 kB) free!