[1572547016] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174224 kB) free!
[1572547101] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168276 kB) free!
[1572547136] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174420 kB) free!
[1572547221] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164300 kB) free!
[1572547256] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187792 kB) free!
[1572547294] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180024 kB) free!
[1572547341] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165948 kB) free!
[1572547346] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182020 kB) free!
[1572547415] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196152 kB) free!
[1572547462] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17233888 kB) free!
[1572547466] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188784 kB) free!
[1572547535] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181008 kB) free!
[1572547582] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14616500 kB) free!
[1572547586] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172060 kB) free!
[1572547702] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29268124 kB) free.
[1572548302] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15713768 kB) free!
[1572548422] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29913548 kB) free.
[1572549622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16252144 kB) free!
[1572549742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29907016 kB) free.
[1572550815] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4556444 kB) free!
[1572550942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15231312 kB) free!
[1572551063] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28763884 kB) free.
[1572553462] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15382820 kB) free!
[1572553582] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28557752 kB) free.
[1572555011] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242664 kB) free!
[1572556211] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 85.3% (42012964 kB) free.
[1572556535] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 21.6% (5341084 kB) free.
[1572558060] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1036) < 49.50.64.122 [01/Nov/2019:03:10:25 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1572558180] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572558364] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1332) < 2019/11/01 03:12:22 [error] 24135#0: *1583883 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572558382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15725228 kB) free!
[1572558484] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572558502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29524196 kB) free.
[1572559102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15089096 kB) free!
[1572559223] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29217844 kB) free.
[1572561022] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15501888 kB) free!
[1572561142] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28986788 kB) free.
[1572563542] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15692940 kB) free!
[1572563662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28816348 kB) free.
[1572565462] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15414044 kB) free!
[1572565535] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4937328 kB) free!
[1572565582] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28595220 kB) free.
[1572565655] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4927608 kB) free!
[1572565775] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4919516 kB) free!
[1572567982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15446880 kB) free!
[1572568102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28805916 kB) free.
[1572568702] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14906072 kB) free!
[1572568822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28853468 kB) free.
[1572569187] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.1% (20865704 kB) free.
[1572569375] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198552 kB) free!
[1572569456] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 81.3% (19933840 kB) free.
[1572569541] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.1% (21347244 kB) free.
[1572570022] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14896324 kB) free!
[1572570142] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28746256 kB) free.
[1572571231] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1572571352] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1572571942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16122692 kB) free!
[1572572062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29014412 kB) free.
[1572572662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14515272 kB) free!
[1572572782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26639528 kB) free.
[1572576242] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 72.4% (35681304 kB) free.
[1572576270] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.2% (18562100 kB) free.
[1572576302] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (26915036 kB) free.
[1572584673] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.9% (3429496 kB) free!
[1572584793] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.5% (3332824 kB) free!
[1572584914] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.2% (3265760 kB) free!
[1572586114] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2416640 kB) free!
[1572587182] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12714568 kB) free!
[1572587187] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4899816 kB) free!
[1572587302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31911988 kB) free.
[1572587307] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4685280 kB) free!
[1572587428] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4466592 kB) free!
[1572588223] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 8 users currently logged in
[1572588502] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10387964 kB) free!
[1572588622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30044316 kB) free.
[1572588628] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2254312 kB) free!
[1572589222] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8835088 kB) free!
[1572589342] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.7% (24687916 kB) free!
[1572589463] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19702604 kB) free!
[1572589464] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 36.31, 57.09, 40.48
[1572589584] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 23.50, 45.78, 38.37
[1572589704] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 24.65, 38.75, 36.70
[1572589811] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9323764 kB) free!
[1572589931] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (8851120 kB) free!
[1572590051] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (8364076 kB) free!
[1572590063] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30880280 kB) free.
[1572590283] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.35, 5.92, 5.04
[1572590403] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.57, 5.35, 4.94
[1572590456] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4710964 kB) free!
[1572590576] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4488808 kB) free!
[1572590697] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (4284552 kB) free!
[1572591003] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.25, 5.01, 5.04
[1572591124] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.14, 4.69, 4.92
[1572591251] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3598952 kB) free!
[1572591505] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 32.33, 29.46, 29.87
[1572591897] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2216196 kB) free!
[1572592292] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.3% (19080328 kB) free.
[1572592313] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 67.2% (16473196 kB) free.
[1572592339] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 30.6% (15083640 kB) free.
[1572592402] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.7% (22240828 kB) free.
[1572592463] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15317156 kB) free!
[1572592583] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34681184 kB) free.
[1572592892] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.9% (3665696 kB) free!
[1572593012] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193168 kB) free!
[1572593132] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1572593132] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182968 kB) free!
[1572593183] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (12987580 kB) free!
[1572593303] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33372992 kB) free.
[1572593305] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 20.39, 57.03, 49.36
[1572593426] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 20.45, 44.79, 45.80
[1572593541] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4397772 kB) free!
[1572593546] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 24.61, 38.32, 43.32
[1572593661] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (4290372 kB) free!
[1572593768] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1572593781] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (4179960 kB) free!
[1572593889] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1572593903] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11348388 kB) free!
[1572594009] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1572594023] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31582364 kB) free.
[1572594146] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1572594146] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 21.92, 25.18, 33.92
[1572594341] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166552 kB) free!
[1572594461] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179476 kB) free!
[1572594581] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193440 kB) free!
[1572594609] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1572594623] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9859724 kB) free!
[1572594743] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29934392 kB) free.
[1572594747] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 21.22, 23.35, 28.63
[1572595181] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 40.4% (9911028 kB) free.
[1572595209] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1572595329] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1572595445] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8449024 kB) free!
[1572595449] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1572595565] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8449000 kB) free!
[1572595685] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (8449384 kB) free!
[1572595781] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179952 kB) free!
[1572595901] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163796 kB) free!
[1572595944] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11444640 kB) free!
[1572596021] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185660 kB) free!
[1572596064] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29917296 kB) free.
[1572596144] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1572596181] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2184368 kB) free!
[1572596231] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1572596264] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1572596351] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1572596650] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1572596650] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1572596664] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10677616 kB) free!
[1572596784] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29209552 kB) free.
[1572597384] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8974632 kB) free!
[1572597504] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28506684 kB) free.
[1572597843] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162532 kB) free!
[1572597963] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (162880 kB) free!
[1572598083] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164860 kB) free!
[1572598104] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7956988 kB) free!
[1572598224] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27702424 kB) free.
[1572598672] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.2% (14668976 kB) free.
[1572598682] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.2% (22342452 kB) free.
[1572598686] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.2% (40498024 kB) free.
[1572598704] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22944912 kB) free.
[1572598722] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1572598722] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 92.6% (22701764 kB) free.
[1572598819] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4516836 kB) free!
[1572598824] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7018224 kB) free!
[1572598939] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (4177176 kB) free!
[1572598944] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26677812 kB) free.
[1572599059] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 16.5% (3888884 kB) free!
[1572599314] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (3856156 kB) free!
[1572599434] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.3% (3781456 kB) free!
[1572599544] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5502888 kB) free!
[1572599554] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (3712700 kB) free!
[1572599659] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2257912 kB) free!
[1572599664] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.5% (25699212 kB) free!
[1572599784] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21504692 kB) free!
[1572600074] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4031996 kB) free!
[1572600194] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.1% (3839684 kB) free!
[1572600314] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.2% (3646760 kB) free!
[1572600484] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2019/11/01 14:55:38 [error] 24130#0: *1715480 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572600604] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572600754] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2190988 kB) free!
[1572600820] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4898440 kB) free!
[1572600859] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.7% (22793156 kB) free.
[1572600914] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 85.0% (19084984 kB) free.
[1572600940] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4693212 kB) free!
[1572600985] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1572600985] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3396344 kB) free!
[1572601060] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4508092 kB) free!
[1572601204] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/11/01 15:10:02 [error] 24134#0: *1715995 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572601324] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(5) < 2019/11/01 15:12:00 [error] 24132#0: *1716083 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572601444] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(20) < 2019/11/01 15:14:03 [error] 24135#0: *1716331 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572601660] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (24.3% change >= 20.0% threshold)
[1572601660] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1874392 kB) free!
[1572601664] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1572601784] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1572602260] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (3368960 kB) free!
[1572602513] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (4168708 kB) free!
[1572602634] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (4033628 kB) free!
[1572602645] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572602754] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (3877560 kB) free!
[1572603385] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27337776 kB) free.
[1572603460] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1776536 kB) free!
[1572603618] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.3% (23272344 kB) free.
[1572603735] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.3% (22858052 kB) free.
[1572603797] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.2% (22840980 kB) free.
[1572605185] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13023040 kB) free!
[1572605305] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34776492 kB) free.
[1572605905] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11984924 kB) free!
[1572606025] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34161520 kB) free.
[1572606147] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572606625] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10742236 kB) free!
[1572606745] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32967096 kB) free.
[1572606934] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9638884 kB) free!
[1572607054] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9543088 kB) free!
[1572607127] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1572607174] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (9452472 kB) free!
[1572607345] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10785232 kB) free!
[1572607450] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1572607465] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32954176 kB) free.
[1572608066] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9905340 kB) free!
[1572608186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32139668 kB) free.
[1572608786] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8778060 kB) free!
[1572608906] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31107604 kB) free.
[1572609146] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8674464 kB) free!
[1572609266] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8331400 kB) free!
[1572609386] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (7998720 kB) free!
[1572609506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8929616 kB) free!
[1572609626] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4748096 kB) free!
[1572609746] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28047868 kB) free.
[1572609859] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 11.0% (2602252 kB) free!
[1572609979] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (2152488 kB) free!
[1572610095] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9404488 kB) free!
[1572610099] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1690108 kB) free!
[1572610215] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9404392 kB) free!
[1572610335] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (9404008 kB) free!
[1572610346] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4872872 kB) free!
[1572610466] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28534812 kB) free.
[1572610514] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.2% (3186020 kB) free!
[1572610586] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4691604 kB) free!
[1572610634] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2888112 kB) free!
[1572610754] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.6% (2612924 kB) free!
[1572611066] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4941408 kB) free!
[1572611186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28181692 kB) free.
[1572611354] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1572611354] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (983168 kB) free!
[1572611733] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.2% (40500188 kB) free.
[1572611786] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5582640 kB) free!
[1572611900] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.9% (22845212 kB) free.
[1572611905] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.2% (47424992 kB) free.
[1572611906] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15673296 kB) free!
[1572611954] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 85.1% (19110124 kB) free.
[1572612022] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4764788 kB) free!
[1572612026] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 49.0% (64648620 kB) free.
[1572612098] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 94.9% (46753536 kB) free.
[1572612142] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.4% (2812536 kB) free!
[1572612228] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1572612228] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.0% (21477000 kB) free.
[1572612811] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4715800 kB) free!
[1572612931] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4551668 kB) free!
[1572613051] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4392184 kB) free!
[1572613233] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1572613948] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (4572700 kB) free!
[1572614068] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (4392868 kB) free!
[1572614188] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4237284 kB) free!
[1572614851] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1834748 kB) free!
[1572615233] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (2157944 kB) free!
[1572615353] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (2071380 kB) free!
[1572615388] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2405388 kB) free!
[1572615473] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180184 kB) free!
[1572615633] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24829308 kB) free!
[1572615753] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.6% (48261128 kB) free.
[1572615936] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (251028 kB) free!
[1572616056] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237216 kB) free!
[1572616176] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230000 kB) free!
[1572616687] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4567712 kB) free!
[1572616807] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4500956 kB) free!
[1572616927] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4316652 kB) free!
[1572619353] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13777052 kB) free!
[1572619473] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35396864 kB) free.
[1572619776] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1572619776] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (5435096 kB) free!
[1572620073] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14265564 kB) free!
[1572620194] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35681436 kB) free.
[1572620376] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224112 kB) free!
[1572620793] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15131292 kB) free!
[1572620913] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30912864 kB) free.
[1572621513] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16955464 kB) free!
[1572621633] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13269848 kB) free!
[1572621727] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2358648 kB) free!
[1572621753] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34421456 kB) free.
[1572622353] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15099800 kB) free!
[1572622473] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34901536 kB) free.
[1572622764] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572623073] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16885812 kB) free!
[1572623193] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13255444 kB) free!
[1572623313] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33610720 kB) free.
[1572623913] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14819592 kB) free!
[1572624033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34513708 kB) free.
[1572624633] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17457296 kB) free!
[1572624753] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13962196 kB) free!
[1572624874] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33026708 kB) free.
[1572625473] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15881028 kB) free!
[1572625593] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34304784 kB) free.
[1572626188] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1572626193] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17004468 kB) free!
[1572626313] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34205076 kB) free.
[1572626913] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18557424 kB) free!
[1572627033] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15300144 kB) free!
[1572627153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33776796 kB) free.
[1572627753] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16336236 kB) free!
[1572627873] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34045064 kB) free.
[1572628473] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18478564 kB) free!
[1572628593] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15328948 kB) free!
[1572628713] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33600320 kB) free.
[1572629313] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16904656 kB) free!
[1572629433] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33947308 kB) free.
[1572630034] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19260900 kB) free!
[1572630154] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16182036 kB) free!
[1572630274] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33900796 kB) free.
[1572630874] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18162092 kB) free!
[1572630994] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (15011808 kB) free!
[1572631114] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32702212 kB) free.
[1572631714] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17124832 kB) free!
[1572631776] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572631834] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33827940 kB) free.
[1572632434] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21393020 kB) free!
[1572632554] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.6% (19205056 kB) free!
[1572632674] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16970696 kB) free!