[1567881415] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.4% (2801492 kB) free!
[1567881476] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1567881487] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181696 kB) free!
[1567881535] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186568 kB) free!
[1567881600] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1567881607] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183192 kB) free!
[1567881655] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200664 kB) free!
[1567881727] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (213388 kB) free!
[1567881946] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567882066] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567884697] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.1% (9399376 kB) free.
[1567885297] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4844548 kB) free!
[1567885417] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (3946156 kB) free!
[1567885537] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.6% (3358096 kB) free!
[1567886137] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185852 kB) free!
[1567886750] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.67, 6.01, 5.09
[1567886870] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.88, 6.26, 5.29
[1567886990] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.93, 6.35, 5.46
[1567887590] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.17, 3.02, 3.96
[1567888666] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567888786] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567889596] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (245136 kB) free!
[1567889716] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239240 kB) free!
[1567889836] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227968 kB) free!
[1567890697] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 303 processes
[1567890817] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1567891036] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.9% (44286956 kB) free.
[1567892455] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 46.1% (11294884 kB) free.
[1567892527] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 59.7% (14627992 kB) free.
[1567892614] SERVICE ALERT: sig6-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 264 processes
[1567892734] SERVICE ALERT: sig6-vm2;CPU_Procs;OK;HARD;1;CPU OK: 259 processes
[1567903186] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 407 processes
[1567903306] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 409 processes
[1567904455] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179536 kB) free!
[1567904537] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1567904575] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195992 kB) free!
[1567904576] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 467 processes
[1567904636] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 461 processes
[1567904660] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189232 kB) free!
[1567904696] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171816 kB) free!
[1567904780] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195256 kB) free!
[1567915906] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567916026] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1567920824] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.9% (4881624 kB) free!
[1567920944] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.8% (4861328 kB) free!
[1567921064] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.8% (4841788 kB) free!
[1567925747] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 362 processes
[1567925867] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 361 processes
[1567925987] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 363 processes
[1567926098] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 93.7% (22958468 kB) free.
[1567926112] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.7% (22380748 kB) free.
[1567926124] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.0% (22058608 kB) free.
[1567926145] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39774820 kB) free.
[1567926164] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.4% (23147936 kB) free.
[1567926184] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14649984 kB) free.
[1567929116] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182940 kB) free!
[1567929236] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173264 kB) free!
[1567929356] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195812 kB) free!
[1567933356] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8733928 kB) free!
[1567933476] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.4% (8067420 kB) free!
[1567933597] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.1% (7445752 kB) free!
[1567934139] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1567934197] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4745024 kB) free!
[1567934259] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1567938839] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.4% (23053856 kB) free.
[1567938877] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.5% (40678468 kB) free.
[1567939662] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 19.20, 16.35, 7.43
[1567939782] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.86, 11.43, 6.70
[1567939790] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.04, 14.94, 9.08
[1567939902] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.09, 7.89, 5.98
[1567939910] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.36, 10.63, 8.22
[1567940030] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.79, 7.30, 7.29
[1567940045] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (4196488 kB) free!
[1567940165] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.4% (3546368 kB) free!
[1567940285] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.3% (3521632 kB) free!
[1567940502] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.42, 1.55, 3.42
[1567940630] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.79, 2.18, 4.38
[1567940885] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170248 kB) free!
[1567947884] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9266008 kB) free!
[1567948004] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8843440 kB) free!
[1567948124] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (8391524 kB) free!
[1567951085] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22826808 kB) free.
[1567951724] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4394624 kB) free!
[1567954430] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 135.13, 67.98, 27.64
[1567954550] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 139.97, 92.00, 41.33
[1567954670] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 118.67, 103.35, 51.83
[1567954933] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4481492 kB) free!
[1567955042] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.5% (40654812 kB) free.
[1567955053] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4159488 kB) free!
[1567955173] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (3775708 kB) free!
[1567955179] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 96.07, 80.51, 43.07
[1567955299] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 28.21, 64.55, 42.12
[1567955419] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.55, 43.48, 37.12
[1567955773] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 23.5% (5748292 kB) free.
[1567955870] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.73, 5.20, 16.17
[1567956019] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.97, 6.78, 19.97
[1567957070] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 0.65, 4.83
[1567957685] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189048 kB) free!
[1567957805] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196584 kB) free!
[1567957819] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.63, 0.84, 3.60
[1567957925] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163664 kB) free!
[1567959120] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.4% (23295392 kB) free.
[1567963573] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4877872 kB) free!
[1567963693] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4873680 kB) free!
[1567963813] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4854720 kB) free!
[1567966447] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9717864 kB) free!
[1567966567] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9718584 kB) free!
[1567966687] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (9717292 kB) free!