[1580063531] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.5% (30318280 kB) free.
[1580063778] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.4% (2795224 kB) free!
[1580063898] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191812 kB) free!
[1580064010] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4580908 kB) free!
[1580064018] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174136 kB) free!
[1580064131] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4583452 kB) free!
[1580064251] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4584504 kB) free!
[1580064798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178564 kB) free!
[1580064918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196908 kB) free!
[1580065038] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188608 kB) free!
[1580066239] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.8% (10820372 kB) free.
[1580067439] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (3883860 kB) free!
[1580067559] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (1764740 kB) free!
[1580067679] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195560 kB) free!
[1580067810] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3877632 kB) free!
[1580068160] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (26762416 kB) free.
[1580068528] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2020/01/27 01:16:03 [error] 18367#0: *44173749 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580068648] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580070019] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.6% (10451800 kB) free.
[1580074451] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.2% (4960292 kB) free.
[1580074648] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(278) < 2020/01/27 03:01:31 [error] 18366#0: *44189118 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580074681] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(216) < 103.67.237.219 [27/Jan/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1580074768] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580074771] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2147) < 172.105.59.173 [27/Jan/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1580074802] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580074892] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580075051] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4912884 kB) free!
[1580075171] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4902336 kB) free!
[1580075291] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (4893660 kB) free!
[1580081419] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.9% (2673828 kB) free!
[1580081539] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (157372 kB) free!
[1580081660] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (140412 kB) free!
[1580081891] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 67.4% (16520052 kB) free.
[1580082079] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 46.5% (11478056 kB) free.
[1580085861] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.6% (40196060 kB) free.
[1580085896] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.4% (22163692 kB) free.
[1580090106] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1580090227] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580094680] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4494708 kB) free!
[1580094800] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (4482088 kB) free!
[1580094920] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4470900 kB) free!
[1580097920] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2472836 kB) free!
[1580100307] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.3% (21291104 kB) free.
[1580101690] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 114.39, 214.77, 118.25
[1580101810] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 25.68, 147.35, 105.23
[1580101930] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.41, 101.96, 93.70
[1580103130] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 19.52, 14.72, 34.31
[1580103730] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 10.13, 12.64, 24.15
[1580104330] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 18.54, 45.94, 37.75
[1580104450] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 16.20, 35.44, 34.87
[1580104570] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1580104570] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 17.19, 29.47, 32.75
[1580105170] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 13.19, 56.27, 59.44
[1580105368] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2020/01/27 11:30:49 [error] 18366#0: *44308226 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1580105488] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580105771] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 16.56, 22.41, 39.49
[1580106303] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (1483864 kB) free!
[1580106371] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 12.20, 14.22, 26.95
[1580106423] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174436 kB) free!
[1580106543] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174108 kB) free!
[1580108707] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (1604680 kB) free!
[1580108827] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (1482572 kB) free!
[1580108947] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1440484 kB) free!
[1580108958] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1580108958] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.3% (21554288 kB) free.
[1580108993] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.9% (22287772 kB) free.
[1580109088] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2020/01/27 12:38:41 [error] 18360#0: *44340228 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580109208] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580109597] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186192 kB) free!
[1580109717] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197956 kB) free!
[1580109837] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199184 kB) free!
[1580109874] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1580109874] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 60.7% (14879880 kB) free.
[1580110480] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172736 kB) free!
[1580110600] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167156 kB) free!
[1580110720] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184268 kB) free!
[1580110962] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 336 processes
[1580111082] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1580112091] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(674) < 103.255.146.164 [27/Jan/2020:13:30:06 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1580112211] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580114572] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.6% (22208140 kB) free.
[1580115784] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (351880 kB) free!
[1580115809] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(58) < 2020/01/27 14:30:55 [error] 18364#0: *44385710 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580115904] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (271476 kB) free!
[1580115929] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/01/27 14:35:18 [error] 18369#0: *44387855 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.7849005561080498 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1580116024] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173924 kB) free!
[1580116049] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580116168] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (4106388 kB) free!
[1580116288] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (4005232 kB) free!
[1580116300] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.5% (22188216 kB) free.
[1580116408] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (3896192 kB) free!
[1580117079] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21660988 kB) free.
[1580117515] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174808 kB) free!
[1580117635] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172204 kB) free!
[1580117755] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190424 kB) free!
[1580117772] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1580118667] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.5% (21191508 kB) free.
[1580119267] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171896 kB) free!
[1580119387] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187940 kB) free!
[1580119479] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 52.3% (12810676 kB) free.
[1580119490] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.8% (4616984 kB) free!
[1580119547] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21782100 kB) free.
[1580120089] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195744 kB) free!
[1580120209] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198288 kB) free!
[1580120329] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178232 kB) free!
[1580122134] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 58.9% (14428044 kB) free.
[1580122737] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (972820 kB) free!
[1580122857] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (786248 kB) free!
[1580122946] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1011 processes
[1580122977] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (666052 kB) free!
[1580123011] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(80) < 203.212.70.170 [27/Jan/2020:16:29:55 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1580123066] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 359 processes
[1580123131] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580123249] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2020/01/27 16:28:28 [error] 18366#0: *44434715 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580123370] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580123691] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4828204 kB) free!
[1580123811] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4738500 kB) free!
[1580123931] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4648864 kB) free!
[1580125133] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187308 kB) free!
[1580126667] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8759332 kB) free!
[1580126787] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8758616 kB) free!
[1580126907] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (8759084 kB) free!
[1580129580] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.8% (21178432 kB) free.
[1580129613] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.4% (21433692 kB) free.
[1580129660] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9797396 kB) free!
[1580129780] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9735256 kB) free!
[1580129900] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (9678856 kB) free!
[1580134419] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1580135300] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4750752 kB) free!
[1580135589] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1580135731] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [27/Jan/2020:20:05:31 +0530] TCP 502 0 0 7.014 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "7.014, 0.000"
[1580135852] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(2) < 172.105.59.173 [27/Jan/2020:20:06:31 +0530] TCP 502 0 0 67.023 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "7.020, 60.003, 0.000"
[1580135972] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580136507] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4583752 kB) free!
[1580139789] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4429332 kB) free!
[1580139909] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (4419876 kB) free!
[1580140029] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4408856 kB) free!
[1580141226] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580141346] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580142882] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1580143002] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1580143578] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 66.5% (32748532 kB) free.
[1580143618] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.6% (15076256 kB) free.
[1580143772] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 203.212.70.170 [27/Jan/2020:22:19:12 +0530] TCP 502 0 0 180.001 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 60.000, 60.000, 60.001"
[1580143892] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580144229] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187128 kB) free!
[1580144291] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (252984 kB) free!
[1580144411] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (238416 kB) free!
[1580144531] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180024 kB) free!
[1580145366] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580145486] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1580145606] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1580146206] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 356 processes
[1580146806] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1580147406] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1580147406] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1580148007] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1580148127] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1580148731] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217100 kB) free!
[1580148852] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215572 kB) free!
[1580148972] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240156 kB) free!
[1580149160] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.3% (18896464 kB) free!
[1580149280] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.3% (18904180 kB) free!
[1580149327] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580149400] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.3% (18909692 kB) free!
[1580149447] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580149567] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes