[1573670108] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21876444 kB) free!
[1573670153] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 95.4% (46999028 kB) free.
[1573670205] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175268 kB) free!
[1573670228] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18679372 kB) free!
[1573670325] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190276 kB) free!
[1573670348] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37640188 kB) free.
[1573670434] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170024 kB) free!
[1573670445] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181652 kB) free!
[1573670554] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178656 kB) free!
[1573670585] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (203052 kB) free!
[1573670674] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190776 kB) free!
[1573670706] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187768 kB) free!
[1573670796] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179200 kB) free!
[1573670826] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199472 kB) free!
[1573670916] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191400 kB) free!
[1573670948] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21357136 kB) free!
[1573671036] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172028 kB) free!
[1573671068] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.9% (18383464 kB) free!
[1573671188] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36935708 kB) free.
[1573671788] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20073764 kB) free!
[1573671909] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37729628 kB) free.
[1573671982] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 11.5% (2824632 kB) free!
[1573672042] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 63.5% (15552504 kB) free.
[1573672417] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(193) < 2019/11/14 00:42:26 [error] 23635#0: *2733375 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:462/17, bytes from/to upstream:17/462
[1573672509] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21901500 kB) free!
[1573672538] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/11/14 00:43:40 [error] 23640#0: *2733635 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:202/17, bytes from/to upstream:17/202
[1573672629] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18701492 kB) free!
[1573672658] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573672749] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36963164 kB) free.
[1573672837] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 53.0% (13080208 kB) free.
[1573673225] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.6% (13935688 kB) free!
[1573673258] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/14 00:53:11 [error] 23642#0: *2735333 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:248/17, bytes from/to upstream:17/248
[1573673349] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20376600 kB) free!
[1573673378] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573673469] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37726296 kB) free.
[1573673978] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/14 01:07:28 [error] 23640#0: *2737958 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:131/17, bytes from/to upstream:17/131
[1573674037] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193492 kB) free!
[1573674069] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21542912 kB) free!
[1573674095] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8883040 kB) free!
[1573674098] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573674157] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (203084 kB) free!
[1573674189] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (18581368 kB) free!
[1573674215] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8616804 kB) free!
[1573674277] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195908 kB) free!
[1573674309] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36153896 kB) free.
[1573674335] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (8329308 kB) free!
[1573674698] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/14 01:15:43 [error] 23646#0: *2739466 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:233/17, bytes from/to upstream:17/233
[1573674819] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573674909] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18751976 kB) free!
[1573674936] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4908872 kB) free!
[1573675029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36385680 kB) free.
[1573675418] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/14 01:32:31 [error] 23641#0: *2742596 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:792/17, bytes from/to upstream:17/792
[1573675538] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573675629] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19018892 kB) free!
[1573675749] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36161620 kB) free.
[1573676077] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.7% (21401208 kB) free.
[1573676138] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/14 01:40:02 [error] 23638#0: *2743816 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1573676226] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.5% (23174072 kB) free.
[1573676258] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573676349] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19213976 kB) free!
[1573676445] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1573676445] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 62.8% (15391184 kB) free.
[1573676469] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35830084 kB) free.
[1573676674] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.2% (22835264 kB) free.
[1573676677] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179888 kB) free!
[1573676736] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 69.8% (34379868 kB) free.
[1573676797] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193772 kB) free!
[1573676918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186036 kB) free!
[1573677070] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18907404 kB) free!
[1573677190] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35876736 kB) free.
[1573677790] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19188048 kB) free!
[1573677910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36384036 kB) free.
[1573678025] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12849956 kB) free!
[1573678510] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19726456 kB) free!
[1573678613] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250456 kB) free!
[1573678625] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1573678625] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.0% (18418200 kB) free!
[1573678630] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37012752 kB) free.
[1573679230] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19144124 kB) free!
[1573679258] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/14 02:29:18 [error] 23637#0: *2753020 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:190/17, bytes from/to upstream:17/190
[1573679318] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.8% (18950216 kB) free.
[1573679350] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36096928 kB) free.
[1573679378] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573679813] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.4% (47027956 kB) free.
[1573679918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168840 kB) free!
[1573680038] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181656 kB) free!
[1573680158] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (438324 kB) free!
[1573680550] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21462700 kB) free!
[1573680670] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37791940 kB) free.
[1573680759] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.0% (20726400 kB) free.
[1573681270] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20114940 kB) free!
[1573681390] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36902960 kB) free.
[1573681569] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1041.30, 440.17, 166.47
[1573681689] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1129.37, 672.72, 284.87
[1573681809] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1141.12, 828.26, 388.89
[1573682081] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(22) < 2019/11/14 03:21:06 [error] 24131#0: *6244066 recv() failed (104: Connection reset by peer) while proxying connection, client: 121.241.242.109, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:44/0, bytes from/to upstream:0/44
[1573682201] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573682409] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.82, 119.62, 213.87
[1573682445] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573682590] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22305368 kB) free!
[1573682711] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (19102968 kB) free!
[1573682831] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36495460 kB) free.
[1573683192] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1573683312] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1573683401] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(19) < 2019/11/14 03:46:40 [error] 24130#0: *6245211 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573683431] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19504444 kB) free!
[1573683432] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1573683521] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(26) < 2019/11/14 03:48:39 [error] 24130#0: *6245412 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573683551] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36628008 kB) free.
[1573683641] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(7) < 2019/11/14 03:49:12 [error] 24130#0: *6245305 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573684033] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1573684151] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22118028 kB) free!
[1573684241] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573684271] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (18990264 kB) free!
[1573684391] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36597760 kB) free.
[1573684451] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 338 processes
[1573684571] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1573684691] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1573684991] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19724748 kB) free!
[1573685111] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36486260 kB) free.
[1573685711] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19833228 kB) free!
[1573685831] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37727316 kB) free.
[1573686431] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22880028 kB) free!
[1573686551] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19700660 kB) free!
[1573686671] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37616900 kB) free.
[1573687271] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22149264 kB) free!
[1573687392] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18788060 kB) free!
[1573687512] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37479424 kB) free.
[1573688112] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20809872 kB) free!
[1573688232] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17346948 kB) free!
[1573688352] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36368512 kB) free.
[1573688952] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19411676 kB) free!
[1573689072] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37208012 kB) free.
[1573689672] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20840924 kB) free!
[1573689792] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17379576 kB) free!
[1573689912] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35859740 kB) free.
[1573690025] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573690512] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18576524 kB) free!
[1573690632] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36946884 kB) free.
[1573690841] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/14 05:45:47 [error] 24130#0: *6267419 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573690957] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1573690961] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573691077] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573691232] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19921716 kB) free!
[1573691352] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26472068 kB) free.
[1573691560] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191760 kB) free!
[1573691680] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189580 kB) free!
[1573691800] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197116 kB) free!
[1573691952] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20672172 kB) free!
[1573692046] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177916 kB) free!
[1573692072] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (17062752 kB) free!
[1573692166] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165432 kB) free!
[1573692193] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36287592 kB) free.
[1573692274] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166932 kB) free!
[1573692286] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191312 kB) free!
[1573692394] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197936 kB) free!
[1573692425] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 37.2% (49097780 kB) free.
[1573692427] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166032 kB) free!
[1573692514] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168692 kB) free!
[1573692547] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176432 kB) free!
[1573692667] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188460 kB) free!
[1573692792] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18700408 kB) free!
[1573692886] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 53.5% (13105532 kB) free.
[1573692912] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30349372 kB) free.
[1573693114] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.1% (21828248 kB) free.
[1573693267] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 78.0% (19107752 kB) free.
[1573693512] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19856160 kB) free!
[1573693632] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16214352 kB) free!
[1573693752] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35686000 kB) free.
[1573694352] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17363284 kB) free!
[1573694472] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35939152 kB) free.
[1573695072] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18093876 kB) free!
[1573695192] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (14149096 kB) free!
[1573695312] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33813168 kB) free.
[1573695912] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13923684 kB) free!
[1573696032] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34960084 kB) free.
[1573696632] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15172544 kB) free!
[1573696752] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34642524 kB) free.
[1573697353] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15998792 kB) free!
[1573697376] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4280268 kB) free!
[1573697473] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12184744 kB) free!
[1573697496] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4126680 kB) free!
[1573697593] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33337120 kB) free.
[1573697616] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.8% (3995880 kB) free!
[1573697800] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573698193] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13037384 kB) free!
[1573698313] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33720500 kB) free.
[1573698913] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13934832 kB) free!
[1573699033] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9916116 kB) free!
[1573699153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31611520 kB) free.
[1573699416] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1480668 kB) free!
[1573699753] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10964576 kB) free!
[1573699873] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32655380 kB) free.
[1573700473] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12626804 kB) free!
[1573700593] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17972952 kB) free!
[1573700713] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30363168 kB) free.
[1573701313] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11654236 kB) free!
[1573701433] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31846228 kB) free.
[1573702034] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13772840 kB) free!
[1573702154] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15317216 kB) free!
[1573702274] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29743080 kB) free.
[1573702417] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 32.2% (7236088 kB) free.
[1573702778] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(187) < 2019/11/14 09:08:08 [error] 23639#0: *2810219 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1573702874] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8867412 kB) free!
[1573702898] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573702994] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31074648 kB) free.
[1573703594] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10813772 kB) free!
[1573703714] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29649320 kB) free.
[1573704314] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11064972 kB) free!
[1573704434] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6662080 kB) free!
[1573704554] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29925204 kB) free.
[1573705154] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8084960 kB) free!
[1573705274] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4227460 kB) free!
[1573705394] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27332320 kB) free.
[1573705994] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5671936 kB) free!
[1573706114] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28548732 kB) free.
[1573706686] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4741484 kB) free!
[1573706715] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7056372 kB) free!
[1573706806] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4620680 kB) free!
[1573706835] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2942500 kB) free!
[1573706926] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (4478392 kB) free!
[1573706955] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26778236 kB) free.
[1573707068] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4699360 kB) free!
[1573707188] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4535208 kB) free!
[1573707217] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4287496 kB) free!
[1573707308] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4397340 kB) free!
[1573707337] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4093980 kB) free!
[1573707457] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1573707457] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.2% (3870740 kB) free!
[1573707555] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5310560 kB) free!
[1573707675] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27144516 kB) free.
[1573708057] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2195064 kB) free!
[1573708275] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4201540 kB) free!
[1573708395] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384756 kB) free!
[1573708508] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1843036 kB) free!
[1573708515] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29049108 kB) free.
[1573708562] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 31.2% (7007032 kB) free.
[1573708609] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.3% (21888264 kB) free.
[1573708648] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.4% (22795012 kB) free.
[1573708699] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 58.6% (14364008 kB) free.
[1573709115] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5771388 kB) free!
[1573709235] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1630220 kB) free!
[1573709355] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27384620 kB) free.
[1573709762] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4048608 kB) free!
[1573709882] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (3741244 kB) free!
[1573709955] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4631084 kB) free!
[1573710003] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (3450056 kB) free!
[1573710075] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (568720 kB) free!
[1573710195] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25685604 kB) free!
[1573710337] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.4% (7572088 kB) free!
[1573710457] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (8208364 kB) free!
[1573710491] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1573710577] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (8208392 kB) free!
[1573710603] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2108044 kB) free!
[1573710611] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 338 processes
[1573710731] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1573710796] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3251348 kB) free!
[1573711996] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18530644 kB) free!
[1573712561] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2019/11/14 11:52:39 [error] 24131#0: *6334189 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573712681] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(6) < 2019/11/14 11:54:25 [error] 24130#0: *6334287 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573712801] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(25) < 2019/11/14 11:56:12 [error] 24130#0: *6334407 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573713003] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 32.3% (7263112 kB) free.
[1573713196] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1573713196] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2845084 kB) free!
[1573713401] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573713412] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (4067176 kB) free!
[1573713532] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3620356 kB) free!
[1573713652] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1573713652] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.4% (3523524 kB) free!
[1573713796] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18852296 kB) free!
[1573714177] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3733916 kB) free!
[1573714252] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 23.2% (5675780 kB) free.
[1573714654] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (4080344 kB) free!
[1573714774] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (4030160 kB) free!
[1573714894] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2834880 kB) free!
[1573714996] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33139056 kB) free.
[1573715201] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/14 12:30:18 [error] 24134#0: *6340922 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573715303] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (4314416 kB) free!
[1573715322] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573715424] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4086060 kB) free!
[1573715452] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (3777588 kB) free!
[1573715494] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (2000100 kB) free!
[1573715544] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (3886792 kB) free!
[1573715572] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3507288 kB) free!
[1573715596] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9917492 kB) free!
[1573715692] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.4% (3291772 kB) free!
[1573715716] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22556524 kB) free!
[1573715813] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9361880 kB) free!
[1573715836] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31885356 kB) free.
[1573715914] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4760004 kB) free!
[1573715933] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9119264 kB) free!
[1573716034] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4624772 kB) free!
[1573716053] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (8861436 kB) free!
[1573716154] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4505940 kB) free!
[1573716292] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452252 kB) free!
[1573716436] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9351268 kB) free!
[1573716556] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33102684 kB) free.
[1573716744] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1573716744] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2125040 kB) free!
[1573717156] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11944660 kB) free!
[1573717276] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (7976500 kB) free!
[1573717397] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32277056 kB) free.
[1573717803] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4282248 kB) free!
[1573717923] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (3980040 kB) free!
[1573717954] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2332692 kB) free!
[1573717996] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10931156 kB) free!
[1573718043] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (3670576 kB) free!
[1573718116] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6604848 kB) free!
[1573718236] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31294876 kB) free.
[1573718453] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3948808 kB) free!
[1573718643] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2071068 kB) free!
[1573718836] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8921388 kB) free!
[1573718956] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4477848 kB) free!
[1573719076] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30267644 kB) free.
[1573719131] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1573719251] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1573719676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8062560 kB) free!
[1573719796] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3910976 kB) free!
[1573719916] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29178360 kB) free.
[1573720516] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7534460 kB) free!
[1573720636] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3221900 kB) free!
[1573720756] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20092300 kB) free!
[1573721356] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7578824 kB) free!
[1573721957] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17636116 kB) free!
[1573721975] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.0% (22457880 kB) free.
[1573722011] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.9% (22523384 kB) free.
[1573722043] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 32.7% (7338440 kB) free.
[1573722077] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 49.1% (24204540 kB) free.
[1573722105] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 29.1% (14356360 kB) free.
[1573722133] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 61.7% (15131616 kB) free.
[1573722170] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.2% (22593776 kB) free.
[1573722557] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29326988 kB) free.
[1573722583] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183972 kB) free!
[1573722703] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180588 kB) free!
[1573722824] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1573722824] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164372 kB) free!
[1573723157] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4887612 kB) free!
[1573723252] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (3865096 kB) free!
[1573723277] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (748908 kB) free!
[1573723372] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3600864 kB) free!
[1573723397] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29926076 kB) free.
[1573723492] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3342684 kB) free!
[1573723997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5515076 kB) free!
[1573724093] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 32.3% (7241724 kB) free.
[1573724117] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1379524 kB) free!
[1573724237] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30906388 kB) free.
[1573724837] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6837824 kB) free!
[1573724957] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2502192 kB) free!
[1573725077] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31078840 kB) free.
[1573725677] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6557032 kB) free!
[1573725703] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.0% (22219504 kB) free.
[1573725740] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.6% (23147048 kB) free!
[1573725744] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40520644 kB) free.
[1573726033] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1573726154] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 336 processes
[1573726274] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1573726313] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (282884 kB) free!
[1573726433] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164476 kB) free!
[1573726553] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169284 kB) free!
[1573726721] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2019/11/14 15:40:23 [error] 24129#0: *6377122 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573726841] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(17) < 2019/11/14 15:49:37 [error] 24130#0: *6378347 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573726946] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1573726961] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573727014] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9555812 kB) free!
[1573727134] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9334876 kB) free!
[1573727254] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (9132712 kB) free!
[1573727414] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573727544] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21806764 kB) free!
[1573727546] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (4234892 kB) free!
[1573727664] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17649960 kB) free!
[1573727666] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (4073216 kB) free!
[1573727784] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.9% (47350520 kB) free.
[1573727786] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (3922996 kB) free!
[1573728293] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4295196 kB) free!
[1573728357] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1164_n"
[1573728384] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22188732 kB) free!
[1573728414] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (3933336 kB) free!
[1573728479] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573728504] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.5% (17847252 kB) free!
[1573728534] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.9% (3560328 kB) free!
[1573728625] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45526324 kB) free.
[1573728986] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2339028 kB) free!
[1573729134] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1623340 kB) free!
[1573729224] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19460784 kB) free!
[1573729282] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8237112 kB) free!
[1573729306] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.0% (22699464 kB) free.
[1573729344] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15149804 kB) free!
[1573729368] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1573729368] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 60.8% (14894952 kB) free.
[1573729402] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (7765708 kB) free!
[1573729407] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 31.6% (7090336 kB) free.
[1573729464] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31586292 kB) free.
[1573729522] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.7% (7761764 kB) free!
[1573729906] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171276 kB) free!
[1573730026] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184808 kB) free!
[1573730064] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5379500 kB) free!
[1573730146] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174332 kB) free!
[1573730184] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1157056 kB) free!
[1573730254] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4722904 kB) free!
[1573730304] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32851736 kB) free.
[1573730615] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2924656 kB) free!
[1573730735] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.0% (2463012 kB) free!
[1573730855] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2043012 kB) free!
[1573730904] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6205648 kB) free!
[1573731024] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1700880 kB) free!
[1573731098] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(110) < 2019/11/14 17:01:37 [error] 23646#0: *2879238 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:229/2888, bytes from/to upstream:2888/229
[1573731144] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8463276 kB) free!
[1573731218] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(161) < 2019/11/14 17:03:32 [error] 23642#0: *2880333 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:56/17, bytes from/to upstream:17/56
[1573731315] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2844448 kB) free!
[1573731338] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573731435] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3254104 kB) free!
[1573731555] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3670288 kB) free.
[1573731614] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (4121616 kB) free!
[1573731674] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1573731734] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3947852 kB) free!
[1573731854] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (3777776 kB) free!
[1573731938] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/14 17:08:53 [error] 23648#0: *2881088 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: rq-1-1gbps.server, request: "GET /favicon.ico HTTP/1.1", host: "rq-1-1gbps.server:8091"
[1573732058] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573732122] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.1% (22981380 kB) free.
[1573732163] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 92.5% (45594848 kB) free.
[1573732248] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 32.5% (7297104 kB) free.
[1573732277] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.3% (49125528 kB) free.
[1573733054] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2125756 kB) free!
[1573733123] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3038660 kB) free!
[1573733331] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180552 kB) free!
[1573733451] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173176 kB) free!
[1573733452] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (3692500 kB) free!
[1573733571] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172592 kB) free!
[1573733572] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3325500 kB) free!
[1573733693] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.3% (2984656 kB) free!
[1573734293] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1419600 kB) free!
[1573734555] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3224064 kB) free!
[1573734675] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3733564 kB) free.
[1573734686] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18783832 kB) free!
[1573734806] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48640248 kB) free.
[1573734893] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 31.8% (7142408 kB) free.
[1573735373] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4577780 kB) free!
[1573735375] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4795824 kB) free!
[1573735407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24010988 kB) free!
[1573735493] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4486340 kB) free!
[1573735495] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4672636 kB) free!
[1573735527] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19715432 kB) free!
[1573735613] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4366452 kB) free!
[1573735615] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (4539860 kB) free!
[1573735647] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15553540 kB) free!
[1573736475] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2078900 kB) free!
[1573736596] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2619360 kB) free!
[1573736716] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (3184328 kB) free!
[1573736847] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37021684 kB) free.
[1573737316] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3596936 kB) free.
[1573737415] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1910400 kB) free!
[1573737447] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14407220 kB) free!
[1573737567] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44945612 kB) free.
[1573737712] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9355068 kB) free!
[1573737832] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9119712 kB) free!
[1573737953] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (8892240 kB) free!
[1573738013] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2119256 kB) free!
[1573738167] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19903244 kB) free!
[1573738287] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15431908 kB) free!
[1573738407] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (11140976 kB) free!
[1573739007] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21250100 kB) free!
[1573739093] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4312264 kB) free!
[1573739213] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (3968868 kB) free!
[1573739333] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.2% (3636604 kB) free!
[1573739607] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33247168 kB) free.
[1573739753] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4944884 kB) free!
[1573739933] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2063736 kB) free!
[1573740207] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10301680 kB) free!
[1573740316] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2527860 kB) free!
[1573740327] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40879644 kB) free.
[1573740436] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3114812 kB) free!
[1573740556] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3745284 kB) free.
[1573740927] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16585164 kB) free!
[1573741048] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12371152 kB) free!
[1573741168] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8493060 kB) free!
[1573741768] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18831236 kB) free!
[1573741874] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1573741994] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1573742368] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30159076 kB) free.
[1573742968] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7776880 kB) free!
[1573743088] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38082084 kB) free.
[1573743688] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17477168 kB) free!
[1573743771] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1573743808] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13617272 kB) free!
[1573743928] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (9887964 kB) free!
[1573744528] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20002548 kB) free!
[1573744756] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (2465412 kB) free!
[1573744876] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2891428 kB) free!
[1573744961] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/14 20:51:16 [error] 24132#0: *6435265 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:47/0, bytes from/to upstream:0/47
[1573744996] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3394284 kB) free.
[1573745081] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573745128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30747928 kB) free.
[1573745728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11110356 kB) free!
[1573745848] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37015992 kB) free.
[1573745934] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 32.8% (7356188 kB) free.
[1573746448] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19836332 kB) free!
[1573746568] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16384284 kB) free!
[1573746689] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12824736 kB) free!
[1573747289] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22404996 kB) free!
[1573747386] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22862132 kB) free.
[1573747423] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.5% (22188240 kB) free.
[1573747466] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 72.7% (35849324 kB) free.
[1573747493] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 35.1% (17286676 kB) free.
[1573747533] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 61.5% (15071524 kB) free.
[1573747571] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.3% (22863796 kB) free.
[1573747889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31094964 kB) free.
[1573747993] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (2191536 kB) free!
[1573748113] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (2170604 kB) free!
[1573748234] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2151308 kB) free!
[1573748260] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (2489396 kB) free!
[1573748380] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.3% (2486868 kB) free!
[1573748489] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13567352 kB) free!
[1573748500] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (2485428 kB) free!
[1573748609] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31797496 kB) free.
[1573748858] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2019/11/14 21:52:39 [error] 23645#0: *2945161 "/usr/share/nginx/html/nagios/index.html" is not found (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /nagios/ HTTP/1.1", host: "testnagios.1gbps.server.com:8091"
[1573748978] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573749209] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14410516 kB) free!
[1573749329] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10930968 kB) free!
[1573749449] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7429488 kB) free!
[1573750049] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16859404 kB) free!
[1573750649] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26554460 kB) free.
[1573751249] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8765600 kB) free!
[1573751369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31957332 kB) free.
[1573751934] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573751969] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15092232 kB) free!
[1573752089] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11761572 kB) free!
[1573752209] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (8572052 kB) free!
[1573752342] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573752809] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19531844 kB) free!
[1573753410] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29767116 kB) free.
[1573753625] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.0% (18431392 kB) free!
[1573753745] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.2% (20014268 kB) free!
[1573753865] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (14175272 kB) free!
[1573754010] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14169120 kB) free!
[1573754130] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11220120 kB) free!
[1573754250] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (8553864 kB) free!
[1573754850] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18200536 kB) free!
[1573755281] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(141) < 2019/11/14 23:41:41 [error] 24130#0: *6466906 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573755401] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573755450] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28914568 kB) free.
[1573755546] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(11) < 49.50.86.152 [14/Nov/2019:23:40:42 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1573755666] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573756050] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15036580 kB) free!
[1573756170] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12356936 kB) free!