[1597516242] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3211944 kB) free!
[1597516362] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2761108 kB) free!
[1597516927] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193104 kB) free!
[1597516944] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 9.12, 32.10, 41.01
[1597517047] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199184 kB) free!
[1597517168] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 51.8% (12773308 kB) free.
[1597517705] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.91, 5.85, 734.67
[1597517767] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171088 kB) free!
[1597517887] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194776 kB) free!
[1597517924] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.78, 1.92, 475.54
[1597518007] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.1% (20992312 kB) free.
[1597518145] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1597518145] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 67.76, 46.65, 42.43
[1597518305] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.94, 1.53, 385.36
[1597518745] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 37.38, 39.55, 40.41
[1597520195] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597520255] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193692 kB) free!
[1597520315] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.7% (17076672 kB) free.
[1597521016] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7376504 kB) free!
[1597521332] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (767768 kB) free!
[1597521452] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166448 kB) free!
[1597521572] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166004 kB) free!
[1597521607] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229428 kB) free!
[1597521727] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231080 kB) free!
[1597521763] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3822144 kB) free.
[1597521847] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (221160 kB) free!
[1597522772] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 59.9% (14693840 kB) free.
[1597523647] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 54.6% (26926768 kB) free.
[1597523972] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191608 kB) free!
[1597524007] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164472 kB) free!
[1597524092] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195200 kB) free!
[1597524127] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176860 kB) free!
[1597524212] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197276 kB) free!
[1597524247] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193224 kB) free!
[1597524847] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.2% (21262008 kB) free.
[1597525344] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 65.8% (32432160 kB) free.
[1597525413] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1597525413] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.1% (15706336 kB) free.
[1597526013] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (763240 kB) free!
[1597526133] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.2% (23088976 kB) free.
[1597526231] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 72.2% (35596856 kB) free.
[1597528945] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.68, 50.34, 46.77
[1597529324] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16201.63, 13591.18, 7195.67
[1597529444] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16203.82, 14455.73, 8287.93
[1597529545] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 39.48, 40.85, 43.50
[1597529564] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16205.32, 15034.47, 9247.84
[1597530103] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597530145] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 54.92, 49.49, 45.61
[1597530223] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1597530745] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 30.94, 40.15, 43.41
[1597532565] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.70, 17.17, 533.29
[1597534345] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 66.84, 50.46, 42.24
[1597534945] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 46.23, 47.39, 43.68
[1597536333] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1597537945] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.29, 48.76, 43.50
[1597538545] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 51.80, 46.87, 44.80
[1597539616] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (30336300 kB) free.
[1597539771] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 321 processes
[1597539891] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1597540345] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 45.44, 68.45, 59.36
[1597542933] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200176 kB) free!
[1597543053] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182500 kB) free!
[1597543173] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184992 kB) free!
[1597543773] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.1% (23064256 kB) free.
[1597544692] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 321 processes
[1597544812] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1597547706] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1597547780] SERVICE ALERT: pservercluster2;Memory;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1597547805] SERVICE ALERT: pservercluster2;Total Processes;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1597547826] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1597547900] SERVICE ALERT: pservercluster2;Memory;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1597547925] SERVICE ALERT: pservercluster2;Total Processes;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1597547946] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1597548020] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 78.9% (19328516 kB) free.
[1597548046] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 388 processes
[1597548546] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.03, 2994.40, 5837.58
[1597548608] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548645] SERVICE ALERT: sservercluster2;HDD_Root status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1597548698] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548738] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548775] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548775] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548825] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548828] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548867] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548868] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548900] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597548905] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548906] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548923] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548949] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597548950] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597548958] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548984] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597548997] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549005] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549022] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549029] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597549036] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549053] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549074] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597549080] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549080] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549114] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549127] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549135] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549152] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549166] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549183] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549204] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549210] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549245] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549265] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549282] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549340] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597549546] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16152.91, 10760.47, 4793.30
[1597549666] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16214.55, 12568.27, 6179.28
[1597549786] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16222.70, 13777.81, 7397.17
[1597550545] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1597550946] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.52, 2.87, 445.80
[1597551306] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 32894 MB (64% inode=98%):
[1597551358] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 197623 MB (89% inode=99%):
[1597551429] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.44, 1.47, 1.51
[1597551485] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=3.27% system=8.88% iowait=0.01% idle=87.84%
[1597551528] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1597551566] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1597551584] SERVICE ALERT: sservercluster2;Open-Files;OK;HARD;1;OK: 8960 open files (0% of max 2424092)
[1597551604] SERVICE ALERT: sservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1597551645] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.0% (20844340 kB) free.
[1597551682] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1597551740] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1597551868] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 329 processes
[1597552265] SERVICE ALERT: sservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597552280] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597552400] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597552963] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3216496 kB) free!
[1597553083] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3621436 kB) free.
[1597553386] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.22, 1.67, 304.88
[1597555186] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16244.66, 12141.87, 5850.42
[1597555306] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16259.61, 13505.06, 7112.79
[1597555426] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16262.07, 14417.40, 8222.18
[1597556229] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15994.50, 11534.39, 5383.42
[1597556349] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16019.04, 13019.49, 6673.43
[1597556469] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16022.39, 14013.26, 7807.06
[1597557283] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3177748 kB) free!
[1597557403] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3140556 kB) free!
[1597557523] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (2988348 kB) free!
[1597559026] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1597559026] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.37, 1.60, 265.94
[1597559469] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1597559469] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.85, 3.98, 637.75
[1597560069] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.67, 1.87, 334.88
[1597566312] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.8% (6471172 kB) free!
[1597566432] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.6% (6405812 kB) free!
[1597566523] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1159844 kB) free!
[1597566552] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.3% (6325324 kB) free!
[1597567123] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.8% (2398000 kB) free!
[1597567152] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 21.1% (6908592 kB) free.
[1597567426] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1597567752] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.9% (6529508 kB) free!
[1597567872] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.7% (6461188 kB) free!
[1597567992] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.6% (6407776 kB) free!
[1597569431] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9383504 kB) free!
[1597569551] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9262060 kB) free!
[1597569672] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (9152112 kB) free!
[1597571469] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1597573123] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1060456 kB) free!
[1597573236] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(10) < 103.255.146.164 [16/Aug/2020:15:45:07 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1597573356] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597573723] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1597573723] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (1784780 kB) free!
[1597573869] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14934.85, 6539.19, 2484.42
[1597573989] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15983.81, 9718.29, 4141.04
[1597574109] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16125.09, 11845.40, 5596.76
[1597576872] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4893228 kB) free!
[1597577109] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.36, 6.60, 813.76
[1597577709] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.63, 2.32, 427.10
[1597577923] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1076892 kB) free!
[1597578523] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.8% (1921168 kB) free!
[1597580923] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1464764 kB) free!
[1597581523] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.2% (2315636 kB) free!
[1597583346] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16197.32, 11898.08, 5621.56
[1597583466] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16216.98, 13328.16, 6906.64
[1597583560] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1597583581] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597583592] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (3181544 kB) free!
[1597583596] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1597583610] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597583665] SERVICE ALERT: pservercluster2;CPU_Procs;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1597583685] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597583685] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597583701] SERVICE ALERT: pservercluster2;Open-Files;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1597583730] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;(17) < Aug 16 18:44:41 pservercluster2 xinetd[711]: nrpe: fork failed: Cannot allocate memory (errno = 12)
[1597583785] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1597583805] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597583805] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1597583821] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 9216 open files (0% of max 2424156)
[1597583850] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597583930] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.79% system=4.44% iowait=0.00% idle=93.75%
[1597584792] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 10.3% (3374920 kB) free!
[1597585392] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (3215580 kB) free!
[1597586596] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.80, 1.79, 413.30
[1597589626] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16238.29, 11999.11, 5692.27
[1597589746] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16255.84, 13408.48, 6973.44
[1597589866] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16258.30, 14351.52, 8099.26
[1597590616] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.2% (17460656 kB) free!
[1597590736] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.2% (17458840 kB) free!
[1597590856] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.2% (17453520 kB) free!
[1597591944] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9837060 kB) free!
[1597591992] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1597591992] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 44.4% (14546548 kB) free.
[1597592064] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (9806756 kB) free!
[1597592184] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (9776856 kB) free!
[1597592923] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1597593466] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.85, 1.09, 258.22
[1597595792] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1505 processes
[1597595912] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1498 processes
[1597597992] SERVICE FLAPPING ALERT: mrq4;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1597600345] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 45.36, 42.14, 43.48
[1597600945] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 52.69, 49.83, 46.97