[1569868608] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569868672] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195228 kB) free!
[1569868732] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1569868797] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193820 kB) free!
[1569868927] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1569869754] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569870354] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 596 processes
[1569872154] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569872754] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1569872754] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 449 processes
[1569876132] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 16.06, 8.22, 5.46
[1569876252] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 165.44, 67.94, 27.68
[1569876372] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 185.60, 107.90, 47.20
[1569877055] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.2% (8482008 kB) free!
[1569877655] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (323028 kB) free!
[1569878172] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 1.62, 13.06
[1569878255] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.8% (47212028 kB) free.
[1569879697] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.5% (12133208 kB) free.
[1569879727] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 59.7% (14640068 kB) free.
[1569879972] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 1.41, 3.64
[1569883422] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1569883542] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1569884154] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569885481] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569885601] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569890401] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 409 processes
[1569890521] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 407 processes
[1569892555] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 599 processes
[1569892675] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 601 processes
[1569892795] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569892953] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.9% (11813368 kB) free.
[1569893198] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 519 processes
[1569893258] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 522 processes
[1569893318] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 527 processes
[1569893396] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 598 processes
[1569893497] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185316 kB) free!
[1569893521] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1569893530] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196284 kB) free!
[1569893553] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191960 kB) free!
[1569893617] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191328 kB) free!
[1569893641] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1569893651] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190976 kB) free!
[1569893673] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187528 kB) free!
[1569893738] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196304 kB) free!
[1569893774] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191060 kB) free!
[1569893793] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188804 kB) free!
[1569895738] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569895858] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1569902394] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.35, 6.81, 5.39
[1569902514] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.21, 6.93, 5.61
[1569902634] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.99, 6.53, 5.62
[1569911034] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 17 users currently logged in
[1569911154] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 17 users currently logged in
[1569911255] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8681904 kB) free!
[1569911274] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 17 users currently logged in
[1569911375] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.6% (8178656 kB) free!
[1569911397] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 639 processes
[1569911448] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569911495] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.9% (7819248 kB) free!
[1569911508] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1569911568] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1569911628] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1569912095] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4876588 kB) free!
[1569912228] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1569919230] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4404416 kB) free!
[1569919350] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (4051256 kB) free!
[1569919470] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 15.2% (3590744 kB) free!
[1569919934] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.4% (23052404 kB) free.
[1569919990] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.6% (22938896 kB) free.
[1569920036] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39964696 kB) free.
[1569920070] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1385396 kB) free!
[1569920097] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.9% (47266588 kB) free.
[1569920145] SERVICE ALERT: sig3-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1569920146] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.4% (22956952 kB) free.
[1569920199] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.3% (23117588 kB) free.
[1569920266] SERVICE ALERT: sig3-vm2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1569921168] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 29.3% (38620644 kB) free.
[1569921234] SERVICE ALERT: rq-2;CPU_Procs;OK;HARD;1;CPU OK: 240 processes
[1569921346] SERVICE ALERT: rq-2;CPU Load;OK;HARD;1;OK - load average: 1.07, 1.03, 0.63
[1569921347] SERVICE ALERT: rq-2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1569921347] SERVICE ALERT: rq-2;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /rq2vns2 1068147 MB (99% inode=99%):
[1569921347] SERVICE ALERT: rq-2;Total Processes;OK;HARD;1;PROCS OK: 239 processes
[1569921347] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1569921348] SERVICE ALERT: rq-2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 44762 MB (87% inode=99%):
[1569921354] SERVICE ALERT: rq-2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 278827 MB (69% inode=99%):
[1569921528] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1569921768] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;1;CRITICAL - 9.3% (12276536 kB) free!
[1569922128] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1569923548] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (560356 kB) free!
[1569923668] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (507000 kB) free!
[1569923788] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168656 kB) free!
[1569926628] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1569926640] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 330 processes
[1569927171] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 82.23, 70.06, 36.05
[1569927240] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 330 processes
[1569927291] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.60, 47.08, 31.75
[1569927411] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.51, 31.86, 28.03
[1569927474] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1569927529] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 55 users currently logged in
[1569928011] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.09, 5.26, 15.21
[1569929211] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.35, 1.16, 4.93
[1569932205] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (4473252 kB) free!
[1569932325] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (4264724 kB) free!
[1569932445] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (4087224 kB) free!
[1569932704] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4786228 kB) free!
[1569932825] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4749520 kB) free!
[1569932945] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.1% (4686820 kB) free!
[1569933645] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2330464 kB) free!
[1569933793] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (4282824 kB) free!
[1569933913] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4171620 kB) free!
[1569934011] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.30, 8.37, 6.94
[1569934033] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (4021104 kB) free!
[1569934131] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.59, 5.79, 6.17
[1569934251] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1569934251] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 4.08, 5.50
[1569934851] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.62, 1.16, 3.23
[1569935034] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 339 processes
[1569935154] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1569935519] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 41.68, 11.25, 4.68
[1569935639] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 94.72, 41.68, 16.63
[1569935645] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4770028 kB) free!
[1569935704] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8970588 kB) free!
[1569935759] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 18.56, 32.15, 16.35
[1569935765] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (3032224 kB) free!
[1569935824] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.7% (8726752 kB) free!
[1569935833] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2217588 kB) free!
[1569935885] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4170888 kB) free!
[1569935944] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.1% (8432876 kB) free!
[1569936359] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.19, 5.34, 9.15
[1569937559] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.94, 1.14, 3.37
[1569937744] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4133736 kB) free!
[1569939651] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.94, 11.86, 5.92
[1569939771] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 25.89, 18.99, 9.28
[1569939891] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.84, 13.24, 8.35
[1569939959] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.77, 18.64, 12.77
[1569939972] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.44, 8.68, 8.10
[1569940079] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.83, 26.26, 16.86
[1569940092] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.68, 6.32, 7.30
[1569940199] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1569940199] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 62.61, 36.34, 21.33
[1569940212] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.35, 4.87, 6.64
[1569940812] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.77, 2.29, 4.36
[1569941412] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15.35, 40.25, 22.61
[1569941532] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.54, 27.54, 20.09
[1569941652] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.40, 18.89, 17.83
[1569942234] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.27, 2.91, 4.07
[1569942710] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 634 processes
[1569942766] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 629 processes
[1569942830] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 615 processes
[1569942852] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1569942852] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.85, 39.66, 28.96
[1569942891] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.30, 23.54, 21.83
[1569942950] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 337 processes
[1569943366] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1569943452] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.87, 5.96, 15.54
[1569943966] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1569944034] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 16.56, 10.40, 6.56
[1569944086] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1569944091] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.42, 9.03, 16.64
[1569944154] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.21, 8.71, 6.42
[1569944206] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 349 processes
[1569944274] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.10, 6.29, 5.81
[1569944399] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.09, 4.74, 14.82
[1569944874] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.30, 3.70, 4.48
[1569944945] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2454572 kB) free!
[1569945252] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.71, 0.82, 2.97
[1569945599] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.22, 1.23, 4.91
[1569945891] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.81, 0.66, 3.00
[1569951645] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 31.4% (7700752 kB) free.
[1569952703] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 325 processes
[1569952824] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 330 processes
[1569952944] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 325 processes
[1569953544] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 322 processes