[1556735956] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10499916 kB) free.
[1556736556] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8627548 kB) free!
[1556736676] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7119912 kB) free!
[1556736796] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10913436 kB) free.
[1556737396] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8851300 kB) free!
[1556737516] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7389176 kB) free!
[1556737636] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10912568 kB) free.
[1556738088] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556738688] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556738808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556738837] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8420904 kB) free!
[1556738929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556738957] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (6976656 kB) free!
[1556739077] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10903532 kB) free.
[1556739529] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1556739829] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2818672 kB) free!
[1556739906] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 28.91, 128.05, 83.02
[1556740026] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.12, 85.79, 72.99
[1556740146] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.10, 57.61, 64.22
[1556740277] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9609600 kB) free!
[1556740397] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8220756 kB) free!
[1556740517] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6607484 kB) free!
[1556740729] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556740849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556740969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1556741117] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10902000 kB) free.
[1556741346] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.37, 1.45, 17.94
[1556741888] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2458340 kB) free!
[1556742317] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6514356 kB) free!
[1556742437] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10911408 kB) free.
[1556743146] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.75, 0.58, 3.02
[1556743637] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6747080 kB) free!
[1556743757] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10915212 kB) free.
[1556743969] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1556744163] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 243.79, 208.00, 112.79
[1556744283] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 51.24, 153.95, 105.16
[1556744289] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (3230856 kB) free!
[1556744403] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.47, 103.58, 92.61
[1556744957] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7239896 kB) free!
[1556745077] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10909328 kB) free.
[1556745677] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6405720 kB) free!
[1556745797] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10918416 kB) free.
[1556745829] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2448564 kB) free!
[1556746203] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.34, 1.76, 14.68
[1556746369] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556746489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556746997] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7062368 kB) free!
[1556747117] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10892704 kB) free.
[1556747618] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 247.55, 129.03, 51.87
[1556747689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556747738] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 254.35, 170.87, 76.56
[1556747809] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556747858] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 130.99, 166.86, 87.48
[1556748003] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.63, 1.59, 3.49
[1556748058] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2641188 kB) free!
[1556748318] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7163188 kB) free!
[1556748438] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10909000 kB) free.
[1556749009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556749038] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6267620 kB) free!
[1556749129] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556749158] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10905932 kB) free.
[1556749658] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.09, 0.70, 12.86
[1556750329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556750358] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6552600 kB) free!
[1556750450] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1556750458] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2453964 kB) free!
[1556750478] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10898724 kB) free.
[1556750858] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.20, 0.26, 3.72
[1556751049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556751169] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556751546] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 404.12, 232.93, 97.27
[1556751666] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 409.56, 291.64, 135.23
[1556751678] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7073444 kB) free!
[1556751786] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 91.60, 224.30, 130.73
[1556751798] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10908960 kB) free.
[1556751829] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2858268 kB) free!
[1556752369] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556752489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556752998] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7215660 kB) free!
[1556753119] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10902328 kB) free.
[1556753586] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.37, 0.99, 19.16
[1556753689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556753718] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6127872 kB) free!
[1556753810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556753838] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556753838] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10113484 kB) free.
[1556753930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556754438] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7703896 kB) free!
[1556754558] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5931860 kB) free!
[1556754678] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10110732 kB) free.
[1556755089] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2184292 kB) free!
[1556755278] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8000132 kB) free!
[1556755386] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.51, 3.22
[1556755398] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6370568 kB) free!
[1556755518] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10377632 kB) free.
[1556756118] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8322260 kB) free!
[1556756238] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6863976 kB) free!
[1556756330] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556756358] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10849972 kB) free.
[1556756930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556756959] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8711476 kB) free!
[1556757051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1556757079] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7306964 kB) free!
[1556757171] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556757199] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10822588 kB) free.
[1556757799] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9348600 kB) free!
[1556757829] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2450756 kB) free!
[1556757919] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7972848 kB) free!
[1556758039] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.2% (6484116 kB) free!
[1556758971] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556759571] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556759691] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556759811] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556759839] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10791140 kB) free.
[1556760412] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556760439] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8967688 kB) free!
[1556760532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556760559] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7521800 kB) free!
[1556760652] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556760679] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10785968 kB) free.
[1556761280] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9143792 kB) free!
[1556761400] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7695240 kB) free!
[1556761520] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10772772 kB) free.
[1556762120] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9312292 kB) free!
[1556762240] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7893340 kB) free!
[1556762344] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.8% (22908472 kB) free.
[1556762360] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6141740 kB) free!
[1556762405] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.6% (22458700 kB) free.
[1556762451] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.4% (14531196 kB) free.
[1556762511] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.0% (20835644 kB) free.
[1556762951] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201588 kB) free!
[1556763072] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189788 kB) free!
[1556763192] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174788 kB) free!
[1556763716] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (381736 kB) free!
[1556763836] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (473808 kB) free!
[1556763956] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (542140 kB) free!
[1556764160] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1556764160] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10761392 kB) free.
[1556764602] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1556764722] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1556764760] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8702580 kB) free!
[1556764880] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7246432 kB) free!
[1556764986] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 282.85, 158.20, 65.59
[1556765000] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10749608 kB) free.
[1556765106] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 287.85, 201.35, 92.63
[1556765226] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 174.91, 202.63, 107.21
[1556765452] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1556765601] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9096484 kB) free!
[1556765721] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7663136 kB) free!
[1556765841] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10738128 kB) free.
[1556766052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1556766172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556766292] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556766337] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.5% (20955132 kB) free.
[1556766388] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1556766388] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22661144 kB) free.
[1556766441] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9229460 kB) free!
[1556766561] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7792672 kB) free!
[1556766681] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6046324 kB) free!
[1556766892] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556766946] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (389640 kB) free!
[1556767012] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556767026] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 1.24, 16.37
[1556767066] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (403236 kB) free!
[1556767133] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556767186] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1556767186] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (428952 kB) free!
[1556767600] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (185040 kB) free!
[1556767720] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170588 kB) free!
[1556767840] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179316 kB) free!
[1556768481] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10681764 kB) free.
[1556768826] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.81, 1.13, 3.35
[1556768858] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 126.98, 181.57, 105.73
[1556768978] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.30, 121.57, 92.94
[1556769081] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8760440 kB) free!
[1556769098] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.67, 81.46, 81.71
[1556769201] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7281384 kB) free!
[1556769322] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10659688 kB) free.
[1556769921] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9063760 kB) free!
[1556770041] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7647884 kB) free!
[1556770161] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6020572 kB) free!
[1556770898] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.82, 0.79, 12.35
[1556771961] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10565968 kB) free.
[1556772098] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.07, 2.43, 4.60
[1556772544] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1556772561] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9076532 kB) free!
[1556772681] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7603580 kB) free!
[1556772801] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6127640 kB) free!
[1556773240] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.9% (18229948 kB) free.
[1556773335] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 63.7% (15608032 kB) free.
[1556773733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1556773800] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.04, 21.31, 20.03
[1556773857] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (165924 kB) free!
[1556773882] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201820 kB) free!
[1556773935] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182928 kB) free!
[1556773977] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (179472 kB) free!
[1556774002] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167460 kB) free!
[1556774056] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178092 kB) free!
[1556774097] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (191976 kB) free!
[1556774122] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165864 kB) free!
[1556774176] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184472 kB) free!
[1556774333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556774400] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 20.47, 18.72, 18.99
[1556774453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1556774573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556774602] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10464116 kB) free.
[1556774697] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.1% (1804692 kB) free!
[1556774959] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 281 processes
[1556775058] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.5% (15802816 kB) free.
[1556775079] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 276 processes
[1556775202] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8949260 kB) free!
[1556775297] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1197396 kB) free!
[1556775322] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7435428 kB) free!
[1556775322] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 63.9% (15778232 kB) free.
[1556775442] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.9% (5882060 kB) free!
[1556775714] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1556776373] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556776754] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.32, 5.42, 5.01
[1556776874] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.24, 5.41, 5.06
[1556776914] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1556776974] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1556776994] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.17, 5.96, 5.29
[1556777094] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1556777214] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556777242] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10270776 kB) free.
[1556777814] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556777842] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8362772 kB) free!
[1556777934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556777962] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6664980 kB) free!
[1556778005] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (4195524 kB) free!
[1556778054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1556778082] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10200596 kB) free.
[1556778125] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188172 kB) free!
[1556778245] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 57.2% (14008080 kB) free.
[1556778682] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8740288 kB) free!
[1556778795] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.48, 4.61, 4.91
[1556778802] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7112844 kB) free!
[1556778923] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5361960 kB) free!
[1556779854] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1556779995] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.69, 5.05, 5.07
[1556780116] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.55, 5.24, 5.13
[1556780236] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.76, 5.35, 5.18
[1556780455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556780575] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556780695] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556780723] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4517768 kB) free!
[1556781323] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.9% (8352644 kB) free!
[1556781429] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46634 MB (20% inode=99%):
[1556781549] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45608 MB (20% inode=99%):
[1556781669] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45574 MB (20% inode=99%):
[1556782495] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556783095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556783215] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556783336] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556783458] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (4174132 kB) free!
[1556783578] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.6% (3090692 kB) free!
[1556783698] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2519548 kB) free!
[1556783935] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1556784055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1556784175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556784298] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2176532 kB) free!
[1556784436] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1556784436] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.64, 4.34, 4.70
[1556784670] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 66659 MB (29% inode=99%):
[1556784897] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 39.9% (6481752 kB) free.
[1556785523] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4425680 kB) free!
[1556785576] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (3000092 kB) free!
[1556786122] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1556786123] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.8% (8301164 kB) free!
[1556786176] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (539032 kB) free!
[1556786236] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.77, 5.89, 5.35
[1556786356] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.77, 5.42, 5.24
[1556786476] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.20, 5.29, 5.21
[1556786767] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1556786826] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 439.46, 295.38, 132.50
[1556786946] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 343.64, 322.16, 162.91
[1556787066] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 46.56, 215.65, 143.19
[1556787076] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.50, 4.77, 4.83
[1556787297] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (3144444 kB) free!
[1556787324] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4032508 kB) free!
[1556787417] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (2995204 kB) free!
[1556787522] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1556787537] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (2841836 kB) free!
[1556787643] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1556787698] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.48, 7.02, 6.57
[1556787818] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.50, 5.56, 6.09
[1556787924] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.0% (8372692 kB) free!
[1556787938] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.97, 4.06, 5.49
[1556788144] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1556788538] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.91, 3.08
[1556788738] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (1471052 kB) free!
[1556789124] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4002780 kB) free!
[1556789466] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.72, 0.84, 11.36
[1556789724] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.5% (8137668 kB) free!
[1556790544] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1556790666] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.23, 0.49, 3.50
[1556791144] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1556791375] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1556791831] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1556791937] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.1% (18408172 kB) free.
[1556791951] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1556791975] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556792096] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556792125] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 39.4% (19428500 kB) free.
[1556792167] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1556792216] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556792287] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1556792338] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1556792338] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.5% (14381636 kB) free.
[1556792407] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1556793066] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.26, 9.70, 5.26
[1556793186] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.99, 6.67, 4.69
[1556793735] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4224024 kB) free!
[1556793795] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 14.1% (3462800 kB) free!
[1556793855] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 14.1% (3463308 kB) free!
[1556793915] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 14.1% (3465404 kB) free!
[1556793922] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4777088 kB) free!
[1556794042] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4737920 kB) free!
[1556794162] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4714096 kB) free!
[1556794704] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 592 processes
[1556794716] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3249648 kB) free!
[1556794764] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 596 processes
[1556794836] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3180944 kB) free!
[1556794956] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.1% (3103556 kB) free!
[1556795115] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.4% (2301260 kB) free!
[1556795547] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (3643420 kB) free!
[1556795667] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3636052 kB) free!
[1556795787] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (3516304 kB) free!
[1556796112] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.2% (22847060 kB) free.
[1556796190] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.7% (18310312 kB) free.
[1556796259] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.4% (13393744 kB) free.
[1556796311] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.7% (23623660 kB) free.
[1556796362] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 59.7% (14634120 kB) free.
[1556798364] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 596 processes
[1556798424] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 597 processes
[1556798476] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556798484] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 599 processes
[1556798544] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 0 warn out of 600 processes
[1556798766] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4821000 kB) free!
[1556798826] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3660728 kB) free!
[1556798886] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 13.3% (3256404 kB) free!
[1556798946] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 13.2% (3241048 kB) free!
[1556799023] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (127752 kB) free!
[1556799143] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (133340 kB) free!
[1556799144] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 614 processes
[1556799151] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1556799264] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (120816 kB) free!
[1556799271] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1556799391] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1556799864] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 40.4% (2970104 kB) free.
[1556799991] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1556800738] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1556801725] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1556801817] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1556802417] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1556802538] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556802658] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1556803060] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 59.6% (14616068 kB) free.
[1556806258] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1556806859] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1556806979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556806995] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556807099] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1556807699] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556807819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556807939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556808042] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1556808162] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1556809739] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1556810339] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556810459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556810579] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556810939] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3119980 kB) free!
[1556811059] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.0% (3086488 kB) free!
[1556811162] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1556811179] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (3061088 kB) free!
[1556811282] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1556812379] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1556812980] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556813100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556813220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556815020] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1556815620] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556815740] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556815861] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556816517] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4863692 kB) free!
[1556816637] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4855984 kB) free!
[1556816757] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4851928 kB) free!
[1556817661] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1556818261] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1556818381] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556818501] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1556820301] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1556820901] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556821021] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556821141] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1556821379] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1611188 kB) free!
[1556821741] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes