[1563474757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563474803] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7478384 kB) free!
[1563474878] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563474923] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11939776 kB) free.
[1563475096] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172624 kB) free!
[1563475219] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185264 kB) free!
[1563475306] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 66.5% (16416692 kB) free.
[1563475339] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201176 kB) free!
[1563475367] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.5% (20723496 kB) free.
[1563475392] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187556 kB) free!
[1563475433] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39780596 kB) free.
[1563475477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563475498] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.9% (20802292 kB) free.
[1563475523] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7160248 kB) free!
[1563475554] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.9% (14777976 kB) free.
[1563475597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563475643] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11949660 kB) free.
[1563475911] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168576 kB) free!
[1563476031] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175232 kB) free!
[1563476151] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185172 kB) free!
[1563476197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563476244] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6987480 kB) free!
[1563476317] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563476364] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11947416 kB) free.
[1563476917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563476964] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7352356 kB) free!
[1563477037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563477084] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11940812 kB) free.
[1563477637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563477684] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7246744 kB) free!
[1563477706] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563477758] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563477804] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11947072 kB) free.
[1563478312] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 538 processes
[1563478357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563478372] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 531 processes
[1563478404] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7048952 kB) free!
[1563478417] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180152 kB) free!
[1563478477] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563478477] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182404 kB) free!
[1563478524] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11938300 kB) free.
[1563478537] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (182904 kB) free!
[1563478597] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (177984 kB) free!
[1563478967] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1563478977] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.61, 4.23, 4.70
[1563479077] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563479087] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1563479124] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7280860 kB) free!
[1563479197] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563479244] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11945108 kB) free.
[1563479797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563479844] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7631268 kB) free!
[1563479917] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563479964] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11949064 kB) free.
[1563480517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563480564] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7548012 kB) free!
[1563480637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563480684] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11940012 kB) free.
[1563481238] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563481284] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7815592 kB) free!
[1563481358] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563481404] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11950028 kB) free.
[1563481958] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563482004] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8011816 kB) free!
[1563482078] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563482124] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11937080 kB) free.
[1563482204] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (153168 kB) free!
[1563482324] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176868 kB) free!
[1563482444] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156744 kB) free!
[1563482678] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563482724] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8147728 kB) free!
[1563482798] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563482844] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11947536 kB) free.
[1563482918] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563483444] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8267872 kB) free!
[1563483519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563483564] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11946472 kB) free.
[1563483639] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563484164] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8191488 kB) free!
[1563484239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563484284] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11940516 kB) free.
[1563484359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563484884] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8228068 kB) free!
[1563484959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563485005] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11945920 kB) free.
[1563485079] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563485463] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 340.11, 92.37, 31.49
[1563485583] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 46.20, 61.94, 27.72
[1563485604] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8164148 kB) free!
[1563485679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563485703] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.32, 41.51, 24.38
[1563485724] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11938612 kB) free.
[1563485799] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563486054] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4906936 kB) free!
[1563486174] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4899376 kB) free!
[1563486294] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4890224 kB) free!
[1563486303] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.10, 5.68, 12.85
[1563486324] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8287484 kB) free!
[1563486400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563486444] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11944884 kB) free.
[1563486520] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563486951] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563487044] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8322652 kB) free!
[1563487120] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563487164] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11949144 kB) free.
[1563487240] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563487503] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.16, 0.25, 3.66
[1563487764] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8316796 kB) free!
[1563487840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563487884] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11939128 kB) free.
[1563487960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563488484] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8447268 kB) free!
[1563488560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563488604] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11946088 kB) free.
[1563488680] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563489204] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8619192 kB) free!
[1563489280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563489325] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6865440 kB) free!
[1563489400] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563489445] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11944076 kB) free.
[1563490001] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563490045] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7087328 kB) free!
[1563490121] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563490165] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11945036 kB) free.
[1563490721] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563490765] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7378540 kB) free!
[1563490841] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563490885] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11933788 kB) free.
[1563491441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563491485] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7121500 kB) free!
[1563491561] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563491605] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11936344 kB) free.
[1563492161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563492205] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6818588 kB) free!
[1563492281] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563492325] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11927360 kB) free.
[1563492882] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563492925] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6968636 kB) free!
[1563493002] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563493045] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11935216 kB) free.
[1563493602] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563493646] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7214540 kB) free!
[1563493722] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563493766] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11460328 kB) free.
[1563494322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563494442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563494562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563494966] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8810660 kB) free!
[1563495086] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7140368 kB) free!
[1563495163] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563495206] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11548572 kB) free.
[1563495763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563495777] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14540.14, 12420.36, 6714.13
[1563495883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563495887] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1563495897] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14541.21, 13122.13, 7663.16
[1563496003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563496007] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1563496017] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14541.98, 13591.90, 8497.20
[1563496406] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8720176 kB) free!
[1563496526] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7163468 kB) free!
[1563496603] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563496603] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563496646] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11523748 kB) free.
[1563496890] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1255.02, 949.08, 459.14
[1563497010] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1255.73, 1050.59, 555.75
[1563497130] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 220.14, 747.17, 506.67
[1563497203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563497323] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563497443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563497444] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 50.6% (8219224 kB) free.
[1563497703] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1124.45, 474.87, 178.73
[1563497823] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1219.35, 726.15, 306.73
[1563497846] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (8996216 kB) free!
[1563497943] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1232.11, 894.23, 419.17
[1563497967] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7629484 kB) free!
[1563498044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563498044] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (154080 kB) free!
[1563498087] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11855920 kB) free.
[1563498164] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (182640 kB) free!
[1563498284] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169868 kB) free!
[1563498884] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 59.5% (9663560 kB) free.
[1563499244] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563499364] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563499887] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8078996 kB) free!
[1563499964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563500007] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11859288 kB) free.
[1563500084] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563500607] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7965956 kB) free!
[1563500727] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11850376 kB) free.
[1563500730] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.09, 2.98, 13.29
[1563500943] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.25, 0.29, 19.06
[1563501284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563501327] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8081092 kB) free!
[1563501404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563501447] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11856576 kB) free.
[1563501525] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563502017] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 0.41, 14.70
[1563502306] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (8646052 kB) free!
[1563502426] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (5226024 kB) free!
[1563502530] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.91, 2.91, 4.44
[1563502546] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (14423692 kB) free.
[1563502648] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8328360 kB) free!
[1563502724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563502743] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.36, 3.03
[1563502768] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11858184 kB) free.
[1563502844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563502964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563503147] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (9400828 kB) free!
[1563503217] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.01, 0.62, 4.42
[1563503267] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (6007308 kB) free!
[1563503387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (15236232 kB) free.
[1563503987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (12934584 kB) free!
[1563504107] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (9541216 kB) free!
[1563504164] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563504227] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (6285716 kB) free!
[1563504358] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181872 kB) free!
[1563504478] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176812 kB) free!
[1563504568] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8527328 kB) free!
[1563504598] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (187472 kB) free!
[1563504688] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6774432 kB) free!
[1563504808] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11855208 kB) free.
[1563504827] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (15553816 kB) free.
[1563504838] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239412 kB) free!
[1563504958] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (249956 kB) free!
[1563505078] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (274280 kB) free!
[1563505365] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563505408] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6966604 kB) free!
[1563505427] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11856304 kB) free!
[1563505485] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563505528] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11844796 kB) free.
[1563505547] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (8490716 kB) free!
[1563505667] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (5051944 kB) free!
[1563506085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563506128] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6740748 kB) free!
[1563506205] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563506248] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11855004 kB) free.
[1563506267] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563506267] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (14974472 kB) free.
[1563506805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563506867] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11555944 kB) free!
[1563506925] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563506987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (8204676 kB) free!
[1563507107] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (4797876 kB) free!
[1563507449] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8255964 kB) free!
[1563507525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563507569] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11854696 kB) free.
[1563507645] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563507707] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (16116896 kB) free.
[1563508169] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8365896 kB) free!
[1563508246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563508289] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11844676 kB) free.
[1563508366] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563508607] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1563508727] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1563508889] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8216496 kB) free!
[1563508908] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (8687216 kB) free!
[1563508966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563509009] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11852784 kB) free.
[1563509028] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (4615492 kB) free!
[1563509086] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563509148] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (15346828 kB) free.
[1563509609] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8256092 kB) free!
[1563509644] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.64, 5.88, 5.21
[1563509686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563509729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11843108 kB) free.
[1563509748] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (11168408 kB) free!
[1563509764] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.61, 5.92, 5.29
[1563509806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563509868] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (7107232 kB) free!
[1563509884] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.73, 6.09, 5.42
[1563509988] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (3010200 kB) free!
[1563510329] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8177296 kB) free!
[1563510406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563510449] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11848868 kB) free.
[1563510526] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563510588] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (13730692 kB) free.
[1563510884] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173820 kB) free!
[1563511004] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (175340 kB) free!
[1563511049] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8343708 kB) free!
[1563511124] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180876 kB) free!
[1563511126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563511127] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 13.0% (2119116 kB) free!
[1563511169] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11849640 kB) free.
[1563511188] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (8404612 kB) free!
[1563511247] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563511247] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 11.2% (1815916 kB) free!
[1563511308] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (4213184 kB) free!
[1563511367] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174940 kB) free!
[1563511368] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1547420 kB) free!
[1563511428] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (261152 kB) free!
[1563511487] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180272 kB) free!
[1563511607] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189144 kB) free!
[1563511769] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8260444 kB) free!
[1563511846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563511890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11841056 kB) free.
[1563511966] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563512028] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (12573184 kB) free!
[1563512490] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8378912 kB) free!
[1563512566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563512610] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11719024 kB) free.
[1563512686] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563513018] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.91, 5.15, 5.10
[1563513069] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.9% (20445948 kB) free.
[1563513115] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 83.5% (20470236 kB) free.
[1563513138] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.23, 5.00, 5.06
[1563513184] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39807136 kB) free.
[1563513210] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8340428 kB) free!
[1563513228] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (3661812 kB) free!
[1563513244] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 81.9% (13310396 kB) free.
[1563513258] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.94, 5.33, 5.20
[1563513286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563513317] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.2% (14830008 kB) free.
[1563513330] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11838308 kB) free.
[1563513354] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 64.5% (10477368 kB) free.
[1563513406] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563513412] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 62.3% (15279972 kB) free.
[1563513473] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 50.2% (32982784 kB) free.
[1563513736] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179884 kB) free!
[1563513856] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189016 kB) free!
[1563513930] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8389556 kB) free!
[1563513959] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (156012 kB) free!
[1563513976] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184484 kB) free!
[1563514006] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563514050] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11842884 kB) free.
[1563514079] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (175092 kB) free!
[1563514127] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563514199] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (147672 kB) free!
[1563514650] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8534668 kB) free!
[1563514677] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11230136 kB) free!
[1563514726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563514770] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6744212 kB) free!
[1563514797] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (6968304 kB) free!
[1563514846] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563514879] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189772 kB) free!
[1563514890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11835036 kB) free.
[1563514917] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (21091968 kB) free.
[1563514999] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192760 kB) free!
[1563515119] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186004 kB) free!
[1563515446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563515566] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563516091] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8315380 kB) free!
[1563516117] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (10711000 kB) free!
[1563516166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563516211] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11842008 kB) free.
[1563516237] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (6177852 kB) free!
[1563516247] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 11.9% (1929860 kB) free!
[1563516286] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563516323] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (128420 kB) free!
[1563516357] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (20428120 kB) free.
[1563516367] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (1395728 kB) free!
[1563516389] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1563516417] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.95, 6.78, 5.34
[1563516443] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (120096 kB) free!
[1563516487] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (934128 kB) free!
[1563516509] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1563516537] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.08, 7.19, 5.66
[1563516563] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (120856 kB) free!
[1563516629] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1563516657] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.22, 7.23, 5.87
[1563516811] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8151628 kB) free!
[1563516886] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563516917] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.1% (2286388 kB) free!
[1563516931] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11830608 kB) free.
[1563516957] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12327416 kB) free!
[1563517007] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563517007] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563517038] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (153672 kB) free!
[1563517077] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (7839084 kB) free!
[1563517111] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4869124 kB) free!
[1563517113] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2104800 kB) free!
[1563517158] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174384 kB) free!
[1563517197] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (3511544 kB) free!
[1563517231] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.7% (4825904 kB) free!
[1563517233] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1897280 kB) free!
[1563517351] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.6% (4795836 kB) free!
[1563517353] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.5% (1707028 kB) free!
[1563517463] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.35, 3.95, 4.77
[1563517551] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1563517797] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (11715140 kB) free!
[1563517954] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 85.4% (13888076 kB) free.
[1563517990] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (222348 kB) free!
[1563518110] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (254752 kB) free!
[1563518131] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7146696 kB) free!
[1563518231] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214132 kB) free!
[1563518251] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11833312 kB) free.
[1563518364] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6830364 kB) free.
[1563518806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563518851] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7285348 kB) free!
[1563518926] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563518971] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11836348 kB) free.
[1563519252] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.7% (20909892 kB) free.
[1563519301] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1563519301] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 80.6% (19766908 kB) free.
[1563519355] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.7% (38760208 kB) free.
[1563519402] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1563519402] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.9% (13468872 kB) free.
[1563519454] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.6% (14398616 kB) free.
[1563519497] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1563519497] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 63.6% (10342756 kB) free.
[1563519526] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563519571] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7667624 kB) free!
[1563519597] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (4012248 kB) free!
[1563519647] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563519691] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11825348 kB) free.
[1563519858] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (2296948 kB) free!
[1563519907] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166128 kB) free!
[1563519960] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.9% (5357168 kB) free!
[1563519978] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171308 kB) free!
[1563520027] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180008 kB) free!
[1563520058] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (174592 kB) free!
[1563520080] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (249596 kB) free!
[1563520098] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1563520098] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173924 kB) free!
[1563520147] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189408 kB) free!
[1563520178] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171328 kB) free!
[1563520200] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1563520200] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237288 kB) free!
[1563520246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563520291] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7512544 kB) free!
[1563520298] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1563520298] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (185560 kB) free!
[1563520366] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563520412] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11834664 kB) free.
[1563520463] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.49, 5.59, 5.12
[1563520583] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.77, 5.66, 5.20
[1563520697] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (716252 kB) free!
[1563520703] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.57, 5.70, 5.27
[1563520797] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (15603864 kB) free.
[1563520817] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 62.6% (10183124 kB) free.
[1563520966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563521011] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7436320 kB) free!
[1563521086] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563521131] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11822712 kB) free.
[1563521151] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1563521303] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1563521303] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.78, 4.32, 4.82
[1563521397] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12801864 kB) free!
[1563521429] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1563521469] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (2449652 kB) free!
[1563521517] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (8317440 kB) free!
[1563521589] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190752 kB) free!
[1563521637] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (4081316 kB) free!
[1563521686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563521709] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198148 kB) free!
[1563521731] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7144188 kB) free!
[1563521806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563521851] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11829968 kB) free.
[1563522309] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (2709644 kB) free!
[1563522406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563522451] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6950900 kB) free!
[1563522527] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1563522571] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11830052 kB) free.
[1563522837] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (15411416 kB) free.
[1563523126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563523171] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6789492 kB) free!
[1563523246] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563523291] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11819432 kB) free.
[1563523846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563523891] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7016268 kB) free!
[1563523966] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563524011] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11829444 kB) free.
[1563524037] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (10402756 kB) free!
[1563524157] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (6206236 kB) free!
[1563524208] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.4% (2348304 kB) free!
[1563524277] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (1919584 kB) free!
[1563524328] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 12.6% (2049148 kB) free!
[1563524449] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.0% (1787008 kB) free!
[1563524566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563524612] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6801896 kB) free!
[1563524686] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563524709] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171872 kB) free!
[1563524732] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11820264 kB) free.
[1563524877] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (15519324 kB) free.
[1563524964] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.3% (754652 kB) free!
[1563525049] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (474936 kB) free!
[1563525084] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (522864 kB) free!
[1563525204] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (286056 kB) free!
[1563525286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563525309] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1563525309] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 68.2% (16722460 kB) free.
[1563525407] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563525527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563525657] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.27, 2.96, 4.28
[1563525909] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.7% (3366144 kB) free!
[1563525910] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4739116 kB) free!
[1563525930] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.37, 6.28, 5.64
[1563526030] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.1% (3451428 kB) free!
[1563526030] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4671464 kB) free!
[1563526050] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.90, 5.83, 5.55
[1563526077] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (10177880 kB) free!
[1563526103] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.10, 6.26, 5.56
[1563526150] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4387984 kB) free!
[1563526150] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4564632 kB) free!
[1563526170] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.83, 5.20, 5.36
[1563526197] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (5896844 kB) free!
[1563526217] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1563526224] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.60, 5.75, 5.47
[1563526317] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1845628 kB) free!
[1563526344] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.43, 4.91, 5.19
[1563526770] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.75, 4.12, 4.69
[1563526944] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.34, 3.89, 4.50
[1563527517] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (15571468 kB) free.
[1563527950] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (234756 kB) free!
[1563528717] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (11777108 kB) free!
[1563528751] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2409612 kB) free!
[1563528804] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.4% (6794784 kB) free.
[1563528837] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (7279936 kB) free!
[1563528957] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2989704 kB) free!
[1563529150] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185648 kB) free!
[1563529150] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 50.6% (12411028 kB) free.
[1563529170] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.07, 5.35, 5.04
[1563529290] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.57, 5.46, 5.12
[1563529410] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.31, 5.38, 5.14
[1563529532] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7199124 kB) free!
[1563529652] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5917844 kB) free!
[1563529750] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186648 kB) free!
[1563529772] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4493880 kB) free!
[1563529870] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (210276 kB) free!
[1563529990] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (359884 kB) free!
[1563530157] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (15870832 kB) free.
[1563530328] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563530373] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11628012 kB) free.
[1563530544] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.27, 6.66, 5.32
[1563530590] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (2757272 kB) free!
[1563530610] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1563530610] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.10, 4.33, 5.00
[1563530664] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.86, 6.89, 5.59
[1563530757] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (11897004 kB) free!
[1563530784] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.24, 6.94, 5.76
[1563530877] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (7414100 kB) free!
[1563530899] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563530947] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563530997] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3150992 kB) free!
[1563531001] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563531057] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 479.90, 211.17, 81.08
[1563531099] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563531177] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 71.72, 147.49, 73.84
[1563531190] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 26.2% (6411580 kB) free.
[1563531210] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.49, 5.10, 5.06
[1563531297] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.37, 99.00, 65.00
[1563531330] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.20, 5.20, 5.11
[1563531450] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.86, 5.15, 5.11
[1563531754] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.3% (2649008 kB) free!
[1563531791] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (3782488 kB) free!
[1563531875] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.3% (2492988 kB) free!
[1563531911] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 64.4% (15793376 kB) free.
[1563531984] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 2.76, 3.49, 4.38
[1563531995] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.4% (2341496 kB) free!
[1563532128] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563532197] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (15523656 kB) free.
[1563532248] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563532368] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563532497] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 2.66, 18.61
[1563532511] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (254260 kB) free!
[1563532595] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 21.0% (3418944 kB) free.
[1563532631] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (214344 kB) free!
[1563532751] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (298224 kB) free!
[1563533195] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2603784 kB) free!
[1563533315] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.6% (2378604 kB) free!
[1563533351] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 61.6% (15094816 kB) free.
[1563533397] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (11322576 kB) free!
[1563533435] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.3% (2156756 kB) free!
[1563533517] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (7028068 kB) free!
[1563533637] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (3002840 kB) free!
[1563533927] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 362 processes
[1563534035] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563534035] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1331984 kB) free!
[1563534047] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1563534290] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.6% (21466184 kB) free.
[1563534297] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1563534297] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.22, 1.72, 3.70
[1563534328] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.7% (22639560 kB) free.
[1563534372] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 71.6% (17551160 kB) free.
[1563534411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 49.2% (32387064 kB) free.
[1563534464] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 30.1% (14830648 kB) free.
[1563534498] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.7% (13934284 kB) free.
[1563534550] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.3% (13373364 kB) free.
[1563534551] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (365264 kB) free!
[1563534592] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.3% (20670104 kB) free.
[1563534630] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 65.2% (15974628 kB) free.
[1563534677] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.8% (14431388 kB) free.
[1563534937] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (4017716 kB) free!
[1563534980] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192148 kB) free!
[1563535057] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (3948836 kB) free!
[1563535100] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188660 kB) free!
[1563535177] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (3853740 kB) free!
[1563535197] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (559228 kB) free!
[1563535220] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189900 kB) free!
[1563535238] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (317400 kB) free!
[1563535277] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162228 kB) free!
[1563535317] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (480708 kB) free!
[1563535358] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172572 kB) free!
[1563535398] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176504 kB) free!
[1563535404] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 12.5% (921068 kB) free!
[1563535437] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1563535437] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (401736 kB) free!
[1563535478] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (764524 kB) free!
[1563535518] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162784 kB) free!
[1563535525] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (663760 kB) free!
[1563535584] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.56, 5.93, 5.11
[1563535645] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (360216 kB) free!
[1563535668] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.1% (4974444 kB) free!
[1563535704] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.02, 6.02, 5.25
[1563535789] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4332352 kB) free!
[1563535825] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.18, 5.99, 5.32
[1563535909] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3672864 kB) free!
[1563536078] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (2795408 kB) free!
[1563536220] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11668624 kB) free!
[1563536340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (7503340 kB) free!
[1563536373] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8890596 kB) free!
[1563536378] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179384 kB) free!
[1563536460] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (3433644 kB) free!
[1563536493] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7587128 kB) free!
[1563536568] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1563536568] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563536613] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11752144 kB) free.
[1563536678] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 61.5% (15063652 kB) free.
[1563537025] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.82, 3.92, 4.45
[1563537060] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (18806296 kB) free.
[1563537168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563537288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563537408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563537813] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9526756 kB) free!
[1563537878] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (2114224 kB) free!
[1563537933] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8347944 kB) free!
[1563537951] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1563537956] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1563537998] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (1902136 kB) free!
[1563538053] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6729456 kB) free!
[1563538118] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (2581348 kB) free!
[1563538706] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.5% (2848348 kB) free!
[1563538718] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203476 kB) free!
[1563538826] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.9% (2584088 kB) free!
[1563538946] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.8% (2244636 kB) free!
[1563539460] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11600060 kB) free!
[1563539546] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (270728 kB) free!
[1563539580] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (6602452 kB) free!
[1563539700] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (2066136 kB) free!
[1563539846] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6832604 kB) free.
[1563540351] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1563540471] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1563540591] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1563540625] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.26, 6.20, 5.24
[1563540745] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.12, 6.07, 5.31
[1563540865] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.29, 6.22, 5.47
[1563541008] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563541050] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.31, 3.92, 4.62
[1563541053] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11738304 kB) free.
[1563541118] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (4541916 kB) free!
[1563541608] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563541653] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7262876 kB) free!
[1563541718] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 55.3% (13556076 kB) free.
[1563541728] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563541773] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11741568 kB) free.
[1563541791] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1563542065] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.40, 3.81, 4.48
[1563542100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (18458768 kB) free.
[1563542156] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 9 users currently logged in
[1563542276] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 9 users currently logged in
[1563542318] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187628 kB) free!
[1563542328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563542373] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7180180 kB) free!
[1563542396] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1563542438] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (4117700 kB) free!
[1563542449] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563542494] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11397732 kB) free.
[1563542558] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 50.4% (12358632 kB) free.
[1563543048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563543168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563543288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563543450] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.59, 5.51, 5.25
[1563543570] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.84, 4.96, 5.08
[1563543620] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.8% (2637904 kB) free!
[1563543690] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.56, 4.50, 4.90
[1563543758] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (1023484 kB) free!
[1563543878] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (244964 kB) free!
[1563543900] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (877892 kB) free!
[1563543999] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1288844 kB) free!
[1563544020] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (280996 kB) free!
[1563544140] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (246052 kB) free!
[1563544220] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1563544220] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 46.4% (11360992 kB) free.
[1563544497] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563544761] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563545088] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563545340] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (14777024 kB) free.
[1563545493] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7884820 kB) free!
[1563545613] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11734816 kB) free.
[1563545884] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.28, 2.98, 4.17
[1563546213] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8084156 kB) free!
[1563546237] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563546289] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563546333] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11726096 kB) free.
[1563546409] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563546543] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 973.51, 547.80, 226.21
[1563546663] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 598.88, 599.95, 287.51
[1563546783] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 81.31, 401.61, 252.71
[1563546933] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8005848 kB) free!
[1563547009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563547053] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11736092 kB) free.
[1563547129] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563547290] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 11647.88, 9466.78, 4853.24
[1563547410] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 11650.37, 10189.38, 5677.42
[1563547530] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1574.12, 6818.93, 4989.39
[1563547653] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8237944 kB) free!
[1563547729] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563547774] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11733040 kB) free.
[1563547849] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563548373] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8174224 kB) free!
[1563548449] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563548493] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11729768 kB) free.
[1563548569] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563548940] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12136928 kB) free!
[1563549060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (8214180 kB) free!
[1563549093] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8306832 kB) free!
[1563549170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563549180] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (4378980 kB) free!
[1563549183] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 0.55, 19.43
[1563549213] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11733448 kB) free.
[1563549290] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563549813] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8536384 kB) free!
[1563549890] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563549933] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6867892 kB) free!
[1563550010] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563550053] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11731640 kB) free.
[1563550494] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 10.0% (2455340 kB) free!
[1563550610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563550653] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7305944 kB) free!
[1563550730] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563550773] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11730656 kB) free.
[1563550946] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563550983] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.27, 0.36, 3.13
[1563551097] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 445.09, 331.23, 144.24
[1563551199] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (3175960 kB) free!
[1563551217] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 63.77, 223.06, 127.27
[1563551330] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563551337] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.45, 150.69, 112.37
[1563551373] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7266008 kB) free!
[1563551450] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563551493] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11726016 kB) free.
[1563552051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563552094] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7462880 kB) free!
[1563552171] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563552214] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11729788 kB) free.
[1563552771] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563552814] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7671256 kB) free!
[1563552865] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563552891] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563552930] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.92, 3.05, 17.80
[1563552934] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11721972 kB) free.
[1563553137] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 2.03, 18.62
[1563553491] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563553534] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7442228 kB) free!
[1563553611] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563553654] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11726808 kB) free.
[1563554211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563554254] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7174400 kB) free!
[1563554331] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563554374] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11730680 kB) free.
[1563554931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563554937] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.32, 2.99
[1563554974] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7054008 kB) free!
[1563555052] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563555094] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11719016 kB) free.
[1563555330] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.74, 2.91, 4.05
[1563555620] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563555651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563555694] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7110380 kB) free!
[1563555771] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563555814] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11722880 kB) free.
[1563556371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563556415] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6921748 kB) free!
[1563556491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563556535] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11712592 kB) free.
[1563556980] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (17351208 kB) free.
[1563557091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563557135] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7151068 kB) free!
[1563557199] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563557211] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563557211] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563557255] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11722008 kB) free.
[1563557811] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563557838] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 88.4% (21674228 kB) free.
[1563557855] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7369908 kB) free!
[1563557900] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.7% (21899964 kB) free.
[1563557931] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563557966] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39780408 kB) free.
[1563557975] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11723092 kB) free.
[1563558026] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.1% (14001276 kB) free.
[1563558075] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.5% (20710192 kB) free.
[1563558124] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.5% (14702672 kB) free.
[1563558532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563558575] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7504936 kB) free!
[1563558652] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563558695] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11715256 kB) free.
[1563558780] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10856980 kB) free!
[1563558900] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (7157764 kB) free!
[1563559020] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (17335368 kB) free.
[1563559252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563559295] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7317928 kB) free!
[1563559372] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563559415] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11720648 kB) free.
[1563559972] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563560015] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7201344 kB) free!
[1563560092] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563560135] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11713632 kB) free.
[1563560220] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (12989020 kB) free!
[1563560337] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.20, 8.23, 8.13
[1563560341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (9530424 kB) free!
[1563560458] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.39, 5.65, 7.20
[1563560461] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (5654680 kB) free!
[1563560578] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1563560578] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.25, 3.86, 6.35
[1563560692] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563560735] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7554952 kB) free!
[1563560812] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563560856] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11718804 kB) free.