[1554229836] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (2803504 kB) free!
[1554229957] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (966456 kB) free!
[1554230128] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1554230469] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.6% (8690668 kB) free!
[1554230557] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.6% (5724044 kB) free!
[1554230728] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554230848] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554230968] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554231157] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10187320 kB) free.
[1554231669] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2161776 kB) free!
[1554231757] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3440836 kB) free!
[1554231877] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (2047556 kB) free!
[1554231997] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10184216 kB) free.
[1554232169] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554232269] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9491240 kB) free!
[1554232597] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5515364 kB) free!
[1554232717] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (4403508 kB) free!
[1554232769] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554232837] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3342852 kB) free!
[1554232889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554233009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554233437] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.2% (8947828 kB) free!
[1554233469] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3149780 kB) free!
[1554234037] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (728196 kB) free!
[1554234069] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.7% (9689248 kB) free!
[1554234209] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554234638] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5388104 kB) free!
[1554234809] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554234929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554235050] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554235238] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10178792 kB) free.
[1554235838] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3510988 kB) free!
[1554235958] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (2311596 kB) free!
[1554236078] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10186232 kB) free.
[1554236250] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554236678] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5166888 kB) free!
[1554236798] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (3972128 kB) free!
[1554236850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554236918] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2754660 kB) free!
[1554236970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554237090] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1554237518] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.8% (7787136 kB) free!
[1554237690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554237810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554237931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554238118] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10180924 kB) free.
[1554238718] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5221016 kB) free!
[1554238839] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (3963432 kB) free!
[1554238959] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2708696 kB) free!
[1554239131] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1554239559] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7924400 kB) free!
[1554239731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554239851] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554239971] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554240069] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3576304 kB) free!
[1554240159] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10182552 kB) free.
[1554240759] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4764756 kB) free!
[1554240879] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3475784 kB) free!
[1554240999] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2097812 kB) free!
[1554241171] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554241269] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.5% (8621448 kB) free!
[1554241599] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6856392 kB) free!
[1554241771] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554241892] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554242012] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554242199] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10159000 kB) free.
[1554242469] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4113324 kB) free!
[1554242799] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4182324 kB) free!
[1554242919] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (2903132 kB) free!
[1554243039] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (1142724 kB) free!
[1554243069] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.2% (6501588 kB) free!
[1554243212] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554243640] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5920884 kB) free!
[1554243812] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554243932] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554244052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554244240] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10155336 kB) free.
[1554244840] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3501092 kB) free!
[1554244869] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4491112 kB) free!
[1554244960] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (2143328 kB) free!
[1554245080] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10160948 kB) free.
[1554245253] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554245680] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5082816 kB) free!
[1554245800] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (3827448 kB) free!
[1554245853] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554245920] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2559632 kB) free!
[1554245973] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554246069] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9866024 kB) free!
[1554246093] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1554246520] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7333408 kB) free!
[1554246693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554246813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554246933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554247120] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10101432 kB) free.
[1554247269] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9894192 kB) free.
[1554247720] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4304140 kB) free!
[1554247841] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (3044808 kB) free!
[1554247961] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1396632 kB) free!
[1554248134] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1554248561] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6117272 kB) free!
[1554248734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554248854] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554248974] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554249161] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10092748 kB) free.
[1554249669] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3047932 kB) free!
[1554249761] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3558500 kB) free!
[1554249789] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9695624 kB) free!
[1554249881] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (2211264 kB) free!
[1554249909] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.5% (7150288 kB) free!
[1554250001] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10100340 kB) free.
[1554250276] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1516796 kB) free!
[1554250510] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4303956 kB) free!
[1554250601] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2368624 kB) free!
[1554250721] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (419920 kB) free!
[1554250774] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554250841] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.7% (9693052 kB) free!
[1554251374] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554251442] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (1133944 kB) free!
[1554251495] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554251615] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554251710] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.3% (7562764 kB) free!
[1554252310] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4376004 kB) free!
[1554252573] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24467 MB (10% inode=99%):
[1554253242] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7436384 kB) free!
[1554253510] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.0% (8354404 kB) free!
[1554253842] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9931580 kB) free.
[1554254442] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (1357792 kB) free!
[1554254562] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10135692 kB) free.
[1554254710] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2674344 kB) free!
[1554255153] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1587452 kB) free!
[1554255162] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2043648 kB) free!
[1554255283] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226556 kB) free!
[1554255310] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.4% (8554260 kB) free!
[1554255403] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.0% (9351060 kB) free!
[1554256003] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (896272 kB) free!
[1554256510] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4171168 kB) free!
[1554257015] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554257203] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.8% (8295096 kB) free!
[1554257615] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554257710] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.3% (7063200 kB) free!
[1554257735] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554257803] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10083636 kB) free.
[1554257855] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554258310] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3931672 kB) free!
[1554258403] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4865308 kB) free!
[1554258455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554258523] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3539620 kB) free!
[1554258576] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554258644] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2243808 kB) free!
[1554258696] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554259510] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.9% (7341424 kB) free!
[1554259844] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.2% (6482932 kB) free!
[1554260110] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4300064 kB) free!
[1554261044] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (628756 kB) free!
[1554261096] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1554261310] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.0% (7392888 kB) free!
[1554261696] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554261816] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554261910] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4852648 kB) free!
[1554261937] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1554262244] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.9% (5860756 kB) free!
[1554262536] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554262656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554262776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554263110] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.8% (7766716 kB) free!
[1554263444] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10090580 kB) free.
[1554264044] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2422948 kB) free!
[1554264164] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (426576 kB) free!
[1554264285] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.1% (9420552 kB) free!
[1554264310] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2144060 kB) free!
[1554264885] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (984920 kB) free!
[1554264910] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.7% (8242616 kB) free!
[1554266110] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (1724352 kB) free!
[1554266685] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7218720 kB) free!
[1554266711] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.3% (8526100 kB) free!
[1554267285] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10247096 kB) free.
[1554267885] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (2566524 kB) free!
[1554267911] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (2604200 kB) free!
[1554268006] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (1150496 kB) free!
[1554268126] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.6% (9654036 kB) free!
[1554268511] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.3% (8517052 kB) free!
[1554268726] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1281044 kB) free!
[1554269711] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3731000 kB) free!
[1554269922] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24207 MB (10% inode=99%):
[1554270526] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7162688 kB) free!
[1554270911] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.7% (7749632 kB) free!
[1554271126] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10034872 kB) free.
[1554271726] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2463660 kB) free!
[1554271846] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (548292 kB) free!
[1554271966] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.8% (9760904 kB) free!
[1554272111] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2080272 kB) free!
[1554272567] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1399072 kB) free!
[1554272711] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.8% (8775992 kB) free!
[1554273576] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1554273767] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5092932 kB) free!
[1554273912] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3303620 kB) free!
[1554274173] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 37395 MB (16% inode=99%):
[1554274176] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1554274276] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.0% (14467732 kB) free.
[1554274284] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.4% (3528508 kB) free!
[1554274776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554274896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554274967] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10152800 kB) free.
[1554275017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554275112] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.0% (7382984 kB) free!
[1554275567] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2043672 kB) free!
[1554275687] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (221640 kB) free!
[1554275712] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4938172 kB) free!
[1554275807] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.4% (9547380 kB) free!
[1554276407] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (1218548 kB) free!
[1554276912] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.4% (8101356 kB) free!
[1554277284] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2447728 kB) free!
[1554277608] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4979024 kB) free!
[1554278017] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1554278112] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4472696 kB) free!
[1554278617] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554278737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554278808] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10116896 kB) free.
[1554278858] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1554278858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554279312] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.2% (8972016 kB) free!
[1554279408] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (1684328 kB) free!
[1554279473] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46930 MB (20% inode=99%):
[1554279528] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10216580 kB) free.
[1554279593] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 58709 MB (26% inode=99%):
[1554280128] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (2104760 kB) free!
[1554280243] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 340 processes
[1554280248] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10220272 kB) free.
[1554280512] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4732252 kB) free!
[1554280848] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2457032 kB) free!
[1554280968] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (442960 kB) free!
[1554281088] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.3% (9494528 kB) free!
[1554281688] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (1043312 kB) free!
[1554281713] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.8% (8756424 kB) free!
[1554282522] SERVICE ALERT: pservercluster1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24570 MB (11% inode=99%):
[1554283488] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7029972 kB) free!
[1554283513] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1579904 kB) free!
[1554284113] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.5% (8610552 kB) free!
[1554284688] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1185772 kB) free!
[1554285043] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 338 processes
[1554285913] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3879888 kB) free!
[1554286315] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554286435] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554286488] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.3% (8040008 kB) free!
[1554286555] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554287088] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10136380 kB) free.
[1554287113] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.7% (8235704 kB) free!
[1554287688] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (2241728 kB) free!
[1554287808] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10129680 kB) free.
[1554287858] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1554288313] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2719768 kB) free!
[1554288408] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2672928 kB) free!
[1554288459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554288529] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (868220 kB) free!
[1554288579] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1554288592] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554288649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9897984 kB) free.
[1554288712] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554288832] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554288914] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.6% (9185916 kB) free!
[1554289179] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554289249] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (1902812 kB) free!
[1554289299] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554289369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10116636 kB) free.
[1554289419] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554289432] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554289722] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24277 MB (10% inode=99%):
[1554289969] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2639672 kB) free!
[1554290089] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (781340 kB) free!
[1554290114] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4877412 kB) free!
[1554290209] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.8% (9780904 kB) free!
[1554290809] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1399376 kB) free!
[1554291314] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.9% (9320924 kB) free!
[1554291522] SERVICE ALERT: pservercluster1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24279 MB (11% inode=99%):
[1554292009] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5838660 kB) free!
[1554292122] SERVICE FLAPPING ALERT: pservercluster1;HDD_Home status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1554292122] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24265 MB (10% inode=99%):
[1554292443] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 332 processes
[1554292514] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3151568 kB) free!
[1554293032] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554293114] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.8% (9250452 kB) free!
[1554293152] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554293209] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (341924 kB) free!
[1554293272] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554293714] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2444040 kB) free!
[1554293847] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 63.46, 28.47, 11.01
[1554293967] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 67.53, 41.62, 17.95
[1554294087] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 68.09, 50.43, 24.05
[1554294276] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2852848 kB) free!
[1554294314] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.9% (8825504 kB) free!
[1554294358] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 636 processes
[1554294687] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.15, 10.20, 16.37
[1554294819] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554295009] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7262984 kB) free!
[1554295419] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554295514] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2496660 kB) free!
[1554295539] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554295887] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.41, 4.67
[1554296115] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.7% (9236028 kB) free!
[1554296140] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554296209] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2023972 kB) free!
[1554296260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554296380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554296869] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 51.31, 34.20, 15.97
[1554296989] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 51.26, 39.86, 20.25
[1554297034] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5458 MB (10% inode=97%):
[1554297109] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.52, 42.88, 23.77
[1554297112] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 720 processes
[1554297232] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 485 processes
[1554297315] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3528744 kB) free!
[1554297355] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554297358] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 323 processes
[1554297409] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.7% (5770408 kB) free!
[1554297634] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7040 MB (13% inode=98%):
[1554297709] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.21, 6.89, 13.09
[1554298515] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.0% (7864876 kB) free!
[1554298609] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10126132 kB) free.
[1554298909] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.23, 1.39, 4.54
[1554299115] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4935576 kB) free!
[1554299209] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2502316 kB) free!
[1554299272] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554299329] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (605228 kB) free!
[1554299380] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1554299449] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.7% (9711464 kB) free!
[1554299981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554300037] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.92, 14.99, 6.05
[1554300049] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1281712 kB) free!
[1554300101] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554300157] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.03, 21.74, 9.61
[1554300221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554300277] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 35.29, 26.25, 12.73
[1554300315] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.4% (8553388 kB) free!
[1554300867] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2748296 kB) free!
[1554300877] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 7.20, 10.25
[1554301249] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5067132 kB) free!
[1554301515] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2523604 kB) free!
[1554302077] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.61, 0.62, 3.16
[1554302115] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.5% (9101856 kB) free!
[1554302449] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10011420 kB) free.
[1554303049] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1789024 kB) free!
[1554303169] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10130020 kB) free.
[1554303316] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3036096 kB) free!
[1554303522] SERVICE FLAPPING ALERT: pservercluster1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554303770] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (2334492 kB) free!
[1554303867] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1518164 kB) free!
[1554303876] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2203604 kB) free!
[1554303890] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (271984 kB) free!
[1554303916] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.5% (9099532 kB) free!
[1554304010] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.9% (9302244 kB) free!
[1554304610] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (807032 kB) free!
[1554305116] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3893160 kB) free!
[1554305393] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46877 MB (20% inode=99%):
[1554305513] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 58659 MB (26% inode=99%):
[1554305810] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5099572 kB) free!
[1554306221] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1554306316] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.6% (7689172 kB) free!
[1554306821] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554306916] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4508716 kB) free!
[1554306941] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554307010] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (476620 kB) free!
[1554307061] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554307662] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554307782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554307902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554308116] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.5% (8626944 kB) free!
[1554308210] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5124988 kB) free!
[1554309316] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2251568 kB) free!
[1554309411] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10001444 kB) free.
[1554309916] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.1% (8899632 kB) free!
[1554310011] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (1832576 kB) free!
[1554310131] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10092844 kB) free.
[1554310731] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2603584 kB) free!
[1554310851] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (701288 kB) free!
[1554310902] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1554310971] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.6% (9673904 kB) free!
[1554311116] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3266948 kB) free!
[1554311503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554311571] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (1495048 kB) free!
[1554311623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554311716] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.0% (8878836 kB) free!
[1554311743] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554312171] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6362952 kB) free!
[1554312343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554312463] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554312583] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554312771] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10070964 kB) free.
[1554312916] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3702108 kB) free!
[1554313371] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3745972 kB) free!
[1554313491] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (2482040 kB) free!
[1554313612] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (472580 kB) free!
[1554313783] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554314116] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.6% (7211212 kB) free!
[1554314212] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5318312 kB) free!
[1554314383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554314504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554314624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554314716] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3895016 kB) free!
[1554314812] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10069780 kB) free.
[1554315224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554315344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554315412] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3003596 kB) free!
[1554315464] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554315532] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (1395812 kB) free!
[1554315547] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 82.44, 38.28, 15.19
[1554315652] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10061864 kB) free.
[1554315667] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 87.05, 54.63, 23.98
[1554315787] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 87.25, 65.44, 31.65
[1554315916] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.5% (7146976 kB) free!