[1591468797] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 24.99, 40.25, 44.55
[1591470597] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 44.35, 45.28, 45.40
[1591471797] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 40.90, 38.11, 41.51
[1591471903] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 15.9% (3899456 kB) free!
[1591471963] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (744584 kB) free!
[1591472026] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (194980 kB) free!
[1591472086] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (179752 kB) free!
[1591472686] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 70.9% (17374856 kB) free.
[1591472916] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1591474045] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243104 kB) free!
[1591474165] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (218628 kB) free!
[1591474285] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214916 kB) free!
[1591474716] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.71, 5.77, 5.06
[1591474797] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1591474797] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 68.21, 55.83, 45.88
[1591474836] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.28, 5.53, 5.05
[1591474956] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.37, 5.42, 5.06
[1591475485] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 25.1% (12392096 kB) free.
[1591475515] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 219843 MB (20% inode=99%):
[1591475556] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.50, 4.79, 4.91
[1591475635] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 219843 MB (20% inode=99%):
[1591475755] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 219843 MB (20% inode=99%):
[1591475997] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 28.53, 36.64, 40.85
[1591478084] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 84.6% (41692300 kB) free.
[1591478289] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.7% (43222612 kB) free.
[1591479282] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(451) < 2020/06/07 03:01:32 [error] 12483#0: *93663549 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1591479363] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2830) < 203.212.70.170 [07/Jun/2020:03:03:32 +0530] TCP 502 0 0 180.000 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1591479387] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(351) < 103.16.101.72 [07/Jun/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1591479402] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591479483] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591479507] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591481397] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 27.60, 23.72, 27.55
[1591481997] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 35.53, 31.50, 30.09
[1591482117] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 33.37, 31.25, 30.12
[1591482237] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 39.00, 34.03, 31.26
[1591483437] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 21.71, 23.33, 27.80
[1591484839] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (204868 kB) free!
[1591486082] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.5% (9428356 kB) free.
[1591486682] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187040 kB) free!
[1591486802] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195748 kB) free!
[1591486922] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1272940 kB) free!
[1591487349] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 65.0% (15919996 kB) free.
[1591488837] SERVICE ALERT: mrq4;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 59.14, 33.54, 27.89
[1591488957] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 36.08, 32.56, 28.22
[1591489557] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 33.10, 34.56, 31.80
[1591489677] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 29.48, 32.24, 31.23
[1591489797] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 40.88, 34.29, 32.07
[1591490697] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 72.0% (17650984 kB) free.
[1591490755] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 230016 MB (21% inode=99%):
[1591491955] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 222703 MB (20% inode=99%):
[1591492075] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 222036 MB (20% inode=99%):
[1591492195] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 221426 MB (20% inode=99%):
[1591492795] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 286418 MB (26% inode=99%):
[1591492929] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 34.8% (45932780 kB) free.
[1591493397] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 59.96, 42.64, 38.46
[1591493997] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 39.70, 40.07, 38.94
[1591494367] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1591494458] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1591494487] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1591494549] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196900 kB) free!
[1591494578] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1591494597] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 56.07, 49.05, 43.11
[1591494607] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1591494669] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177252 kB) free!
[1591494789] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199580 kB) free!
[1591495197] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 24.24, 30.69, 35.95
[1591496194] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 78.0% (19107056 kB) free.
[1591496397] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 24.17, 26.79, 28.98
[1591497597] SERVICE ALERT: mrq4;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.53, 58.05, 47.10
[1591497717] SERVICE ALERT: mrq4;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 76.80, 59.38, 48.55
[1591497837] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 86.74, 67.11, 52.55
[1591500156] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.56, 5.73, 5.18
[1591500276] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.62, 4.87, 4.93
[1591502077] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.27, 6.06, 5.42
[1591502197] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.57, 6.05, 5.51
[1591502317] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.87, 5.24, 5.27
[1591506517] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.18, 4.37, 4.97
[1591508317] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.00, 5.55, 5.20
[1591508437] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.71, 6.39, 5.55
[1591508557] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.27, 6.31, 5.63
[1591508637] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1591509394] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.5% (4532040 kB) free!
[1591509514] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.8% (4357336 kB) free!
[1591509634] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.1% (4189916 kB) free!
[1591509671] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4445728 kB) free!
[1591509791] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4114400 kB) free!
[1591509897] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4734556 kB) free!
[1591509911] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.2% (3865760 kB) free!
[1591510017] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4601792 kB) free!
[1591510137] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4462640 kB) free!
[1591510834] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2274572 kB) free!
[1591511111] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (691128 kB) free!
[1591511337] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2400188 kB) free!
[1591511712] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.7% (16101208 kB) free.
[1591513957] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.87, 4.33, 4.74
[1591514557] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.00, 5.73, 5.28
[1591514677] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.74, 6.32, 5.56
[1591514797] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1591514797] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.62, 5.71, 5.44
[1591518608] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1591518728] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1591518848] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1591519089] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9873168 kB) free!
[1591519209] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9785772 kB) free!
[1591519329] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.7% (9687180 kB) free!
[1591519448] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1591519598] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 6.54, 4.76, 4.76
[1591521248] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1591521368] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 339 processes
[1591521398] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.43, 6.53, 5.50
[1591521488] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1591521518] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.75, 6.87, 5.75
[1591521638] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.29, 6.53, 5.75
[1591521912] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3607552 kB) free!
[1591522032] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3426248 kB) free!
[1591522152] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3287372 kB) free!
[1591522683] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 185.255.8.58 [07/Jun/2020:15:00:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1591522752] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1591522752] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.3% (16010264 kB) free.
[1591522803] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591523529] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3778912 kB) free!
[1591524056] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1881.10, 571.93, 200.39
[1591524176] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 667.52, 680.02, 295.32
[1591524489] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1591524489] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1591525838] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 5.21, 4.18, 4.95
[1591527907] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1297.58, 1008.15, 419.33
[1591528027] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 176.49, 675.13, 368.69
[1591528147] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 24.76, 452.12, 324.12
[1591528204] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 172.105.59.173 [07/Jun/2020:16:35:28 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1591528324] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591528747] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.60, 61.99, 170.55
[1591528924] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(23) < 172.105.59.173 [07/Jun/2020:16:45:09 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591529044] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591529644] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(44) < 172.105.59.173 [07/Jun/2020:16:57:37 +0530] TCP 502 0 0 180.002 "10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 60.001, 0.000"
[1591529764] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591529889] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1591530009] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1591530129] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1591530638] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.26, 4.93, 5.02
[1591530730] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1591530758] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.01, 4.46, 4.83
[1591532914] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.8% (6491548 kB) free!
[1591533034] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 21.9% (7178004 kB) free.
[1591533758] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1591534152] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591534834] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 20.0% (6554476 kB) free!
[1591534954] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.9% (6516896 kB) free!
[1591535074] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.8% (6487212 kB) free!
[1591537958] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.31, 5.14, 5.04
[1591538078] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.68, 5.03, 5.01
[1591538198] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.42, 4.87, 4.95
[1591538684] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9386056 kB) free!
[1591538799] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.04, 6.28, 5.46
[1591538804] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (9262536 kB) free!
[1591538919] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.09, 6.16, 5.51
[1591538924] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (9172924 kB) free!
[1591539039] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.99, 5.87, 5.49
[1591540239] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.59, 4.64, 4.95
[1591542130] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591542274] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (2475580 kB) free!
[1591544074] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 47.0% (15395856 kB) free.
[1591544211] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1591544332] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1591552124] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4860928 kB) free!
[1591554237] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 36.25, 36.19, 41.76