[1572462104] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1572462224] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1572466938] SERVICE FLAPPING ALERT: sigqscluster2;Total Processes;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1572469773] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 92.9% (45770500 kB) free.
[1572473306] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1572473426] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572483147] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.5% (21929804 kB) free.
[1572483549] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.8% (20784560 kB) free.
[1572483654] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.5% (21194124 kB) free.
[1572497240] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461912 kB) free!
[1572497458] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24339040 kB) free!
[1572497578] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37742500 kB) free.
[1572500133] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 92.6% (22694296 kB) free.
[1572500201] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22087948 kB) free.
[1572500340] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39983740 kB) free.
[1572500806] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163620 kB) free!
[1572500926] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172464 kB) free!
[1572500949] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (251952 kB) free!
[1572501046] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186220 kB) free!
[1572501069] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233304 kB) free!
[1572501189] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (221320 kB) free!
[1572502149] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4869456 kB) free!
[1572502269] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4599392 kB) free!
[1572502389] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4353028 kB) free!
[1572502978] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24160008 kB) free!
[1572503106] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (19054188 kB) free!
[1572503226] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14627180 kB) free!
[1572503555] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1572503624] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1572503675] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1572503744] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 386 processes
[1572504063] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.09, 133.94, 130.80
[1572504183] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.17, 90.71, 115.34
[1572504304] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.09, 61.68, 101.71
[1572504789] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1363700 kB) free!
[1572504875] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1572504960] SERVICE ALERT: primary-server;Log Status;CRITICAL;SOFT;1;NRPE: Unable to read output
[1572504995] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1572505080] SERVICE ALERT: primary-server;Log Status;CRITICAL;SOFT;2;NRPE: Unable to read output
[1572505103] SERVICE ALERT: primary-server;Log Status;CRITICAL;HARD;3;NRPE: Unable to read output
[1572505115] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1572505505] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.28, 5.46, 30.69
[1572505704] SERVICE ALERT: primary-server;Log Status;UNKNOWN;HARD;3;NRPE: Unable to read output
[1572506105] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.74, 6.05, 18.92
[1572506107] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.5% (25667152 kB) free!
[1572506226] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1572506227] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (29669404 kB) free.
[1572506915] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (255052 kB) free!
[1572506973] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9504068 kB) free!
[1572507035] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (263936 kB) free!
[1572507054] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (4591164 kB) free!
[1572507093] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9352016 kB) free!
[1572507155] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (285664 kB) free!
[1572507174] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.3% (4248616 kB) free!
[1572507213] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (8967548 kB) free!
[1572507294] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4128900 kB) free!
[1572508107] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.1% (21009768 kB) free.
[1572508166] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.8% (22248836 kB) free.
[1572508353] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39424988 kB) free.
[1572508439] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.4% (14688816 kB) free.
[1572508567] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 94.9% (46765468 kB) free.
[1572508695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 45.6% (60101332 kB) free.
[1572508711] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182876 kB) free!
[1572508738] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.3% (22368892 kB) free.
[1572508780] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.4% (21322464 kB) free.
[1572508960] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229856 kB) free!
[1572509080] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240520 kB) free!
[1572509200] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1572509200] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225796 kB) free!
[1572509380] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186504 kB) free!
[1572509500] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187148 kB) free!
[1572509620] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1572509620] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192300 kB) free!
[1572510747] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4586124 kB) free!
[1572510867] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4508736 kB) free!
[1572510936] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.9% (22676988 kB) free.
[1572510978] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 57.4% (28280524 kB) free.
[1572510988] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4452340 kB) free!
[1572511545] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191908 kB) free!
[1572511581] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236792 kB) free!
[1572511665] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183976 kB) free!
[1572511701] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225956 kB) free!
[1572511785] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199792 kB) free!
[1572511821] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215920 kB) free!
[1572512752] SERVICE ALERT: primary-server;Log Status;WARNING;HARD;3;NRPE: Unable to read output
[1572512868] SERVICE FLAPPING ALERT: primary-server;Log Status;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1572512868] SERVICE ALERT: primary-server;Log Status;OK;HARD;1;Log check data initialized...
[1572513253] SERVICE ALERT: primary-server;Log Status;CRITICAL;SOFT;1;(1786) < 103.209.99.7 [31/Oct/2019:14:44:12 +0530] TCP 200 17 43 0.639 "10.147.212.142:13000" "43" "17" "0.000"
[1572513314] SERVICE ALERT: primary-server;Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572513362] SERVICE FLAPPING ALERT: primary-server;Log Status;DISABLED; Flap detection has been disabled
[1572513478] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1572513598] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1572513988] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2359572 kB) free!
[1572515394] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1572515524] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2019/10/31 15:21:56 [error] 24133#0: *1450523 recv() failed (104: Connection reset by peer) while proxying connection, client: 139.59.21.208, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:121/0, bytes from/to upstream:0/121
[1572515644] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(8) < 2019/10/31 15:23:53 [error] 24131#0: *1450639 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572515705] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1045.40, 811.21, 400.38
[1572515764] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(8) < 2019/10/31 15:25:40 [error] 24130#0: *1450711 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572515825] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1046.37, 889.08, 478.73
[1572515945] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 237.88, 670.40, 451.41
[1572516380] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1572516380] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21229320 kB) free.
[1572516410] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.6% (6711796 kB) free!
[1572516413] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 55.2% (27181460 kB) free.
[1572516438] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.5% (23171252 kB) free.
[1572516721] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.75, 52.59, 197.88
[1572516986] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196236 kB) free!
[1572517013] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2596960 kB) free!
[1572517106] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194676 kB) free!
[1572517133] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231684 kB) free!
[1572517226] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176964 kB) free!
[1572517253] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237228 kB) free!
[1572517564] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572518300] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23974340 kB) free!
[1572518420] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28227236 kB) free.
[1572518870] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1572518990] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1572519021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (21959816 kB) free!
[1572519110] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1572519141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39557076 kB) free.
[1572520341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20914976 kB) free!
[1572520461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37953872 kB) free.
[1572520627] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (22817868 kB) free!
[1572520748] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.1% (30481344 kB) free.
[1572521661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19220260 kB) free!
[1572521781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36661088 kB) free.
[1572522981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17750280 kB) free!
[1572523101] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36016744 kB) free.
[1572523147] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26237428 kB) free!
[1572523267] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.4% (33438296 kB) free.
[1572523310] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1572523867] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25026060 kB) free!
[1572523987] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.4% (32129848 kB) free.
[1572524342] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4885084 kB) free!
[1572524462] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4799232 kB) free!
[1572524583] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (4712568 kB) free!
[1572524587] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.1% (23796116 kB) free!
[1572524707] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (28302968 kB) free.
[1572524901] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18619776 kB) free!
[1572525021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36408968 kB) free.
[1572526221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16629988 kB) free!
[1572526341] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34844264 kB) free.
[1572526983] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2419092 kB) free!
[1572527373] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4819804 kB) free!
[1572527426] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1572527493] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4747984 kB) free!
[1572527541] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15474580 kB) free!
[1572527613] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (4673860 kB) free!
[1572527661] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33752412 kB) free.
[1572527764] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 42.3% (20857792 kB) free.
[1572527811] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.7% (22235356 kB) free.
[1572527843] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.4% (23293252 kB) free.
[1572528110] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1572528861] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14900260 kB) free!
[1572528981] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32992912 kB) free.
[1572529643] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187988 kB) free!
[1572529763] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194664 kB) free!
[1572529883] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194076 kB) free!
[1572530181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14126316 kB) free!
[1572530301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32230036 kB) free.
[1572530613] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2358424 kB) free!
[1572532101] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15154704 kB) free!
[1572532222] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32360488 kB) free.
[1572532575] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1572533422] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14269868 kB) free!
[1572533542] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31716020 kB) free.
[1572534742] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12004048 kB) free!
[1572534862] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30941604 kB) free.
[1572536062] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12571496 kB) free!
[1572536182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30293928 kB) free.
[1572536775] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9287168 kB) free!
[1572536895] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9287228 kB) free!
[1572537015] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (9288584 kB) free!
[1572537078] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.8% (23137344 kB) free.
[1572537138] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 77.1% (18894288 kB) free.
[1572537382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13346516 kB) free!
[1572537502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29753620 kB) free.
[1572537971] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9692324 kB) free!
[1572538091] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9637224 kB) free!
[1572538211] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9582960 kB) free!
[1572539302] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14625444 kB) free!
[1572539422] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30155584 kB) free.
[1572540622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13891264 kB) free!
[1572540742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29611368 kB) free.
[1572541942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14578304 kB) free!
[1572542062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29839856 kB) free.
[1572543262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15392124 kB) free!
[1572543307] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7101636 kB) free!
[1572543382] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30174820 kB) free.
[1572543427] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1104808 kB) free!
[1572543547] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2769396 kB) free!
[1572543982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13866800 kB) free!
[1572544102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28768900 kB) free.
[1572544702] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13527252 kB) free!
[1572544822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28475120 kB) free.
[1572546022] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16503168 kB) free!
[1572546142] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13726760 kB) free!
[1572546262] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28560412 kB) free.