[1563215408] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563215812] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.85, 3.73, 6.96
[1563215947] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7490196 kB) free!
[1563216008] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563216067] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194664 kB) free!
[1563216067] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10890160 kB) free.
[1563216128] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563216187] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166068 kB) free!
[1563216267] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8394360 kB) free!
[1563216307] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177100 kB) free!
[1563216387] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (8204388 kB) free!
[1563216507] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (8197836 kB) free!
[1563216668] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7371460 kB) free!
[1563216728] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563216788] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10898448 kB) free.
[1563216848] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563217388] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7210672 kB) free!
[1563217449] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563217508] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10888860 kB) free.
[1563217569] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563218108] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7007300 kB) free!
[1563218160] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563218169] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563218228] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10896792 kB) free.
[1563218289] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563218814] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21222548 kB) free.
[1563218828] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7027260 kB) free!
[1563218866] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.8% (20783536 kB) free.
[1563218889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563218948] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10894980 kB) free.
[1563219009] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563219055] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1563219175] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1563219190] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.1% (13505328 kB) free.
[1563219243] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180616 kB) free!
[1563219262] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.4% (21420836 kB) free.
[1563219303] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179080 kB) free!
[1563219319] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 62.9% (15414136 kB) free.
[1563219365] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.6% (14720312 kB) free.
[1563219412] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.47, 3.35
[1563219548] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7083664 kB) free!
[1563219609] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563219668] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10890988 kB) free.
[1563219729] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563220268] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7232444 kB) free!
[1563220329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563220388] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10895904 kB) free.
[1563220450] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563220612] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.68, 11.19, 8.38
[1563220732] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.55, 7.64, 7.42
[1563220852] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 5.25, 6.57
[1563220989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7428176 kB) free!
[1563221049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563221109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10888680 kB) free.
[1563221169] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563221709] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7481304 kB) free!
[1563221769] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563221829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10895112 kB) free.
[1563221889] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563222429] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7348256 kB) free!
[1563222489] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563222549] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10898900 kB) free.
[1563222609] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1563222616] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1179136 kB) free!
[1563222874] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.1% (2784692 kB) free!
[1563222960] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 41.59, 32.42, 19.69
[1563222994] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 28.3% (4602000 kB) free.
[1563223149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7167064 kB) free!
[1563223209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563223269] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10890472 kB) free.
[1563223330] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563223495] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.29, 49.42, 39.23
[1563223618] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174704 kB) free!
[1563223738] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (468736 kB) free!
[1563223858] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (439712 kB) free!
[1563223869] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6926412 kB) free!
[1563223929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563223989] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10897300 kB) free.
[1563224049] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563224160] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.79, 8.42, 13.99
[1563224452] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 40.85, 24.85, 12.98
[1563224589] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6589772 kB) free!
[1563224709] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10884276 kB) free.
[1563225249] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563225295] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.00, 14.41, 17.27
[1563225309] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6179872 kB) free!
[1563225369] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563225430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10892984 kB) free.
[1563225652] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 3.81, 7.96
[1563225969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563226029] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6048816 kB) free!
[1563226089] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563226149] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10892960 kB) free.
[1563226690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563226749] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6090964 kB) free!
[1563226810] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563226869] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10888380 kB) free.
[1563227410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563227469] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6003376 kB) free!
[1563227530] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563227589] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10891596 kB) free.
[1563227695] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.27, 1.29, 3.58
[1563228130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563228250] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563228295] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.60, 4.43, 6.55
[1563228360] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.37, 25.80, 21.91
[1563228415] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 65.67, 29.30, 15.29
[1563228535] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.12, 20.28, 13.67
[1563228616] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563228789] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7379368 kB) free!
[1563228850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563228909] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10882404 kB) free.
[1563228970] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563229135] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 39.60, 23.29, 15.11
[1563229509] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7099660 kB) free!
[1563229571] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563229630] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10892928 kB) free.
[1563229691] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563229735] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.43, 3.54, 8.15
[1563229852] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 102.93, 74.29, 39.28
[1563230160] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1563230160] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.92, 7.24, 14.14
[1563230229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6823420 kB) free!
[1563230349] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10894448 kB) free.
[1563230891] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563230907] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4133400 kB) free!
[1563230949] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6510820 kB) free!
[1563231011] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563231069] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10885084 kB) free.
[1563231186] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22501676 kB) free.
[1563231237] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39757020 kB) free.
[1563231287] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 69.8% (11341220 kB) free.
[1563231611] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563231652] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.00, 11.99, 19.81
[1563231669] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6566648 kB) free!
[1563231731] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563231789] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10892616 kB) free.
[1563231960] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.20, 24.29, 20.40
[1563232331] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563232389] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6756832 kB) free!
[1563232452] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563232509] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10886864 kB) free.
[1563232735] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.72, 26.41, 17.32
[1563233051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563233109] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6865940 kB) free!
[1563233171] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563233229] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10890400 kB) free.
[1563233291] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563233335] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.48, 4.01, 9.32
[1563233829] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6945432 kB) free!
[1563233891] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563233950] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10895900 kB) free.
[1563234011] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563234472] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1068 processes
[1563234549] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6974708 kB) free!
[1563234593] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1068 processes
[1563234611] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563234652] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.39, 6.88, 6.60
[1563234669] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10884744 kB) free.
[1563234713] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1065 processes
[1563234731] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563235252] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 4.79, 7.92
[1563235269] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6724948 kB) free!
[1563235313] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 386 processes
[1563235389] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10766816 kB) free.
[1563235735] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.53, 1.41, 4.58
[1563235932] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563235989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9452440 kB) free!
[1563236052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563236109] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8059228 kB) free!
[1563236172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563236229] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6616360 kB) free!
[1563236829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10859488 kB) free.
[1563236935] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.95, 14.02, 8.33
[1563237055] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.98, 9.47, 7.35
[1563237175] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 6.47, 6.51
[1563237429] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9462600 kB) free!
[1563237549] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8171168 kB) free!
[1563237669] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6850332 kB) free!
[1563237775] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1563237895] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.5% (1874264 kB) free!
[1563237895] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1563237960] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.66, 7.22, 12.97
[1563238015] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 50.0% (8127956 kB) free.
[1563238252] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.56, 26.08, 17.42
[1563238270] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563238270] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10831008 kB) free.
[1563238852] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.70, 4.13, 9.45
[1563238870] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9395840 kB) free!
[1563238990] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8102572 kB) free!
[1563239110] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6693092 kB) free!
[1563239575] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.55, 0.89, 3.92
[1563239710] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10798444 kB) free.
[1563239773] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563240175] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.27, 5.38, 7.03
[1563240295] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.31, 3.72, 6.22
[1563240310] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6632688 kB) free!
[1563240415] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 43.24, 23.03, 13.13
[1563240430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10807516 kB) free.
[1563240973] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563241015] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.99, 3.80, 7.24
[1563241030] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6690576 kB) free!
[1563241093] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563241150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10806952 kB) free.
[1563241693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563241750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6686640 kB) free!
[1563241813] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563241870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10803144 kB) free.
[1563242413] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563242470] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6813388 kB) free!
[1563242533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563242591] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10803236 kB) free.
[1563242654] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563243052] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 65.49, 22.89, 12.09
[1563243190] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6899888 kB) free!
[1563243253] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563243310] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10795600 kB) free.
[1563243373] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563243493] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563243652] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 3.42, 6.50
[1563243960] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563244253] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.71, 24.11, 27.28
[1563244853] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.52, 8.63, 19.33
[1563245047] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (226444 kB) free!
[1563245167] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231368 kB) free!
[1563245170] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173896 kB) free!
[1563245215] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.67, 9.48, 8.89
[1563245287] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243780 kB) free!
[1563245290] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (170696 kB) free!
[1563245410] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168500 kB) free!
[1563245815] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.87, 4.30, 8.07
[1563245815] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (2497136 kB) free!
[1563245893] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1563245893] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563245935] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.8% (2569904 kB) free!
[1563246055] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2542484 kB) free!
[1563246310] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6102700 kB) free!
[1563246360] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 93.01, 38.16, 22.58
[1563246430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10796652 kB) free.
[1563246653] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.56, 4.63
[1563246655] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174320 kB) free!
[1563247030] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6501192 kB) free!
[1563247150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10801040 kB) free.
[1563247253] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 31.46, 23.24, 11.79
[1563247373] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.72, 15.71, 10.42
[1563247493] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.83, 10.59, 9.18
[1563247560] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.96, 8.80, 14.81
[1563247693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563247750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6667880 kB) free!
[1563247813] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563247871] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10790756 kB) free.
[1563249013] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563249070] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6208196 kB) free!
[1563249133] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563249190] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10793852 kB) free.
[1563249293] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.72, 3.80
[1563249360] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1563249360] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7.66, 61.31, 62.78
[1563249733] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563249790] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6409196 kB) free!
[1563249854] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563249893] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.64, 4.34, 6.61
[1563249910] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10788408 kB) free.
[1563250013] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.55, 3.06, 5.86
[1563250133] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.72, 2.29, 5.24
[1563250453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563250510] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6520584 kB) free!
[1563250573] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563250615] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 1.23, 4.98
[1563250630] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10790904 kB) free.
[1563251173] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563251215] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.59, 16.62, 9.87
[1563251230] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6451872 kB) free!
[1563251293] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563251333] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.43, 0.99, 4.03
[1563251335] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.16, 11.27, 8.73
[1563251350] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10795588 kB) free.
[1563251455] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.81, 7.75, 7.75
[1563251760] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.86, 4.42, 13.16
[1563251893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563251933] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.03, 11.02, 8.42
[1563251950] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6249188 kB) free!
[1563252013] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563252053] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.57, 7.50, 7.45
[1563252055] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 1.55, 4.37
[1563252070] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10469664 kB) free.
[1563252173] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 5.18, 6.60
[1563252671] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563252677] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.7% (4818532 kB) free!
[1563252773] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.25, 1.03, 3.67
[1563252797] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.6% (4795772 kB) free!
[1563252917] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.4% (4760324 kB) free!
[1563252925] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 364 processes
[1563252962] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.39, 5.75, 5.44
[1563253045] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1563253082] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.71, 5.54, 5.41
[1563253202] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.04, 5.61, 5.44
[1563253373] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.97, 18.73, 10.61
[1563253493] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.90, 12.71, 9.39
[1563253613] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.55, 8.60, 8.29
[1563253813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563253870] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6448096 kB) free!
[1563253933] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563253990] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10789216 kB) free.
[1563254001] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 12.8% (2086172 kB) free!
[1563254121] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 11.3% (1842180 kB) free!
[1563254213] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 1.53, 4.56
[1563254241] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1615816 kB) free!
[1563254455] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.46, 7.67, 6.62
[1563254533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563254575] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.02, 5.50, 5.96
[1563254590] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6055752 kB) free!
[1563254653] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563254695] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.74, 3.99, 5.36
[1563254710] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10774664 kB) free.
[1563254814] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.15, 15.28, 9.41
[1563254934] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.64, 10.37, 8.32
[1563254959] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1563255054] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 7.05, 7.35
[1563255079] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1563255199] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1563255253] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563255373] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563255493] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563255654] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.43, 1.33, 4.06
[1563255895] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.69, 1.59, 4.70
[1563255895] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1563255910] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9707668 kB) free!
[1563256015] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1563256030] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8366260 kB) free!
[1563256150] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (6983800 kB) free!
[1563256399] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1563256750] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10791932 kB) free.
[1563257160] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.61, 3.43, 4.34
[1563257351] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8565988 kB) free!
[1563257355] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 11.8% (871384 kB) free!
[1563257471] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7265748 kB) free!
[1563257475] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (468864 kB) free!
[1563257591] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5965640 kB) free!
[1563257595] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (118324 kB) free!
[1563257980] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (12871184 kB) free!
[1563258094] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.3% (39549116 kB) free.
[1563258100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (22088904 kB) free.
[1563258191] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563258191] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10776092 kB) free.
[1563258199] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1563258319] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1563258340] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 81.9% (13310192 kB) free.
[1563258401] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 81.1% (13176956 kB) free.
[1563258439] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1563258444] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 68.5% (11131288 kB) free.
[1563258445] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1563258484] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.6% (6810084 kB) free.
[1563258565] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 373 processes
[1563258685] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 378 processes
[1563258795] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168160 kB) free!
[1563258915] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183824 kB) free!
[1563259006] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (186304 kB) free!
[1563259035] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170444 kB) free!
[1563259126] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166888 kB) free!
[1563259246] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181600 kB) free!
[1563259285] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1563259299] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9519388 kB) free!
[1563259391] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8737020 kB) free!
[1563259419] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8313584 kB) free!
[1563259511] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7430724 kB) free!
[1563259539] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (7255968 kB) free!
[1563259631] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5598704 kB) free!
[1563259693] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563259885] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 371 processes
[1563260006] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 367 processes
[1563260126] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 376 processes
[1563260139] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (249972 kB) free!
[1563260293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563260414] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563260534] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563261439] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1563261439] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1563262031] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10612476 kB) free.
[1563262517] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2428296 kB) free!
[1563262526] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1563262526] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1563262632] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9485960 kB) free!
[1563262639] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1563262752] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8152068 kB) free!
[1563262872] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6761888 kB) free!
[1563263126] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1563263145] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.8% (2400804 kB) free!
[1563263246] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 367 processes
[1563263265] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 12.9% (2093084 kB) free!
[1563263366] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 367 processes
[1563263385] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.0% (1789160 kB) free!
[1563263472] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10733152 kB) free.
[1563263660] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.7% (21486592 kB) free.
[1563263717] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.8% (20686568 kB) free.
[1563263760] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1563263780] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1563263780] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39790512 kB) free.
[1563263830] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 82.7% (13442168 kB) free.
[1563263839] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1563263879] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4711636 kB) free!
[1563263985] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1563263985] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (317932 kB) free!
[1563263999] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4510244 kB) free!
[1563264072] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9311972 kB) free!
[1563264120] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4388232 kB) free!
[1563264192] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (8005988 kB) free!
[1563264312] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6576580 kB) free!
[1563264324] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (1968908 kB) free!
[1563264436] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151316 kB) free!
[1563264445] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185660 kB) free!
[1563264556] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (148424 kB) free!
[1563264565] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179824 kB) free!
[1563264566] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1563264677] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1563264677] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157324 kB) free!
[1563264912] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10726300 kB) free.
[1563265088] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.4% (1134428 kB) free!
[1563265208] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.0% (954000 kB) free!
[1563265328] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.5% (774324 kB) free!
[1563265615] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1563265735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 351 processes
[1563265855] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1563265900] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (12864692 kB) free!
[1563265928] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1563265928] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (117248 kB) free!
[1563266021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (22825416 kB) free.
[1563266366] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 369 processes
[1563266455] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1563266486] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 375 processes
[1563266606] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1563266695] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563267055] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563267120] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2426404 kB) free!
[1563267913] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7278892 kB) free!
[1563268033] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10618556 kB) free.
[1563268334] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563268633] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7167472 kB) free!
[1563268639] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1563268753] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10731672 kB) free.
[1563269353] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7050868 kB) free!
[1563269473] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10716704 kB) free.
[1563269528] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6831444 kB) free.
[1563269685] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'sig4-vm2' was out of bounds)
[1563269786] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (242432 kB) free!
[1563269815] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1563269862] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1563269865] HOST ALERT: sig4-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.191)
[1563269884] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4-vm2' was out of bounds)
[1563269906] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234012 kB) free!
[1563269936] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;2;(Return code of 255 for service 'Total Processes' on host 'sig4-vm2' was out of bounds)
[1563269953] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig4-vm2' was out of bounds)
[1563269985] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1563270027] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232004 kB) free!
[1563270073] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6779536 kB) free!
[1563270117] HOST ALERT: sig4-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1563270134] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563270193] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10727300 kB) free.
[1563270254] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563270416] HOST ALERT: sig4-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.191)
[1563270535] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 192 processes
[1563270539] HOST FLAPPING ALERT: sig4-vm2;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1563270539] HOST ALERT: sig4-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1563270582] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1563270604] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 89.5% (6536736 kB) free.
[1563270734] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4-vm2' was out of bounds)
[1563270734] SERVICE ALERT: sig4-vm2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig4-vm2' was out of bounds)
[1563270737] HOST ALERT: sig4-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.191)
[1563270755] SERVICE ALERT: sig4-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4-vm2' was out of bounds)
[1563270854] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563270854] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4-vm2' was out of bounds)
[1563270854] SERVICE ALERT: sig4-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig4-vm2' was out of bounds)
[1563270974] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563270982] HOST ALERT: sig4-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1563271094] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563271145] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1563271148] HOST ALERT: sig4-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.191)
[1563271185] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1563271207] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4-vm2' was out of bounds)
[1563271268] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig4-vm2' was out of bounds)
[1563271397] HOST ALERT: sig4-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1563271454] SERVICE ALERT: sig4-vm2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1563271454] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 208 processes
[1563271476] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 118987 MB (92% inode=99%):
[1563271696] HOST ALERT: sig4-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.191)
[1563271760] HOST ALERT: sig4-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1563271865] SERVICE FLAPPING ALERT: sig4-vm2;Total Processes;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563271865] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 207 processes
[1563271873] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.01, 0.05
[1563271905] SERVICE FLAPPING ALERT: sig4-vm2;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1563271905] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1563271927] SERVICE FLAPPING ALERT: sig4-vm2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1563271927] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 88.2% (6444172 kB) free.
[1563271993] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8393808 kB) free!
[1563272113] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7014076 kB) free!
[1563272234] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10716256 kB) free.
[1563272295] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563272833] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9715808 kB) free!
[1563272895] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563272953] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8378908 kB) free!
[1563273015] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563273065] HOST FLAPPING ALERT: sig4-vm2;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563273073] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7087108 kB) free!
[1563273135] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563273735] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563274273] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10217360 kB) free.
[1563274335] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563274455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 337 processes
[1563274464] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 85.5% (13897716 kB) free.
[1563274480] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 81.0% (13172992 kB) free.
[1563274497] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39480704 kB) free.
[1563274515] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 81.7% (20021624 kB) free.
[1563274524] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.3% (20793872 kB) free.
[1563274576] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563274639] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1563275021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (11083892 kB) free!
[1563275066] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.9% (3885976 kB) free!
[1563275141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (17788544 kB) free.
[1563275176] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563275186] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.7% (3607944 kB) free!
[1563275306] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (3345840 kB) free!
[1563275455] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1563275473] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9062376 kB) free!
[1563275575] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1563275593] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7229440 kB) free!
[1563275606] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563275713] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11663108 kB) free.
[1563275776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563275896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563275906] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2215832 kB) free!
[1563276016] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563276128] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (1153020 kB) free!
[1563276248] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.8% (941336 kB) free!
[1563276328] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184048 kB) free!
[1563276341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (11048648 kB) free!
[1563276368] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (667932 kB) free!
[1563276394] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2719316 kB) free!
[1563276448] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191444 kB) free!
[1563276461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (16405232 kB) free.
[1563276514] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2622224 kB) free!
[1563276568] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175224 kB) free!
[1563276634] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 23.4% (3795476 kB) free.
[1563276654] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.72, 92.62, 53.45
[1563276774] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.90, 62.63, 47.21
[1563276894] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.60, 42.28, 41.62
[1563276914] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8799016 kB) free!
[1563277034] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12424756 kB) free.
[1563277234] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3197940 kB) free!
[1563277354] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.0% (3091492 kB) free!
[1563277468] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (751732 kB) free!
[1563277474] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.1% (2934804 kB) free!
[1563277588] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (184320 kB) free!
[1563277708] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159212 kB) free!
[1563278094] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.09, 1.66, 12.11
[1563278706] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2018144 kB) free!
[1563278826] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243244 kB) free!
[1563278946] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235760 kB) free!
[1563279274] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1546396 kB) free!
[1563279295] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.22, 1.09, 4.11
[1563279886] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563279895] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 86.84, 215.79, 109.44
[1563279960] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 57.44, 28.03, 12.36
[1563280015] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.95, 144.87, 96.35
[1563280080] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.34, 19.67, 11.19
[1563280135] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1563280135] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.94, 97.46, 84.86
[1563280200] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.45, 13.97, 10.13
[1563280454] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 211 processes
[1563280474] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.0% (13334984 kB) free.
[1563280485] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.4% (2495352 kB) free!
[1563280515] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.1% (3460440 kB) free!
[1563280568] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.8% (6750464 kB) free.
[1563280575] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 209 processes
[1563280605] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.1% (2291204 kB) free!
[1563280635] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (1411908 kB) free!
[1563280639] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563280695] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 207 processes
[1563280725] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.0% (2108748 kB) free!
[1563280755] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1402224 kB) free!
[1563281234] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9690292 kB) free!
[1563281325] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1195360 kB) free!
[1563281354] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8351436 kB) free!
[1563281400] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.06, 3.15, 4.97
[1563281474] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12397940 kB) free.
[1563281695] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 596.30, 540.56, 278.72
[1563281815] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 84.68, 363.32, 245.52
[1563281935] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.78, 1.72, 13.43
[1563281935] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.61, 244.72, 216.34
[1563282461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12392152 kB) free!
[1563282582] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (19482844 kB) free.
[1563282705] SERVICE FLAPPING ALERT: sig4-vm2;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563282727] SERVICE FLAPPING ALERT: sig4-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563283135] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.09, 0.89, 4.35
[1563283265] SERVICE FLAPPING ALERT: sig4-vm2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563283274] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8634548 kB) free!
[1563283394] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12392004 kB) free.
[1563284335] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.11, 0.23, 16.76
[1563285194] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9005712 kB) free!
[1563285314] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7421140 kB) free!
[1563285434] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11854824 kB) free.
[1563285535] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.15, 0.26, 4.83
[1563285546] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1563285581] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (11680072 kB) free!
[1563285701] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (18566236 kB) free.
[1563286034] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563286108] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1563286301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (11151040 kB) free!
[1563286421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (18044384 kB) free.
[1563286568] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (1365564 kB) free!
[1563286634] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9227960 kB) free!
[1563286688] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.0% (1031864 kB) free!
[1563286735] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 9.50, 287.46, 226.02
[1563286754] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7601252 kB) free!
[1563286808] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.8% (791164 kB) free!
[1563286816] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563286855] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.34, 192.66, 198.72
[1563286874] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11945632 kB) free.
[1563286975] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.08, 129.23, 174.74
[1563287021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (10464980 kB) free!
[1563287141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (17385812 kB) free.
[1563287409] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.8% (6823500 kB) free.
[1563287416] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563287536] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563287657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563287741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (10226716 kB) free!
[1563287861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (16834832 kB) free.
[1563288074] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8826348 kB) free!
[1563288194] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12349536 kB) free.
[1563288461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (7725596 kB) free!
[1563288581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (20587280 kB) free.
[1563288857] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563289375] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.28, 1.64, 14.51
[1563289994] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8790688 kB) free!
[1563290057] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563290114] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12336560 kB) free.
[1563290177] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563290576] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.94, 0.86, 4.69
[1563290714] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8665132 kB) free!
[1563290777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563290834] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12343388 kB) free.
[1563290897] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563291002] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.58, 4.41, 4.88
[1563291176] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1460.11, 431.96, 151.92
[1563291296] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 223.67, 314.19, 144.05
[1563291416] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.46, 210.37, 126.63
[1563291434] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8398368 kB) free!
[1563291498] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563291554] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12344568 kB) free.
[1563291618] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563292109] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 22.7% (3692800 kB) free.
[1563292147] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 58.2% (28690652 kB) free.
[1563292154] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8042472 kB) free!
[1563292169] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.3% (22025072 kB) free.
[1563292201] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.6% (21214900 kB) free.
[1563292229] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.3% (20181720 kB) free.
[1563292275] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 35.4% (17445776 kB) free.
[1563292285] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.2% (13363012 kB) free.
[1563292709] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (164592 kB) free!
[1563292818] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563292829] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (183200 kB) free!
[1563292938] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563292949] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164496 kB) free!
[1563293347] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (259428 kB) free!
[1563293467] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (254716 kB) free!
[1563293474] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9712900 kB) free!
[1563293538] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563293587] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236628 kB) free!
[1563293594] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8392372 kB) free!
[1563293658] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563293714] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12436520 kB) free.
[1563293816] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.89, 1.26, 13.88
[1563293981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12676568 kB) free!
[1563294101] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (19696656 kB) free.
[1563294258] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1563294314] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8116524 kB) free!
[1563294379] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563294434] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12445688 kB) free.
[1563294458] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162936 kB) free!
[1563294578] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173288 kB) free!
[1563294685] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (149056 kB) free!
[1563294698] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155744 kB) free!
[1563294805] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (172496 kB) free!
[1563294925] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180472 kB) free!
[1563294978] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563295016] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.89, 0.78, 4.34
[1563295034] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7801264 kB) free!
[1563295098] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563295155] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12445480 kB) free.
[1563295525] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 73.5% (11945408 kB) free.
[1563295698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563295818] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563296354] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9200076 kB) free!
[1563296418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563296474] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7593664 kB) free!
[1563296538] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563296594] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12447120 kB) free.
[1563297138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563297194] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7737712 kB) free!
[1563297259] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563297314] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12438220 kB) free.
[1563297858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563297914] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7831596 kB) free!
[1563297978] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563298034] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12442468 kB) free.
[1563298578] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563298616] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 27.26, 295.72, 190.82
[1563298634] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7399272 kB) free!
[1563298698] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563298736] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.17, 198.09, 167.77
[1563298754] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12443400 kB) free.
[1563298809] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563298856] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.06, 132.76, 147.50
[1563299298] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563299418] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563299955] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8842564 kB) free!
[1563300018] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563300075] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12438060 kB) free.
[1563300138] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563300675] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8648368 kB) free!
[1563300739] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563300795] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12443568 kB) free.
[1563300859] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563301395] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8308648 kB) free!
[1563301459] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563301515] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12430712 kB) free.