[1569782288] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569782344] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181796 kB) free!
[1569782408] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 121424 MB (54% inode=99%):
[1569782446] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186832 kB) free!
[1569782458] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196016 kB) free!
[1569782464] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191536 kB) free!
[1569782566] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194160 kB) free!
[1569782578] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199364 kB) free!
[1569782584] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182896 kB) free!
[1569782687] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (231304 kB) free!
[1569782698] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180144 kB) free!
[1569785451] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 80.37, 21.35, 7.77
[1569785571] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 26.12, 26.11, 11.64
[1569785691] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.91, 17.61, 10.28
[1569786113] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 450 processes
[1569786891] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.33, 1.40, 3.47
[1569787073] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9324356 kB) free!
[1569787193] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8816704 kB) free!
[1569787313] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (8490044 kB) free!
[1569788087] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.3% (5510896 kB) free.
[1569788687] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (3757180 kB) free!
[1569788807] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.1% (3476564 kB) free!
[1569788928] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.3% (3278468 kB) free!
[1569790128] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1569790128] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188548 kB) free!
[1569790812] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230636 kB) free!
[1569790932] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237800 kB) free!
[1569791052] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (423116 kB) free!
[1569791652] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.7% (44181868 kB) free.
[1569793128] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.0% (3214188 kB) free!
[1569793728] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202676 kB) free!
[1569793984] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 54.9% (13456396 kB) free.
[1569794098] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 46.1% (11289452 kB) free.
[1569804454] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 565 processes
[1569804515] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 574 processes
[1569805128] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569805314] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (465308 kB) free!
[1569806014] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 578 processes
[1569806075] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 1 warn out of 568 processes
[1569806135] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 569 processes
[1569806195] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 564 processes
[1569810191] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (212896 kB) free!
[1569810300] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184628 kB) free!
[1569810311] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190672 kB) free!
[1569810420] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189068 kB) free!
[1569810431] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191444 kB) free!
[1569810540] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192024 kB) free!
[1569812513] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 594 processes
[1569812633] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 594 processes
[1569812753] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 593 processes
[1569816011] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 391 processes
[1569816131] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 387 processes
[1569816251] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 382 processes
[1569816361] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1569816481] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569820193] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.29, 5.37, 5.06
[1569820313] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.54, 4.70, 4.85
[1569820566] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.3% (20809444 kB) free.
[1569820631] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.1% (23070996 kB) free.
[1569820668] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 65 users currently logged in
[1569820728] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 65 users currently logged in
[1569820788] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 65 users currently logged in
[1569820848] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1569820913] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.82, 8.57, 6.86
[1569821033] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.44, 8.41, 7.00
[1569821122] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.3% (40070412 kB) free.
[1569821153] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.41, 8.55, 7.20
[1569821175] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188580 kB) free!
[1569821295] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176572 kB) free!
[1569821340] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.2% (23344084 kB) free.
[1569821416] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.0% (20987288 kB) free.
[1569821981] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 305 processes
[1569822015] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176168 kB) free!
[1569822048] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 68 users currently logged in
[1569822101] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1569822135] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173776 kB) free!
[1569822249] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177568 kB) free!
[1569822310] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.6% (18897500 kB) free.
[1569822377] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.4% (14692912 kB) free.
[1569822701] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 313 processes
[1569822821] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1569822913] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163504 kB) free!
[1569823033] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183056 kB) free!
[1569823135] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.0% (23440164 kB) free.
[1569824345] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (4261764 kB) free!
[1569824465] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (4163988 kB) free!
[1569824586] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1569824586] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (434148 kB) free!
[1569825239] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.14, 11.95, 6.97
[1569825291] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.67, 16.39, 8.06
[1569825348] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1569825359] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.80, 8.82, 6.42
[1569825411] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.29, 11.27, 7.20
[1569825479] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.46, 6.70, 5.94
[1569825492] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.52, 14.18, 11.62
[1569825531] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.96, 7.84, 6.44
[1569825612] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.25, 10.15, 10.45
[1569825732] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.64, 8.03, 9.63
[1569825953] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569826079] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.47, 2.46, 4.01
[1569826131] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.97, 2.13, 3.99
[1569826932] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.33, 1.72, 3.77
[1569827052] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9129300 kB) free!
[1569827153] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569827172] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8660700 kB) free!
[1569827292] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.6% (8160368 kB) free!
[1569827753] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569828048] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1569828492] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3146780 kB) free!
[1569828703] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 22386 MB (16% inode=99%):
[1569828823] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 22388 MB (16% inode=99%):
[1569828943] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 22388 MB (16% inode=99%):
[1569828953] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1569828953] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 597 processes
[1569829995] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4698048 kB) free!
[1569830115] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4461660 kB) free!
[1569830235] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (4224004 kB) free!
[1569831435] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1884664 kB) free!
[1569831948] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1569832404] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.3% (46981948 kB) free.
[1569832460] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.5% (23312348 kB) free.
[1569832496] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.5% (22682624 kB) free.
[1569833539] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 12.2% (2882560 kB) free!
[1569833659] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 10.8% (2550960 kB) free!
[1569833743] SERVICE ALERT: sig2-vm2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 14190 MB (10% inode=99%):
[1569833779] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2214044 kB) free!
[1569833891] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.5% (22989336 kB) free.
[1569833982] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4873628 kB) free!
[1569834103] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (4743980 kB) free!
[1569834223] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4639512 kB) free!
[1569834263] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178352 kB) free!
[1569834364] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22008976 kB) free.
[1569834445] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.6% (22936872 kB) free.
[1569834943] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 46840 MB (33% inode=99%):
[1569834976] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191336 kB) free!
[1569835096] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199184 kB) free!
[1569835216] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179532 kB) free!
[1569836754] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.44, 4.25, 4.70
[1569837110] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1569837166] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.2% (23242580 kB) free.
[1569837230] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1569837648] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 57 users currently logged in
[1569837954] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.26, 5.26, 5.03
[1569838074] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.39, 4.59, 4.82
[1569838674] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.29, 5.21, 5.01
[1569838794] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.34, 5.19, 5.02
[1569838914] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.01, 5.23, 5.06
[1569839753] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569840114] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.82, 3.90, 4.52
[1569840830] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 331 processes
[1569840906] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4467100 kB) free!
[1569840950] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 332 processes
[1569840981] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (1319732 kB) free!
[1569841026] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192292 kB) free!
[1569841070] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 331 processes
[1569841101] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176600 kB) free!
[1569841146] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191136 kB) free!
[1569841221] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (183232 kB) free!
[1569841314] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.06, 5.99, 5.49
[1569841434] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.51, 5.61, 5.42
[1569841554] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1569841554] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.35, 5.10, 5.25
[1569841670] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1569842147] SERVICE ALERT: webservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: /home 66655 MB (29% inode=99%):
[1569842270] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 331 processes
[1569842391] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1569842511] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 332 processes
[1569843022] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 82.9% (13475124 kB) free.
[1569843111] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1569843111] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1569843174] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185760 kB) free!
[1569843295] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191528 kB) free!
[1569843322] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.8% (5791460 kB) free!
[1569843415] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187188 kB) free!
[1569843442] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (2991004 kB) free!
[1569843562] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2869352 kB) free!
[1569844162] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (5257712 kB) free!
[1569844762] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223260 kB) free!
[1569846415] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2650808 kB) free!
[1569846562] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1569846562] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (7142152 kB) free!
[1569846711] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1569846831] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1569847015] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1723388 kB) free!
[1569847162] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 28.3% (13946128 kB) free.
[1569847659] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22435648 kB) free.
[1569847704] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22943920 kB) free.
[1569848153] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569848304] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569848363] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.5% (5189452 kB) free!
[1569848424] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.5% (22972140 kB) free.
[1569848483] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (4225180 kB) free!
[1569848603] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4225268 kB) free!
[1569848615] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9817904 kB) free!
[1569848735] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9680016 kB) free!
[1569848753] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1569848855] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.4% (9546856 kB) free!
[1569849354] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.74, 3.35, 4.37
[1569850056] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4705540 kB) free!
[1569850176] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4643632 kB) free!
[1569850296] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4420248 kB) free!
[1569850554] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.31, 7.05, 5.42
[1569850674] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.70, 7.09, 5.66
[1569850794] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.48, 5.62, 5.29
[1569851394] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 6.74, 4.70, 4.76
[1569851855] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4908280 kB) free!
[1569852531] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 118.63, 41.35, 15.41
[1569852651] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 144.84, 76.92, 31.59
[1569852771] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 148.14, 100.71, 45.80
[1569853679] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 105.54, 61.96, 26.61
[1569853799] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 106.61, 76.78, 36.33
[1569853896] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (811408 kB) free!
[1569853919] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 107.40, 86.93, 44.95
[1569854571] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 1.23, 11.30
[1569854632] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569855119] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.03, 4.02, 18.68
[1569855771] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.83, 0.91, 3.79
[1569856713] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569856919] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.28, 1.11, 3.67
[1569858513] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4739888 kB) free!
[1569858633] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4701364 kB) free!
[1569858753] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4653004 kB) free!
[1569859403] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569861753] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2056756 kB) free!
[1569862794] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)