[1555957981] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11083424 kB) free.
[1555958096] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10461344 kB) free.
[1555958197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555958317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555958438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555958581] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.6% (6715896 kB) free!
[1555958696] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9322136 kB) free!
[1555958701] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.0% (5433232 kB) free!
[1555958816] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7989520 kB) free!
[1555958821] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3772388 kB) free!
[1555958936] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6687148 kB) free!
[1555959536] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4941348 kB) free!
[1555960013] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.6% (15052792 kB) free.
[1555960136] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10167604 kB) free.
[1555960737] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8811824 kB) free!
[1555960857] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7450968 kB) free!
[1555960977] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6071048 kB) free!
[1555961461] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182928 kB) free!
[1555961577] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10557520 kB) free.
[1555961581] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183484 kB) free!
[1555961701] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184700 kB) free!
[1555961813] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (174972 kB) free!
[1555961933] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180764 kB) free!
[1555962053] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159376 kB) free!
[1555962177] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9456668 kB) free!
[1555962297] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8238804 kB) free!
[1555962417] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (6980652 kB) free!
[1555962653] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (14943460 kB) free.
[1555963617] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10585344 kB) free.
[1555964217] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8635912 kB) free!
[1555964337] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7297260 kB) free!
[1555964457] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5980816 kB) free!
[1555965058] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10609384 kB) free.
[1555965658] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9634592 kB) free!
[1555965778] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8247304 kB) free!
[1555965898] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6865184 kB) free!
[1555967098] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10665800 kB) free.
[1555968898] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9157068 kB) free!
[1555969018] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7774980 kB) free!
[1555969138] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6408160 kB) free!
[1555969238] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1555970438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555970558] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1555970678] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1555971538] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10761252 kB) free.
[1555971878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1555971998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1555972599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555972719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1555974519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1555974538] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6706480 kB) free!
[1555974639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555974658] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10831984 kB) free.
[1555974759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555975258] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9681144 kB) free!
[1555975378] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8361520 kB) free!
[1555975498] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.1% (6968148 kB) free!
[1555976698] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10132972 kB) free.
[1555977298] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8827264 kB) free!
[1555977418] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7498696 kB) free!
[1555977538] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5828800 kB) free!
[1555978138] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10839196 kB) free.
[1555978738] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9198248 kB) free!
[1555978858] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7735156 kB) free!
[1555978978] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6229896 kB) free!
[1555979578] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10684596 kB) free.
[1555980159] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1555980759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1555980778] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9219640 kB) free!
[1555980880] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555980899] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7833072 kB) free!
[1555981000] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555981019] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6245132 kB) free!
[1555981619] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10905564 kB) free.
[1555982819] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8847604 kB) free!
[1555982939] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7331344 kB) free!
[1555983059] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5459484 kB) free!
[1555983659] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10665140 kB) free.
[1555984259] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9136412 kB) free!
[1555984379] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7671612 kB) free!
[1555984499] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.9% (5873172 kB) free!
[1555984600] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1555985100] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10891584 kB) free.
[1555985201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555985321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555985327] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4631812 kB) free!
[1555985373] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4783544 kB) free!
[1555985441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555985447] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4542328 kB) free!
[1555985493] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4737800 kB) free!
[1555985567] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4456200 kB) free!
[1555985613] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4691116 kB) free!
[1555985700] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9251760 kB) free!
[1555985820] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7900048 kB) free!
[1555985940] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6533852 kB) free!
[1555986540] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4705100 kB) free!
[1555987140] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10146108 kB) free.
[1555987740] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8308848 kB) free!
[1555987860] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6760856 kB) free!
[1555987980] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10794868 kB) free.
[1555988567] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2454480 kB) free!
[1555988580] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9506028 kB) free!
[1555988700] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8026344 kB) free!
[1555988820] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6578520 kB) free!
[1555990020] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10709828 kB) free.
[1555990553] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4872708 kB) free!
[1555990620] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9434964 kB) free!
[1555990673] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4821632 kB) free!
[1555990740] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7949124 kB) free!
[1555990793] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (4770696 kB) free!
[1555990860] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6397372 kB) free!
[1555991013] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2306672 kB) free!
[1555991460] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10521248 kB) free.
[1555992060] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9448308 kB) free!
[1555992180] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8160784 kB) free!
[1555992300] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6790808 kB) free!
[1555993500] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10566768 kB) free.
[1555994700] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8599872 kB) free!
[1555994820] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7172836 kB) free!
[1555994940] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.6% (5713640 kB) free!
[1555995041] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1555995540] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10471688 kB) free.
[1555995641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555995705] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (219832 kB) free!
[1555995761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555995825] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (159828 kB) free!
[1555995881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1555995945] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161144 kB) free!
[1555996140] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9140848 kB) free!
[1555996193] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2215332 kB) free!
[1555996260] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7582404 kB) free!
[1555996380] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6001700 kB) free!
[1555996980] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10391228 kB) free.
[1555997145] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.2% (9948368 kB) free.
[1555997580] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9289880 kB) free!
[1555997700] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7621068 kB) free!
[1555997820] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5959816 kB) free!
[1555998421] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4168160 kB) free!
[1555999021] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10051312 kB) free.
[1555999481] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1555999621] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8057016 kB) free!
[1555999741] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6394504 kB) free!
[1555999853] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.6% (2704460 kB) free!
[1555999861] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4539332 kB) free!
[1555999973] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.0% (2595112 kB) free!
[1556000081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1556000093] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.2% (2465712 kB) free!
[1556000201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1556000221] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10476824 kB) free.
[1556000321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1556000461] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.5% (8628360 kB) free!
[1556000587] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1556000707] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1556000821] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.4% (6600136 kB) free!
[1556000827] SERVICE ALERT: primary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1556000941] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.9% (5852960 kB) free!
[1556001061] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.2% (5045940 kB) free!
[1556001293] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1327324 kB) free!
[1556001427] SERVICE ALERT: primary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1556001661] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (385560 kB) free!
[1556003461] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10025716 kB) free.
[1556004061] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8730580 kB) free!
[1556004182] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7190100 kB) free!
[1556004302] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5475440 kB) free!
[1556004902] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4652824 kB) free!
[1556005502] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10347276 kB) free.
[1556006102] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8430144 kB) free!
[1556006222] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6714020 kB) free!
[1556006342] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4915624 kB) free!
[1556006921] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1556006942] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10428372 kB) free.
[1556007521] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1556007542] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9441480 kB) free!
[1556007591] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.58, 6.88, 7.01
[1556007642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1556007662] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7910440 kB) free!
[1556007711] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.16, 4.99, 6.30
[1556007762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1556007782] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.1% (6435760 kB) free!
[1556007831] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.73, 3.55, 5.61
[1556008431] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.97, 1.20, 3.36
[1556008962] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1556008983] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10348564 kB) free.
[1556009562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1556009583] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9036524 kB) free!
[1556009682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1556009703] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7624540 kB) free!
[1556009802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1556009823] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6238204 kB) free!
[1556010061] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.6% (13127188 kB) free.
[1556010423] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4468204 kB) free!
[1556010661] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9158300 kB) free!
[1556010781] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8512492 kB) free!
[1556010901] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1556010901] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.0% (7876716 kB) free!
[1556011002] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1556011023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10310104 kB) free.
[1556011501] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4606212 kB) free!
[1556011603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1556011623] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8630188 kB) free!
[1556011723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1556011743] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7012196 kB) free!
[1556011843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1556011863] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5342036 kB) free!
[1556012463] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3613548 kB) free!
[1556013043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1556013063] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.7% (9684612 kB) free!
[1556013643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1556013763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556013883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556015084] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1556015684] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556015745] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.1% (2778316 kB) free!
[1556015804] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556015865] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.2% (2631948 kB) free!
[1556015924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1556015985] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.2% (2473184 kB) free!
[1556017124] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1556017185] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (971196 kB) free!
[1556017724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1556017844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1556017964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1556018463] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4793936 kB) free!
[1556019164] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1556019663] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.3% (9507356 kB) free!
[1556019764] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1556019884] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556020004] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556021204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1556022404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1556022525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1556022645] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1556022663] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4002184 kB) free!
[1556022901] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556023185] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 80.3% (13046792 kB) free.
[1556023245] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1556023263] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.0% (9349300 kB) free!
[1556023502] SERVICE ALERT: sig2-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 279 processes
[1556023623] SERVICE ALERT: sig2-vm2;CPU_Procs;OK;HARD;1;CPU OK: 275 processes
[1556023845] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1556023965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556024085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1556025285] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1556025663] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4172392 kB) free!
[1556025885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556026005] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1556026126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1556026863] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.6% (9184848 kB) free!
[1556027326] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1556027728] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.59, 5.54, 6.08
[1556027848] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.18, 3.75, 5.36
[1556027926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1556027968] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.38, 2.65, 4.76
[1556028046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1556028166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1556028663] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4556056 kB) free!
[1556029366] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1556029863] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.0% (8888276 kB) free!
[1556029967] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1556030087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1556030207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556031407] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1556032007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1556032127] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1556032247] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1556032263] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3600420 kB) free!
[1556033447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1556033463] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11063156 kB) free.
[1556033827] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1556033947] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1556034047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1556034063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7194756 kB) free!
[1556034067] SERVICE ALERT: primary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1556034167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1556034183] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5539832 kB) free!
[1556034287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556034303] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3835460 kB) free!
[1556035487] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1556035504] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.6% (8674156 kB) free!
[1556036087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1556036207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556036328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1556037304] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4602040 kB) free!
[1556037528] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1556038128] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1556038248] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1556038368] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1556038504] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.2% (8486120 kB) free!
[1556040304] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3846408 kB) free!
[1556040901] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.7% (8709684 kB) free!
[1556042101] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2784492 kB) free!
[1556042105] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7061184 kB) free!
[1556043168] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1556043769] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1556043889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1556044009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes