[1590863894] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 18.4% (6022992 kB) free!
[1590863928] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17850052 kB) free!
[1590864014] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 18.4% (6024524 kB) free!
[1590864074] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (2792256 kB) free!
[1590864134] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.4% (6039152 kB) free!
[1590864194] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3180908 kB) free!
[1590864277] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590864315] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3648144 kB) free.
[1590864528] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406156 kB) free!
[1590865128] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19425420 kB) free!
[1590865728] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384520 kB) free!
[1590865862] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.29420.927> (10.147.212.143:49108 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590865983] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590866114] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173088 kB) free!
[1590866234] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (457088 kB) free!
[1590866328] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18112244 kB) free!
[1590866354] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (632736 kB) free!
[1590866410] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.6491.6097> (10.147.212.38:48754 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1590866532] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590866928] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (429552 kB) free!
[1590867529] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17636760 kB) free!
[1590867900] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 61.0% (14943064 kB) free.
[1590868129] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404744 kB) free!
[1590868154] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2159404 kB) free!
[1590868729] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16677496 kB) free!
[1590868786] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.4% (13692112 kB) free!
[1590868906] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12896120 kB) free!
[1590868959] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3996220 kB) free!
[1590869026] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (12146236 kB) free!
[1590869079] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (3994724 kB) free!
[1590869155] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590869199] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3992720 kB) free!
[1590869329] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384184 kB) free!
[1590869354] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (602496 kB) free!
[1590869929] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14698204 kB) free!
[1590869954] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (2076380 kB) free!
[1590870226] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (13704080 kB) free!
[1590870529] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439244 kB) free!
[1590870999] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 24.2% (11926172 kB) free.
[1590873152] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.3% (44990660 kB) free.
[1590874001] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2715552 kB) free!
[1590874129] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26076336 kB) free!
[1590874195] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.3% (44488008 kB) free.
[1590874334] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(429) < 2020/05/31 03:01:42 [error] 22132#0: *92178138 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590874334] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.12857.6383> (10.147.212.181:35863 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1590874454] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590874455] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590874602] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172628 kB) free!
[1590874689] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(326) < 103.16.101.72 [31/May/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590874729] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3351536 kB) free!
[1590874776] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3543) < 172.105.59.173 [31/May/2020:03:02:42 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1590874810] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590874896] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590875329] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20050444 kB) free!
[1590875682] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1119 processes
[1590875929] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1680740 kB) free!
[1590876134] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.06, 24.95, 21.32
[1590876134] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 20.3% (6645676 kB) free.
[1590876530] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17100080 kB) free!
[1590876554] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4195144 kB) free.
[1590877130] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409720 kB) free!
[1590878330] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24189904 kB) free!
[1590878930] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382720 kB) free!
[1590879530] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20744456 kB) free!
[1590880130] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397516 kB) free!
[1590880144] SERVICE ALERT: mrq4;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1233 processes
[1590880184] SERVICE ALERT: mrq4;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 287.61, 258.12, 143.86
[1590880267] SERVICE ALERT: mrq4;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 963 processes
[1590880301] SERVICE ALERT: mrq4;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1590880305] SERVICE ALERT: mrq4;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 292.73, 267.81, 161.12
[1590880334] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.55, 13.76, 19.82
[1590880389] SERVICE ALERT: mrq4;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1228 processes
[1590880422] SERVICE ALERT: mrq4;CPU_Procs;OK;HARD;1;CPU OK: 1332 processes
[1590880425] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 283.15, 272.68, 175.90
[1590880602] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.6% (9468496 kB) free.
[1590881083] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1545 processes
[1590881203] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1546 processes
[1590881323] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1544 processes
[1590881330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33693200 kB) free.
[1590881802] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (224584 kB) free!
[1590881922] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (223568 kB) free!
[1590881930] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12357392 kB) free!
[1590882042] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1590882042] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (234004 kB) free!
[1590882050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8492856 kB) free!
[1590882170] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4691432 kB) free!
[1590882494] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 228 processes
[1590882554] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1590882646] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.8% (18832072 kB) free.
[1590882770] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18487756 kB) free!
[1590882842] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1590882854] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.7663.6666> (10.147.212.192:47442 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1590882962] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1590882966] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.6% (14718692 kB) free.
[1590882974] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590883334] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.02, 0.37, 4.19
[1590883370] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34038996 kB) free.
[1590883699] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 777 processes
[1590883819] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 856 processes
[1590883940] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 928 processes
[1590883970] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15148144 kB) free!
[1590884091] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11578344 kB) free!
[1590884211] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (8097392 kB) free!
[1590884811] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17662484 kB) free!
[1590885411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28812388 kB) free.
[1590886008] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 72.1% (17684424 kB) free.
[1590886011] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8051912 kB) free!
[1590886131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35450460 kB) free.
[1590886393] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) <                  {amqp_error,frame_error,
[1590886426] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 37.4% (49333656 kB) free.
[1590886513] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590886731] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14673856 kB) free!
[1590886851] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10661480 kB) free!
[1590886971] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (6678440 kB) free!
[1590887571] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17497548 kB) free!
[1590888171] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29010000 kB) free.
[1590888366] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.6% (2605084 kB) free!
[1590888486] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.4% (2567700 kB) free!
[1590888606] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (2566876 kB) free!
[1590888771] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7978360 kB) free!
[1590888891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.5% (25653268 kB) free!
[1590889011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32175956 kB) free.
[1590889207] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2360492 kB) free!
[1590889611] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11496188 kB) free!
[1590889731] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7713132 kB) free!
[1590889852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34335224 kB) free.
[1590890451] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14316196 kB) free!
[1590890571] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10325372 kB) free!
[1590890691] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6462692 kB) free!
[1590891291] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17551816 kB) free!
[1590891300] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 12.7% (3106212 kB) free!
[1590891360] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 12.7% (3114756 kB) free!
[1590891420] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 12.7% (3114616 kB) free!
[1590891480] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 12.7% (3115968 kB) free!
[1590891891] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28710704 kB) free.
[1590892080] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.9% (5113724 kB) free.
[1590892380] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 12.4% (3042188 kB) free!
[1590892440] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 10.8% (2654728 kB) free!
[1590892491] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8061040 kB) free!
[1590892500] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 9.2% (2261132 kB) free!
[1590892560] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 7.9% (1931236 kB) free!
[1590892611] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4167112 kB) free!
[1590892731] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32328816 kB) free.
[1590892819] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184120 kB) free!
[1590892842] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1590892846] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182044 kB) free!
[1590892939] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188064 kB) free!
[1590892966] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189544 kB) free!
[1590893059] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198436 kB) free!
[1590893086] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174956 kB) free!
[1590893133] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1590893155] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590893253] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1590893275] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590893331] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11010128 kB) free!
[1590893373] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1590893395] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590893451] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (7033556 kB) free!
[1590893571] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3066288 kB) free!
[1590893973] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590893995] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1590894171] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13522596 kB) free!
[1590895371] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2709336 kB) free!
[1590896572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26774708 kB) free.
[1590897172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4027476 kB) free!
[1590897259] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 61.9% (15159524 kB) free.
[1590897292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400664 kB) free!
[1590897412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29953748 kB) free.
[1590898009] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177468 kB) free!
[1590898012] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5516348 kB) free!
[1590898129] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191236 kB) free!
[1590898132] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (928916 kB) free!
[1590898249] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178532 kB) free!
[1590898252] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (506620 kB) free!
[1590899452] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16236976 kB) free!
[1590900052] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377036 kB) free!
[1590900134] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 18.7% (6115552 kB) free!
[1590900255] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 17.9% (5869548 kB) free!
[1590900360] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 15.3% (3741944 kB) free!
[1590900375] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 16.9% (5522544 kB) free!
[1590901252] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28517732 kB) free.
[1590901575] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2767160 kB) free!
[1590901852] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3457516 kB) free!
[1590901972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385380 kB) free!
[1590902093] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446344 kB) free!
[1590902693] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14331552 kB) free!
[1590902762] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1590902882] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1590902954] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (1300808 kB) free!
[1590902995] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9363408 kB) free!
[1590903074] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3665060 kB) free.
[1590903115] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8953924 kB) free!
[1590903235] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (8523020 kB) free!
[1590903259] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.4% (4267984 kB) free!
[1590903293] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (450412 kB) free!
[1590903379] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.9% (3904988 kB) free!
[1590903499] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.5% (3545780 kB) free!
[1590903674] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2995708 kB) free!
[1590903794] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3665332 kB) free.
[1590903883] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2881584 kB) free!
[1590903893] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13261988 kB) free!
[1590903960] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 7.9% (1938612 kB) free!
[1590904003] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.1% (2494552 kB) free!
[1590904099] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1590904099] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2001124 kB) free!
[1590904123] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2143356 kB) free!
[1590904394] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.0% (2433000 kB) free!
[1590904436] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4636360 kB) free!
[1590904493] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392928 kB) free!
[1590904514] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3287008 kB) free.
[1590905114] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2898440 kB) free!
[1590905234] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3426004 kB) free.
[1590905282] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 359 processes
[1590905402] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 335 processes
[1590905522] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1590905552] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8326292 kB) free!
[1590905672] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7978120 kB) free!
[1590905792] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (7608648 kB) free!
[1590905834] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3035876 kB) free!
[1590905954] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (2374460 kB) free!
[1590906074] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (2893832 kB) free!
[1590906293] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15126092 kB) free!
[1590906893] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (427516 kB) free!
[1590906992] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4156872 kB) free!
[1590907123] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.4% (16030208 kB) free.
[1590907426] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26308872 kB) free!
[1590907546] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.4% (32107644 kB) free.
[1590907753] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1590907875] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38774 MB (75% inode=98%):
[1590909293] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32263972 kB) free.
[1590909674] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1034096 kB) free!
[1590909893] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4146152 kB) free!
[1590910014] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394468 kB) free!
[1590910134] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396460 kB) free!
[1590910274] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (2753440 kB) free!
[1590910734] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16343064 kB) free!
[1590910874] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1590910874] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (878696 kB) free!
[1590911334] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396172 kB) free!
[1590911474] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (2023600 kB) free!
[1590912074] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (626064 kB) free!
[1590912123] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1590912240] SERVICE ALERT: sig6-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 40.90, 22.96, 18.30
[1590912243] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1590912346] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (24884060 kB) free!
[1590912360] SERVICE ALERT: sig6-vm2;CPU Load;OK;HARD;1;OK - load average: 18.50, 21.62, 18.48
[1590912466] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (30643112 kB) free.
[1590912674] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (2465744 kB) free!
[1590912843] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1590912963] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1590913066] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.9% (18324404 kB) free!
[1590913186] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (28933756 kB) free.
[1590913723] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4485232 kB) free!
[1590913843] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4202560 kB) free!
[1590913963] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.8% (3778220 kB) free!
[1590914386] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25602840 kB) free!
[1590914506] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.0% (31630992 kB) free.
[1590914563] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1590914563] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1855752 kB) free!
[1590915106] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.6% (24533272 kB) free!
[1590915226] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (27695640 kB) free.
[1590915499] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590915826] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.3% (25397364 kB) free!
[1590915946] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (30956104 kB) free.
[1590916274] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (511268 kB) free!
[1590916546] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (24947312 kB) free!
[1590916666] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.6% (31044956 kB) free.
[1590917266] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.3% (25391024 kB) free!
[1590917386] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (31005856 kB) free.
[1590917474] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1697360 kB) free!
[1590917497] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 185.255.8.58 [31/May/2020:15:00:02 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590917563] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16111608 kB) free.
[1590917617] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590917987] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.3% (21545904 kB) free!
[1590918107] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.3% (32076696 kB) free.
[1590918674] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (403760 kB) free!
[1590919274] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2775208 kB) free!
[1590919874] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (895404 kB) free!
[1590920474] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (1897864 kB) free!
[1590920507] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.4% (24273236 kB) free!
[1590920627] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (30562588 kB) free.
[1590921227] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25620972 kB) free!
[1590921347] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.7% (19384976 kB) free!
[1590921467] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.1% (22533660 kB) free!
[1590922067] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.1% (34383404 kB) free.
[1590922734] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590923017] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(795) < 203.212.70.170 [31/May/2020:16:33:02 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1590923055] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1382) < 2020/05/31 16:34:13 [error] 22134#0: *92311136 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:37/0, bytes from/to upstream:0/37
[1590923138] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(675) < 203.212.70.170 [31/May/2020:16:35:12 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590923181] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(60) < 2020/05/31 16:34:46 [error] 22123#0: *92311122 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590923258] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(42) < 203.212.70.170 [31/May/2020:16:36:23 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590923267] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.8% (26090860 kB) free!
[1590923301] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590923387] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.1% (19880176 kB) free!
[1590923410] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(582) < 172.105.57.57 [31/May/2020:16:33:49 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590923474] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1126412 kB) free!
[1590923507] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (27264672 kB) free.
[1590923530] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590924107] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.0% (22469144 kB) free!
[1590924227] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.6% (29786844 kB) free.
[1590924458] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590924674] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (2916796 kB) free!
[1590924827] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.7% (21995428 kB) free!
[1590924947] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.5% (15097676 kB) free!
[1590924963] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1590925067] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (27730312 kB) free.
[1590925083] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1590925667] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24422228 kB) free!
[1590925787] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.3% (33314992 kB) free.
[1590925874] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1002884 kB) free!
[1590926563] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1590926987] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.3% (24122988 kB) free!
[1590927074] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2702340 kB) free!
[1590927107] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (30947960 kB) free.
[1590927163] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.4% (3468576 kB) free!
[1590927283] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3275100 kB) free!
[1590927403] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.6% (3045992 kB) free!
[1590928003] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2107620 kB) free!
[1590928307] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.3% (21452148 kB) free!
[1590928392] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (2529040 kB) free!
[1590928427] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.9% (26250692 kB) free!
[1590928547] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.4% (18971912 kB) free!
[1590928603] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1590928603] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16115472 kB) free.
[1590929147] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1590929147] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.6% (32482576 kB) free.
[1590929592] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (785456 kB) free!
[1590929884] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 342 processes
[1590930004] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 338 processes
[1590930124] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1590930674] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3656040 kB) free.
[1590931547] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.9% (22314392 kB) free!
[1590931667] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (28346664 kB) free.
[1590931874] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2814068 kB) free!
[1590931994] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3774700 kB) free.
[1590932867] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.8% (26119656 kB) free!
[1590932987] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (31814724 kB) free.
[1590933794] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3240540 kB) free!
[1590933914] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3212776 kB) free!
[1590934034] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (539648 kB) free!
[1590934101] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/31 19:34:05 [error] 22135#0: *92341039 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590934222] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590935235] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (2551068 kB) free!
[1590935835] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (760228 kB) free!
[1590937035] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.9% (1763648 kB) free!
[1590937788] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.2% (21386984 kB) free!
[1590937908] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.3% (16195824 kB) free!
[1590938025] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 34.78, 32.62, 37.85
[1590938028] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7679516 kB) free!
[1590938235] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.8% (4525928 kB) free.
[1590939381] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 46.9% (15338032 kB) free.
[1590940003] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590940035] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2950564 kB) free!
[1590940134] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30595824 kB) free.
[1590940155] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3598048 kB) free.
[1590940734] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1743840 kB) free!
[1590940755] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2212916 kB) free!
[1590940854] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (429788 kB) free!
[1590940875] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (2706956 kB) free!
[1590940974] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (438448 kB) free!
[1590940995] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3276296 kB) free.
[1590941574] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16258204 kB) free!
[1590941625] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 12.20, 19.93, 27.03
[1590942174] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1590942174] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388716 kB) free!
[1590943395] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3194876 kB) free!
[1590943515] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3509076 kB) free.
[1590944028] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1590944625] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 40.48, 34.11, 28.85
[1590944745] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 33.34, 32.36, 28.79
[1590945315] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (2773900 kB) free!
[1590945435] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2886956 kB) free!
[1590945555] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (2930652 kB) free!
[1590945774] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32904912 kB) free.
[1590945828] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.7% (14139344 kB) free!
[1590946155] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3400488 kB) free.
[1590946374] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6242888 kB) free!
[1590946494] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1921928 kB) free!
[1590946614] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382572 kB) free!
[1590947028] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12394900 kB) free!
[1590947214] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15693668 kB) free!
[1590947355] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3210296 kB) free!
[1590947475] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3116236 kB) free!
[1590947595] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3016764 kB) free!
[1590947628] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1590947628] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.2% (16035148 kB) free!
[1590947814] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (453392 kB) free!
[1590948828] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12733620 kB) free!
[1590949014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31771828 kB) free.
[1590949395] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3549972 kB) free.
[1590949428] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.8% (15500092 kB) free!
[1590949614] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5793584 kB) free!
[1590949734] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1575128 kB) free!