[1567103846] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (1197148 kB) free!
[1567103915] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179312 kB) free!
[1567103971] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182300 kB) free!
[1567104035] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183064 kB) free!
[1567104091] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181556 kB) free!
[1567104155] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172184 kB) free!
[1567104745] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1567107043] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 469 processes
[1567107103] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 464 processes
[1567107117] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (194200 kB) free!
[1567107163] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 467 processes
[1567107341] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 293 processes
[1567107417] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 24.5% (6008872 kB) free.
[1567107455] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (5071524 kB) free!
[1567107573] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (398988 kB) free!
[1567107575] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (239704 kB) free!
[1567107693] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (414584 kB) free!
[1567107695] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (1176292 kB) free!
[1567107813] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1480580 kB) free!
[1567107941] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 288 processes
[1567109375] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.06, 5.80, 5.40
[1567109495] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.49, 5.77, 5.44
[1567109615] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.97, 5.33, 5.31
[1567110173] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567110215] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.52, 4.13, 4.95
[1567110745] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.7% (3122524 kB) free!
[1567111303] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1567111346] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192520 kB) free!
[1567111423] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567112978] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.1% (43915464 kB) free.
[1567113097] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 94868 MB (42% inode=99%):
[1567114304] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 89.00, 67.29, 32.86
[1567114424] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 75.80, 71.65, 38.74
[1567114542] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1411 processes
[1567114544] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.41, 48.40, 34.21
[1567114662] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1407 processes
[1567114782] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 335 processes
[1567114891] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 65.2% (15980064 kB) free.
[1567114955] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 47.8% (11712884 kB) free.
[1567115144] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 7.11, 18.31
[1567116944] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.27, 1.14, 3.49
[1567120423] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567120543] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567125095] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 96.3% (33739952 kB) free.
[1567125814] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.7% (23337936 kB) free!
[1567126294] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190960 kB) free!
[1567126355] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (815052 kB) free!
[1567126417] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193052 kB) free!
[1567126475] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (807148 kB) free!
[1567126538] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192060 kB) free!
[1567126595] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (801588 kB) free!
[1567126895] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1567126964] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 501 processes
[1567127014] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 40.9% (53922124 kB) free.
[1567127024] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 510 processes
[1567127084] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 495 processes
[1567133743] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567133863] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567138895] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 310 processes
[1567139015] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1567145263] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567145383] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567146276] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.3% (22788616 kB) free.
[1567146297] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.3% (22612228 kB) free.
[1567146315] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40735576 kB) free.
[1567146332] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.2% (22606096 kB) free.
[1567148085] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180164 kB) free!
[1567148205] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179792 kB) free!
[1567148325] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173372 kB) free!
[1567151124] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8949400 kB) free!
[1567151244] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8631596 kB) free!
[1567151364] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (8326604 kB) free!
[1567152564] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2780860 kB) free!
[1567154015] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.27, 7.34, 5.38
[1567154135] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.61, 8.01, 5.86
[1567154255] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.54, 7.91, 6.09
[1567156055] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.70, 3.84, 4.69
[1567157255] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.35, 8.13, 5.88
[1567157375] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.56, 8.63, 6.34
[1567157380] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22120544 kB) free.
[1567157424] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 52.8% (26020040 kB) free.
[1567157495] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.55, 8.56, 6.59
[1567157978] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9041804 kB) free!
[1567157984] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195552 kB) free!
[1567158035] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217124 kB) free!
[1567158098] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8859900 kB) free!
[1567158104] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191116 kB) free!
[1567158155] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (211792 kB) free!
[1567158218] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (8675572 kB) free!
[1567158224] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168560 kB) free!
[1567158275] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1567158275] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227372 kB) free!
[1567158695] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1567158695] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.61, 4.22, 4.71
[1567159295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1567159295] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.81, 7.93, 6.70
[1567159415] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1567159415] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.69, 7.08, 6.53
[1567159535] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1567159536] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.41, 6.83, 6.50
[1567160135] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1567160619] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4182156 kB) free!
[1567163377] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 315 processes
[1567163497] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 315 processes
[1567163617] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1567165897] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47116 MB (20% inode=99%):
[1567166017] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46743 MB (20% inode=99%):
[1567166137] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46737 MB (20% inode=99%):
[1567166706] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4721496 kB) free!
[1567166737] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (4224984 kB) free!
[1567166740] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 172.85, 138.22, 70.25
[1567166744] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 69.40, 24.46, 9.59
[1567166826] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (4305532 kB) free!
[1567166857] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 26.4% (6466832 kB) free.
[1567166860] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 63.86, 118.32, 71.90
[1567166864] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 85.66, 45.58, 19.14
[1567166946] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3887284 kB) free!
[1567166980] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 9.29, 79.42, 63.27
[1567166984] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 87.68, 59.63, 27.49
[1567167336] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.34, 2.44, 3.92
[1567167546] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 21.1% (5177372 kB) free.
[1567167875] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1567168146] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4897304 kB) free!
[1567168180] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.81, 2.24, 18.03
[1567168184] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.06, 4.32, 15.68
[1567168266] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4722916 kB) free!
[1567168386] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (4636556 kB) free!
[1567169497] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 63.7% (31393772 kB) free.
[1567169538] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 35.8% (17658748 kB) free.
[1567169573] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.3% (19821628 kB) free.
[1567169615] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 314 processes
[1567169735] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1567169857] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4647940 kB) free!
[1567169977] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (4578644 kB) free!
[1567169980] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.43, 1.21, 3.48
[1567169984] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.35, 1.59, 3.62
[1567170097] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4510632 kB) free!
[1567170173] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172788 kB) free!
[1567170293] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (591996 kB) free!
[1567170413] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184560 kB) free!
[1567170707] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8955884 kB) free!
[1567170827] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8615980 kB) free!
[1567170947] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (8336032 kB) free!
[1567172147] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2500924 kB) free!
[1567173186] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1567173186] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2405292 kB) free!
[1567173336] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1567173583] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 404 processes
[1567173704] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567176097] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2442880 kB) free!
[1567176937] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.05, 6.58, 5.97
[1567177057] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.95, 6.10, 5.88
[1567177177] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.96, 5.54, 5.71
[1567177777] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.36, 4.00, 4.63
[1567180386] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1567182783] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.0% (21479232 kB) free.
[1567182817] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.4% (22641556 kB) free.
[1567182849] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 60.5% (29829676 kB) free.
[1567182880] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.4% (22654644 kB) free.
[1567183983] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.7% (3369748 kB) free!
[1567184103] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.5% (3322500 kB) free!
[1567184223] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (3155936 kB) free!
[1567184823] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175852 kB) free!
[1567186903] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567187023] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 402 processes
[1567187143] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567189543] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567189663] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes