[1556821861] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1556821981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556823338] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 268.87, 190.11, 87.76
[1556823459] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 60.28, 146.39, 84.94
[1556823579] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.30, 98.02, 74.67
[1556823781] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1556823957] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.0% (5391484 kB) free.
[1556824381] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556824445] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4881180 kB) free!
[1556824501] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556824565] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4871568 kB) free!
[1556824621] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556824685] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (4869732 kB) free!
[1556825379] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.17, 0.42, 10.91
[1556825586] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 250.85, 108.10, 41.06
[1556825706] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 270.42, 162.80, 69.23
[1556825826] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 272.79, 199.32, 93.96
[1556826003] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 248.38, 182.79, 87.23
[1556826123] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 248.69, 204.61, 106.82
[1556826243] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 108.67, 181.04, 111.04
[1556826249] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 314 processes
[1556826579] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.16, 0.23, 3.17
[1556826849] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 2 warn out of 314 processes
[1556827621] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556827626] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.10, 0.70, 14.52
[1556828043] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.72, 2.59, 17.67
[1556828221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556828341] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1556828826] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.28, 4.18
[1556828941] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556829061] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1556829357] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4876448 kB) free!
[1556829478] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4869172 kB) free!
[1556829598] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4854148 kB) free!
[1556829843] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.60, 1.67, 3.97
[1556830261] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556830382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556830502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556831702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1556834102] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1556834223] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1556835422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556835542] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1556836142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556836262] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1556837463] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556837583] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556840583] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556840703] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556841304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556841424] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1556842179] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 438.36, 269.25, 115.43
[1556842198] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2424764 kB) free!
[1556842299] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 441.72, 326.49, 155.06
[1556842419] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 64.98, 223.56, 138.39
[1556842624] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556842744] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556842798] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4648116 kB) free!
[1556843344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556843464] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1556843584] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1556843584] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1556843998] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.0% (5392768 kB) free.
[1556844185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556844305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556844425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556844485] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2117364 kB) free!
[1556844670] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47032 MB (20% inode=99%):
[1556844790] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46952 MB (20% inode=99%):
[1556844910] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46949 MB (20% inode=99%):
[1556845026] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 217.84, 151.00, 69.01
[1556845146] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 87.26, 136.58, 74.56
[1556845267] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.86, 91.44, 65.54
[1556846019] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 0.46, 12.82
[1556846467] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.27, 1.99, 18.31
[1556847219] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.11, 0.25, 3.71
[1556847925] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8537488 kB) free!
[1556848045] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12694816 kB) free.
[1556848198] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4854792 kB) free!
[1556848267] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.53, 0.43, 3.02
[1556848318] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4844176 kB) free!
[1556848438] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1556848438] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (4830840 kB) free!
[1556849245] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8593292 kB) free!
[1556849365] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12707112 kB) free.
[1556850565] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8534520 kB) free!
[1556850685] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12803596 kB) free.
[1556851482] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1556851603] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1556851886] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8402112 kB) free!
[1556852006] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12703412 kB) free.
[1556852826] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556853206] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8855692 kB) free!
[1556853249] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 314 processes
[1556853326] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13255904 kB) free.
[1556853426] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556853546] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556853666] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556854526] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9189808 kB) free!
[1556854602] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1556854646] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13247224 kB) free.
[1556854722] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1556855846] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9252812 kB) free!
[1556855966] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13208296 kB) free.
[1556856238] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1556856249] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 313 processes
[1556857167] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9578444 kB) free!
[1556857287] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13165420 kB) free.
[1556857914] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1556857944] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 563 processes
[1556858004] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 561 processes
[1556858038] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1974720 kB) free!
[1556858064] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 561 processes
[1556858487] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9510408 kB) free!
[1556858607] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7821416 kB) free!
[1556858727] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12536376 kB) free.
[1556859316] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4799404 kB) free!
[1556859436] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4741864 kB) free!
[1556859556] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4689648 kB) free!
[1556859714] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1556859834] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1556859927] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8407472 kB) free!
[1556859954] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1556860047] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12867296 kB) free.
[1556860401] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (219776 kB) free!
[1556860521] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (213172 kB) free!
[1556860641] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (506788 kB) free!
[1556860761] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1556860881] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1556860979] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 49.1% (7988024 kB) free.
[1556861110] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 75643 MB (33% inode=99%):
[1556861247] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8576420 kB) free!
[1556861367] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12952080 kB) free.
[1556862219] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 244.24, 101.41, 38.18
[1556862340] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 266.91, 157.33, 66.34
[1556862460] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 269.80, 194.69, 91.07
[1556862567] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9416100 kB) free!
[1556862687] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7549308 kB) free!
[1556862807] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12006164 kB) free.
[1556862838] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2797108 kB) free!
[1556863041] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.1% (3696080 kB) free!
[1556863698] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.8% (23397736 kB) free.
[1556863760] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.8% (21773836 kB) free.
[1556863808] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.9% (22770648 kB) free.
[1556863903] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.4% (18728360 kB) free.
[1556864007] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8187896 kB) free!
[1556864127] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12776812 kB) free.
[1556864260] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.00, 1.13, 14.16
[1556865179] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3181688 kB) free!
[1556865299] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (3025752 kB) free!
[1556865327] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8630124 kB) free!
[1556865419] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.7% (2870732 kB) free!
[1556865447] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12699628 kB) free.
[1556865620] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191156 kB) free!
[1556865708] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (203120 kB) free!
[1556865740] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191548 kB) free!
[1556865828] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (778240 kB) free!
[1556865843] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 272.94, 232.19, 125.59
[1556865860] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1556865860] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192104 kB) free!
[1556865948] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1556865948] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200552 kB) free!
[1556865963] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 48.46, 165.20, 114.43
[1556866060] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.62, 0.95, 3.71
[1556866083] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.54, 111.32, 100.84
[1556866619] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1626004 kB) free!
[1556866647] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9043636 kB) free!
[1556866768] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7106924 kB) free!
[1556866888] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11701676 kB) free.
[1556867309] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (315004 kB) free!
[1556867364] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191804 kB) free!
[1556867429] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (265460 kB) free!
[1556867484] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196472 kB) free!
[1556867488] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9857388 kB) free!
[1556867549] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199260 kB) free!
[1556867604] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194172 kB) free!
[1556867608] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8194556 kB) free!
[1556867728] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12547352 kB) free.
[1556867883] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.13, 2.21, 16.17
[1556868837] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.2% (18688148 kB) free.
[1556868902] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.4% (22881156 kB) free.
[1556868928] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8932028 kB) free!
[1556868952] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1556868952] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.1% (14640544 kB) free.
[1556868996] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.7% (22471296 kB) free.
[1556869048] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7036560 kB) free!
[1556869052] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1556869052] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.0% (23435280 kB) free.
[1556869168] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11744652 kB) free.
[1556869683] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.11, 1.96, 3.89
[1556869768] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9770644 kB) free!
[1556869867] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.64, 11.76, 7.46
[1556869888] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8138560 kB) free!
[1556869987] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.25, 7.89, 6.57
[1556870008] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12450356 kB) free.
[1556870107] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.36, 5.42, 5.83
[1556870466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1556870707] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 1.06, 3.25
[1556871067] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556871187] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556871208] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8244348 kB) free!
[1556871307] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556871328] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12367836 kB) free.
[1556871846] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (4000728 kB) free!
[1556871907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1556871966] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3936444 kB) free!
[1556872027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556872086] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (3956276 kB) free!
[1556872147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556872528] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8207900 kB) free!
[1556872648] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12310260 kB) free.
[1556873286] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2368944 kB) free!
[1556873848] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8801984 kB) free!
[1556873886] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (2926616 kB) free!
[1556873968] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7061620 kB) free!
[1556874088] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11754016 kB) free.
[1556874688] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9529852 kB) free!
[1556874808] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7855168 kB) free!
[1556874929] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12160244 kB) free.
[1556876128] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8533068 kB) free!
[1556876248] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6636896 kB) free!
[1556876258] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4471104 kB) free!
[1556876368] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11464140 kB) free.
[1556876378] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4432508 kB) free!
[1556876498] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (4391624 kB) free!
[1556876706] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1556876755] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.9% (23170316 kB) free.
[1556876756] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1556876823] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.7% (18551728 kB) free.
[1556877568] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8447128 kB) free!
[1556877688] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6522760 kB) free!
[1556877808] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11414760 kB) free.
[1556878367] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4799416 kB) free!
[1556878408] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9544092 kB) free!
[1556878427] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4797888 kB) free!
[1556878487] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.6% (4799872 kB) free!
[1556878528] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7803548 kB) free!
[1556878548] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.6% (4800672 kB) free!
[1556878628] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3600656 kB) free!
[1556878648] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11958856 kB) free.
[1556878748] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.8% (3375368 kB) free!
[1556878807] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1556878868] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166108 kB) free!
[1556879204] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (2260600 kB) free!
[1556879324] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (2221356 kB) free!
[1556879445] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2173412 kB) free!
[1556879706] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (1110620 kB) free!
[1556879826] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (1088268 kB) free!
[1556879848] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8592364 kB) free!
[1556879946] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (1045400 kB) free!
[1556879968] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6922320 kB) free!
[1556880007] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1556880088] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11720428 kB) free.
[1556880608] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1556880688] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9691720 kB) free!
[1556880808] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8048692 kB) free!
[1556880928] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6051736 kB) free!
[1556881208] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1556881208] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1556881294] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.7% (21493660 kB) free.
[1556881343] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.3% (22385420 kB) free.
[1556881392] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.6% (18534868 kB) free.
[1556881439] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 59.4% (14567932 kB) free.
[1556881528] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10425128 kB) free.
[1556881556] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (2863356 kB) free!
[1556881676] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.1% (2771504 kB) free!
[1556881796] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.6% (2699876 kB) free!
[1556882128] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8398032 kB) free!
[1556882162] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185940 kB) free!
[1556882248] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6527592 kB) free!
[1556882282] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166392 kB) free!
[1556882368] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11264120 kB) free.
[1556882402] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180088 kB) free!
[1556882968] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9134804 kB) free!
[1556883089] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7545000 kB) free!
[1556883120] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22846928 kB) free.
[1556883177] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.2% (14667504 kB) free.
[1556883209] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11770304 kB) free.
[1556883269] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3141856 kB) free!
[1556883389] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.0% (3090332 kB) free!
[1556883509] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.7% (3045712 kB) free!
[1556884147] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1556884171] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 79.6% (12935136 kB) free.
[1556884409] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8654216 kB) free!
[1556884529] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7061740 kB) free!
[1556884554] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1556884649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11725564 kB) free.
[1556884747] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556884867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556884987] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556885249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9653980 kB) free!
[1556885369] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8076900 kB) free!
[1556885489] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6110184 kB) free!
[1556886089] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1556886089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10582812 kB) free.
[1556886689] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8681960 kB) free!
[1556886809] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7014508 kB) free!
[1556886930] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11484764 kB) free.
[1556887083] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 337.32, 185.99, 77.07
[1556887203] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 343.97, 238.60, 109.56
[1556887323] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 344.99, 273.82, 138.11
[1556887529] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9499892 kB) free!
[1556887554] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1556887649] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7859976 kB) free!
[1556887674] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1556887769] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11700948 kB) free.
[1556887794] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1556888260] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 194.24, 66.78, 24.00
[1556888380] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 237.65, 125.54, 50.72
[1556888500] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 243.67, 164.89, 74.21
[1556888547] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (4304744 kB) free!
[1556888667] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 27.3% (6688296 kB) free.
[1556888969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8156124 kB) free!
[1556889089] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6302020 kB) free!
[1556889209] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11040460 kB) free.
[1556889723] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.48, 1.79, 15.51
[1556889796] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1556889809] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9272060 kB) free!
[1556889907] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 27.68, 125.53, 82.08
[1556889929] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7804948 kB) free!
[1556890027] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.02, 84.10, 72.17
[1556890049] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11628444 kB) free.
[1556890148] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.75, 56.37, 63.45
[1556890194] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1556890300] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.30, 0.94, 14.73
[1556890987] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1556891250] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8400220 kB) free!
[1556891348] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.14, 1.27, 17.67
[1556891370] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6847688 kB) free!
[1556891490] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11389076 kB) free.
[1556891500] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.28, 0.42, 4.36
[1556891523] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.72, 1.71, 3.70
[1556891587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556891708] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556891828] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1556891828] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556892090] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9557980 kB) free!
[1556892210] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8037516 kB) free!
[1556892330] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6098132 kB) free!
[1556892608] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556892729] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1556892930] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10304316 kB) free.
[1556893148] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.34, 2.86
[1556893530] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8550360 kB) free!
[1556893650] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7125772 kB) free!
[1556893770] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547948 kB) free.
[1556894228] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1556894828] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556894948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556894970] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7933940 kB) free!
[1556895068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556895091] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6176552 kB) free!
[1556895211] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11059192 kB) free.
[1556895811] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8862908 kB) free!
[1556895931] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7238464 kB) free!
[1556896051] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11487008 kB) free.
[1556896651] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9756276 kB) free!
[1556896771] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8171864 kB) free!
[1556896891] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6371464 kB) free!
[1556897440] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5533 MB (10% inode=97%):
[1556897491] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10433232 kB) free.
[1556897523] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 256.67, 213.61, 112.73
[1556897643] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 74.45, 171.22, 110.39
[1556897763] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.56, 115.10, 97.20
[1556898040] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6871 MB (13% inode=98%):
[1556898091] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8748132 kB) free!
[1556898211] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7216376 kB) free!
[1556898331] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11443604 kB) free.
[1556898932] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9414748 kB) free!
[1556899052] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7889816 kB) free!
[1556899172] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5943596 kB) free!
[1556899563] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1556899563] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.69, 2.07, 15.56
[1556899772] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10240288 kB) free.
[1556899776] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (158636 kB) free!
[1556899897] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (184904 kB) free!
[1556899996] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195380 kB) free!
[1556900017] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178956 kB) free!
[1556900116] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (396076 kB) free!
[1556900236] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (453140 kB) free!
[1556900372] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8285328 kB) free!
[1556900492] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6564644 kB) free!
[1556900612] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11145444 kB) free.
[1556900617] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 35.1% (5699676 kB) free.
[1556901212] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8984600 kB) free!
[1556901217] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (148020 kB) free!
[1556901332] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7403868 kB) free!
[1556901337] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166220 kB) free!
[1556901363] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.76, 1.70, 3.64
[1556901452] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11420636 kB) free.
[1556901457] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161760 kB) free!
[1556902467] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4908944 kB) free!
[1556902587] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4899320 kB) free!
[1556902652] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8204796 kB) free!
[1556902707] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (4891016 kB) free!
[1556902773] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6439244 kB) free!
[1556902868] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1556902893] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10945636 kB) free.
[1556903468] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556903493] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9042180 kB) free!
[1556903588] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556903613] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7511672 kB) free!
[1556903708] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556903733] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11404764 kB) free.
[1556904333] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9821888 kB) free!
[1556904453] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8341116 kB) free!
[1556904573] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.5% (6670320 kB) free!
[1556905057] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1556905057] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.5% (14213888 kB) free.
[1556905173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10579716 kB) free.
[1556905773] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8627912 kB) free!
[1556905893] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7073992 kB) free!
[1556906013] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11417292 kB) free.
[1556906614] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9367172 kB) free!
[1556906734] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7809084 kB) free!
[1556906854] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11418556 kB) free.
[1556907308] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556907908] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556908028] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556908054] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8167852 kB) free!
[1556908148] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556908174] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6344536 kB) free!