[1590431438] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3583088 kB) free.
[1590431643] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10033968 kB) free!
[1590431763] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6016212 kB) free!
[1590431884] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2210684 kB) free!
[1590432484] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15876932 kB) free!
[1590433084] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32118544 kB) free.
[1590433684] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8956804 kB) free!
[1590433804] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5141408 kB) free!
[1590433838] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (1030968 kB) free!
[1590433924] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1475876 kB) free!
[1590433958] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180416 kB) free!
[1590434078] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1368048 kB) free!
[1590434524] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13903940 kB) free!
[1590435127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27370108 kB) free.
[1590435225] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195436 kB) free!
[1590435288] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1678816 kB) free!
[1590435288] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193872 kB) free!
[1590435350] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (176616 kB) free!
[1590435411] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (199912 kB) free!
[1590435729] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6284932 kB) free!
[1590435839] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1590435849] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2599712 kB) free!
[1590435959] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1590435969] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32597928 kB) free.
[1590436011] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 80.7% (19787952 kB) free.
[1590436475] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9440960 kB) free!
[1590436569] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10073552 kB) free!
[1590436595] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242168 kB) free!
[1590436689] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6320720 kB) free!
[1590436715] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233368 kB) free!
[1590436809] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2815436 kB) free!
[1590437266] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (394280 kB) free!
[1590437386] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382612 kB) free!
[1590437409] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13917948 kB) free!
[1590437506] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1221788 kB) free!
[1590438609] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5696312 kB) free!
[1590439115] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.7% (6236696 kB) free!
[1590439210] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15778936 kB) free!
[1590439810] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26831296 kB) free.
[1590440410] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7363372 kB) free!
[1590440530] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33205064 kB) free.
[1590440902] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 85.7% (42212696 kB) free.
[1590441130] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13326336 kB) free!
[1590441250] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9558920 kB) free!
[1590441353] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.3% (43519740 kB) free.
[1590441370] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5913548 kB) free!
[1590441970] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14085816 kB) free!
[1590442233] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.9% (9769092 kB) free.
[1590442327] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(428) < 2020/05/26 03:01:32 [error] 22130#0: *90857505 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590442422] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1590442447] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590442559] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(336) < 49.50.86.152 [26/May/2020:03:00:37 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590442604] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (3106272 kB) free!
[1590442680] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590442716] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.9% (44293468 kB) free.
[1590442766] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2205) < 172.105.59.173 [26/May/2020:03:02:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590442833] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181832 kB) free!
[1590442886] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590442953] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179108 kB) free!
[1590443073] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184760 kB) free!
[1590443770] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (11471644 kB) free!
[1590444288] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.4% (4123652 kB) free.
[1590444370] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21116444 kB) free!
[1590444970] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32029572 kB) free.
[1590445570] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10916692 kB) free!
[1590445690] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7430112 kB) free!
[1590445810] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29633396 kB) free.
[1590446073] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2782684 kB) free!
[1590446204] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1857820 kB) free!
[1590446411] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8375664 kB) free!
[1590446531] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17428948 kB) free!
[1590446651] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12899832 kB) free!
[1590446673] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1590446673] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (246272 kB) free!
[1590447251] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21584600 kB) free!
[1590447851] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412772 kB) free!
[1590448451] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16378812 kB) free!
[1590449651] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1810508 kB) free!
[1590450288] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1590450362] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 51.6% (12741612 kB) free.
[1590450851] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22795032 kB) free!
[1590451003] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 72.7% (17808748 kB) free.
[1590451004] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 46.6% (11418932 kB) free.
[1590451451] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29142828 kB) free.
[1590452051] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9332696 kB) free!
[1590452171] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5453436 kB) free!
[1590452291] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28135168 kB) free.
[1590452891] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6855988 kB) free!
[1590453012] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28774312 kB) free.
[1590453611] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8721024 kB) free!
[1590453731] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4654388 kB) free!
[1590453851] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28431624 kB) free.
[1590454451] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7871400 kB) free!
[1590454571] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28256252 kB) free.
[1590455162] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183604 kB) free!
[1590455171] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9812568 kB) free!
[1590455282] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (204380 kB) free!
[1590455291] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6015244 kB) free!
[1590455402] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185196 kB) free!
[1590455411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28069720 kB) free.
[1590456011] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7983192 kB) free!
[1590456131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4245708 kB) free!
[1590456251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26433108 kB) free.
[1590456851] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6260580 kB) free!
[1590456971] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27643244 kB) free.
[1590457571] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8102228 kB) free!
[1590457691] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4343056 kB) free!
[1590457812] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26918616 kB) free.
[1590457906] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 76.3% (100617596 kB) free.
[1590458073] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590458411] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6380112 kB) free!
[1590458531] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2645748 kB) free!
[1590458651] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (25039632 kB) free!
[1590459247] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/26 07:35:05 [error] 22135#0: *90888624 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590459251] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4845036 kB) free!
[1590459367] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590459783] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (187348 kB) free!
[1590460451] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16365380 kB) free!
[1590461651] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1288824 kB) free!
[1590462851] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15310832 kB) free!
[1590463604] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4770184 kB) free!
[1590463724] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4653820 kB) free!
[1590463844] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4521836 kB) free!
[1590464051] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395360 kB) free!
[1590464199] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.39, 7.66, 5.68
[1590464319] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.37, 9.67, 6.75
[1590464439] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.94, 8.83, 6.78
[1590465044] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2162532 kB) free!
[1590465251] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15794616 kB) free!
[1590465851] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399496 kB) free!
[1590466452] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14547496 kB) free!
[1590466488] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3177084 kB) free!
[1590466608] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3787660 kB) free.
[1590467167] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/26 09:50:04 [error] 22135#0: *90906007 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590467287] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590467652] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386328 kB) free!
[1590467808] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1851032 kB) free!
[1590467928] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3074536 kB) free!
[1590468048] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3005952 kB) free!
[1590468222] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1590468342] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590468403] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185284 kB) free!
[1590468441] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1590468462] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590468523] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190140 kB) free!
[1590468561] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1590468643] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170140 kB) free!
[1590468681] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1590468852] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18568248 kB) free!
[1590469062] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590470052] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5653048 kB) free!
[1590470652] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18604604 kB) free!
[1590470776] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 38.4% (9408316 kB) free.
[1590470862] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590470983] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590471044] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (3119048 kB) free!
[1590471103] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 357 processes
[1590471252] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394864 kB) free!
[1590471644] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1590471644] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2030624 kB) free!
[1590471742] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (3726520 kB) free!
[1590471862] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3355672 kB) free!
[1590471982] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.3% (2986144 kB) free!
[1590472248] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1280272 kB) free!
[1590472582] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1047848 kB) free!
[1590472848] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (2594008 kB) free!
[1590473486] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(379) < 103.255.146.164 [26/May/2020:11:32:39 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590473606] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590474252] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19142100 kB) free!
[1590474703] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1590474703] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 358 processes
[1590474852] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414192 kB) free!
[1590474976] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.3% (4474576 kB) free!
[1590474982] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.7% (15866524 kB) free.
[1590475096] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.8% (4113840 kB) free!
[1590475216] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.2% (3724656 kB) free!
[1590475239] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 40.48, 25.64, 17.36
[1590475243] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 23.4% (5735272 kB) free.
[1590475303] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590475816] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1955108 kB) free!
[1590475843] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199132 kB) free!
[1590475963] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183612 kB) free!
[1590476084] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181396 kB) free!
[1590476217] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.1399.804> (10.147.212.158:51486 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590476337] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590476448] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1590476448] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (564796 kB) free!
[1590476527] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.39, 23.23, 22.40
[1590477048] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1695712 kB) free!
[1590477127] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 14.68, 17.40, 19.92
[1590477252] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27471340 kB) free.
[1590477359] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1590477479] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1590477502] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8494300 kB) free!
[1590477622] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8201252 kB) free!
[1590477648] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1355492 kB) free!
[1590477742] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (7903152 kB) free!
[1590477852] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (425672 kB) free!
[1590477972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (433480 kB) free!
[1590478079] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1590478092] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414568 kB) free!
[1590478199] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1590478327] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.23, 22.53, 20.91
[1590478554] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8658136 kB) free!
[1590478644] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (6136172 kB) free!
[1590478674] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8379912 kB) free!
[1590478794] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.4% (8081720 kB) free!
[1590478927] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1590478927] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 21.61, 17.74, 18.65
[1590479244] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1400192 kB) free!
[1590479292] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33191440 kB) free.
[1590479448] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (2120264 kB) free!
[1590479504] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1590479527] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.33, 21.28, 20.05
[1590479542] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3636340 kB) free!
[1590479892] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5376024 kB) free!
[1590479994] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4781904 kB) free!
[1590480012] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (874704 kB) free!
[1590480127] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 17.37, 16.04, 17.32
[1590480132] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (451348 kB) free!
[1590480732] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15183744 kB) free!
[1590480733] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 810 processes
[1590480853] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 811 processes
[1590480973] SERVICE ALERT: qservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 811 processes
[1590480982] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3970280 kB) free!
[1590481102] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3444380 kB) free!
[1590481222] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (2925072 kB) free!
[1590481332] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396968 kB) free!
[1590481822] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1590481822] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (438724 kB) free!
[1590482887] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/26 14:10:13 [error] 22132#0: *90972922 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: rq-1-1gbps.server, request: "GET /favicon.ico HTTP/1.1", host: "rq-1-1gbps.server:8091", referrer: "http://rq-1-1gbps.server:8091/rabbitmq/"
[1590482937] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.21069.842> (10.147.212.158:35880 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590483007] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590483044] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590483057] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590483973] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 856 processes
[1590484327] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (5454080 kB) free!
[1590484447] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 14.0% (4917768 kB) free!
[1590484567] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 12.5% (4383784 kB) free!
[1590484573] SERVICE ALERT: qservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 831 processes
[1590484799] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1590484919] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1590485168] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1984864 kB) free!
[1590485422] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.0% (16164316 kB) free.
[1590485606] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 185.255.8.58 [26/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590485726] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590485768] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.0% (31888472 kB) free.
[1590486048] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (928020 kB) free!
[1590486067] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4713972 kB) free!
[1590486187] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4464416 kB) free!
[1590486308] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 17.0% (4195636 kB) free!
[1590486373] SERVICE FLAPPING ALERT: qservercluster2;Total Processes;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1590486373] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 907 processes
[1590486648] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.0% (2280088 kB) free!
[1590486908] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.9% (22938548 kB) free.
[1590487248] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1198860 kB) free!
[1590487304] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1590487848] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2168540 kB) free!
[1590487919] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 340 processes
[1590488039] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 343 processes
[1590488159] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1590488407] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/26 15:45:42 [error] 22135#0: *91005209 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590488528] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590489057] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.12929.817> (10.147.212.158:43606 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590489132] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15520808 kB) free!
[1590489177] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590489359] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1590489373] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 552 processes
[1590489732] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387252 kB) free!
[1590490239] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.24, 14.39, 19.69
[1590490264] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < fatal_frame_error
[1590490304] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590490385] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590490848] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1378504 kB) free!
[1590491527] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590491650] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4912700 kB) free!
[1590491770] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4911828 kB) free!
[1590491890] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4909868 kB) free!
[1590492127] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/26 16:47:58 [error] 22135#0: *91026486 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/dashboard"
[1590492132] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31046272 kB) free.
[1590492248] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590492359] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1590492479] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1590492599] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590492732] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (427724 kB) free!
[1590492847] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/26 17:00:20 [error] 22135#0: *91031072 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590492852] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (412848 kB) free!
[1590492967] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590492972] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433656 kB) free!
[1590493248] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (2978172 kB) free!
[1590493567] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/26 17:14:32 [error] 22135#0: *91035978 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net"
[1590493687] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590493822] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1590494583] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (1803856 kB) free!
[1590495022] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4007452 kB) free!
[1590495142] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.0% (3814912 kB) free!
[1590495262] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.1% (3606600 kB) free!
[1590496462] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16078820 kB) free.
[1590496848] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (713368 kB) free!
[1590497173] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46605 MB (20% inode=99%):
[1590497294] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46593 MB (20% inode=99%):
[1590497414] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46587 MB (20% inode=99%):
[1590497772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31816728 kB) free.
[1590498048] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.2% (4413732 kB) free.
[1590498372] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (411108 kB) free!
[1590498373] SERVICE FLAPPING ALERT: qservercluster2;Total Processes;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1590498493] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408756 kB) free!
[1590498613] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433128 kB) free!
[1590498648] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.2% (1660860 kB) free!
[1590498768] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3393972 kB) free.
[1590499368] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181160 kB) free!
[1590499489] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3040784 kB) free!
[1590499609] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.9% (4535480 kB) free.
[1590499799] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1590499919] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590500809] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1070380 kB) free!
[1590500929] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2420876 kB) free!
[1590500994] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 33.8% (16672572 kB) free.
[1590501013] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (25415716 kB) free!
[1590501049] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3636300 kB) free.
[1590501142] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 29.4% (14481420 kB) free.
[1590501613] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416696 kB) free!
[1590501614] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 187142 MB (83% inode=99%):
[1590502213] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30410924 kB) free.
[1590502813] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (984772 kB) free!
[1590502849] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (803308 kB) free!
[1590502933] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405284 kB) free!
[1590502969] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.3% (1995704 kB) free!
[1590503053] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418004 kB) free!
[1590503090] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (3235892 kB) free!
[1590503439] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.90, 4.69, 4.94
[1590503690] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3620800 kB) free.
[1590505453] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31775972 kB) free.
[1590505919] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1590506039] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1590506054] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1744236 kB) free!
[1590506174] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (454420 kB) free!
[1590506294] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400904 kB) free!
[1590506690] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2985100 kB) free!
[1590506810] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3282844 kB) free.
[1590506894] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13480748 kB) free!
[1590507410] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (2918768 kB) free!
[1590507494] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (449696 kB) free!
[1590507530] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3256128 kB) free!
[1590507650] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (3238256 kB) free!
[1590508694] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16994660 kB) free!
[1590509294] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4874456 kB) free!
[1590509894] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25897924 kB) free!
[1590510494] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393968 kB) free!
[1590510651] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3959160 kB) free.
[1590511094] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17491288 kB) free!
[1590511251] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3249796 kB) free!
[1590511371] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3680088 kB) free.
[1590511694] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420128 kB) free!
[1590512571] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2839476 kB) free!
[1590512691] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3112048 kB) free!
[1590512811] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3332384 kB) free.
[1590512894] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30825928 kB) free.
[1590513178] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.26244.823> (10.147.212.158:47122 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590513298] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590513494] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5323948 kB) free!
[1590513614] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1176856 kB) free!
[1590513704] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1590513735] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403900 kB) free!
[1590514794] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9785012 kB) free!
[1590514914] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9762940 kB) free!
[1590514935] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31571828 kB) free.
[1590515034] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.8% (9741324 kB) free!
[1590515535] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5320872 kB) free!
[1590515655] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1341776 kB) free!
[1590515775] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401676 kB) free!
[1590515811] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2216544 kB) free!
[1590515931] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.1% (1963276 kB) free!
[1590516051] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1683312 kB) free!
[1590516975] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26261004 kB) free!
[1590517575] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3646008 kB) free!