[1551379260] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44059 MB (19% inode=99%):
[1551379272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551379380] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52251 MB (23% inode=99%):
[1551379392] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551379619] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42487 MB (20% inode=99%):
[1551379739] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46084 MB (22% inode=99%):
[1551379980] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46967 MB (20% inode=99%):
[1551379992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551380100] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46967 MB (20% inode=99%):
[1551380112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551380144] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3850696 kB) free!
[1551380220] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52277 MB (23% inode=99%):
[1551380232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551380744] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.9% (7318816 kB) free!
[1551380820] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44085 MB (19% inode=99%):
[1551380832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551380939] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42559 MB (20% inode=99%):
[1551380941] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52277 MB (23% inode=99%):
[1551380953] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551381060] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46158 MB (22% inode=99%):
[1551381552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551381672] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551381792] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551381944] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4597576 kB) free!
[1551382259] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42288 MB (20% inode=99%):
[1551382379] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46497 MB (22% inode=99%):
[1551382392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551382512] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551383112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551383144] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.0% (7392780 kB) free!
[1551383232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551383579] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42746 MB (20% inode=99%):
[1551383699] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46874 MB (22% inode=99%):
[1551383833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551383953] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551384540] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44079 MB (19% inode=99%):
[1551384553] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551384660] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44079 MB (19% inode=99%):
[1551384673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551384780] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52271 MB (23% inode=99%):
[1551384793] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551384944] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4447628 kB) free!
[1551385380] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43343 MB (19% inode=99%):
[1551385393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551385500] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52015 MB (23% inode=99%):
[1551385513] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551385545] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.6% (6221640 kB) free!
[1551386100] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46575 MB (20% inode=99%):
[1551386113] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551386220] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42479 MB (18% inode=99%):
[1551386233] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551386340] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52271 MB (23% inode=99%):
[1551386834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551386940] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44078 MB (19% inode=99%):
[1551386954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551387060] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52270 MB (23% inode=99%):
[1551387074] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551387660] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44078 MB (19% inode=99%):
[1551387674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551387780] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44078 MB (19% inode=99%):
[1551387794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551387900] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52270 MB (23% inode=99%):
[1551387914] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551388500] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44078 MB (19% inode=99%):
[1551388514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551388620] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52270 MB (23% inode=99%):
[1551388634] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1551388634] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551389221] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44078 MB (19% inode=99%):
[1551389235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551389341] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52270 MB (23% inode=99%):
[1551389355] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551389941] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44078 MB (19% inode=99%):
[1551389955] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551390061] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52270 MB (23% inode=99%):
[1551390075] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1551390675] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551390795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551390915] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551391515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551391545] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4737192 kB) free!
[1551391636] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551392145] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.3% (7551140 kB) free!
[1551392235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551392355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551392475] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551393061] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44071 MB (19% inode=99%):
[1551393075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551393181] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52263 MB (23% inode=99%):
[1551393195] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551393795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551393915] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551394036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551394635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551394755] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551395355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551395475] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551395595] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551396181] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44070 MB (19% inode=99%):
[1551396195] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551396301] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52262 MB (23% inode=99%):
[1551396315] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551396901] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44069 MB (19% inode=99%):
[1551396916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551397021] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52261 MB (23% inode=99%):
[1551397036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551397636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551397756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551397876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551398476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551398596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551399196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551399317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551399437] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551400021] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44069 MB (19% inode=99%):
[1551400037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551400141] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52261 MB (23% inode=99%):
[1551400157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551400741] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44069 MB (19% inode=99%):
[1551400757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551400861] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52261 MB (23% inode=99%):
[1551400877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551401477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551401597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551401717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551402318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551402438] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551403038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551403158] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551403278] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551403545] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551403861] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44073 MB (19% inode=99%):
[1551403878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551403981] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52265 MB (23% inode=99%):
[1551403998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551404581] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44072 MB (19% inode=99%):
[1551404598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551404701] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44072 MB (19% inode=99%):
[1551404719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551404821] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52264 MB (23% inode=99%):
[1551404839] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551405421] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44072 MB (19% inode=99%):
[1551405439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551405541] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52264 MB (23% inode=99%):
[1551405559] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551406141] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44068 MB (19% inode=99%):
[1551406159] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551406261] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44068 MB (19% inode=99%):
[1551406279] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551406381] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52260 MB (23% inode=99%):
[1551406879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551406981] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44068 MB (19% inode=99%):
[1551406999] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551407102] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52260 MB (23% inode=99%):
[1551407120] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551407701] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43812 MB (19% inode=99%):
[1551407719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551407821] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52004 MB (23% inode=99%):
[1551407839] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551408439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551408559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551408679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551409021] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3241292 kB) free!
[1551409141] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3232332 kB) free!
[1551409261] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.8% (3222060 kB) free!
[1551409660] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 320 processes
[1551411080] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1551411421] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46605 MB (20% inode=99%):
[1551411541] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46605 MB (20% inode=99%):
[1551411661] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51981 MB (23% inode=99%):
[1551411680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551411800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551411920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551414321] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1551414321] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1551414921] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1551415041] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551415161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551415261] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44007 MB (19% inode=99%):
[1551415381] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52198 MB (23% inode=99%):
[1551416745] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2234180 kB) free!
[1551417661] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1278944 kB) free!
[1551418762] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1551419145] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.7% (6735624 kB) free!
[1551419362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1551419482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1551419602] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551420602] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.15, 5.33, 5.05
[1551420722] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.90, 5.50, 5.14
[1551420842] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.41, 5.09, 5.04
[1551421060] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 322 processes
[1551421442] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.16, 4.12, 4.52
[1551421545] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3316276 kB) free!
[1551421773] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.6% (2694156 kB) free!
[1551421862] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.5% (2685396 kB) free!
[1551421893] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.6% (2540316 kB) free!
[1551422013] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.8% (2403304 kB) free!
[1551422100] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43172 MB (20% inode=99%):
[1551422220] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49894 MB (23% inode=99%):
[1551422462] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1285084 kB) free!
[1551422581] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46730 MB (20% inode=99%):
[1551422613] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1459060 kB) free!
[1551422701] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46728 MB (20% inode=99%):
[1551422745] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1551422745] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.3% (6060036 kB) free!
[1551422821] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42632 MB (18% inode=99%):
[1551423345] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4460084 kB) free!
[1551424553] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47125 MB (20% inode=99%):
[1551424673] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47125 MB (20% inode=99%):
[1551424710] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 150.47, 97.24, 43.08
[1551424793] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47061 MB (20% inode=99%):
[1551424830] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 20.62, 65.18, 37.90
[1551424950] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.32, 43.80, 33.37
[1551425221] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51044 MB (22% inode=99%):
[1551425462] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1551425462] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.2% (2627968 kB) free!
[1551425550] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 6.38, 17.76
[1551425993] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49802 MB (22% inode=99%):
[1551426345] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.3% (8533600 kB) free!
[1551426661] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9861164 kB) free!
[1551426662] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (671168 kB) free!
[1551426781] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.1% (8414652 kB) free!
[1551426901] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.6% (14582944 kB) free.
[1551427350] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.78, 0.48, 2.95
[1551427402] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551427621] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46417 MB (20% inode=99%):
[1551427732] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46689 MB (20% inode=99%):
[1551427741] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46419 MB (20% inode=99%):
[1551427852] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46689 MB (20% inode=99%):
[1551427861] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42835 MB (19% inode=99%):
[1551427972] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46688 MB (20% inode=99%):
[1551428003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551428123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551428243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551428393] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46939 MB (20% inode=99%):
[1551428513] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46917 MB (20% inode=99%):
[1551428633] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46909 MB (20% inode=99%):
[1551429061] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1551429061] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51014 MB (22% inode=99%):
[1551429062] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.8% (1750940 kB) free!
[1551429233] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1551429233] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49024 MB (21% inode=99%):
[1551429346] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4007492 kB) free!
[1551429663] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1354052 kB) free!
[1551429833] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4759336 kB) free!
[1551429953] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4703956 kB) free!
[1551430073] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4661876 kB) free!
[1551430501] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9829724 kB) free!
[1551430620] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43017 MB (20% inode=99%):
[1551430622] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8497396 kB) free!
[1551430740] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46078 MB (22% inode=99%):
[1551430742] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.2% (14397112 kB) free.
[1551431146] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.0% (9387360 kB) free!
[1551431243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551431461] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46908 MB (20% inode=99%):
[1551431581] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46908 MB (20% inode=99%):
[1551431633] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46780 MB (20% inode=99%):
[1551431775] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (30.4% change >= 20.0% threshold)
[1551431775] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1551431775] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1551431786] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (30.4% change >= 20.0% threshold)
[1551431786] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1551431786] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1551431836] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1551431836] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50487 MB (22% inode=99%):
[1551431869] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46768 MB (20% inode=99%):
[1551431989] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (27.4% change >= 20.0% threshold)
[1551431989] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46736 MB (20% inode=99%):
[1551432049] SERVICE ALERT: sig5-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (137564 kB) free!
[1551432169] SERVICE ALERT: sig5-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (119204 kB) free!
[1551432253] SERVICE ALERT: sig5-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 280 processes
[1551432289] SERVICE ALERT: sig5-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (154328 kB) free!
[1551432373] SERVICE ALERT: sig5-vm1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 280 processes
[1551432443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551432493] SERVICE ALERT: sig5-vm1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 280 processes
[1551432563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551432683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551433093] SERVICE ALERT: sig5-vm1;CPU_Procs;OK;HARD;1;CPU OK: 276 processes
[1551433634] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (28.0% change >= 20.0% threshold)
[1551433634] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551433634] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (26.1% change >= 20.0% threshold)
[1551433634] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (27.8% change >= 20.0% threshold)
[1551433644] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45865 MB (20% inode=99%):
[1551433673] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452688 kB) free!
[1551433701] SERVICE ALERT: sig5-vm2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1551433756] SERVICE ALERT: sig5-vm1;Memory;OK;HARD;1;OK - 95.5% (7022068 kB) free.
[1551433764] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45865 MB (20% inode=99%):
[1551433768] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1551433821] SERVICE ALERT: sig5-vm2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1551433835] SERVICE ALERT: sig5-vm2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1551433867] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1551433867] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551433867] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1551433867] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1551433878] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1551433878] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551433878] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1551433878] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1551433888] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1551433888] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44905 MB (19% inode=99%):
[1551433888] SERVICE ALERT: sig5-vm2;CPU_Procs;OK;HARD;1;CPU OK: 275 processes
[1551433941] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.02, 0.19
[1551433955] SERVICE ALERT: sig5-vm2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1551434037] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (129808 kB) free!
[1551434157] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (126260 kB) free!
[1551434277] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (129980 kB) free!
[1551434342] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9531992 kB) free!
[1551434462] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8167876 kB) free!
[1551434582] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.3% (13941952 kB) free.
[1551434908] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4883216 kB) free!
[1551435022] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1551435022] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1551435022] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1551435022] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1551435088] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50460 MB (22% inode=99%):
[1551435088] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551435152] SERVICE ALERT: sig6-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 275 processes
[1551435229] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1551435272] SERVICE ALERT: sig6-vm1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 275 processes
[1551435283] SERVICE ALERT: sig6-vm1;HDD_Root status;CRITICAL;SOFT;1;NRPE: Command 'check_/dev/xvda3' not defined
[1551435349] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1551435392] SERVICE ALERT: sig6-vm1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 275 processes
[1551435403] SERVICE ALERT: sig6-vm1;HDD_Root status;CRITICAL;SOFT;2;NRPE: Command 'check_/dev/xvda3' not defined
[1551435413] SERVICE ALERT: sig6-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (152636 kB) free!
[1551435469] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1551435523] SERVICE ALERT: sig6-vm1;HDD_Root status;CRITICAL;HARD;3;NRPE: Command 'check_/dev/xvda3' not defined
[1551435533] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 19.7% (1446964 kB) free!
[1551435534] SERVICE ALERT: sig6-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (115056 kB) free!
[1551435653] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.4% (1427736 kB) free!
[1551435654] SERVICE ALERT: sig6-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (121172 kB) free!
[1551435655] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 19.7% (1452328 kB) free!
[1551435688] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551435773] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 19.2% (1410664 kB) free!
[1551435775] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 19.4% (1424528 kB) free!
[1551435808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551435895] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 19.1% (1406992 kB) free!
[1551435928] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1551435992] SERVICE ALERT: sig6-vm1;CPU_Procs;OK;HARD;1;CPU OK: 273 processes
[1551436186] SERVICE ALERT: sig6-vm1;Memory;OK;HARD;1;OK - 95.6% (7028560 kB) free.
[1551436234] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1551436234] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (32.5% change >= 20.0% threshold)
[1551436234] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (25.7% change >= 20.0% threshold)
[1551436234] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1551436303] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 72.8% (35873376 kB) free.
[1551436346] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.3% (22879648 kB) free.
[1551436393] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 95.3% (7007620 kB) free.
[1551436421] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 95.4% (7014588 kB) free.
[1551436449] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.8% (7047804 kB) free.
[1551436505] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.3% (14831464 kB) free.
[1551436546] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1551436546] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.9% (14935460 kB) free.
[1551436564] SERVICE ALERT: sig6-vm1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 97885 MB (81% inode=96%):
[1551436598] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1551436598] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (31.7% change >= 20.0% threshold)
[1551436598] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1551436598] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1551436888] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46355 MB (20% inode=99%):
[1551437008] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45843 MB (20% inode=99%):
[1551437128] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41747 MB (18% inode=99%):
[1551437340] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42145 MB (20% inode=99%):
[1551437460] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 45992 MB (22% inode=99%):
[1551437582] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9404476 kB) free!
[1551437702] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7981720 kB) free!
[1551437822] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.4% (13999636 kB) free.
[1551437854] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'sig2-vm2' was out of bounds)
[1551437854] HOST ALERT: sig2-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.180)
[1551437923] HOST ALERT: sig2-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1551437974] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 94.8% (6970164 kB) free.
[1551438328] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50441 MB (22% inode=99%):
[1551438328] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551438660] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43420 MB (20% inode=99%):
[1551438780] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47123 MB (22% inode=99%):
[1551438793] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4-vm2' was out of bounds)
[1551438804] HOST ALERT: sig4-vm2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1551438819] HOST ALERT: sig4-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1551438913] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 259 processes
[1551439528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551439649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551439693] SERVICE ALERT: sig5-vm1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 272 processes
[1551439769] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551439813] SERVICE ALERT: sig5-vm1;CPU_Procs;OK;HARD;1;CPU OK: 272 processes
[1551439832] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 19.9% (1464884 kB) free!
[1551439952] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 19.8% (1455708 kB) free!
[1551439962] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (40.0% change >= 20.0% threshold)
[1551439962] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1551440027] SERVICE ALERT: sig6-vm2;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig6-vm2' was out of bounds)
[1551440072] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 19.7% (1446028 kB) free!
[1551440093] SERVICE ALERT: sig6-vm2;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig6-vm2' was out of bounds)
[1551440147] SERVICE ALERT: sig6-vm2;CPU Load;CRITICAL;SOFT;2;(Return code of 255 for service 'CPU Load' on host 'sig6-vm2' was out of bounds)
[1551440159] SERVICE ALERT: sig6-vm2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig6-vm2' was out of bounds)
[1551440202] SERVICE ALERT: sig6-vm1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1551440213] SERVICE ALERT: sig6-vm2;CPU_Procs;CRITICAL;SOFT;2;(Return code of 255 for service 'CPU_Procs' on host 'sig6-vm2' was out of bounds)
[1551440225] SERVICE ALERT: sig6-vm2;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6-vm2' was out of bounds)
[1551440233] HOST ALERT: sig6-vm1;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1551440267] SERVICE ALERT: sig6-vm2;CPU Load;CRITICAL;HARD;3;(Return code of 255 for service 'CPU Load' on host 'sig6-vm2' was out of bounds)
[1551440279] SERVICE ALERT: sig6-vm2;Current Users;CRITICAL;SOFT;2;(Return code of 255 for service 'Current Users' on host 'sig6-vm2' was out of bounds)
[1551440282] HOST ALERT: sig6-vm1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1551440291] SERVICE ALERT: sig6-vm2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6-vm2' was out of bounds)
[1551440322] SERVICE ALERT: sig6-vm1;CPU_Procs;OK;HARD;1;CPU OK: 260 processes
[1551440333] SERVICE ALERT: sig6-vm2;CPU_Procs;OK;HARD;1;CPU OK: 270 processes
[1551440345] SERVICE ALERT: sig6-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 104128 MB (86% inode=96%):
[1551440357] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (124564 kB) free!
[1551440399] SERVICE ALERT: sig6-vm2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1551440411] SERVICE ALERT: sig6-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1551440477] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (123712 kB) free!
[1551440580] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43479 MB (20% inode=99%):
[1551440585] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4485804 kB) free!
[1551440598] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (122720 kB) free!
[1551440700] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47563 MB (22% inode=99%):
[1551440706] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4457532 kB) free!
[1551440826] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4441636 kB) free!
[1551440867] SERVICE ALERT: sig6-vm2;CPU Load;OK;HARD;1;OK - load average: 0.01, 0.00, 0.00
[1551441328] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46334 MB (20% inode=99%):
[1551441449] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46334 MB (20% inode=99%):
[1551441569] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42238 MB (18% inode=99%):
[1551442022] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9665892 kB) free!
[1551442142] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8281520 kB) free!
[1551442262] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.1% (13830364 kB) free.
[1551443369] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49912 MB (22% inode=99%):
[1551443369] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551443700] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43395 MB (20% inode=99%):
[1551443820] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47488 MB (22% inode=99%):
[1551443969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551444089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551444209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551444798] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 84.8% (6239232 kB) free.
[1551445059] SERVICE ALERT: sig5-vm1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1551445089] HOST ALERT: sig5-vm1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1551445119] HOST ALERT: sig5-vm1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1551445169] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46323 MB (20% inode=99%):
[1551445179] SERVICE ALERT: sig5-vm1;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1551445289] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46323 MB (20% inode=99%):
[1551445409] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42227 MB (18% inode=99%):
[1551445862] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9632936 kB) free!
[1551445982] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8253948 kB) free!
[1551446103] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.9% (13729524 kB) free.
[1551446332] SERVICE ALERT: sig6-vm1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1551446362] HOST ALERT: sig6-vm1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1551446396] HOST ALERT: sig6-vm1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1551446452] SERVICE ALERT: sig6-vm1;CPU_Procs;OK;HARD;1;CPU OK: 262 processes
[1551446924] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.4% (22890652 kB) free.
[1551446993] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.4% (7014572 kB) free.
[1551447209] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551447209] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49902 MB (22% inode=99%):
[1551448620] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42901 MB (20% inode=99%):
[1551448740] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40853 MB (19% inode=99%):
[1551448860] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44687 MB (21% inode=99%):
[1551449009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551449130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551449250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551449753] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2987744 kB) free!
[1551449873] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.1% (2937968 kB) free!
[1551449993] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.4% (2833056 kB) free!
[1551450593] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1552252 kB) free!
[1551450809] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45837 MB (20% inode=99%):
[1551450902] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9659488 kB) free!
[1551450929] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45837 MB (20% inode=99%):
[1551451022] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8217760 kB) free!
[1551451049] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43917 MB (19% inode=99%):
[1551451142] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6538744 kB) free!
[1551451193] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 11.3% (1837892 kB) free!
[1551451742] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.4% (13522980 kB) free.
[1551451793] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1551451793] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1256004 kB) free!
[1551453449] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49795 MB (22% inode=99%):
[1551453450] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551454051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551454171] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551454291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1551454481] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.2% (14826224 kB) free.
[1551455269] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1551455849] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47108 MB (20% inode=99%):
[1551455942] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9581072 kB) free!
[1551455969] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45060 MB (20% inode=99%):
[1551456062] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8218924 kB) free!
[1551456089] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45060 MB (20% inode=99%):
[1551456182] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13382160 kB) free.
[1551457860] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 37065 MB (17% inode=99%):
[1551457980] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44229 MB (21% inode=99%):
[1551458489] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50080 MB (22% inode=99%):
[1551458491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551459091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551459211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551459331] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551459824] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 20.0% (1468916 kB) free!
[1551459944] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 19.9% (1466668 kB) free!
[1551460064] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 19.9% (1464940 kB) free!
[1551460889] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45630 MB (20% inode=99%):
[1551460982] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9828456 kB) free!
[1551461009] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46710 MB (20% inode=99%):
[1551461102] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.6% (8655620 kB) free!
[1551461129] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44662 MB (19% inode=99%):
[1551461222] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7437144 kB) free!
[1551463481] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1551463622] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.0% (13314264 kB) free.
[1551463972] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 540 processes
[1551464032] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 551 processes