[1589135824] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2939748 kB) free!
[1589135944] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3448724 kB) free.
[1589135960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8641292 kB) free!
[1589136080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33165752 kB) free.
[1589136680] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12128964 kB) free!
[1589136800] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8306480 kB) free!
[1589136920] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33066824 kB) free.
[1589137320] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (187632 kB) free!
[1589137520] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11380944 kB) free!
[1589137640] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7641900 kB) free!
[1589137744] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1589137761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31935540 kB) free.
[1589137920] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 74.2% (18187364 kB) free.
[1589138344] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1041256 kB) free!
[1589138360] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10308912 kB) free!
[1589138464] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (1028912 kB) free!
[1589138480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32936056 kB) free.
[1589138584] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (516812 kB) free!
[1589139080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12551668 kB) free!
[1589139200] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8705516 kB) free!
[1589139320] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32858600 kB) free.
[1589139920] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12135124 kB) free!
[1589140040] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8516092 kB) free!
[1589140160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32099476 kB) free.
[1589140760] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10643080 kB) free!
[1589140880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32752292 kB) free.
[1589140891] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1516888 kB) free!
[1589141011] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (494732 kB) free!
[1589141062] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224856 kB) free!
[1589141131] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1854656 kB) free!
[1589141182] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215764 kB) free!
[1589141302] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (247788 kB) free!
[1589141480] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12851764 kB) free!
[1589141600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9125968 kB) free!
[1589141720] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32515804 kB) free.
[1589142006] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.62, 6.52, 5.76
[1589142126] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.22, 6.76, 5.94
[1589142246] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.43, 6.92, 6.09
[1589142320] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11454744 kB) free!
[1589142441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32677528 kB) free.
[1589143040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13641960 kB) free!
[1589143102] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.3% (5054052 kB) free!
[1589143160] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10100424 kB) free!
[1589143280] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32629188 kB) free.
[1589143880] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12620668 kB) free!
[1589144000] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9158176 kB) free!
[1589144120] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30577140 kB) free.
[1589144584] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (2384296 kB) free!
[1589144720] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11007820 kB) free!
[1589144840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32601980 kB) free.
[1589144903] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4945832 kB) free!
[1589144974] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 80.7% (39778720 kB) free.
[1589145137] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 82.2% (40498704 kB) free.
[1589145184] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1000844 kB) free!
[1589145440] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12813328 kB) free!
[1589145560] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9361376 kB) free!
[1589145680] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30275768 kB) free.
[1589145736] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 37.4% (9155300 kB) free.
[1589145784] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1589145784] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3454696 kB) free.
[1589145846] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.64, 2.86, 4.53
[1589146236] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2338) < 139.59.21.208 [11/May/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1589146261] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1589146280] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10934876 kB) free!
[1589146357] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(34) < 172.105.59.173 [11/May/2020:03:02:32 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1589146384] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3126328 kB) free!
[1589146400] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589146400] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32591864 kB) free.
[1589146407] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(383) < 172.105.57.57 [11/May/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589146477] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589146504] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3486368 kB) free.
[1589146527] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589146542] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(504) < 2020/05/11 03:01:42 [error] 22135#0: *87580503 no live upstreams while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589146662] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589146703] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1589146703] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.1% (44407332 kB) free.
[1589146936] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (204928 kB) free!
[1589147000] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12512044 kB) free!
[1589147056] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199200 kB) free!
[1589147120] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9215516 kB) free!
[1589147176] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (2961712 kB) free!
[1589147241] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30085368 kB) free.
[1589147704] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3175684 kB) free!
[1589147776] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (220088 kB) free!
[1589147824] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3489416 kB) free.
[1589147840] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10802036 kB) free!
[1589147960] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32493484 kB) free.
[1589148560] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11966536 kB) free!
[1589148680] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32505692 kB) free.
[1589148976] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1589148976] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.0% (10057260 kB) free.
[1589149280] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13314340 kB) free!
[1589149400] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (19153796 kB) free!
[1589149520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30007664 kB) free.
[1589150120] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10678228 kB) free!
[1589150176] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (218644 kB) free!
[1589150240] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30605620 kB) free.
[1589150296] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199820 kB) free!
[1589150416] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (216464 kB) free!
[1589150840] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11149440 kB) free!
[1589150960] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30219284 kB) free.
[1589152160] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14139640 kB) free!
[1589152216] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.2% (9616896 kB) free.
[1589152280] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32545380 kB) free.
[1589152881] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14919516 kB) free!
[1589153001] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11547040 kB) free!
[1589153121] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30767456 kB) free.
[1589153416] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4441868 kB) free!
[1589153537] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (153696 kB) free!
[1589153657] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179420 kB) free!
[1589153721] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11469524 kB) free!
[1589153841] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31038216 kB) free.
[1589154441] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11101308 kB) free!
[1589154476] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.5% (11228556 kB) free.
[1589154493] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 55.4% (13567724 kB) free.
[1589154561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29014164 kB) free.
[1589154757] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.4% (18229924 kB) free.
[1589155161] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3372452 kB) free!
[1589155281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29340852 kB) free.
[1589155881] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22293872 kB) free!
[1589156001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29130412 kB) free.
[1589156904] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1589157201] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8731324 kB) free!
[1589157321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31924300 kB) free.
[1589157424] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589157921] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9431100 kB) free!
[1589158042] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33773712 kB) free.
[1589158076] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199500 kB) free!
[1589158196] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179876 kB) free!
[1589158316] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (215700 kB) free!
[1589158532] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 36.2% (47666176 kB) free.
[1589158641] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15114304 kB) free!
[1589158761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33628884 kB) free.
[1589159361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13880340 kB) free!
[1589159481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32135920 kB) free.
[1589160681] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14713144 kB) free!
[1589160801] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33172612 kB) free.
[1589161401] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13182580 kB) free!
[1589161521] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31434232 kB) free.
[1589162721] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13732776 kB) free!
[1589162841] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32441812 kB) free.
[1589163441] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12474884 kB) free!
[1589163561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30762828 kB) free.
[1589164457] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589164762] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12275884 kB) free!
[1589164882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31586376 kB) free.
[1589165482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10214628 kB) free!
[1589165602] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29669396 kB) free.
[1589166494] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4345524 kB) free!
[1589166614] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4229972 kB) free!
[1589166734] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (4121960 kB) free!
[1589166802] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10054036 kB) free!
[1589166846] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.75, 6.38, 5.34
[1589166922] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28360484 kB) free.
[1589166966] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.79, 6.65, 5.55
[1589167086] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.32, 8.05, 6.20
[1589168122] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14125656 kB) free!
[1589168242] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36159416 kB) free.
[1589168534] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1807576 kB) free!
[1589168842] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13074740 kB) free!
[1589168962] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34219076 kB) free.
[1589169562] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2571952 kB) free!
[1589169682] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35079780 kB) free.
[1589170062] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/11 09:37:20 [error] 22135#0: *87625789 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589170182] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589170282] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6903548 kB) free!
[1589170403] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2724036 kB) free!
[1589170523] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33394908 kB) free.
[1589171123] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9205024 kB) free!
[1589171243] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31995680 kB) free.
[1589171677] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(530) < 103.255.146.164 [11/May/2020:10:02:12 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589171797] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589171843] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9530600 kB) free!
[1589171963] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5260308 kB) free!
[1589172083] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28763428 kB) free.
[1589172290] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (2690352 kB) free!
[1589172410] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.5% (2686652 kB) free!
[1589172424] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3172736 kB) free!
[1589172531] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (2691432 kB) free!
[1589172544] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3120008 kB) free!
[1589172665] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2540004 kB) free!
[1589172683] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3542228 kB) free!
[1589172803] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29253880 kB) free.
[1589173253] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1589173261] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1589173357] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193268 kB) free!
[1589173373] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1589173381] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1589173403] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5195792 kB) free!
[1589173477] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184232 kB) free!
[1589173523] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (649880 kB) free!
[1589173597] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179388 kB) free!
[1589173643] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27540764 kB) free.
[1589174243] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1594316 kB) free!
[1589174363] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393572 kB) free!
[1589174483] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29419360 kB) free.
[1589175065] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1302784 kB) free!
[1589175084] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2885016 kB) free!
[1589175093] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (1757416 kB) free!
[1589175134] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (2794428 kB) free!
[1589175204] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400792 kB) free!
[1589175213] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (1132924 kB) free!
[1589175324] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32133748 kB) free.
[1589175333] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (530344 kB) free!
[1589175665] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2534152 kB) free!
[1589175735] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1589175735] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1374676 kB) free!
[1589175781] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1589175827] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 47.0% (11525964 kB) free.
[1589175883] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.6075.362> (10.147.212.158:58214 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1589175901] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1589175924] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6008264 kB) free!
[1589176003] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589176021] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1589176044] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1521936 kB) free!
[1589176164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29569488 kB) free.
[1589176532] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.0% (21142884 kB) free!
[1589176597] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(729) < 103.255.146.164 [11/May/2020:11:26:34 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589176652] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.3% (22755892 kB) free!
[1589176717] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589176764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3338868 kB) free!
[1589176772] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (31726420 kB) free.
[1589176884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384920 kB) free!
[1589177004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30724720 kB) free.
[1589177372] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.2% (24004580 kB) free!
[1589177492] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28099544 kB) free.
[1589177604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5381004 kB) free!
[1589177724] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30726292 kB) free.
[1589178324] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5246348 kB) free!
[1589178421] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1589178444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (448368 kB) free!
[1589178536] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4517372 kB) free!
[1589178564] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376028 kB) free!
[1589178656] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (4108480 kB) free!
[1589178692] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25892516 kB) free!
[1589178776] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3684564 kB) free!
[1589178812] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.2% (18715328 kB) free!
[1589178932] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28101820 kB) free.
[1589178933] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16072248 kB) free.
[1589179376] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1599732 kB) free!
[1589179427] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.1% (3946852 kB) free!
[1589179547] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.0% (3684192 kB) free!
[1589179668] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.0% (3421484 kB) free!
[1589179765] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22626260 kB) free!
[1589179865] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1589179865] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (356308 kB) free!
[1589179976] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22855500 kB) free.
[1589180132] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.8% (19468432 kB) free!
[1589180252] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.7% (24667760 kB) free!
[1589180268] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2303216 kB) free!
[1589180365] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (459476 kB) free!
[1589180372] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.9% (18267084 kB) free!
[1589180465] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1678772 kB) free!
[1589180537] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8458528 kB) free!
[1589180657] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8193480 kB) free!
[1589180777] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (7903912 kB) free!
[1589180917] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [11/May/2020:12:31:26 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, 10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 60.000, 0.000"
[1589181037] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(7) < 103.13.97.241 [11/May/2020:12:39:07 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589181158] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589181421] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1589181541] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1589181565] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1589181565] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24312524 kB) free!
[1589181572] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.4% (29585328 kB) free.
[1589181574] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9550084 kB) free!
[1589181661] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1589181665] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (832880 kB) free!
[1589181694] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9288680 kB) free!
[1589181814] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9076772 kB) free!
[1589182165] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (430720 kB) free!
[1589182172] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.3% (21533192 kB) free!
[1589182265] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.8% (1910152 kB) free!
[1589182292] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.9% (27564092 kB) free.
[1589182577] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3936212 kB) free!
[1589182892] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24380624 kB) free!
[1589183012] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (31015960 kB) free.
[1589183257] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1589183365] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25705764 kB) free!
[1589183377] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1589183465] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1195200 kB) free!
[1589183612] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.9% (23553192 kB) free!
[1589183732] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (28805084 kB) free.
[1589183965] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (428052 kB) free!
[1589184214] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4208320 kB) free!
[1589184332] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.0% (23669776 kB) free!
[1589184452] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.2% (17446124 kB) free!
[1589184565] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13663528 kB) free!
[1589184572] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.2% (23996684 kB) free!
[1589184577] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 327 processes
[1589184665] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (1793856 kB) free!
[1589184698] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1589185533] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.6% (3059380 kB) free!
[1589185653] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.3% (2762776 kB) free!
[1589185765] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408384 kB) free!
[1589185773] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.0% (2479392 kB) free!
[1589185865] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (254568 kB) free!
[1589186365] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17030880 kB) free!
[1589186373] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1589186373] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (181748 kB) free!
[1589186965] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405144 kB) free!
[1589186972] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1589186972] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12467272 kB) free!
[1589187065] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.6% (2046268 kB) free!
[1589187136] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589187572] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (27277832 kB) free.
[1589187665] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159720 kB) free!
[1589188166] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29490288 kB) free.
[1589188172] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.8% (20777364 kB) free!
[1589188265] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (2365840 kB) free!
[1589188292] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12482996 kB) free!
[1589188412] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (7042620 kB) free!
[1589188766] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (403684 kB) free!
[1589188865] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1345864 kB) free!
[1589188886] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (423852 kB) free!
[1589189006] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28120140 kB) free.
[1589189557] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(568) < 103.255.146.164 [11/May/2020:15:01:14 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589189576] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.6% (3594944 kB) free!
[1589189606] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (385988 kB) free!
[1589189677] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589189696] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 12.4% (3056312 kB) free!
[1589189726] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393252 kB) free!
[1589189816] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.5% (2601180 kB) free!
[1589189846] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386408 kB) free!
[1589189974] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16258892 kB) free.
[1589190213] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.8% (19468560 kB) free!
[1589190416] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1589190416] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (728688 kB) free!
[1589190446] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16156460 kB) free!
[1589190661] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1589190665] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2529076 kB) free!
[1589190877] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [11/May/2020:15:16:10 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.000, 0.000"
[1589190998] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589191016] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22878180 kB) free.
[1589191046] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35944612 kB) free.
[1589191597] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 203.212.70.170 [11/May/2020:15:33:50 +0530] TCP 502 0 0 60.002 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.001, 0.000"
[1589191646] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6547468 kB) free!
[1589191717] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589191766] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1909780 kB) free!
[1589191862] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1589191886] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386700 kB) free!
[1589191982] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1589192013] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (9623068 kB) free!
[1589192102] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589192317] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(202) < 103.255.146.164 [11/May/2020:15:42:55 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589192437] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589192498] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1589192618] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1589193086] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29890736 kB) free.
[1589193302] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1589193302] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1589193686] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2838596 kB) free!
[1589193807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399492 kB) free!
[1589193927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37262868 kB) free.
[1589194181] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/11 16:17:50 [error] 22135#0: *87699775 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589194301] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589194413] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.4% (17688440 kB) free!
[1589194527] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9275656 kB) free!
[1589194647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4621932 kB) free!
[1589194767] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376296 kB) free!
[1589194865] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.3% (4761040 kB) free.
[1589195367] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14573804 kB) free!
[1589195465] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.0% (1795232 kB) free!
[1589195585] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.5% (4470484 kB) free.
[1589195967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34350800 kB) free.
[1589196186] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.6% (2372720 kB) free!
[1589196306] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.0% (4231104 kB) free.
[1589196567] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5967936 kB) free!
[1589196687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1360928 kB) free!
[1589196807] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393408 kB) free!
[1589196906] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2409568 kB) free!
[1589197026] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.0% (4880396 kB) free.
[1589197502] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1589197626] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (1563460 kB) free!
[1589197746] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3140232 kB) free!
[1589197776] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.32590.5586> (10.147.212.53:35250 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1589197866] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.5% (4961964 kB) free.
[1589197896] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589198007] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32308432 kB) free.
[1589198374] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3627128 kB) free!
[1589198466] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (2001480 kB) free!
[1589198494] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.5% (3264972 kB) free!
[1589198586] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3289640 kB) free.
[1589198607] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4151168 kB) free!
[1589198613] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (27023276 kB) free.
[1589198614] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (2866596 kB) free!
[1589198727] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374696 kB) free!
[1589198847] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403428 kB) free!
[1589199213] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7074924 kB) free!
[1589199214] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1073528 kB) free!
[1589199333] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.8% (19459432 kB) free!
[1589199453] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (10955248 kB) free!
[1589199786] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (983592 kB) free!
[1589199907] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180992 kB) free!
[1589200027] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1438304 kB) free!
[1589200047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32851752 kB) free.
[1589200053] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.1% (13288964 kB) free!
[1589200415] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.6% (16305936 kB) free.
[1589200627] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3803208 kB) free.
[1589200647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4827876 kB) free!
[1589200653] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (11506904 kB) free!
[1589200767] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (413808 kB) free!
[1589200887] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405284 kB) free!
[1589201227] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (921388 kB) free!
[1589201253] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.3% (20178052 kB) free!
[1589201347] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (2686252 kB) free!
[1589201467] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3845168 kB) free.
[1589201487] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13451056 kB) free!
[1589202067] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1068800 kB) free!
[1589202087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35654028 kB) free.
[1589202101] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/11 18:25:28 [error] 22135#0: *87720723 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589202187] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2664760 kB) free!
[1589202221] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589202307] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4060492 kB) free.
[1589202416] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589202687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6595180 kB) free!
[1589202807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2108332 kB) free!
[1589202907] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (1495552 kB) free!
[1589202927] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380380 kB) free!
[1589203027] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3206272 kB) free!
[1589203147] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3973512 kB) free.
[1589203527] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13386372 kB) free!
[1589203747] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (1545096 kB) free!
[1589203867] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3235892 kB) free!
[1589203987] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3971920 kB) free.
[1589204286] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.81, 2.70, 4.31
[1589204588] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (1887496 kB) free!
[1589204708] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3673268 kB) free.
[1589204727] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (434552 kB) free!
[1589205218] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1589205302] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1589205308] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (224444 kB) free!
[1589205338] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 330 processes
[1589205428] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.0% (1952960 kB) free!
[1589205458] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1589205548] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3736548 kB) free.
[1589205814] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.6% (5705928 kB) free!
[1589206148] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (220816 kB) free!
[1589206268] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.8% (1762720 kB) free!
[1589206388] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3633856 kB) free.
[1589207014] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4722348 kB) free!
[1589207128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34134948 kB) free.
[1589207253] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.8% (30111272 kB) free.
[1589207728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3757748 kB) free!
[1589207848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (441704 kB) free!
[1589207853] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.1% (22574308 kB) free!
[1589207887] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.08, 6.34, 5.44
[1589207968] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396204 kB) free!
[1589207973] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.7% (15361428 kB) free!
[1589208007] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.81, 5.56, 5.28
[1589208093] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.7% (24597940 kB) free!
[1589208127] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.55, 5.16, 5.16
[1589208568] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14867792 kB) free!
[1589208727] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.54, 3.86, 4.48
[1589208788] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.7% (1903556 kB) free!
[1589208908] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3645256 kB) free.
[1589209168] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37838780 kB) free.
[1589209293] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (29724972 kB) free.
[1589209327] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.45, 5.41, 5.01
[1589209447] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.86, 5.31, 5.03
[1589209567] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.68, 5.04, 4.96
[1589209768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7762040 kB) free!
[1589209888] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3244568 kB) free!
[1589209893] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.2% (23980112 kB) free!
[1589210008] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408560 kB) free!
[1589210013] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.3% (32068872 kB) free.
[1589210102] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1589210108] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (1400496 kB) free!
[1589210228] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (2714764 kB) free!
[1589210348] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.3% (3946320 kB) free.
[1589210608] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16960548 kB) free!
[1589210613] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.6% (24548228 kB) free!
[1589210733] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.2% (31947276 kB) free.
[1589211089] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 26963 MB (20% inode=99%):
[1589211208] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38894908 kB) free.
[1589211209] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 26963 MB (20% inode=99%):
[1589211215] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589211329] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 26963 MB (20% inode=99%):
[1589211333] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.3% (24116312 kB) free!
[1589211453] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.0% (31605676 kB) free.
[1589211548] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (457856 kB) free!
[1589211668] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.5% (2037140 kB) free!
[1589211788] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3371904 kB) free.
[1589211808] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11429592 kB) free!
[1589211928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (7008168 kB) free!
[1589212048] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2862908 kB) free!
[1589212053] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24421300 kB) free!
[1589212173] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.1% (30492844 kB) free.
[1589212648] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21382092 kB) free!
[1589213249] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387912 kB) free!
[1589213374] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (25957556 kB) free!
[1589213494] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.0% (31685092 kB) free.
[1589213849] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15839812 kB) free!
[1589214188] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3091240 kB) free!
[1589214309] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589214309] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3422844 kB) free.
[1589214449] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35266860 kB) free.
[1589215049] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9504488 kB) free!
[1589215169] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5387544 kB) free!
[1589215289] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1452852 kB) free!
[1589215508] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2096784 kB) free!
[1589215628] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2845584 kB) free!
[1589215748] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (3256432 kB) free!
[1589215889] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16050492 kB) free!
[1589216348] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3567988 kB) free.
[1589216489] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (36982920 kB) free.
[1589217089] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10906496 kB) free!
[1589217209] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6738644 kB) free!
[1589217329] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2714764 kB) free!
[1589217929] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17388864 kB) free!
[1589218237] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [11/May/2020:22:59:31 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1589218357] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589218529] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34151816 kB) free.
[1589218748] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3184600 kB) free!
[1589218868] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3547432 kB) free.
[1589218957] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 203.212.70.170 [11/May/2020:23:12:33 +0530] TCP 502 0 0 60.001 "10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 60.000, 0.000"
[1589219078] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589219129] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9364380 kB) free!
[1589219249] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5297520 kB) free!
[1589219369] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1415672 kB) free!
[1589219970] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15142808 kB) free!
[1589220570] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30106344 kB) free.
[1589220669] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (2258132 kB) free!
[1589220694] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589220789] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2414156 kB) free!
[1589220909] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (2491664 kB) free!
[1589221170] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7523716 kB) free!
[1589221290] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3743648 kB) free!
[1589221410] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36820348 kB) free.