[1568399529] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7658328 kB) free!
[1568399649] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.5% (41586356 kB) free.
[1568400249] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21917204 kB) free!
[1568400369] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18068360 kB) free!
[1568400489] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14201628 kB) free!
[1568401089] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 24.8% (32725592 kB) free.
[1568401689] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13403052 kB) free!
[1568401809] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9541228 kB) free!
[1568401929] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 32.0% (42203556 kB) free.
[1568402529] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24325316 kB) free!
[1568402649] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20497708 kB) free!
[1568402769] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16663072 kB) free!
[1568403369] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 26.7% (35208712 kB) free.
[1568403969] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15707824 kB) free!
[1568404090] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11809604 kB) free!
[1568404210] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (7936948 kB) free!
[1568404810] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 20.1% (26473308 kB) free.
[1568405156] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.38, 6.32, 5.31
[1568405276] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.24, 6.11, 5.37
[1568405396] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.92, 6.35, 5.54
[1568405410] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7066664 kB) free!
[1568405530] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.1% (41025860 kB) free.
[1568406130] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21643192 kB) free!
[1568406250] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 13.5% (17754604 kB) free!
[1568406370] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13865684 kB) free!
[1568406970] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 24.6% (32448500 kB) free.
[1568407570] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12884960 kB) free!
[1568407690] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8949656 kB) free!
[1568407733] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (4004908 kB) free!
[1568407810] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.7% (41763932 kB) free.
[1568408024] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231612 kB) free!
[1568408396] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.43, 2.34, 3.79
[1568408410] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23690708 kB) free!
[1568408530] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19835884 kB) free!
[1568408650] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15959032 kB) free!
[1568408933] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1630112 kB) free!
[1568409224] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 76.5% (37699616 kB) free.
[1568409251] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 26.3% (34604744 kB) free.
[1568409851] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15058324 kB) free!
[1568409971] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11196276 kB) free!
[1568410091] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (7079744 kB) free!
[1568410691] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24336012 kB) free!
[1568411291] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.5% (41540932 kB) free.
[1568411375] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 45.4% (11131468 kB) free.
[1568411435] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 46.5% (11403612 kB) free.
[1568411891] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23326208 kB) free!
[1568412011] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19457052 kB) free!
[1568412131] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15626472 kB) free!
[1568412731] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 26.2% (34564020 kB) free.
[1568413331] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13551768 kB) free!
[1568413451] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9209828 kB) free!
[1568413571] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.5% (41467160 kB) free.
[1568414771] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8899240 kB) free!
[1568414891] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.5% (41460020 kB) free.
[1568416091] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10549880 kB) free!
[1568416211] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6697440 kB) free!
[1568416331] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.4% (41436464 kB) free.
[1568417531] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10138356 kB) free!
[1568417651] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6297520 kB) free!
[1568417771] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.4% (41428192 kB) free.
[1568418371] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25648972 kB) free!
[1568418491] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (21799892 kB) free!
[1568418611] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17947956 kB) free!
[1568419211] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.4% (41412332 kB) free.
[1568419811] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25319508 kB) free!
[1568419931] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 16.4% (21601396 kB) free!
[1568420051] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17790804 kB) free!
[1568420651] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.4% (41403816 kB) free.
[1568421252] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25244660 kB) free!
[1568421372] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21404928 kB) free!
[1568421492] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17576856 kB) free!
[1568422092] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 27.5% (36302872 kB) free.
[1568422199] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 435 processes
[1568422259] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 433 processes
[1568422692] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16727244 kB) free!
[1568422812] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12858900 kB) free!
[1568422932] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (8992588 kB) free!
[1568423532] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 21.0% (27736212 kB) free.
[1568424132] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7907812 kB) free!
[1568424252] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.2% (41068616 kB) free.
[1568424359] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 463 processes
[1568424420] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 464 processes
[1568424578] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179460 kB) free!
[1568424636] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187276 kB) free!
[1568424699] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200208 kB) free!
[1568424756] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196868 kB) free!
[1568424820] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191008 kB) free!
[1568424852] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21852072 kB) free!
[1568424876] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197516 kB) free!
[1568424972] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17975108 kB) free!
[1568425092] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14106852 kB) free!
[1568425692] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 25.1% (33067776 kB) free.
[1568426292] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13800904 kB) free!
[1568426413] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9939500 kB) free!
[1568426533] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6026984 kB) free!
[1568427133] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24931640 kB) free!
[1568427733] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (5600380 kB) free!
[1568428333] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24518124 kB) free!
[1568428933] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5198024 kB) free!
[1568429533] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24368488 kB) free!
[1568430133] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (5028000 kB) free!
[1568430733] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23578464 kB) free!
[1568431333] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3717944 kB) free!
[1568431933] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22586764 kB) free!
[1568432533] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 20.7% (27342244 kB) free.
[1568433133] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20277952 kB) free!
[1568433253] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16307008 kB) free!
[1568433373] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12337996 kB) free!
[1568433973] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 23.9% (31491144 kB) free.
[1568434574] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11523044 kB) free!
[1568434694] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7632688 kB) free!
[1568434814] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3727728 kB) free!
[1568435414] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23969156 kB) free!
[1568436014] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4032352 kB) free!
[1568436614] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24367488 kB) free!
[1568437214] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4554332 kB) free!
[1568437814] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25861212 kB) free!
[1568438414] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (5534472 kB) free!
[1568438503] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1568438623] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1568439014] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 20.1% (26554808 kB) free.
[1568439111] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568439231] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1568439351] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1568439614] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6333216 kB) free!
[1568439734] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2362332 kB) free!
[1568439824] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9363740 kB) free!
[1568439854] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385328 kB) free!
[1568439944] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9006904 kB) free!
[1568440064] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (8674144 kB) free!
[1568440454] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20611484 kB) free!
[1568441064] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1568441084] HOST ALERT: rq-2;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1568441511] SERVICE ALERT: rq-2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'rq-2' was out of bounds)
[1568441616] SERVICE ALERT: rq-2;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'rq-2' was out of bounds)
[1568441616] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'rq-2' was out of bounds)
[1568441616] SERVICE ALERT: rq-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'rq-2' was out of bounds)
[1568441626] SERVICE ALERT: rq-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'rq-2' was out of bounds)
[1568441626] SERVICE ALERT: rq-2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'rq-2' was out of bounds)
[1568441626] SERVICE ALERT: rq-2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'rq-2' was out of bounds)
[1568441629] SERVICE ALERT: rq-2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'rq-2' was out of bounds)
[1568441652] HOST ALERT: rq-2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.204)
[1568441864] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4041964 kB) free!
[1568442004] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1568442124] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 166541 MB (74% inode=99%):
[1568442225] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.48, 10.92, 6.95
[1568442345] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.92, 8.45, 6.52
[1568442351] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1568442464] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.5% (7628936 kB) free!
[1568442465] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.28, 6.64, 6.09
[1568443065] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.83, 3.84, 4.77
[1568443552] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568443665] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1568443665] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4637384 kB) free!
[1568443672] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1568443792] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1568444146] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.68, 5.22, 5.01
[1568444266] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.93, 5.14, 5.00
[1568445466] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.96, 5.59, 5.19
[1568445586] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.79, 6.55, 5.60
[1568445706] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.75, 6.75, 5.80
[1568447392] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1568447392] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1568447994] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568448115] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568449045] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.5% (22671972 kB) free.
[1568449100] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.5% (40676620 kB) free.
[1568449158] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.2% (45422400 kB) free.
[1568449210] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.4% (22883360 kB) free.
[1568449255] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.1% (22966620 kB) free.
[1568449347] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.13, 9.53, 8.09
[1568449467] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.09, 7.03, 7.35
[1568449588] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.64, 5.27, 6.67
[1568450188] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.10, 2.36, 4.47
[1568450465] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164328 kB) free!
[1568450585] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186208 kB) free!
[1568450705] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181292 kB) free!
[1568451706] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.68, 4.03, 4.60
[1568452467] SERVICE FLAPPING ALERT: rq-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568452588] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 125.38, 54.93, 22.15
[1568452654] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1568452708] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 134.78, 81.92, 36.02
[1568452752] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.9% (22928040 kB) free.
[1568452828] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 135.56, 99.74, 48.12
[1568453254] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 587 processes
[1568454563] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1568454628] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1568454628] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.97, 2.32, 12.87
[1568455792] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1568455828] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.75, 1.87, 4.89
[1568456254] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1568457028] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 56.16, 14.39, 6.82
[1568457148] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 32.95, 27.64, 13.20
[1568457245] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 825 processes
[1568457268] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.59, 18.96, 11.77
[1568457365] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 815 processes
[1568457485] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 344 processes
[1568458100] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8646588 kB) free!
[1568458220] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9146280 kB) free!
[1568458340] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (9147408 kB) free!
[1568458468] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.43, 1.91, 4.36
[1568459996] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.76, 9.76, 6.53
[1568460116] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.96, 6.76, 5.82
[1568460236] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 4.67, 5.16
[1568460454] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4394048 kB) free!
[1568460574] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (4316248 kB) free!
[1568460610] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4599516 kB) free!
[1568460694] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (3994572 kB) free!
[1568460730] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4227016 kB) free!
[1568460836] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.93, 1.35, 3.08
[1568460850] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (4033508 kB) free!
[1568461757] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.3% (3276356 kB) free!
[1568461877] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.0% (3216064 kB) free!
[1568461894] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2118008 kB) free!
[1568461997] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (3146664 kB) free!
[1568462597] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1698544 kB) free!
[1568462650] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2071984 kB) free!
[1568462855] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1568463455] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1568463455] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1568464163] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9795040 kB) free!
[1568464283] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9676524 kB) free!
[1568464314] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1568464403] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (9476356 kB) free!
[1568464434] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568465258] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 300 processes
[1568465378] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1568466140] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4249240 kB) free!
[1568466456] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1568467656] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1568468003] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4655144 kB) free!
[1568468256] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1568468563] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.3% (23013472 kB) free.
[1568468600] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.8% (22748684 kB) free.
[1568468641] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40414760 kB) free.
[1568468677] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 47.9% (23615968 kB) free.
[1568468716] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.9% (23024808 kB) free.
[1568469163] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.7% (3370024 kB) free!
[1568469268] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568469283] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (3031448 kB) free!
[1568469403] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1568469403] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1745288 kB) free!
[1568469456] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1568470056] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1568472456] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1568473940] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.8% (22895892 kB) free.
[1568477028] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.68, 16.12, 9.60
[1568477068] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.09, 8.49, 6.14
[1568477148] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.27, 11.11, 8.56
[1568477188] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.86, 5.98, 5.51
[1568477268] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.82, 7.67, 7.60
[1568477308] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.98, 4.33, 4.96
[1568477869] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.89, 1.85, 4.45
[1568480544] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1568480634] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1568480754] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568482056] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1568482108] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.95, 12.18, 7.40
[1568482228] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.14, 8.44, 6.61
[1568482349] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 5.77, 5.86
[1568482723] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1568482949] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.58, 1.34, 3.37
[1568483323] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1568483636] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 34.18, 17.04, 6.77
[1568483756] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.77, 11.47, 5.97
[1568483856] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 514 processes
[1568483876] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.78, 7.74, 5.27
[1568483923] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 335 processes
[1568484043] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 339 processes
[1568484163] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 1 warn out of 342 processes
[1568484476] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.52, 1.50, 3.00
[1568485423] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1568485543] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes