[1597603123] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3466980 kB) free.
[1597603135] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597603202] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597603221] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (240212 kB) free!
[1597603255] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195548 kB) free!
[1597603256] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597603261] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=8.60% system=5.44% iowait=5.23% idle=80.73%
[1597603323] SERVICE ALERT: pservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597603343] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188468 kB) free!
[1597603345] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 8.75, 33.36, 42.48
[1597603376] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187828 kB) free!
[1597603376] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 2 processes with STATE = Z
[1597603447] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176308 kB) free!
[1597603463] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179156 kB) free!
[1597603496] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187444 kB) free!
[1597603568] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.7% (20167968 kB) free.
[1597603631] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597603660] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597603712] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1503 processes
[1597603757] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 9472 open files (0% of max 2424156)
[1597603780] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597603833] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1503 processes
[1597603855] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597603945] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1597603945] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 74.08, 44.95, 40.19
[1597603953] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1497 processes
[1597603975] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 337 processes
[1597604096] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.2% (20138916 kB) free.
[1597604545] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 43.16, 45.51, 42.87
[1597604664] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.0% (22298204 kB) free.
[1597605745] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 62.39, 48.26, 45.42
[1597606353] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1503 processes
[1597606415] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (1916640 kB) free!
[1597606473] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1504 processes
[1597606476] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (183928 kB) free!
[1597606536] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (190616 kB) free!
[1597606594] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1505 processes
[1597606596] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (184464 kB) free!
[1597606773] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.9% (3150088 kB) free!
[1597606893] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199544 kB) free!
[1597606897] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 68.1% (16696060 kB) free.
[1597607013] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174552 kB) free!
[1597607613] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 25.4% (6235996 kB) free.
[1597607656] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381540 kB) free!
[1597608145] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 31.97, 40.07, 42.94
[1597608247] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227680 kB) free!
[1597608367] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (221008 kB) free!
[1597608394] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1493 processes
[1597608488] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (245140 kB) free!
[1597608813] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178224 kB) free!
[1597608933] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.0% (21813424 kB) free.
[1597609345] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 35.18, 45.31, 45.60
[1597609533] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (253556 kB) free!
[1597609653] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182600 kB) free!
[1597609774] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174028 kB) free!
[1597609945] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 36.62, 43.70, 44.24
[1597610288] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 43.7% (21546112 kB) free.
[1597610545] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 74.85, 49.43, 45.15
[1597610664] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175400 kB) free!
[1597610784] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.5% (21455688 kB) free.
[1597610974] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1597610974] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (4290928 kB) free!
[1597611145] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 47.50, 42.10, 42.72
[1597611384] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 71.9% (35420876 kB) free.
[1597611574] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180576 kB) free!
[1597612174] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (3519908 kB) free!
[1597612272] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 71.8% (35399020 kB) free.
[1597612774] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.2% (23087508 kB) free.
[1597614985] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597615105] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1597615174] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (989012 kB) free!
[1597615294] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.3% (20653332 kB) free.
[1597615345] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 53.12, 47.68, 46.54
[1597616545] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 22.59, 33.81, 41.09
[1597617796] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16030.20, 9125.75, 3782.40
[1597617917] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16266.04, 11500.44, 5300.49
[1597618037] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16297.45, 13089.27, 6634.45
[1597618345] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 53.57, 53.13, 45.49
[1597618377] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597618497] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 3 processes with STATE = Z
[1597618945] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 54.47, 43.66, 42.42
[1597619545] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.61, 49.99, 46.91
[1597620745] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 35.19, 37.09, 42.96
[1597621037] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.72, 3.06, 551.28
[1597622545] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 42.32, 46.54, 46.02
[1597623471] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623596] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623660] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623686] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623695] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185872 kB) free!
[1597623698] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623721] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623727] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623785] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623811] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623815] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198448 kB) free!
[1597623823] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623834] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623852] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623872] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623877] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623910] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623930] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623935] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194524 kB) free!
[1597623936] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623943] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623947] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623959] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623967] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623972] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623977] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623993] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597623994] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1504 processes
[1597623997] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624002] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624055] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624072] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624084] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624092] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624097] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624114] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1507 processes
[1597624117] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624118] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624127] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624180] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624197] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624217] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624222] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624234] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1509 processes
[1597624243] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597624516] SERVICE ALERT: pservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.27% system=3.97% iowait=0.45% idle=94.30%
[1597624535] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.0% (23048256 kB) free.
[1597624536] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 321 processes
[1597624543] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1597624684] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1597624717] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 9952 open files (0% of max 2440681)
[1597624727] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39288 MB (76% inode=99%):
[1597624780] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1597624797] SERVICE ALERT: pservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1597624817] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1597624822] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.2% (9917244 kB) free.
[1597624843] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 202036 MB (91% inode=99%):
[1597624921] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.50, 1.23, 1.46
[1597624956] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 103.255.146.164 [17/Aug/2020:06:03:43 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1597625056] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.0% (18467596 kB) free!
[1597625076] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597625289] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625322] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625332] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625385] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625402] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625409] SERVICE ALERT: pservercluster1;Sensors Status;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1597625417] SERVICE ALERT: pservercluster1;CPU_Procs;UNKNOWN;SOFT;1;Could not open pipe: /usr/bin/ps
[1597625422] SERVICE ALERT: pservercluster1;Memory;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1597625442] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 10272 open files (0% of max 2440681)
[1597625452] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1597625505] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;UNKNOWN;SOFT;2;Could not open pipe: /usr/bin/ps
[1597625521] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1597625522] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1597625534] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625542] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625542] SERVICE ALERT: pservercluster1;Memory;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1597625550] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1597625577] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625630] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625642] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625646] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625656] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446888 kB) free!
[1597625662] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625667] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625670] SERVICE ALERT: sservercluster2;Sensors Status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1597625674] SERVICE ALERT: sservercluster2;Total Processes;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1597625701] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625709] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625711] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625721] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625741] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625748] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625758] SERVICE ALERT: sservercluster2;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1597625771] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625790] SERVICE ALERT: sservercluster2;Sensors Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1597625794] SERVICE ALERT: sservercluster2;Total Processes;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1597625826] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625831] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625834] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625846] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625866] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625873] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625878] SERVICE ALERT: sservercluster2;HDD_Home status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1597625890] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625919] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625933] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625951] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625954] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625956] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625971] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625971] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625989] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625991] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597625998] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626003] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626009] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626015] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626043] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626047] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626058] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626071] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626087] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626096] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626114] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626134] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626140] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626167] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626168] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626183] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626196] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626212] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626221] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626240] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626256] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.8% (32754608 kB) free.
[1597626259] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626288] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626292] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626321] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626337] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597626395] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597627176] SERVICE FLAPPING ALERT: pservercluster1;CPU-STATISTICS;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1597627176] SERVICE ALERT: pservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.06% system=3.29% iowait=0.00% idle=95.64%
[1597627191] SERVICE FLAPPING ALERT: pservercluster1;Total Processes;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1597627191] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 322 processes
[1597627198] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1597627198] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1597627334] SERVICE FLAPPING ALERT: pservercluster1;Sensors Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1597627334] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1597627377] SERVICE FLAPPING ALERT: pservercluster1;HDD_Root status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1597627377] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39275 MB (76% inode=99%):
[1597627430] SERVICE FLAPPING ALERT: pservercluster1;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1597627430] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1597627442] SERVICE FLAPPING ALERT: pservercluster1;zombie_procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1597627442] SERVICE ALERT: pservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1597627462] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1597627462] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.8% (13040852 kB) free.
[1597627467] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1597627467] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1597627488] SERVICE FLAPPING ALERT: pservercluster1;HDD_Home status;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1597627488] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 201821 MB (91% inode=99%):
[1597627492] SERVICE FLAPPING ALERT: pservercluster1;Open-Files;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1597627492] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 9344 open files (0% of max 2440681)
[1597627571] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1597627571] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.55, 74.52, 69.99
[1597628177] SERVICE ALERT: pservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597631319] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 329 processes
[1597631352] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.0% (18639972 kB) free.
[1597631355] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.30, 1.62, 1.62
[1597631357] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 32866 MB (64% inode=98%):
[1597631403] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 192599 MB (87% inode=99%):
[1597631545] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.40% system=5.27% iowait=0.01% idle=93.31%
[1597631583] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1597631621] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1597631640] SERVICE ALERT: sservercluster2;Open-Files;OK;HARD;1;OK: 10592 open files (0% of max 2424092)
[1597631659] SERVICE ALERT: sservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1597631737] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1597631795] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1597632322] SERVICE ALERT: sservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597632745] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1597633155] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15961.32, 10968.47, 4970.69
[1597633280] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597633350] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597633388] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1597633400] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16010.77, 13791.22, 7531.61
[1597633475] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=5.79% system=6.56% iowait=0.03% idle=87.62%
[1597633508] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 404 processes
[1597633522] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;SOFT;1;(3) < Aug 17 08:33:03 sservercluster2 xinetd[14876]: nrpe: fork failed: Cannot allocate memory (errno = 12)
[1597633628] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 356 processes
[1597633642] SERVICE ALERT: sservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597634228] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1597634228] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1597634435] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1499 processes
[1597635035] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1502 processes
[1597635155] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1503 processes
[1597635276] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1507 processes
[1597635876] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1597635876] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1495 processes
[1597635935] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1597636637] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16255.67, 12891.58, 6461.04
[1597636757] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16260.85, 14006.61, 7649.34
[1597636878] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16262.92, 14753.22, 8693.74
[1597636975] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 924 processes
[1597637095] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 340 processes
[1597637534] SERVICE FLAPPING ALERT: pservercluster1;Sensors Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1597637577] SERVICE FLAPPING ALERT: pservercluster1;HDD_Root status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1597637630] SERVICE FLAPPING ALERT: pservercluster1;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1597637642] SERVICE FLAPPING ALERT: pservercluster1;zombie_procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1597637662] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1597637667] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1597637771] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1597637982] SERVICE FLAPPING ALERT: pservercluster1;CPU-STATISTICS;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1597637991] SERVICE FLAPPING ALERT: pservercluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1597637998] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1597638200] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.68, 3.10, 600.92
[1597638289] SERVICE FLAPPING ALERT: pservercluster1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1597638292] SERVICE FLAPPING ALERT: pservercluster1;Open-Files;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1597638800] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1597638800] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.78, 2.22, 315.94
[1597638876] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1505 processes
[1597638997] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1505 processes
[1597639117] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1499 processes
[1597639717] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1502 processes
[1597639838] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1496 processes
[1597639878] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.96, 2.33, 441.39
[1597640437] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1502 processes
[1597640557] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1497 processes
[1597643323] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3145996 kB) free!
[1597643443] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3076416 kB) free!
[1597643564] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (2136712 kB) free!
[1597643771] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 7190.93, 1798.24, 602.74
[1597643891] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15026.55, 6579.42, 2499.76
[1597644011] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16082.71, 9778.39, 4166.70
[1597644164] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179188 kB) free!
[1597645628] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1597647164] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (1708928 kB) free!
[1597647757] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1597647764] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1597647764] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1567576 kB) free!
[1597647877] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(48) < 103.16.101.72 [17/Aug/2020:12:29:21 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1597647997] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597648211] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.71, 1.70, 211.86
[1597648597] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(44) < 103.16.101.72 [17/Aug/2020:12:36:52 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1597648717] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597650200] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1597650792] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.9% (6513344 kB) free!
[1597650912] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.6% (6420868 kB) free!
[1597651032] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.3% (6330348 kB) free!
[1597651272] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9690136 kB) free!
[1597651392] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9548728 kB) free!
[1597651513] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.1% (9414536 kB) free!
[1597653412] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 319 processes
[1597653532] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1597653652] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 322 processes
[1597654252] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1597654513] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4814732 kB) free!
[1597654957] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1503 processes
[1597655077] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1500 processes
[1597657852] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 320 processes
[1597657972] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1597659164] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1597659432] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (3237976 kB) free!
[1597660477] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1503 processes
[1597660598] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1504 processes
[1597660718] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1499 processes
[1597661318] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1504 processes
[1597661438] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1500 processes
[1597661572] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 325 processes
[1597661692] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 322 processes
[1597661784] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9525652 kB) free!
[1597661812] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1597661904] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9441720 kB) free!
[1597662024] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9358312 kB) free!
[1597662164] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (1719596 kB) free!
[1597662930] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4913752 kB) free!
[1597663050] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.8% (4861956 kB) free!
[1597663170] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.9% (4874200 kB) free!
[1597663365] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1621936 kB) free!
[1597663965] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (1865116 kB) free!
[1597664212] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 322 processes
[1597664332] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1597665639] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1501 processes
[1597665759] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1496 processes
[1597665765] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1597665765] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1612612 kB) free!
[1597666824] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4512292 kB) free!
[1597666965] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (1904416 kB) free!
[1597667332] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1597667452] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1597668165] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1119728 kB) free!
[1597668765] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (2548824 kB) free!
[1597669138] SERVICE ALERT: sig1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3079708 kB) free!
[1597669258] SERVICE ALERT: sig1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3078452 kB) free!
[1597669365] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1206688 kB) free!
[1597669378] SERVICE ALERT: sig1-vm2;Memory;WARNING;HARD;3;WARNING - 18.8% (3078312 kB) free!
[1597669965] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (2074360 kB) free!
[1597671052] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 325 processes
[1597671165] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1585028 kB) free!
[1597671172] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1597671765] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2852220 kB) free!
[1597673228] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 377 processes
[1597673348] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 429 processes
[1597673391] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1488 processes
[1597673469] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 437 processes
[1597673511] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1483 processes
[1597673565] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (946248 kB) free!
[1597673631] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 327 processes
[1597674069] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1597674166] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.9% (2094512 kB) free!
[1597674669] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1597674766] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (562132 kB) free!
[1597674789] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1597675366] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (3032084 kB) free!
[1597677256] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11601036 kB) free!
[1597677376] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11597108 kB) free!
[1597677496] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (11598736 kB) free!
[1597678632] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 45.1% (14765948 kB) free.
[1597679566] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1571512 kB) free!
[1597680166] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.4% (2334036 kB) free!
[1597681959] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1511 processes
[1597681966] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3431700 kB) free.
[1597682079] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1516 processes
[1597682199] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1521 processes
[1597683517] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(712) < 172.105.51.131 [17/Aug/2020:22:28:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1597683637] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1597683766] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3057360 kB) free!
[1597683886] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2607796 kB) free!
[1597684006] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2158248 kB) free!
[1597684606] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3260736 kB) free.
[1597685206] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3230840 kB) free!
[1597685326] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3208064 kB) free!
[1597685446] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (2961464 kB) free!