[1568659226] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188176 kB) free!
[1568659346] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199468 kB) free!
[1568659466] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192416 kB) free!
[1568661764] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.1% (10396608 kB) free.
[1568662664] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.7% (18071224 kB) free!
[1568662784] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.1% (18591484 kB) free!
[1568662904] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.1% (18589776 kB) free!
[1568662965] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (1259116 kB) free!
[1568663085] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (407556 kB) free!
[1568663205] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184628 kB) free!
[1568663516] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568663636] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568664030] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.28, 7.39, 2.97
[1568664150] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 9.55, 9.69, 4.57
[1568668333] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.7% (46187696 kB) free.
[1568670266] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 49.2% (12062960 kB) free.
[1568670507] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 54.6% (13375196 kB) free.
[1568672036] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1568672156] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568674063] SERVICE ALERT: sig4-vm2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 12072 MB (9% inode=99%):
[1568682405] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.3% (12417412 kB) free.
[1568683143] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 497 processes
[1568683204] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 495 processes
[1568683264] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 496 processes
[1568684805] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (1393140 kB) free!
[1568684908] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190900 kB) free!
[1568684925] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194368 kB) free!
[1568685028] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169292 kB) free!
[1568685046] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192040 kB) free!
[1568685148] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186292 kB) free!
[1568685267] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182936 kB) free!
[1568685356] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1568685391] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165688 kB) free!
[1568685476] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568685511] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195520 kB) free!
[1568686550] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 118.54, 31.97, 11.73
[1568686670] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 108.89, 56.93, 23.33
[1568686790] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 107.60, 73.63, 33.53
[1568686864] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 492 processes
[1568686925] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 485 processes
[1568686950] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 106.67, 81.15, 40.39
[1568687070] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 106.46, 89.55, 48.42
[1568687190] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 106.76, 95.26, 55.50
[1568687548] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2641168 kB) free!
[1568688590] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.90, 1.28, 11.23
[1568689348] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2397160 kB) free!
[1568689592] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.43, 3.20, 18.80
[1568689791] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.97, 0.76, 3.63
[1568691392] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.97, 1.75, 4.67
[1568695676] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1568695742] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.6% (23098696 kB) free.
[1568695790] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.2% (23081096 kB) free.
[1568695796] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1568695852] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 93.7% (22970604 kB) free.
[1568695964] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 89.2% (117636256 kB) free.
[1568696104] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40430692 kB) free.
[1568696170] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.0% (23049168 kB) free.
[1568698801] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4837800 kB) free!
[1568698921] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4712216 kB) free!
[1568699041] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (4542224 kB) free!
[1568699348] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176272 kB) free!
[1568699468] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178644 kB) free!
[1568699588] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182788 kB) free!
[1568700472] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22938192 kB) free.
[1568700841] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2352472 kB) free!
[1568701326] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1568701326] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.2% (22844604 kB) free.
[1568703112] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568703232] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1568703352] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1568703698] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 308 processes
[1568703818] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1568703824] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1568703884] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1568703944] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1568703992] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.16, 22.46, 14.84
[1568704004] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1568704112] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.98, 15.29, 13.14
[1568704191] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.62, 18.18, 18.22
[1568704232] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 10.49, 11.64
[1568704311] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.46, 12.66, 16.19
[1568704431] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.85, 8.72, 14.32
[1568705746] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.13, 6.44, 5.26
[1568705774] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4558112 kB) free!
[1568705866] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.13, 7.02, 5.63
[1568705894] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (4323268 kB) free!
[1568705986] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.30, 6.29, 5.53
[1568706014] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4117324 kB) free!
[1568706032] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.78, 1.25, 3.14
[1568706231] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.89, 1.28, 3.24
[1568706733] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9230492 kB) free!
[1568706854] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8797608 kB) free!
[1568706974] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.1% (8424740 kB) free!
[1568707676] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (4378704 kB) free!
[1568707796] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4250708 kB) free!
[1568707814] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2068216 kB) free!
[1568707916] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4146720 kB) free!
[1568708152] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1568708174] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4579980 kB) free!
[1568708204] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 52 users currently logged in
[1568708534] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.9% (47265516 kB) free.
[1568708611] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.9% (22780936 kB) free.
[1568709217] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 549 processes
[1568709717] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1568709717] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1270828 kB) free!
[1568710271] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.6% (23091432 kB) free.
[1568711214] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4052980 kB) free!
[1568711334] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.2% (3583748 kB) free!
[1568711455] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.8% (3258116 kB) free!
[1568711533] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568711986] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.44, 4.91, 4.94
[1568712055] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1676256 kB) free!
[1568712423] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.8% (23056808 kB) free.
[1568712944] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1568713064] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1568713324] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 492 processes
[1568713346] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.4% (7582088 kB) free!
[1568713384] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 489 processes
[1568713466] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.9% (7333580 kB) free!
[1568713516] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9770316 kB) free!
[1568713586] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1568713586] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.4% (7097156 kB) free!
[1568713636] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9756240 kB) free!
[1568713756] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 20.7% (10178012 kB) free.
[1568713933] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4460204 kB) free!
[1568714053] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4385068 kB) free!
[1568714173] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4319448 kB) free!
[1568714357] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9316460 kB) free!
[1568714474] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.6% (40227912 kB) free.
[1568714550] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.3% (47452224 kB) free.
[1568714600] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.8% (22512168 kB) free.
[1568716275] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.9% (2678380 kB) free!
[1568716395] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.3% (2548208 kB) free!
[1568716515] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2425728 kB) free!
[1568716786] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.84, 6.57, 5.02
[1568716906] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.65, 6.42, 5.14
[1568717026] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.40, 6.44, 5.29
[1568717076] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.4% (22800512 kB) free.
[1568720016] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4519664 kB) free!
[1568720018] SERVICE ALERT: smscapp1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 552 processes
[1568720084] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4932704 kB) free!
[1568720137] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4427328 kB) free!
[1568720138] SERVICE ALERT: smscapp1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 551 processes
[1568720204] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (4787540 kB) free!
[1568720257] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (4320292 kB) free!
[1568720258] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 551 processes
[1568720324] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4683572 kB) free!
[1568720858] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 545 processes
[1568720924] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204960 kB) free!
[1568721428] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.1% (3323640 kB) free!
[1568721548] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.0% (2839776 kB) free!
[1568721594] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (5926992 kB) free!
[1568721668] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.3% (2432948 kB) free!
[1568721715] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.2% (5683692 kB) free!
[1568721835] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (5384200 kB) free!
[1568721903] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22840840 kB) free.
[1568722047] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1568722047] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.2% (22901076 kB) free.
[1568722162] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 94.3% (33026004 kB) free.
[1568722221] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.7% (22234688 kB) free.
[1568723263] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 115460 MB (89% inode=99%):
[1568723996] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1568724116] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568724754] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568725426] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.03, 3.84, 4.82
[1568726554] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9722532 kB) free!
[1568726674] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9488652 kB) free!
[1568726794] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9251336 kB) free!
[1568727032] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 67.70, 16.99, 6.38
[1568727152] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.02, 30.47, 13.30
[1568727204] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4327304 kB) free!
[1568727272] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.68, 20.79, 11.84
[1568727324] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (4197388 kB) free!
[1568727444] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (4121792 kB) free!
[1568727685] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9085468 kB) free!
[1568727805] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8991560 kB) free!
[1568727925] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (8898724 kB) free!
[1568728472] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.07, 1.25, 3.90
[1568728504] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173120 kB) free!
[1568728525] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239128 kB) free!
[1568728624] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179376 kB) free!
[1568728744] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188788 kB) free!
[1568729194] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215552 kB) free!
[1568729258] SERVICE ALERT: smscapp1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 552 processes
[1568729379] SERVICE ALERT: smscapp1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 551 processes
[1568729499] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 551 processes
[1568729844] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2234412 kB) free!
[1568730099] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1568730099] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 549 processes
[1568730699] SERVICE ALERT: smscapp1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 552 processes
[1568730819] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 548 processes
[1568731426] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.29, 7.10, 5.80
[1568731435] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 330 processes
[1568731546] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.21, 7.69, 6.19
[1568731658] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568731666] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.93, 6.55, 5.94
[1568732635] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 326 processes
[1568734064] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 328 processes
[1568734066] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.72, 2.92, 4.15
[1568734184] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176120 kB) free!
[1568734184] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1568734304] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174852 kB) free!
[1568734424] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175368 kB) free!
[1568734745] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1568735037] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.36, 118.68, 76.54
[1568735157] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 9.45, 81.23, 67.92
[1568735277] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.72, 55.90, 60.27
[1568735497] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: /home 46201 MB (20% inode=99%):
[1568735617] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: /home 46201 MB (20% inode=99%):
[1568735737] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: /home 46201 MB (20% inode=99%):
[1568736477] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 1.43, 16.99
[1568737677] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.33, 4.94
[1568737916] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568738036] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1568741619] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)