[1560105223] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.8% (2885376 kB) free!
[1560105299] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 357 processes
[1560105467] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (1915236 kB) free!
[1560105587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10214868 kB) free.
[1560105816] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.7% (2229800 kB) free!
[1560105882] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.46, 5.17, 18.59
[1560106499] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 387 processes
[1560107016] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1545368 kB) free!
[1560107344] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.2% (5479572 kB) free.
[1560107358] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 114.03, 282.64, 165.96
[1560107478] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.78, 189.28, 145.90
[1560107598] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.30, 126.72, 128.23
[1560107616] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.3% (2637744 kB) free!
[1560107682] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.23, 1.46, 3.93
[1560107987] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2392864 kB) free!
[1560108051] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1560108107] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10562108 kB) free.
[1560108299] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 371 processes
[1560109398] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 0.55, 18.69
[1560109416] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (579980 kB) free!
[1560109499] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 396 processes
[1560109851] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560109971] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560110091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560110617] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.6% (4644088 kB) free.
[1560111107] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3114700 kB) free!
[1560111198] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.25, 0.33, 2.98
[1560111227] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11236228 kB) free.
[1560111291] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560111827] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3448604 kB) free!
[1560111891] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560111947] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11224016 kB) free.
[1560112011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560112131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560112417] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (966544 kB) free!
[1560112537] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167656 kB) free!
[1560112657] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1348252 kB) free!
[1560113623] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2453768 kB) free!
[1560113699] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 380 processes
[1560113747] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2464632 kB) free!
[1560113867] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11226788 kB) free.
[1560114299] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 376 processes
[1560114457] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (1690688 kB) free!
[1560115067] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (2557784 kB) free!
[1560115132] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560115187] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11223844 kB) free.
[1560115323] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 197.15, 145.76, 69.33
[1560115443] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 34.41, 104.42, 63.78
[1560115563] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.08, 70.03, 56.10
[1560115732] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560115787] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3167956 kB) free!
[1560115852] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1560115907] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10991260 kB) free.
[1560116099] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 400 processes
[1560116699] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 393 processes
[1560116763] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.32, 1.59, 15.65
[1560117052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560117172] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560117544] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4924160 kB) free!
[1560117664] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4915244 kB) free!
[1560117707] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5127200 kB) free!
[1560117772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560117785] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4910372 kB) free!
[1560117827] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (2805068 kB) free!
[1560117882] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560117892] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560117948] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11179560 kB) free.
[1560117963] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.27, 4.50
[1560118493] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560118547] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5371464 kB) free!
[1560118563] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 505.47, 291.21, 123.63
[1560118613] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560118667] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (3025192 kB) free!
[1560118683] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 512.23, 364.69, 170.88
[1560118733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1560118787] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11192664 kB) free.
[1560118803] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1560118803] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 114.24, 280.39, 165.00
[1560119257] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (3288856 kB) free.
[1560119333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560119387] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5521860 kB) free!
[1560119453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560119507] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (3192628 kB) free!
[1560119573] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560119627] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11188696 kB) free.
[1560119699] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 381 processes
[1560120173] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1560120227] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3537684 kB) free!
[1560120293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560120347] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11179724 kB) free.
[1560120414] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560120947] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4305404 kB) free!
[1560121013] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560121067] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11185788 kB) free.
[1560121133] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560121203] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.15, 2.98, 16.96
[1560121499] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 341 processes
[1560121667] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6449308 kB) free!
[1560121733] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560121788] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4578792 kB) free!
[1560121853] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560121908] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11187516 kB) free.
[1560121973] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560122023] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (3641340 kB) free!
[1560122082] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.91, 12.85, 12.44
[1560122202] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.80, 9.18, 11.15
[1560122322] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.62, 6.63, 9.97
[1560122403] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.40, 4.90
[1560122508] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6631556 kB) free!
[1560122573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560122628] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4812904 kB) free!
[1560122694] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560122699] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 385 processes
[1560122748] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2510520 kB) free!
[1560122814] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560123348] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6743440 kB) free!
[1560123414] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560123522] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.38, 1.53, 3.79
[1560123534] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560123654] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1560125099] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 406 processes
[1560125623] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2330084 kB) free!
[1560126055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560126175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560126295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560126457] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1560127122] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 132.22, 110.54, 59.08
[1560127165] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 44.6% (10921860 kB) free.
[1560127242] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 33.47, 85.39, 56.53
[1560127362] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.07, 57.76, 49.90
[1560128099] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 386 processes
[1560128562] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1560128562] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.71, 2.73, 14.92
[1560128699] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 375 processes
[1560129094] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1560130362] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.58, 1.62, 3.50
[1560130399] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 93.99, 28.76, 10.24
[1560130519] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 131.07, 64.49, 25.57
[1560130639] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 135.48, 88.22, 38.98
[1560131839] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.33, 2.73, 14.73
[1560132295] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1560133039] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.25, 4.22
[1560133203] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560133413] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4878564 kB) free!
[1560133499] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 379 processes
[1560133533] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4861332 kB) free!
[1560133653] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4848060 kB) free!
[1560134100] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 402 processes
[1560138295] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1560138348] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9896440 kB) free.
[1560138785] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2276496 kB) free!
[1560138948] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9874504 kB) free!
[1560139068] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7612684 kB) free!
[1560139188] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5664280 kB) free!
[1560139203] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 509.96, 322.51, 140.15
[1560139323] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 513.29, 385.81, 185.46
[1560139443] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 114.58, 294.58, 177.84
[1560139586] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1560139706] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 361 processes
[1560139826] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 6 crit, 0 warn out of 365 processes
[1560140100] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 435 processes
[1560140257] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.9% (2748464 kB) free!
[1560140377] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.2% (3933416 kB) free.
[1560140426] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1560140562] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560141026] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1560141146] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1560141295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560141301] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 442 processes
[1560141415] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560141453] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2019320 kB) free!
[1560141535] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560141577] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.0% (1939132 kB) free!
[1560141660] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1560141697] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (3332004 kB) free.
[1560141780] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1560141843] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 0.82, 14.08
[1560141900] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1560142297] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.0% (2912808 kB) free!
[1560142346] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 367 processes
[1560142362] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 473.17, 204.37, 78.37
[1560142417] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.5% (2359188 kB) free!
[1560142467] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1560142482] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 509.41, 307.13, 131.30
[1560142500] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1560142537] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.2% (1814964 kB) free!
[1560142602] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 514.16, 375.87, 177.81
[1560142767] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.81, 6.42, 5.35
[1560142887] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.29, 6.16, 5.39
[1560143007] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.80, 5.74, 5.34
[1560143043] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 0.64, 4.35
[1560143066] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1560143137] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1067224 kB) free!
[1560143186] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1560143607] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 557 processes
[1560143667] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 562 processes
[1560143737] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.9% (2087700 kB) free!
[1560143972] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (1794836 kB) free!
[1560144094] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 9 users currently logged in
[1560144214] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 9 users currently logged in
[1560144334] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1560144386] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 4 crit, 0 warn out of 366 processes
[1560144506] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 365 processes
[1560144573] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (188080 kB) free!
[1560144626] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1560145226] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1560145244] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.5% (2024780 kB) free!
[1560145346] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1560145537] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1560145537] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (148984 kB) free!
[1560145773] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (2370136 kB) free!
[1560145844] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (916384 kB) free!
[1560146137] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.3% (2975308 kB) free!
[1560146373] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1560146373] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (149968 kB) free!
[1560146546] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 366 processes
[1560146667] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1560146737] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (154352 kB) free!
[1560146787] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1560146935] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1560146973] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (2830028 kB) free!
[1560147134] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (1352120 kB) free!
[1560147255] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.1% (1037532 kB) free!
[1560147337] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.3% (2809036 kB) free!
[1560147375] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (731144 kB) free!
[1560147402] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.54, 2.58, 17.34
[1560147573] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165024 kB) free!
[1560147588] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9961612 kB) free.
[1560147937] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (3813464 kB) free.
[1560147987] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 3 warn out of 370 processes
[1560148107] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 4 crit, 1 warn out of 371 processes
[1560148173] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (2405716 kB) free!
[1560148227] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 368 processes
[1560148537] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.2% (2466012 kB) free!
[1560148575] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.7% (6889088 kB) free.
[1560148657] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (3807808 kB) free.
[1560148735] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560148773] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161076 kB) free!
[1560148856] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560148976] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1560148976] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560149202] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.42, 1.51, 3.86
[1560149257] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (163460 kB) free!
[1560149377] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (182184 kB) free!
[1560149388] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (306512 kB) free!
[1560149427] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1560149427] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1560149497] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.9% (2745316 kB) free!
[1560149509] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9654688 kB) free!
[1560149629] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1560149629] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7089340 kB) free!
[1560150027] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 5 crit, 0 warn out of 370 processes
[1560150097] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (299360 kB) free!
[1560150147] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 368 processes
[1560150268] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 3 warn out of 368 processes
[1560150301] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 421 processes
[1560150901] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 397 processes
[1560151173] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (1958020 kB) free!
[1560151298] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (3651720 kB) free.
[1560151773] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (177632 kB) free!
[1560151898] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (1305168 kB) free!
[1560151976] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1560152018] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173604 kB) free!
[1560152029] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9919060 kB) free.
[1560152138] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.3% (2652020 kB) free!
[1560152373] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (2131556 kB) free!
[1560152668] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1560152701] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 414 processes
[1560152738] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (472116 kB) free!
[1560152973] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161572 kB) free!
[1560153301] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 389 processes
[1560153334] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1560153777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560153807] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.09, 4.71, 4.96
[1560153897] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560154017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560154407] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.98, 6.43, 5.75
[1560154429] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233652 kB) free!
[1560154527] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.60, 6.55, 5.88
[1560154538] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (3448636 kB) free.
[1560154549] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10009636 kB) free.
[1560154647] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.27, 6.30, 5.86
[1560155175] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (1210504 kB) free!
[1560155295] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.5% (1068168 kB) free!
[1560155415] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.5% (918228 kB) free!
[1560155738] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (154112 kB) free!
[1560155847] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.78, 4.28, 4.83
[1560155858] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.9% (1766964 kB) free!
[1560155978] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.5% (2994148 kB) free!
[1560156015] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1560156015] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (117808 kB) free!
[1560156447] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.01, 5.68, 5.39
[1560156567] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.91, 5.77, 5.49
[1560156579] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (3298972 kB) free.
[1560156687] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1560156687] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.73, 5.51, 5.43
[1560156949] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234508 kB) free!
[1560157017] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1560157069] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10312416 kB) free.
[1560157179] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3175776 kB) free!
[1560157299] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (401784 kB) free!
[1560157419] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (3066884 kB) free!
[1560158217] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560158338] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560158458] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560158619] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169776 kB) free!
[1560158668] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1560158788] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 368 processes
[1560158869] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (245728 kB) free!
[1560158908] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 367 processes
[1560158973] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (3680632 kB) free.
[1560158989] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10446464 kB) free.
[1560159015] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.3% (6863400 kB) free.
[1560159508] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1560159687] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.61, 4.34, 4.88
[1560159819] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.0% (2762332 kB) free!
[1560160108] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1560160174] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181036 kB) free!
[1560160228] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1560160294] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3025436 kB) free!
[1560160348] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1560160414] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (212908 kB) free!
[1560160419] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179640 kB) free!
[1560160639] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 248.65, 64.39, 22.05
[1560160759] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 474.32, 211.68, 81.16
[1560160879] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 505.75, 310.59, 133.23
[1560160887] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.57, 5.69, 5.04
[1560160949] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1560161007] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.62, 6.23, 5.30
[1560161019] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.8% (4344760 kB) free.
[1560161127] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.61, 7.24, 5.77
[1560161202] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.48, 8.99, 5.42
[1560161322] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.08, 6.57, 4.96
[1560161549] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 370 processes
[1560161619] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (149660 kB) free!
[1560161669] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 369 processes
[1560161700] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1560161739] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (531220 kB) free!
[1560161789] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1560161821] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1560161859] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.5% (2994868 kB) free!
[1560161941] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1560161989] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238348 kB) free!
[1560162058] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560162109] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1560162109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10593164 kB) free.
[1560162389] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 369 processes
[1560162459] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156808 kB) free!
[1560162509] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1560162843] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 481.52, 222.24, 86.10
[1560162927] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 6.52, 4.76, 4.74
[1560162963] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 507.28, 317.91, 137.67
[1560163083] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 511.09, 381.99, 183.00
[1560163109] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 6 warn out of 383 processes
[1560163230] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 6 warn out of 382 processes
[1560163258] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560163279] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 0.64, 16.09
[1560163350] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 380 processes
[1560163378] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560163498] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560163950] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 373 processes
[1560164259] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.6% (2523140 kB) free!
[1560164479] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.21, 0.87, 4.89
[1560164509] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (477116 kB) free!
[1560164550] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1560164614] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (3853192 kB) free.
[1560164629] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10054444 kB) free.
[1560164701] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560164859] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156332 kB) free!
[1560165150] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 367 processes
[1560165214] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (215124 kB) free!
[1560165229] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4151068 kB) free!
[1560165270] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1560165301] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1560165335] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (3605904 kB) free.
[1560165349] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (2519780 kB) free!
[1560165459] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.0% (2916332 kB) free!
[1560165469] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (1606720 kB) free!
[1560165483] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.37, 1.18, 18.13
[1560165870] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 5 crit, 0 warn out of 369 processes
[1560165901] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 353 processes
[1560165934] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (165624 kB) free!
[1560165990] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 371 processes
[1560166054] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (3442404 kB) free.
[1560166059] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160152 kB) free!
[1560166069] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5578728 kB) free!
[1560166110] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1560166669] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (1040064 kB) free!
[1560166711] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 6 warn out of 405 processes
[1560166741] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 15 users currently logged in
[1560166815] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.0% (1246960 kB) free!
[1560166831] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 3 warn out of 409 processes
[1560166936] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (1048924 kB) free!
[1560166951] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 2 warn out of 409 processes
[1560167056] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.7% (859576 kB) free!
[1560167127] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.50, 5.22, 5.01
[1560167159] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.04, 9.00, 6.05
[1560167247] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.64, 5.08, 4.99
[1560167254] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (170812 kB) free!
[1560167269] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1560167269] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5928848 kB) free!
[1560167279] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.44, 9.20, 6.48
[1560167283] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 215.24, 158.25, 77.56
[1560167374] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (3619984 kB) free.
[1560167400] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.75, 8.31, 6.49
[1560167552] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 403 processes
[1560167656] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (119288 kB) free!
[1560167859] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (3441884 kB) free.
[1560167869] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4391720 kB) free!
[1560167974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (159944 kB) free!
[1560168094] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (188312 kB) free!
[1560168152] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 396 processes
[1560168214] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (213832 kB) free!
[1560168417] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3226704 kB) free!
[1560168447] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.40, 5.43, 5.09
[1560168459] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (155224 kB) free!
[1560168469] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6290460 kB) free!
[1560168537] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3112028 kB) free!
[1560168567] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.27, 5.67, 5.21
[1560168579] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (1552704 kB) free!
[1560168657] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.7% (3034384 kB) free!
[1560168687] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.30, 5.30, 5.15
[1560168700] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161876 kB) free!
[1560169083] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1560169083] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 0.94, 12.05
[1560169200] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.81, 3.86, 4.48
[1560169287] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.59, 4.96, 4.86
[1560169498] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1560169722] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 511.46, 414.50, 212.65
[1560169842] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 242.39, 375.35, 224.78
[1560169857] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.9% (13642480 kB) free.
[1560169952] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1560169962] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.11, 251.71, 197.73
[1560170056] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.6% (6886028 kB) free.
[1560170283] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.96, 0.58, 3.63
[1560170487] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.40, 6.44, 5.49
[1560170509] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.4% (4123332 kB) free.
[1560170552] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 6 crit, 0 warn out of 381 processes
[1560170607] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.87, 5.97, 5.43
[1560170672] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 6 crit, 0 warn out of 381 processes
[1560170698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560170701] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1560170727] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.12, 5.79, 5.44
[1560170792] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 6 crit, 0 warn out of 382 processes
[1560170818] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560170938] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560171109] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178292 kB) free!
[1560171229] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.7% (4169324 kB) free.
[1560171301] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1560171301] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 348 processes
[1560171392] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 379 processes
[1560171469] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10040884 kB) free.
[1560171829] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.6% (2203692 kB) free!
[1560171949] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (149260 kB) free!
[1560171993] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 380 processes
[1560172069] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.3% (4107400 kB) free.
[1560172113] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1560172362] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.53, 1.94, 16.79
[1560172669] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (765216 kB) free!
[1560172790] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (150036 kB) free!
[1560172910] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.3% (4096868 kB) free.
[1560173479] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.05, 35.18, 34.07
[1560173510] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.9% (2909712 kB) free!
[1560173600] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.24, 24.32, 30.23
[1560173614] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (3539672 kB) free.
[1560173630] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (156716 kB) free!
[1560173720] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.14, 17.05, 26.85
[1560173750] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166332 kB) free!
[1560173869] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (582840 kB) free!
[1560173938] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1560173989] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10945808 kB) free.
[1560174162] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.02, 2.97, 5.00
[1560174214] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (177052 kB) free!
[1560174320] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1560174320] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.03, 3.20, 14.61
[1560174327] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.44, 4.19, 4.61
[1560174335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (154756 kB) free!
[1560174455] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (4092232 kB) free.
[1560174538] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560174658] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1560174778] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560174927] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.95, 6.70, 5.85
[1560175047] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.25, 5.82, 5.63
[1560175055] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169696 kB) free!
[1560175167] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.95, 4.92, 5.33
[1560175175] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (157996 kB) free!
[1560175189] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243900 kB) free!
[1560175295] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (233516 kB) free!
[1560175309] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10618708 kB) free.
[1560175520] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.50, 1.02, 4.91
[1560175794] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1560175914] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1560176034] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1560176314] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1560176434] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1560176495] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (2513608 kB) free!
[1560176656] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (1196012 kB) free!
[1560176667] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 589 processes
[1560176727] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 584 processes
[1560176776] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.8% (1013484 kB) free!
[1560176787] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 579 processes
[1560176896] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.2% (824408 kB) free!
[1560177096] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165460 kB) free!
[1560177110] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232324 kB) free!
[1560177230] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10209008 kB) free.
[1560177496] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.1% (6696536 kB) free.
[1560177901] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 405 processes
[1560178107] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (136204 kB) free!
[1560178167] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 6.31, 4.96, 4.94
[1560178227] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (118248 kB) free!
[1560178347] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (129176 kB) free!
[1560178430] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (812384 kB) free!
[1560178550] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11002400 kB) free.
[1560179101] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 406 processes
[1560179496] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (2045332 kB) free!
[1560179750] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (1380764 kB) free!
[1560179870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11206384 kB) free.
[1560180096] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172012 kB) free!
[1560180470] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2259108 kB) free!
[1560180590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11246724 kB) free.
[1560180934] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1560181083] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560181190] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1599064 kB) free!
[1560181296] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (1672120 kB) free!
[1560181310] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1560181310] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10860684 kB) free.
[1560181897] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (153948 kB) free!
[1560182368] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.20, 5.03, 5.26
[1560182488] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 0.97, 3.74, 4.77
[1560182511] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (686696 kB) free!
[1560182579] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1560182631] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10228416 kB) free.
[1560182701] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 370 processes
[1560183034] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560183179] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560183299] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560183302] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 343 processes
[1560183351] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1560183419] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560183831] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2440836 kB) free!
[1560183950] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.1% (1804604 kB) free!
[1560183951] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11185888 kB) free.
[1560184347] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 95.3% (7008432 kB) free.
[1560184550] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155700 kB) free!
[1560184551] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2021412 kB) free!
[1560184619] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560184671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10678776 kB) free.
[1560185488] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.70, 6.40, 5.47
[1560185497] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (3008080 kB) free!
[1560185608] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.59, 5.28, 5.16
[1560185728] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.86, 6.03, 5.43
[1560185819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560185871] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (1505784 kB) free!
[1560185939] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560185991] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11565756 kB) free.
[1560186097] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161024 kB) free!
[1560186320] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560186328] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.69, 3.51, 4.30
[1560186540] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560186591] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2957228 kB) free!
[1560186660] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560186711] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (869620 kB) free!
[1560186780] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560186832] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11587764 kB) free.
[1560186902] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1560186928] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.53, 6.18, 5.02
[1560186950] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.0% (2429248 kB) free!
[1560187048] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.17, 5.64, 4.96
[1560187380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560187431] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (1304004 kB) free!
[1560187500] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560187502] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 357 processes
[1560187550] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1560187550] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.0% (3893592 kB) free.
[1560187551] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (253152 kB) free!
[1560187620] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560187671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11495944 kB) free.
[1560188034] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1560188150] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (165720 kB) free!
[1560188220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560188270] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174224 kB) free!
[1560188283] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 133.13, 96.37, 39.13
[1560188340] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560188390] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156868 kB) free!
[1560188403] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 18.50, 64.68, 34.46
[1560188460] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1560188497] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (2986452 kB) free!
[1560188523] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.07, 43.48, 30.35
[1560188871] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1796508 kB) free!
[1560188896] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.9% (1097580 kB) free!
[1560188991] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11161016 kB) free.
[1560189016] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.3% (979392 kB) free!
[1560189097] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172568 kB) free!
[1560189123] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.85, 6.36, 16.16
[1560189136] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.6% (850284 kB) free!
[1560189736] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (638756 kB) free!
[1560190190] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.1% (2129936 kB) free!
[1560190323] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.51, 4.70
[1560190362] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 495.89, 251.76, 101.17
[1560190482] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 511.25, 338.46, 151.20
[1560190502] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 377 processes
[1560190602] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 513.55, 396.49, 195.17
[1560190790] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168924 kB) free!
[1560190791] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3753348 kB) free!
[1560190911] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11701168 kB) free.
[1560191102] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 398 processes