[1558377749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558377791] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7618320 kB) free!
[1558377869] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558377912] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12772068 kB) free.
[1558377989] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558378511] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7954708 kB) free!
[1558378589] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558378631] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12786052 kB) free.
[1558378709] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558379231] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8073256 kB) free!
[1558379309] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558379351] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12787104 kB) free.
[1558379429] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1558379951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8262352 kB) free!
[1558380030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558380071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12775276 kB) free.
[1558380150] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558380602] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1558381271] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6223952 kB) free!
[1558381391] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12778516 kB) free.
[1558381950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558381991] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6396660 kB) free!
[1558382070] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558382112] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12766260 kB) free.
[1558382670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558382711] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7951036 kB) free!
[1558382790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558382831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12778248 kB) free.
[1558382911] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558383431] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8405780 kB) free!
[1558383510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558383551] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12780828 kB) free.
[1558383630] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1558384151] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8508016 kB) free!
[1558384230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1558384271] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6395368 kB) free!
[1558384350] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558384391] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12778088 kB) free.
[1558384950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558384991] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7199644 kB) free!
[1558385070] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558385111] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12764600 kB) free.
[1558385670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558385711] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7508288 kB) free!
[1558385791] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1558385831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12776740 kB) free.
[1558386002] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 267.74, 117.31, 44.74
[1558386122] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 286.59, 174.26, 74.40
[1558386243] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 265.91, 207.56, 98.92
[1558386390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558386432] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7739832 kB) free!
[1558386510] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558386552] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12780452 kB) free.
[1558386607] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2774588 kB) free!
[1558387110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558387152] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7913836 kB) free!
[1558387230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558387272] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12770488 kB) free.
[1558387350] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558387872] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8061264 kB) free!
[1558387950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558387992] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558387992] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12776708 kB) free.
[1558388043] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.22, 0.73, 14.48
[1558388070] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558388592] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8229248 kB) free!
[1558388671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558388712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12758252 kB) free.
[1558388791] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1558389243] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.48, 0.35, 4.20
[1558389312] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8321416 kB) free!
[1558389377] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 259.50, 100.31, 37.85
[1558389391] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558389432] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12766028 kB) free.
[1558389497] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 294.54, 166.36, 69.62
[1558389511] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558389617] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 299.02, 210.51, 97.53
[1558390032] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8431036 kB) free!
[1558390111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558390153] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12768236 kB) free.
[1558390231] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1558390254] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1558390375] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1558390752] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8576444 kB) free!
[1558390831] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558390872] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6479356 kB) free!
[1558390952] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558390992] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12765432 kB) free.
[1558391551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558391592] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6661120 kB) free!
[1558391671] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558391712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12750816 kB) free.
[1558392017] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.53, 1.76, 15.81
[1558392271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558392312] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6821528 kB) free!
[1558392391] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558392432] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12758808 kB) free.
[1558392991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558393032] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6943540 kB) free!
[1558393111] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558393152] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12762340 kB) free.
[1558393711] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558393752] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6938164 kB) free!
[1558393817] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.30, 1.38, 3.53
[1558393832] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1558393872] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12751148 kB) free.
[1558394407] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2435380 kB) free!
[1558394431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558394473] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7164652 kB) free!
[1558394551] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558394593] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12759624 kB) free.
[1558395151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558395193] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7243224 kB) free!
[1558395271] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558395313] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12750284 kB) free.
[1558395871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558395913] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7675328 kB) free!
[1558395991] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1558396033] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12755236 kB) free.
[1558396592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558396633] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7885836 kB) free!
[1558396712] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558396753] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12759252 kB) free.
[1558397312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558397432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558397552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558398553] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8477696 kB) free!
[1558398674] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6292068 kB) free!
[1558398794] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11873184 kB) free.
[1558399893] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 4 crit, 0 warn out of 563 processes
[1558399953] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 560 processes
[1558399994] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9271644 kB) free!
[1558400114] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7714800 kB) free!
[1558400234] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12669520 kB) free.
[1558400553] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 557 processes
[1558400613] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 554 processes
[1558402952] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558403234] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9118760 kB) free!
[1558403354] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7892256 kB) free!
[1558403475] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6247092 kB) free!
[1558403552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558403672] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558403793] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558404681] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 163.03, 193.14, 99.81
[1558404766] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2881832 kB) free!
[1558404801] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 22.34, 129.34, 87.75
[1558404921] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.37, 86.67, 77.15
[1558405275] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12483968 kB) free.
[1558406475] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9452448 kB) free!
[1558406595] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7956764 kB) free!
[1558406715] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12460228 kB) free.
[1558406721] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.48, 0.68, 11.49
[1558407915] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9704528 kB) free!
[1558407921] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.56, 3.58
[1558408035] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8199340 kB) free!
[1558408156] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12468404 kB) free.
[1558408366] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2408188 kB) free!
[1558409043] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 298.39, 193.96, 85.43
[1558409163] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 276.04, 224.16, 109.87
[1558409283] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 37.43, 150.09, 96.58
[1558409407] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2731264 kB) free!
[1558409955] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8190776 kB) free!
[1558410017] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 20.16, 15.28, 8.18
[1558410050] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5581 MB (10% inode=97%):
[1558410075] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5943404 kB) free!
[1558410137] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 19.63, 16.72, 9.58
[1558410195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11768668 kB) free.
[1558410257] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.82, 13.24, 9.23
[1558410366] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2547932 kB) free!
[1558410607] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2019828 kB) free!
[1558410650] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7033 MB (13% inode=98%):
[1558410966] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2205228 kB) free!
[1558411083] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.26, 0.64, 14.15
[1558411395] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8659848 kB) free!
[1558411458] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.02, 2.11, 3.85
[1558411515] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7169484 kB) free!
[1558411635] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12362104 kB) free.
[1558412030] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 358 processes
[1558412150] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1558412270] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558412283] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.35, 4.18
[1558412835] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9495072 kB) free!
[1558412956] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7950136 kB) free!
[1558413076] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5818384 kB) free!
[1558413676] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1558413676] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11713700 kB) free.
[1558414876] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9599660 kB) free!
[1558414996] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8218916 kB) free!
[1558415116] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6903592 kB) free!
[1558416470] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 4 crit, 0 warn out of 359 processes
[1558416590] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558416710] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558416993] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558417310] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 360 processes
[1558417430] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 2 warn out of 365 processes
[1558417516] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12090412 kB) free.
[1558417550] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 2 warn out of 363 processes
[1558417594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558417714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558417834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558418150] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1558419916] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9280128 kB) free!
[1558419950] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 360 processes
[1558420036] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7710340 kB) free!
[1558420071] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 360 processes
[1558420156] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5579856 kB) free!
[1558420191] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 362 processes
[1558420234] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1558420756] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11793036 kB) free.
[1558420791] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1558420791] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1558420834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558420954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558421075] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1558421075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558421568] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.7% (15070280 kB) free.
[1558421626] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.9% (22032480 kB) free.
[1558421678] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (14786156 kB) free.
[1558421735] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22430236 kB) free.
[1558421787] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21945856 kB) free.
[1558422779] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168960 kB) free!
[1558422900] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (150964 kB) free!
[1558422996] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (4091048 kB) free!
[1558423020] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165156 kB) free!
[1558423117] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (4007020 kB) free!
[1558423237] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (3928140 kB) free!
[1558423756] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9770820 kB) free!
[1558423837] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 23.5% (5807608 kB) free.
[1558423876] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8011676 kB) free!
[1558423997] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6313080 kB) free!
[1558424140] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186252 kB) free!
[1558424263] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190624 kB) free!
[1558424383] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166204 kB) free!
[1558424437] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192576 kB) free!
[1558424557] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163316 kB) free!
[1558424597] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11926080 kB) free.
[1558424626] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188644 kB) free!
[1558424677] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1558424677] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192400 kB) free!
[1558424746] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166172 kB) free!
[1558424866] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164376 kB) free!
[1558425797] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9858604 kB) free!
[1558425917] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (8002864 kB) free!
[1558426037] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6299548 kB) free!
[1558426637] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4645596 kB) free!
[1558427237] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11305072 kB) free.
[1558427883] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 276.91, 121.72, 46.70
[1558428003] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 296.68, 180.61, 77.37
[1558428123] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 299.25, 220.08, 104.36
[1558428437] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8458940 kB) free!
[1558428466] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2595212 kB) free!
[1558428557] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6676288 kB) free!
[1558428677] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4304060 kB) free!
[1558429066] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1558429066] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2169016 kB) free!
[1558429278] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10837596 kB) free.
[1558429374] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1558429458] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 296.00, 242.50, 126.02
[1558429495] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1558429578] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 292.68, 258.89, 146.14
[1558429698] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 73.16, 198.33, 138.57
[1558429878] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9572724 kB) free!
[1558429923] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.81, 1.48, 17.26
[1558429998] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7874376 kB) free!
[1558430077] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4528336 kB) free!
[1558430088] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (2899784 kB) free!
[1558430118] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6216952 kB) free!
[1558430164] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 44.3% (21848152 kB) free.
[1558430203] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.3% (21051956 kB) free.
[1558430208] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (2709680 kB) free!
[1558430220] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 75.2% (12224972 kB) free.
[1558430254] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.9% (21053632 kB) free.
[1558430266] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.9% (22271792 kB) free.
[1558430290] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (14797448 kB) free.
[1558430391] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 359 processes
[1558430511] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1558430632] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558430820] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (188324 kB) free!
[1558430866] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (4473688 kB) free!
[1558430940] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167068 kB) free!
[1558430986] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (4385732 kB) free!
[1558431060] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1558431060] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (170940 kB) free!
[1558431106] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (4283552 kB) free!
[1558431231] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 4 crit, 0 warn out of 360 processes
[1558431275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558431351] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558431723] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.58, 0.70, 4.07
[1558431876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558431996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558432007] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4911764 kB) free!
[1558432098] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.14, 3.32, 12.91
[1558432116] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558432127] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4810576 kB) free!
[1558432247] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4707976 kB) free!
[1558432306] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165876 kB) free!
[1558432551] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558432657] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4606264 kB) free!
[1558432778] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4549364 kB) free!
[1558432847] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181216 kB) free!
[1558432908] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1558433299] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.05, 1.96, 4.94
[1558434648] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2748296 kB) free!
[1558434653] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3024328 kB) free!
[1558434773] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2882128 kB) free!
[1558434893] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.9% (2754420 kB) free!
[1558435249] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197456 kB) free!
[1558435630] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1558436121] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.82, 8.84, 4.09
[1558436241] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 20.16, 12.71, 6.09
[1558436365] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 21.98, 15.72, 7.98
[1558436693] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1301340 kB) free!
[1558436916] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558437516] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558437565] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.49, 1.56, 4.06
[1558437636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558437756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558438081] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.8% (14928224 kB) free.
[1558438180] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.5% (22924088 kB) free.
[1558438244] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 86.2% (14018340 kB) free.
[1558438317] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.0% (20103988 kB) free.
[1558438387] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.9% (23407004 kB) free.
[1558438785] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.7% (3352024 kB) free!
[1558438905] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.5% (2564100 kB) free!
[1558439001] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (3086060 kB) free!
[1558439025] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2527600 kB) free!
[1558439123] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178220 kB) free!
[1558439243] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189428 kB) free!
[1558439526] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185060 kB) free!
[1558439626] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177944 kB) free!
[1558439647] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183040 kB) free!
[1558439767] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179980 kB) free!
[1558439899] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2820036 kB) free!
[1558440019] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (2719104 kB) free!
[1558440055] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 325 processes
[1558440139] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.1% (2616912 kB) free!
[1558440952] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 1 warn out of 359 processes
[1558440969] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558441030] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1558441072] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1558441087] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157984 kB) free!
[1558441207] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (187104 kB) free!
[1558441327] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169160 kB) free!
[1558441339] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1459224 kB) free!
[1558441567] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1558441567] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 64.9% (15909136 kB) free.
[1558441643] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.9% (11072172 kB) free.
[1558442026] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.1% (9822708 kB) free.
[1558442843] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.8% (3903424 kB) free!
[1558442963] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171808 kB) free!
[1558443055] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 326 processes
[1558443083] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173124 kB) free!
[1558443226] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187492 kB) free!
[1558443346] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178356 kB) free!
[1558443367] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189288 kB) free!
[1558443466] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170492 kB) free!
[1558443487] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172696 kB) free!
[1558443607] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174608 kB) free!
[1558443655] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 324 processes
[1558444956] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558445338] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.4% (15012336 kB) free.
[1558445407] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.2% (22602836 kB) free.
[1558445467] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (14794684 kB) free.
[1558445556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558445584] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.1% (23076716 kB) free.
[1558445638] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.6% (23100004 kB) free.
[1558445676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558445796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558446565] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 311.33, 177.39, 73.87
[1558446685] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 315.37, 223.33, 103.26
[1558446806] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 137.65, 205.48, 112.56
[1558448606] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.24, 1.04, 16.63
[1558449055] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1558449055] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 323 processes
[1558449723] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 106.88, 178.41, 98.25
[1558449843] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.74, 119.49, 86.38
[1558449963] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.25, 80.08, 75.95
[1558450406] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.70, 1.04, 3.45
[1558450596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558450901] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.75, 6.06, 5.29
[1558451021] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.16, 5.93, 5.33
[1558451141] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.44, 6.15, 5.48
[1558451196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558451316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558451436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558451763] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 0.40, 11.12
[1558451947] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1558452341] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.77, 3.72, 4.61
[1558452637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1558452676] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3183648 kB) free!
[1558452796] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (2983576 kB) free!
[1558452806] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.00, 35.94, 19.87
[1558452917] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.0% (2767584 kB) free!
[1558452926] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.29, 24.40, 17.59
[1558452964] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.41, 3.35
[1558453046] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1558453046] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.10, 16.62, 15.56
[1558453237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558453357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558453448] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1558453477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558453816] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1558454117] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1353336 kB) free!
[1558454246] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.79, 4.67
[1558454450] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3126532 kB) free!
[1558454570] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3013000 kB) free!
[1558454588] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1558454677] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558454690] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.9% (2912908 kB) free!
[1558455277] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558455397] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558455518] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558455655] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1558457090] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1546408 kB) free!
[1558457299] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 37.17, 23.88, 11.42
[1558457419] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 37.07, 28.26, 14.54
[1558457539] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 37.87, 31.44, 17.37
[1558458139] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.86, 7.18, 11.54
[1558459339] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.42, 2.15, 4.66
[1558459718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558460170] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9616192 kB) free!
[1558460290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.5% (15525576 kB) free.
[1558460318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558460439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558460559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558461759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558462359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558462479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558462599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes