[1577298654] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1577299105] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200460 kB) free!
[1577299177] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184216 kB) free!
[1577299225] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (215084 kB) free!
[1577299265] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193588 kB) free!
[1577299297] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164844 kB) free!
[1577299314] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1577299345] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (213828 kB) free!
[1577299386] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196240 kB) free!
[1577299386] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (231620 kB) free!
[1577299417] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (206408 kB) free!
[1577299434] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1577299506] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (208496 kB) free!
[1577299506] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185380 kB) free!
[1577299626] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196820 kB) free!
[1577299708] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1577299828] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1577302026] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.8% (9813072 kB) free.
[1577302627] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.3% (3031540 kB) free!
[1577302747] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (1434264 kB) free!
[1577302867] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202968 kB) free!
[1577303347] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6227404 kB) free!
[1577303467] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.9% (5394548 kB) free!
[1577303587] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4456916 kB) free!
[1577305201] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 310 processes
[1577306617] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.5% (9684600 kB) free.
[1577307237] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 62.0% (30569884 kB) free.
[1577307486] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8624500 kB) free!
[1577307606] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4807648 kB) free!
[1577307726] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.2% (44448340 kB) free.
[1577307745] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (2741216 kB) free!
[1577311690] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577312290] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2019/12/26 03:45:46 [error] 20758#0: *20073809 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577312410] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577313010] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2019/12/26 03:56:15 [error] 20754#0: *20076007 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577313130] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/26 04:00:36 [error] 20756#0: *20076938 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577313250] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577313850] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2019/12/26 04:11:14 [error] 20756#0: *20079184 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577313971] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/26 04:15:35 [error] 20756#0: *20080092 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577314091] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2019/12/26 04:16:37 [error] 20756#0: *20080334 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577314345] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2460460 kB) free!
[1577320418] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174440 kB) free!
[1577320538] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.4% (10152080 kB) free.
[1577320867] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 46.0% (11355344 kB) free.
[1577320945] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.3% (21392388 kB) free.
[1577321106] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.1% (20602952 kB) free.
[1577321138] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (242048 kB) free!
[1577321258] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (237216 kB) free!
[1577321378] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (242380 kB) free!
[1577321632] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.6% (38758024 kB) free.
[1577321666] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.0% (22783068 kB) free.
[1577321837] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 0.8% (188588 kB) free!
[1577323171] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1577323291] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1577323637] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 13.1% (3212860 kB) free!
[1577323891] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1577324011] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1577324237] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 27.6% (6776236 kB) free.
[1577328667] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4749376 kB) free!
[1577328787] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4738772 kB) free!
[1577328907] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4698832 kB) free!
[1577331907] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2416232 kB) free!
[1577335112] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22070940 kB) free.
[1577336719] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < operation none caused a connection exception frame_error: "Malformed UTF-8 in shortstr"
[1577336839] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1577337277] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/26 10:44:36 [error] 28871#0: *27146635 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: rq-1-1gbps.server, request: "GET /favicon.ico HTTP/1.1", host: "rq-1-1gbps.server:8091"
[1577337397] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577339682] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192240 kB) free!
[1577339802] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198176 kB) free!
[1577339922] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189600 kB) free!
[1577340517] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.2% (2515580 kB) free!
[1577340637] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (2407236 kB) free!
[1577340757] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2303292 kB) free!
[1577343726] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8154888 kB) free!
[1577343847] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (1940148 kB) free!
[1577343967] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217080 kB) free!
[1577344437] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3197244 kB) free!
[1577344506] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4445116 kB) free!
[1577344557] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (224360 kB) free!
[1577344626] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3482812 kB) free!
[1577344637] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.9% (4871540 kB) free!
[1577344677] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241584 kB) free!
[1577344746] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (3138036 kB) free!
[1577344757] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.8% (4842308 kB) free!
[1577344878] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.7% (4818040 kB) free!
[1577345346] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170016 kB) free!
[1577345806] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21866968 kB) free.
[1577345852] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.0% (22294648 kB) free.
[1577345877] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 29.7% (14655068 kB) free.
[1577345969] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 94.4% (46544356 kB) free.
[1577346006] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.2% (22360928 kB) free.
[1577347617] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (698668 kB) free!
[1577347737] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (561944 kB) free!
[1577347857] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (413424 kB) free!
[1577351035] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8952956 kB) free!
[1577351155] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8953896 kB) free!
[1577351275] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (8953316 kB) free!
[1577351480] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4405692 kB) free!
[1577351600] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4198880 kB) free!
[1577351720] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.8% (3989092 kB) free!
[1577352593] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8973124 kB) free!
[1577352714] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 53.5% (26351932 kB) free.
[1577352921] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 77.5% (17401364 kB) free.
[1577353828] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1577353945] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4738336 kB) free!
[1577353948] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1577354065] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4673620 kB) free!
[1577354186] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4591908 kB) free!
[1577354686] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.4% (18370904 kB) free.
[1577356243] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (200628 kB) free!
[1577357276] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4601780 kB) free!
[1577357786] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2257604 kB) free!
[1577357924] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(126) < 185.255.8.58 [26/Dec/2019:16:25:31 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1577358044] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577358149] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1577358269] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 337 processes
[1577358338] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 70.6% (34791968 kB) free.
[1577358379] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.6% (22451496 kB) free.
[1577358389] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1577358416] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.2% (14336948 kB) free.
[1577358678] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.9% (2436452 kB) free!
[1577358897] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181940 kB) free!
[1577358987] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3602324 kB) free!
[1577359017] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185224 kB) free!
[1577359055] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198292 kB) free!
[1577359107] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3473856 kB) free!
[1577359137] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197516 kB) free!
[1577359175] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186416 kB) free!
[1577359218] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (995740 kB) free!
[1577359227] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1577359227] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.9% (3396872 kB) free!
[1577359296] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195460 kB) free!
[1577359338] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (909864 kB) free!
[1577359342] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1577359458] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (815788 kB) free!
[1577359462] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1577359582] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1577360797] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(22) < 2019/12/26 17:16:25 [error] 28866#0: *27354166 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577360917] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577361044] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(207) < 203.212.70.170 [26/Dec/2019:17:12:22 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1577361165] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577361627] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2226956 kB) free!
[1577362175] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3125876 kB) free!
[1577362182] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.5% (17887808 kB) free.
[1577362219] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.9% (23007680 kB) free.
[1577362263] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 73.7% (18068660 kB) free.
[1577362291] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 50.3% (24808260 kB) free.
[1577362348] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 69.2% (16961324 kB) free.
[1577362404] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.1% (22819372 kB) free.
[1577363387] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174500 kB) free!
[1577363507] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173324 kB) free!
[1577363627] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170268 kB) free!
[1577364028] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192408 kB) free!
[1577364148] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189400 kB) free!
[1577364157] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175820 kB) free!
[1577364208] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (1055836 kB) free!
[1577364268] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183004 kB) free!
[1577364277] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194452 kB) free!
[1577364328] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (990880 kB) free!
[1577364397] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190124 kB) free!
[1577364448] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (910204 kB) free!
[1577365364] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(705) < 103.255.146.164 [26/Dec/2019:18:32:24 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1577365484] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577365626] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.1% (1965568 kB) free!
[1577365746] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.1% (1966696 kB) free!
[1577365866] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (1969076 kB) free!
[1577366789] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1577366909] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1577369037] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22628448 kB) free!
[1577369157] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46731348 kB) free.
[1577369757] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20295584 kB) free!
[1577369877] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46279848 kB) free.
[1577370477] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19861252 kB) free!
[1577370598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.6% (48263668 kB) free.
[1577371197] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21099532 kB) free!
[1577371317] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.5% (48148468 kB) free.
[1577371917] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21256668 kB) free!
[1577372037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.6% (48284732 kB) free.
[1577372637] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22758780 kB) free!
[1577372757] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.7% (49704436 kB) free.
[1577373357] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23247072 kB) free!
[1577373477] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.6% (49577696 kB) free.
[1577374048] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1577374077] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21753940 kB) free!
[1577374197] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46655528 kB) free.
[1577374698] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1577374773] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1836744 kB) free!
[1577374818] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1577374893] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1821040 kB) free!
[1577375014] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1806276 kB) free!
[1577375205] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.7% (21995836 kB) free.
[1577375227] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 70.9% (17497168 kB) free.
[1577375254] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.7% (21253140 kB) free.
[1577375293] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 74.7% (18310828 kB) free.
[1577375322] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.0% (22555716 kB) free.
[1577375397] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23535248 kB) free!
[1577375510] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (24721336 kB) free!
[1577375517] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.3% (49194312 kB) free.
[1577375630] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.7% (18039988 kB) free!
[1577375750] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (12010924 kB) free!
[1577376118] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23156024 kB) free!
[1577376238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.2% (49075788 kB) free.
[1577376618] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1577376738] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1577376838] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22166800 kB) free!
[1577376958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.5% (49428240 kB) free.
[1577377035] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (4615292 kB) free!
[1577377155] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4590612 kB) free!
[1577377275] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (4563336 kB) free!
[1577377558] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21221952 kB) free!
[1577377678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.3% (49187256 kB) free.
[1577378150] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.4% (18994796 kB) free!
[1577378278] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20788608 kB) free!
[1577378398] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.6% (46975336 kB) free.
[1577378747] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9757896 kB) free!
[1577378868] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9757512 kB) free!
[1577378988] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (9756756 kB) free!
[1577380275] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2473892 kB) free!
[1577381182] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in