[1584556312] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584556430] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584556442] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584556486] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584556551] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1584556572] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584556596] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584556639] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (14011552 kB) free!
[1584556669] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584556726] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584556789] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584556847] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584556922] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4800936 kB) free!
[1584556968] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171640 kB) free!
[1584557042] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.5% (4771872 kB) free!
[1584557088] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191448 kB) free!
[1584557094] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184236 kB) free!
[1584557163] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.3% (4740684 kB) free!
[1584557194] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584557208] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165332 kB) free!
[1584557214] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197768 kB) free!
[1584557239] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417384 kB) free!
[1584557318] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 9.07, 9.88, 9.47
[1584557334] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202500 kB) free!
[1584557839] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15982252 kB) free!
[1584558117] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (203260 kB) free!
[1584558439] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38120652 kB) free.
[1584558950] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1584559008] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4593508 kB) free!
[1584559039] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13878264 kB) free!
[1584559070] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1584559159] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (9992172 kB) free!
[1584559190] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1584559279] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6253692 kB) free!
[1584559318] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 78.6% (19260136 kB) free.
[1584559324] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584559454] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584559584] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584559608] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183672 kB) free!
[1584559790] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584559879] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22720220 kB) free!
[1584559910] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1584560031] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1584560172] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38770 MB (75% inode=98%):
[1584560188] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584560208] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.1% (3734364 kB) free!
[1584560357] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584560479] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3397620 kB) free!
[1584560535] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584560808] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168388 kB) free!
[1584561679] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23027548 kB) free!
[1584561831] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1584561831] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584561889] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219392 kB) free!
[1584562009] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240284 kB) free!
[1584562129] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238096 kB) free!
[1584562279] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385468 kB) free!
[1584562431] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1584562551] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1584562671] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1584562879] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20624900 kB) free!
[1584563057] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (21266692 kB) free!
[1584563177] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.2% (16019836 kB) free!
[1584563272] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1584563297] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (9792592 kB) free!
[1584563480] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (936908 kB) free!
[1584563753] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 7827.80, 4903.21, 2119.03
[1584563873] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3417.79, 4675.73, 2405.24
[1584563929] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 57.5% (28321728 kB) free.
[1584563981] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(122) < 49.50.86.152 [19/Mar/2020:02:03:05 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584563993] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 463.88, 3129.69, 2113.99
[1584564080] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17737812 kB) free!
[1584564102] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584564472] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1584564592] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1584564680] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417788 kB) free!
[1584564701] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 203.212.70.170 [19/Mar/2020:02:16:11 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584564821] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584564957] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2190048 kB) free!
[1584565077] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219916 kB) free!
[1584565192] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584565193] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.65, 58.94, 582.41
[1584565197] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224736 kB) free!
[1584565239] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9488376 kB) free!
[1584565280] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15480724 kB) free!
[1584565312] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584565359] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 10.9% (5353200 kB) free!
[1584565433] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584565479] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (1513544 kB) free!
[1584565793] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.71, 10.22, 306.46
[1584565797] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 78.3% (38599708 kB) free.
[1584565880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31803072 kB) free.
[1584566079] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.6% (43651836 kB) free.
[1584566480] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10922768 kB) free!
[1584566600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7384404 kB) free!
[1584566720] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3924348 kB) free!
[1584567320] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17411684 kB) free!
[1584567920] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400072 kB) free!
[1584568134] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.3% (10127640 kB) free.
[1584569022] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(140) < 92.42.104.166 [19/Mar/2020:03:25:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584569120] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26712676 kB) free.
[1584569142] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584569720] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7111744 kB) free!
[1584569840] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3700000 kB) free!
[1584569960] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29048964 kB) free.
[1584570561] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11585836 kB) free!
[1584570681] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7876164 kB) free!
[1584570801] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4253752 kB) free!
[1584572001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30812152 kB) free.
[1584572208] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584572601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5500604 kB) free!
[1584572721] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1767936 kB) free!
[1584572841] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398400 kB) free!
[1584573441] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13273852 kB) free!
[1584574041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30193500 kB) free.
[1584574641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7299620 kB) free!
[1584574735] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185616 kB) free!
[1584574761] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3407576 kB) free!
[1584574855] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173204 kB) free!
[1584574881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35613412 kB) free.
[1584574975] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177816 kB) free!
[1584575032] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584575481] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11726468 kB) free!
[1584575601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8248168 kB) free!
[1584575722] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4887100 kB) free!
[1584575808] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.2% (12393088 kB) free.
[1584575958] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 44.2% (10830832 kB) free.
[1584576321] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21871264 kB) free!
[1584576921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33172876 kB) free.
[1584577521] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14804252 kB) free!
[1584577641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11180648 kB) free!
[1584577761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36347276 kB) free.
[1584578361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15647876 kB) free!
[1584578481] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11833924 kB) free!
[1584578601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37409428 kB) free.
[1584579201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18330776 kB) free!
[1584579321] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14650468 kB) free!
[1584579441] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (10988048 kB) free!
[1584579963] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2421016 kB) free!
[1584580041] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23132080 kB) free!
[1584580097] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.9% (38077860 kB) free.
[1584580641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32555968 kB) free.
[1584580935] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.9% (39378696 kB) free.
[1584580992] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.0% (20834968 kB) free.
[1584581034] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.5% (21444304 kB) free.
[1584581241] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14011168 kB) free!
[1584581361] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37170800 kB) free.
[1584581808] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196428 kB) free!
[1584581929] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190660 kB) free!
[1584581961] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18009232 kB) free!
[1584582049] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189992 kB) free!
[1584582082] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (14513212 kB) free!
[1584582202] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38970864 kB) free.
[1584582365] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.2% (20280828 kB) free.
[1584582802] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20214276 kB) free!
[1584582922] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16537924 kB) free!
[1584583042] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (21971372 kB) free!
[1584584242] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29414860 kB) free.
[1584585442] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15566476 kB) free!
[1584585562] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11134180 kB) free!
[1584585682] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23049952 kB) free!
[1584586882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26786328 kB) free.
[1584587482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5558436 kB) free!
[1584587602] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33750292 kB) free.
[1584589338] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9436524 kB) free!
[1584589459] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9874348 kB) free!
[1584589569] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4910164 kB) free!
[1584589579] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9875140 kB) free!
[1584589689] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4868976 kB) free!
[1584589809] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4829408 kB) free!
[1584590096] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22124604 kB) free!
[1584590216] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18113800 kB) free!
[1584590336] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45057084 kB) free.
[1584590380] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584590936] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23171092 kB) free!
[1584590980] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 708 processes
[1584591009] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1584591009] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (2008264 kB) free!
[1584591056] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (19082176 kB) free!
[1584591100] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 706 processes
[1584591176] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15128588 kB) free!
[1584591220] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 707 processes
[1584591517] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 78.29, 178.82, 94.12
[1584591637] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.01, 121.26, 83.29
[1584591761] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 9.36, 83.81, 74.16
[1584592159] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4653420 kB) free!
[1584592279] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4469392 kB) free!
[1584592376] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35770128 kB) free.
[1584592399] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4313724 kB) free!
[1584592976] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13375624 kB) free!
[1584593096] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.9% (26294504 kB) free!
[1584593216] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39588416 kB) free.
[1584593816] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16104268 kB) free!
[1584593876] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.6% (2044492 kB) free!
[1584593936] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11782908 kB) free!
[1584593996] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.6% (2042428 kB) free!
[1584594056] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7651936 kB) free!
[1584594116] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (2036652 kB) free!
[1584594117] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.4% (19834852 kB) free.
[1584594156] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 48.4% (63834780 kB) free.
[1584594342] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 182.18.144.246 [19/Mar/2020:10:34:54 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584594379] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3923368 kB) free!
[1584594412] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 104.44, 832.30, 519.16
[1584594462] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584594532] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16.03, 557.66, 456.48
[1584594652] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.46, 373.99, 401.44
[1584594654] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.34, 504.06, 530.35
[1584594774] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.63, 338.13, 466.37
[1584594799] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1815468 kB) free!
[1584595252] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.37, 52.23, 211.48
[1584595662] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(358) < 103.255.146.164 [19/Mar/2020:10:56:11 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584595783] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584595793] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1595.07, 1779.94, 800.77
[1584595824] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2620.12, 2452.45, 1099.02
[1584595913] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 217.59, 1191.87, 703.99
[1584595944] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 356.09, 1642.07, 966.16
[1584595960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10679172 kB) free!
[1584595968] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 992 processes
[1584596022] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.9% (39356792 kB) free.
[1584596034] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 31.25, 798.22, 618.89
[1584596064] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 50.53, 1099.70, 849.36
[1584596067] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.7% (21005764 kB) free.
[1584596080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43291868 kB) free.
[1584596088] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 356 processes
[1584596297] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25793704 kB) free!
[1584596382] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 49.50.86.152 [19/Mar/2020:11:02:02 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584596417] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (26705548 kB) free.
[1584596502] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584596634] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.19, 108.87, 325.27
[1584596639] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.1% (3939052 kB) free!
[1584596664] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.92, 149.89, 446.34
[1584596759] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.0% (3668992 kB) free!
[1584596879] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.1% (3450840 kB) free!
[1584597102] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 103.255.146.164 [19/Mar/2020:11:20:08 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584597117] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181464 kB) free!
[1584597199] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (4120980 kB) free!
[1584597220] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1584597222] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584597237] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168296 kB) free!
[1584597319] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.7% (3855880 kB) free!
[1584597357] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180012 kB) free!
[1584597381] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.7% (21158520 kB) free.
[1584597439] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (3676676 kB) free!
[1584597479] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2137984 kB) free!
[1584597617] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.0% (23768736 kB) free!
[1584597621] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.0% (21561072 kB) free.
[1584597737] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.8% (26161388 kB) free!
[1584597857] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.3% (25454984 kB) free!
[1584597903] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (398556 kB) free!
[1584597970] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4330584 kB) free!
[1584598018] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4813188 kB) free!
[1584598023] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401396 kB) free!
[1584598039] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2077884 kB) free!
[1584598078] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4806092 kB) free!
[1584598090] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (3930904 kB) free!
[1584598138] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.6% (4801104 kB) free!
[1584598143] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (422612 kB) free!
[1584598172] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.9% (20808936 kB) free.
[1584598198] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.6% (4802908 kB) free!
[1584598210] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.6% (3510384 kB) free!
[1584598227] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 48.2% (63550520 kB) free.
[1584598316] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (222856 kB) free!
[1584598810] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1958748 kB) free!
[1584598961] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.22, 14.65, 30.95
[1584599410] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.3% (15116040 kB) free.
[1584599566] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 11.99, 13.75, 22.85
[1584599622] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(110) < 103.255.146.164 [19/Mar/2020:12:01:04 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584599743] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584599789] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3956808 kB) free!
[1584599909] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3509484 kB) free!
[1584600028] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6744176 kB) free!
[1584600029] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.6% (3106484 kB) free!
[1584600148] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1932700 kB) free!
[1584600268] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386756 kB) free!
[1584600386] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (824116 kB) free!
[1584600506] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (1660208 kB) free!
[1584600626] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1478748 kB) free!
[1584600629] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.0% (22714776 kB) free.
[1584600742] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.6% (14080664 kB) free.
[1584600794] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.5% (50763524 kB) free.
[1584600834] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.9% (20692776 kB) free.
[1584601427] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5760952 kB) free!
[1584601547] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.7% (5758220 kB) free!
[1584601667] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (5753564 kB) free!
[1584601998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22318480 kB) free!
[1584602118] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17516384 kB) free!
[1584602238] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12937484 kB) free!
[1584602667] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177812 kB) free!
[1584602787] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181272 kB) free!
[1584602838] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23982864 kB) free!
[1584602907] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173644 kB) free!
[1584603343] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(121) < 103.255.146.164 [19/Mar/2020:12:58:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584603387] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1584603387] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.7% (20035032 kB) free.
[1584603438] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407276 kB) free!
[1584603463] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584603578] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.3% (3496560 kB) free!
[1584603698] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.4% (3277088 kB) free!
[1584603818] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584603818] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (3058836 kB) free!
[1584603840] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (1224208 kB) free!
[1584603961] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (1030544 kB) free!
[1584603987] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176740 kB) free!
[1584604038] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21554692 kB) free!
[1584604067] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1584604067] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3432632 kB) free!
[1584604081] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1725692 kB) free!
[1584604107] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189800 kB) free!
[1584604227] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168044 kB) free!
[1584604418] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1895984 kB) free!
[1584604480] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9556364 kB) free!
[1584604600] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9291912 kB) free!
[1584604638] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394476 kB) free!
[1584604720] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9020012 kB) free!
[1584605238] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17790364 kB) free!
[1584605263] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(146) < 103.255.146.164 [19/Mar/2020:13:31:07 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584605384] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584605429] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.3% (4238812 kB) free!
[1584605549] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.2% (3963256 kB) free!
[1584605669] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.8% (3637736 kB) free!
[1584605838] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46330 MB (20% inode=99%):
[1584605838] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403964 kB) free!
[1584605852] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.7% (20407912 kB) free.
[1584605900] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39815536 kB) free.
[1584605943] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39344380 kB) free.
[1584605958] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46330 MB (20% inode=99%):
[1584605981] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.6% (20969804 kB) free.
[1584606021] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.7% (20993868 kB) free.
[1584606078] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46330 MB (20% inode=99%):
[1584606597] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9478156 kB) free!
[1584606717] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9415184 kB) free!
[1584606832] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1584606837] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (9218792 kB) free!
[1584606869] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1584606869] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1805548 kB) free!
[1584606952] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1584607110] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15504160 kB) free!
[1584607120] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3781172 kB) free!
[1584607230] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36867276 kB) free.
[1584607278] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48409 MB (21% inode=99%):
[1584607658] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183396 kB) free!
[1584607778] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176720 kB) free!
[1584607830] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8809648 kB) free!
[1584607878] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46167 MB (20% inode=99%):
[1584607898] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189680 kB) free!
[1584607950] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3961708 kB) free!
[1584607998] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46171 MB (20% inode=99%):
[1584608070] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397388 kB) free!
[1584608118] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46171 MB (20% inode=99%):
[1584608129] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584608152] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1584608249] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584608272] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1584608322] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1584608369] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1584608384] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(669) < 103.255.146.164 [19/Mar/2020:14:28:00 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584608393] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1584608410] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3104712 kB) free!
[1584608422] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200088 kB) free!
[1584608442] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1584608504] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(118) < 103.255.146.164 [19/Mar/2020:14:30:58 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584608530] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.4% (2778236 kB) free!
[1584608542] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200048 kB) free!
[1584608562] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1584608624] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584608650] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.1% (2487440 kB) free!
[1584608662] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200884 kB) free!
[1584608670] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21477916 kB) free!
[1584608948] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7928056 kB) free!
[1584608969] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1584609068] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (7789624 kB) free!
[1584609188] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.5% (7656680 kB) free!
[1584609250] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1584609250] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (1016396 kB) free!
[1584609270] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389912 kB) free!
[1584609837] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4819308 kB) free!
[1584609869] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 73.7% (18058248 kB) free.
[1584609870] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20421032 kB) free!
[1584609918] SERVICE FLAPPING ALERT: primary-server;HDD_Home status;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1584609918] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 103800 MB (46% inode=99%):
[1584610229] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3899556 kB) free!
[1584610349] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3611392 kB) free!
[1584610451] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.9% (15239236 kB) free.
[1584610469] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3332136 kB) free!
[1584610470] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395952 kB) free!
[1584610813] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1584610933] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1584610988] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4115696 kB) free!
[1584611069] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584611069] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1501620 kB) free!
[1584611070] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19764432 kB) free!
[1584611153] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3163180 kB) free!
[1584611273] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3545324 kB) free.
[1584611369] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584611381] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (3847888 kB) free!
[1584611489] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584611501] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.9% (3650028 kB) free!
[1584611588] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (4958340 kB) free!
[1584611609] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584611621] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.2% (3478716 kB) free!
[1584611624] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(128) < 103.255.146.164 [19/Mar/2020:15:15:54 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584611657] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (8086472 kB) free!
[1584611669] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.7% (3597588 kB) free!
[1584611669] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.0% (22706376 kB) free.
[1584611670] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (422928 kB) free!
[1584611744] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584611789] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2359180 kB) free!
[1584611909] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2123848 kB) free!
[1584612188] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (573660 kB) free!
[1584612209] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1584612209] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1584612257] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.7% (15415788 kB) free!
[1584612270] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17982956 kB) free!
[1584612821] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1373580 kB) free!
[1584612857] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (11975872 kB) free!
[1584612870] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380860 kB) free!
[1584613457] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1584613457] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.6% (20538860 kB) free!
[1584613470] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18115540 kB) free!
[1584614071] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402860 kB) free!
[1584614657] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (9887616 kB) free!
[1584614671] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14793192 kB) free!
[1584615098] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 30.6% (7546484 kB) free.
[1584615188] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 27.6% (13580112 kB) free.
[1584615221] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1584615257] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.6% (16642928 kB) free!
[1584615271] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411332 kB) free!
[1584615341] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1584615461] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1584615473] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.8% (1925140 kB) free!
[1584615593] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2617908 kB) free!
[1584615698] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.4% (2554292 kB) free!
[1584615713] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3379168 kB) free.
[1584615733] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 368 processes
[1584615788] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8383560 kB) free!
[1584615809] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584615818] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (1278144 kB) free!
[1584615853] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1584615908] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8379236 kB) free!
[1584615923] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 0.8% (187076 kB) free!
[1584615929] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584615938] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (1136840 kB) free!
[1584615973] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1584616028] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (8379296 kB) free!
[1584616043] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 1.9% (453724 kB) free!
[1584616049] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584616163] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 1.8% (451520 kB) free!
[1584616313] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2158148 kB) free!
[1584616433] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3006372 kB) free!
[1584616457] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (11507616 kB) free!
[1584616553] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.6% (3997488 kB) free.
[1584616649] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1584617153] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2973936 kB) free!
[1584617273] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3822768 kB) free.
[1584617658] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.0% (18494984 kB) free!
[1584617828] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (252660 kB) free!
[1584617861] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584617873] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (2618576 kB) free!
[1584617993] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (1242744 kB) free!
[1584618114] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.9% (4527008 kB) free.
[1584618262] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584618271] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15692260 kB) free!
[1584618393] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/03/19 17:14:18 [error] 12256#0: *73678186 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: _, request: "GET /favicon.ico HTTP/1.1", host: "rq-1-1gbps.server", referrer: "http://rq-1-1gbps.server/"
[1584618513] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584618858] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12470668 kB) free!
[1584618871] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393320 kB) free!
[1584619362] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1584619451] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4133692 kB) free!
[1584619471] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21451724 kB) free!
[1584619572] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.3% (3882536 kB) free!
[1584619692] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.8% (3558076 kB) free!
[1584620071] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418748 kB) free!
[1584620118] SERVICE FLAPPING ALERT: primary-server;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584620249] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584620369] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584620489] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584620671] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (25084764 kB) free!
[1584620861] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 707 processes
[1584620892] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (640048 kB) free!
[1584620981] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 707 processes
[1584621101] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 707 processes
[1584621271] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (426792 kB) free!
[1584621492] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.5% (15159020 kB) free.
[1584621509] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 58.9% (14446384 kB) free.
[1584621689] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1584621869] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (3874036 kB) free!
[1584621871] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30747956 kB) free.
[1584621973] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1584621989] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.0% (22712340 kB) free.
[1584622093] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 368 processes
[1584622163] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 42.3% (10359548 kB) free.
[1584622213] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1584622471] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3048568 kB) free!
[1584622592] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (423204 kB) free!
[1584622712] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416616 kB) free!
[1584623189] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584623312] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25757616 kB) free!
[1584623622] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584623658] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.3% (13555344 kB) free!
[1584623872] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584623909] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.1% (4190600 kB) free!
[1584623912] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (435076 kB) free!
[1584623993] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2533 processes
[1584624029] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.4% (4017420 kB) free!
[1584624113] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2875940 kB) free!
[1584624149] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.6% (3831600 kB) free!
[1584624176] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584624233] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3753284 kB) free.
[1584624258] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (10928980 kB) free!
[1584624301] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 15.98, 16.51, 15.77
[1584624512] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27063548 kB) free.
[1584624833] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (2621708 kB) free!
[1584624858] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (28927580 kB) free.
[1584624944] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(133) < 103.255.146.164 [19/Mar/2020:19:01:48 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584624953] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3553232 kB) free.
[1584625065] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584625112] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (734900 kB) free!
[1584625232] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (380492 kB) free!
[1584625349] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2158520 kB) free!
[1584625352] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433380 kB) free!
[1584625553] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.2% (1985444 kB) free!
[1584625673] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2770392 kB) free!
[1584625794] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3502968 kB) free.
[1584625952] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19054012 kB) free!
[1584626058] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.8% (23457556 kB) free!
[1584626178] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.6% (15310624 kB) free!
[1584626298] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.8% (26162792 kB) free!
[1584626501] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1584626552] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380080 kB) free!
[1584627152] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23176020 kB) free!
[1584627338] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584627752] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420004 kB) free!
[1584628028] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584628301] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1584628352] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27476936 kB) free.
[1584628421] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584628698] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.9% (30204168 kB) free.
[1584628952] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (396576 kB) free!
[1584629021] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1584629073] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (421132 kB) free!
[1584629141] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584629193] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (453344 kB) free!
[1584629298] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.7% (24593800 kB) free!
[1584629393] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1856368 kB) free!
[1584629419] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.7% (16780928 kB) free!
[1584629513] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2419060 kB) free!
[1584629539] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.7% (26025724 kB) free!
[1584629633] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3002832 kB) free!
[1584629741] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1584629793] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20360084 kB) free!
[1584629861] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1584629864] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21719948 kB) free.
[1584629894] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.7% (19908664 kB) free.
[1584629922] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 30.5% (15026944 kB) free.
[1584629982] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1584629982] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1584629993] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39534080 kB) free.
[1584630012] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.1% (20371612 kB) free.
[1584630049] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 31.8% (15662212 kB) free.
[1584630074] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.5% (20702764 kB) free.
[1584630233] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3401948 kB) free.
[1584630393] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376792 kB) free!
[1584630499] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190556 kB) free!
[1584630598] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9179596 kB) free!
[1584630619] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179936 kB) free!
[1584630718] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8897196 kB) free!
[1584630739] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177240 kB) free!
[1584630838] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (8780200 kB) free!
[1584630993] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25900476 kB) free!
[1584631433] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3188976 kB) free!
[1584631553] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.7% (4011260 kB) free.
[1584631593] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (462288 kB) free!
[1584631665] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(134) < 92.42.104.166 [19/Mar/2020:20:55:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584631785] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584632193] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32008060 kB) free.
[1584632382] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584632489] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584632539] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (30790580 kB) free.
[1584632638] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4143808 kB) free!
[1584632753] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3189932 kB) free!
[1584632793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3196084 kB) free!
[1584632798] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584632873] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3525732 kB) free.
[1584632893] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584632913] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (375328 kB) free!
[1584632920] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584632982] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584633033] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432340 kB) free!
[1584633102] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1584633139] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.6% (24493016 kB) free!
[1584633223] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1584633259] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.8% (30011292 kB) free.
[1584633633] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25557724 kB) free!
[1584633859] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25867064 kB) free!
[1584633979] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.4% (18942748 kB) free!
[1584634073] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (2694392 kB) free!
[1584634099] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.5% (25695500 kB) free!
[1584634193] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3064680 kB) free!
[1584634233] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401480 kB) free!
[1584634313] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3525144 kB) free.
[1584634833] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31705020 kB) free.
[1584635433] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3553712 kB) free!
[1584635553] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (392260 kB) free!
[1584635673] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418244 kB) free!
[1584635899] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (27274284 kB) free.
[1584636223] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584636274] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24249980 kB) free!
[1584636470] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1584636823] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584636874] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (471352 kB) free!
[1584636943] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1584637064] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1584637099] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.8% (26137788 kB) free!
[1584637219] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.8% (32692536 kB) free.
[1584637474] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30066296 kB) free.
[1584637819] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.7% (23392604 kB) free!
[1584637939] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.3% (29385136 kB) free.
[1584638074] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2763980 kB) free!
[1584638194] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (498568 kB) free!
[1584638264] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584638314] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425888 kB) free!
[1584638864] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584638914] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20687596 kB) free!
[1584638984] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1584639104] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1584639139] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.1% (22604148 kB) free!
[1584639260] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27097752 kB) free.
[1584639514] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398172 kB) free!
[1584640459] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.3% (24143796 kB) free!
[1584640579] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (30908832 kB) free.
[1584641314] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14981992 kB) free!
[1584641914] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423364 kB) free!
[1584641922] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4478504 kB) free!
[1584642018] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4441832 kB) free!
[1584642042] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237588 kB) free!
[1584642104] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584642138] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (4323304 kB) free!
[1584642163] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 76.5% (37683068 kB) free.
[1584642258] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4226860 kB) free!
[1584642514] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30021340 kB) free.