[1550169108] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550169228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550169349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550172349] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550172949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550173069] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550181469] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1550185669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550185789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550185909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550188309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550198509] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550198629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550198749] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550199949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550200069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550200189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550201389] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550201989] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550202109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550202229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550203429] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1550203429] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550204029] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550204149] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550204750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550204870] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550205470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550205590] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550205960] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24710 MB (10% inode=99%):
[1550206190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550206310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550206430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550207631] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550208097] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1534832 kB) free!
[1550208231] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550208351] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550208360] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29517 MB (13% inode=99%):
[1550208951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550209071] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550209494] SERVICE ALERT: sig2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1589024 kB) free!
[1550209671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550209791] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550209911] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550210757] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9415420 kB) free!
[1550210877] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13224700 kB) free.
[1550211053] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40605 MB (19% inode=99%):
[1550211173] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47265 MB (22% inode=99%):
[1550212707] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (164732 kB) free!
[1550212827] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (154416 kB) free!
[1550212911] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550212947] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162472 kB) free!
[1550213108] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4338468 kB) free!
[1550213229] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (2817672 kB) free!
[1550213349] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243224 kB) free!
[1550213511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550213573] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43078 MB (20% inode=99%):
[1550213631] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550213693] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46136 MB (22% inode=99%):
[1550213878] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8931428 kB) free!
[1550213998] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12122512 kB) free.
[1550214231] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550214351] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550214472] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550214549] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.4% (5599408 kB) free!
[1550214598] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8195672 kB) free!
[1550214718] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11563524 kB) free.
[1550215318] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7765580 kB) free!
[1550215431] SERVICE ALERT: sig4;Memory;WARNING;SOFT;1;WARNING - 20.0% (3251764 kB) free!
[1550215438] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11165016 kB) free.
[1550215551] SERVICE ALERT: sig4;Memory;WARNING;SOFT;2;WARNING - 20.0% (3243404 kB) free!
[1550215671] SERVICE ALERT: sig4;Memory;WARNING;HARD;3;WARNING - 19.9% (3237216 kB) free!
[1550215997] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550216038] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7006504 kB) free!
[1550216089] SERVICE ALERT: sig3;Memory;WARNING;SOFT;1;WARNING - 19.9% (3228040 kB) free!
[1550216117] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550216158] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10767700 kB) free.
[1550216210] SERVICE ALERT: sig3;Memory;WARNING;SOFT;2;WARNING - 19.8% (3223724 kB) free!
[1550216330] SERVICE ALERT: sig3;Memory;WARNING;HARD;3;WARNING - 19.8% (3217736 kB) free!
[1550216693] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41661 MB (19% inode=99%):
[1550216813] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46255 MB (22% inode=99%):
[1550216949] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2525548 kB) free!
[1550217358] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9556712 kB) free!
[1550217413] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43087 MB (20% inode=99%):
[1550217472] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550217478] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12405988 kB) free.
[1550217533] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41037 MB (19% inode=99%):
[1550217653] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 45118 MB (21% inode=99%):
[1550218072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550218078] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9306252 kB) free!
[1550218192] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550218199] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12357548 kB) free.
[1550218253] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 37388 MB (17% inode=99%):
[1550218275] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1550218373] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48804 MB (23% inode=99%):
[1550218395] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1550218515] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1550218792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550218798] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8933908 kB) free!
[1550218913] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550218918] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12317044 kB) free.
[1550219349] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1550219349] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.5% (5174204 kB) free!
[1550219512] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550219518] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8702208 kB) free!
[1550219632] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550219638] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11874624 kB) free.
[1550219752] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550220238] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8597476 kB) free!
[1550220358] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11930468 kB) free.
[1550220916] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1550220958] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7768872 kB) free!
[1550221078] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11465452 kB) free.
[1550221309] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47193 MB (20% inode=99%):
[1550221429] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47193 MB (20% inode=99%):
[1550221549] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47192 MB (20% inode=99%):
[1550221678] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7703184 kB) free!
[1550221749] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3421132 kB) free!
[1550221798] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11366932 kB) free.
[1550222398] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7229284 kB) free!
[1550222519] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11134848 kB) free.
[1550223118] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7245436 kB) free!
[1550223238] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11072552 kB) free.
[1550223349] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47228 MB (21% inode=99%):
[1550224149] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.4% (5600616 kB) free!
[1550224438] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9129160 kB) free!
[1550224552] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550224558] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11956340 kB) free.
[1550224947] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 53.8% (8745500 kB) free.
[1550224973] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42165 MB (20% inode=99%):
[1550225093] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48835 MB (23% inode=99%):
[1550225114] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24667 MB (10% inode=99%):
[1550225152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550225158] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8609496 kB) free!
[1550225272] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550225278] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11913644 kB) free.
[1550225547] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (266096 kB) free!
[1550225667] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (229260 kB) free!
[1550225787] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (191180 kB) free!
[1550225872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550225878] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8330860 kB) free!
[1550225949] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4515280 kB) free!
[1550225992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550225998] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11527900 kB) free.
[1550226112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550226596] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 321 processes
[1550226598] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7726456 kB) free!
[1550226716] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1550226718] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11161800 kB) free.
[1550227319] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (256020 kB) free!
[1550227439] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13144204 kB) free.
[1550228039] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9296032 kB) free!
[1550228093] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43141 MB (20% inode=99%):
[1550228159] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13077496 kB) free.
[1550228213] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51320 MB (24% inode=99%):
[1550228759] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9226096 kB) free!
[1550228879] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13009848 kB) free.
[1550229479] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (8994564 kB) free!
[1550229599] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12838148 kB) free.
[1550230149] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.9% (5356500 kB) free!
[1550230199] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8760924 kB) free!
[1550230320] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12745876 kB) free.
[1550230919] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8787288 kB) free!
[1550231039] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12697724 kB) free.
[1550231349] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4831172 kB) free!
[1550231639] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8739948 kB) free!
[1550231759] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12588632 kB) free.
[1550231814] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43294 MB (20% inode=99%):
[1550231934] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49429 MB (23% inode=99%):
[1550232359] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8506400 kB) free!
[1550232479] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12452248 kB) free.
[1550232649] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 61.3% (30202200 kB) free.
[1550232741] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.5% (23408088 kB) free.
[1550233008] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46787 MB (20% inode=99%):
[1550233036] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.6% (15219820 kB) free.
[1550233079] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8851140 kB) free!
[1550233128] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51388 MB (22% inode=99%):
[1550233199] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 92.6% (15052656 kB) free.
[1550233199] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12546144 kB) free.
[1550233287] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 93.9% (15265180 kB) free.
[1550233410] SERVICE ALERT: sig3;Memory;OK;HARD;1;OK - 93.8% (15238128 kB) free.
[1550233728] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46851 MB (20% inode=99%):
[1550233799] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8922616 kB) free!
[1550233848] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50434 MB (22% inode=99%):
[1550233919] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12540292 kB) free.
[1550234519] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8935096 kB) free!
[1550234640] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12490868 kB) free.
[1550234804] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 95.2% (15480108 kB) free.
[1550235239] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9054676 kB) free!
[1550235359] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12499092 kB) free.
[1550235914] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 28133 MB (12% inode=99%):
[1550235959] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9380228 kB) free!
[1550236079] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12526232 kB) free.
[1550236248] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46739 MB (20% inode=99%):
[1550236368] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50835 MB (22% inode=99%):
[1550236679] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9430364 kB) free!
[1550236799] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12824172 kB) free.
[1550236913] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1550236968] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46714 MB (20% inode=99%):
[1550237088] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47055 MB (20% inode=99%):
[1550237208] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50253 MB (22% inode=99%):
[1550237399] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9516812 kB) free!
[1550237512] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1550237519] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12798620 kB) free.
[1550237717] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550237808] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46533 MB (20% inode=99%):
[1550237837] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550237928] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50236 MB (22% inode=99%):
[1550237957] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1550238112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1550238119] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9322612 kB) free!
[1550238232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1550238239] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12836144 kB) free.
[1550238444] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173808 kB) free!
[1550238529] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47165 MB (20% inode=99%):
[1550238564] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178556 kB) free!
[1550238649] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50235 MB (22% inode=99%):
[1550238684] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151220 kB) free!
[1550238745] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198180 kB) free!
[1550238832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1550238839] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9413288 kB) free!
[1550238865] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189292 kB) free!
[1550238952] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1550238960] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12815864 kB) free.
[1550238985] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191248 kB) free!
[1550239157] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550239214] SERVICE ALERT: sig3;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1550239249] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47035 MB (20% inode=99%):
[1550239334] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1550239369] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50233 MB (22% inode=99%):
[1550239552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1550239559] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9137044 kB) free!
[1550239645] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.1% (15127268 kB) free.
[1550239672] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1550239679] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.6% (12631268 kB) free.
[1550239708] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.3% (23357264 kB) free.
[1550239969] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46895 MB (20% inode=99%):
[1550240089] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50089 MB (22% inode=99%):
[1550240273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1550240279] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9248980 kB) free!
[1550240393] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1550240399] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12735496 kB) free.
[1550240689] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47019 MB (20% inode=99%):
[1550240809] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50085 MB (22% inode=99%):
[1550240993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1550240999] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9306264 kB) free!
[1550241015] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550241113] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1550241119] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12710540 kB) free.
[1550241135] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550241255] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550241409] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47016 MB (20% inode=99%):
[1550241529] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50085 MB (22% inode=99%):
[1550241573] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 52 users currently logged in
[1550241633] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 52 users currently logged in
[1550241693] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1550241713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1550241719] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9357728 kB) free!
[1550241753] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 52 users currently logged in
[1550241833] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1550241839] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12687064 kB) free.
[1550241879] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550241999] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550242119] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550242129] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47006 MB (20% inode=99%):
[1550242249] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50076 MB (22% inode=99%):
[1550242257] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1550242433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1550242439] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9362268 kB) free!
[1550242553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1550242559] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12660676 kB) free.
[1550242849] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47005 MB (20% inode=99%):
[1550242969] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50587 MB (22% inode=99%):
[1550243114] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 108461 MB (48% inode=99%):
[1550243154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1550243160] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9429832 kB) free!
[1550243274] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1550243280] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.6% (12629728 kB) free.
[1550243334] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40431 MB (19% inode=99%):
[1550243454] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47079 MB (22% inode=99%):
[1550243569] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47002 MB (20% inode=99%):
[1550243689] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51031 MB (22% inode=99%):
[1550243874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1550243880] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9748220 kB) free!
[1550243994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1550244000] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12589228 kB) free.
[1550244004] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2935160 kB) free!
[1550244124] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2876276 kB) free!
[1550244153] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1550244244] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.4% (2829612 kB) free!
[1550244289] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46997 MB (20% inode=99%):
[1550244409] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51094 MB (22% inode=99%):
[1550244594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550244654] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42680 MB (20% inode=99%):
[1550244714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550244775] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40636 MB (19% inode=99%):
[1550244834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550244895] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44601 MB (21% inode=99%):
[1550245009] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46988 MB (20% inode=99%):
[1550245129] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51084 MB (22% inode=99%):
[1550245729] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46739 MB (20% inode=99%):
[1550245849] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50835 MB (22% inode=99%):
[1550246207] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4935464 kB) free!
[1550246327] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4921300 kB) free!
[1550246447] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4899276 kB) free!
[1550246450] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46738 MB (20% inode=99%):
[1550246570] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50834 MB (22% inode=99%):
[1550247170] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46737 MB (20% inode=99%):
[1550247290] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50833 MB (22% inode=99%):
[1550247468] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.3% (22785336 kB) free.
[1550247525] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 92.3% (14995680 kB) free.
[1550247600] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8612216 kB) free!
[1550247720] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12077908 kB) free.
[1550247890] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46939 MB (20% inode=99%):
[1550247895] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43699 MB (20% inode=99%):
[1550248010] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51035 MB (22% inode=99%):
[1550248015] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40630 MB (19% inode=99%):
[1550248135] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44729 MB (21% inode=99%):
[1550248320] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8899632 kB) free!
[1550248434] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550248441] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12429416 kB) free.
[1550248610] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46918 MB (20% inode=99%):
[1550248730] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46797 MB (20% inode=99%):
[1550248850] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49998 MB (22% inode=99%):
[1550249034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550249040] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8957744 kB) free!
[1550249154] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1550249160] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12488964 kB) free.
[1550249450] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46925 MB (20% inode=99%):
[1550249570] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49995 MB (22% inode=99%):
[1550249754] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1550249760] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9177400 kB) free!
[1550249875] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1550249880] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12486808 kB) free.
[1550250170] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46925 MB (20% inode=99%):
[1550250290] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50509 MB (22% inode=99%):
[1550250474] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550250480] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9275896 kB) free!
[1550250594] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550250600] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12475976 kB) free.
[1550250890] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46924 MB (20% inode=99%):
[1550251010] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50506 MB (22% inode=99%):
[1550251194] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550251200] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9545676 kB) free!
[1550251314] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1550251320] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12460120 kB) free.
[1550251610] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46859 MB (20% inode=99%):
[1550251730] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50696 MB (22% inode=99%):
[1550251735] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43727 MB (20% inode=99%):
[1550251855] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40655 MB (19% inode=99%):
[1550251914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550251920] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9679252 kB) free!
[1550251975] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44743 MB (21% inode=99%):
[1550252034] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1550252040] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12435012 kB) free.
[1550252330] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46790 MB (20% inode=99%):
[1550252450] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50945 MB (22% inode=99%):
[1550252634] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550252640] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9874936 kB) free!
[1550252754] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1550252761] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12415208 kB) free.
[1550253050] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46848 MB (20% inode=99%):
[1550253170] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50944 MB (22% inode=99%):
[1550253355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1550253475] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550253595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550253771] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46770 MB (20% inode=99%):
[1550253891] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50866 MB (22% inode=99%):
[1550254491] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46770 MB (20% inode=99%):
[1550254560] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8053140 kB) free!
[1550254611] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50866 MB (22% inode=99%):
[1550254680] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11729512 kB) free.
[1550255211] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46760 MB (20% inode=99%):
[1550255280] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8298024 kB) free!
[1550255331] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50858 MB (22% inode=99%):