[1585333851] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.18, 4.05, 4.61
[1585334002] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 31.9% (7872752 kB) free.
[1585334052] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 24.9% (12256200 kB) free.
[1585334172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6640676 kB) free!
[1585334292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3128764 kB) free!
[1585334388] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187108 kB) free!
[1585334412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28024432 kB) free.
[1585334509] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191288 kB) free!
[1585334602] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191820 kB) free!
[1585334629] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185524 kB) free!
[1585334652] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (1222424 kB) free!
[1585334722] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186624 kB) free!
[1585334772] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241084 kB) free!
[1585334842] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192092 kB) free!
[1585334892] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224628 kB) free!
[1585335012] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6157256 kB) free!
[1585335051] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.15, 6.14, 5.18
[1585335132] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29044956 kB) free.
[1585335171] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.35, 6.12, 5.28
[1585335291] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.05, 6.08, 5.37
[1585335732] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8142024 kB) free!
[1585335844] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166932 kB) free!
[1585335852] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4629048 kB) free!
[1585335904] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179088 kB) free!
[1585335965] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (201768 kB) free!
[1585335972] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27983236 kB) free.
[1585336025] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (200968 kB) free!
[1585336325] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.5% (17038744 kB) free.
[1585336572] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7639428 kB) free!
[1585336692] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5942648 kB) free!
[1585336812] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26441296 kB) free.
[1585337413] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6659908 kB) free!
[1585337533] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28774564 kB) free.
[1585338133] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8870424 kB) free!
[1585338253] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5380844 kB) free!
[1585338373] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27896128 kB) free.
[1585338973] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7048312 kB) free!
[1585339093] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28624920 kB) free.
[1585339693] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8369256 kB) free!
[1585339813] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4658304 kB) free!
[1585339933] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28198612 kB) free.
[1585340533] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6723200 kB) free!
[1585340629] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1585340653] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28336304 kB) free.
[1585340814] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(145) < 92.42.104.166 [28/Mar/2020:01:55:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585340934] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585341253] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7589520 kB) free!
[1585341373] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3662180 kB) free!
[1585341493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28162668 kB) free.
[1585342094] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4561928 kB) free!
[1585342214] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28023068 kB) free.
[1585342814] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6760732 kB) free!
[1585342934] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3098668 kB) free!
[1585343024] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 68.0% (33514164 kB) free.
[1585343054] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27209616 kB) free.
[1585343239] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 34.6% (8476224 kB) free.
[1585343654] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5287564 kB) free!
[1585343774] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27879604 kB) free.
[1585343909] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.0% (42881964 kB) free.
[1585344374] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6098264 kB) free!
[1585344494] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2239748 kB) free!
[1585344614] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25615396 kB) free!
[1585344643] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(345) < 2020/03/28 03:00:22 [error] 22120#0: *77799779 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585344700] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (3245904 kB) free!
[1585344764] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(14) < 2020/03/28 03:01:32 [error] 22126#0: *77800114 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585344884] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585344984] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(282) < 103.16.101.72 [28/Mar/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1585345105] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585345134] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1984) < 172.105.59.173 [28/Mar/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1585345214] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3821480 kB) free!
[1585345255] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585345268] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1585345388] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1585345988] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1585346091] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585346108] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1585346414] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13757512 kB) free!
[1585347100] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2355228 kB) free!
[1585348214] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (5207088 kB) free!
[1585348491] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.47, 3.73, 4.84
[1585349414] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13552344 kB) free!
[1585350891] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.87, 6.17, 5.42
[1585351011] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.52, 6.74, 5.73
[1585351131] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.23, 6.83, 5.88
[1585351214] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26901708 kB) free.
[1585351814] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (2992632 kB) free!
[1585351934] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26615152 kB) free.
[1585352239] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (209708 kB) free!
[1585352359] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (2448024 kB) free!
[1585352479] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165392 kB) free!
[1585352534] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6301292 kB) free!
[1585352654] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (16806936 kB) free!
[1585352774] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23833208 kB) free!
[1585352931] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.90, 4.08, 4.66
[1585353100] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 45.5% (11164240 kB) free.
[1585353229] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 68.8% (16853556 kB) free.
[1585353374] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3401172 kB) free!
[1585355174] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13367652 kB) free!
[1585356531] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.64, 5.31, 5.18
[1585356651] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.49, 5.02, 5.09
[1585356771] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1585356771] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.05, 5.02, 5.07
[1585357371] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.73, 4.13, 4.56
[1585357574] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6501168 kB) free!
[1585359374] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14320404 kB) free!
[1585361774] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1810124 kB) free!
[1585363574] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19427576 kB) free!
[1585365975] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27604588 kB) free.
[1585366575] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7509812 kB) free!
[1585366695] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.3% (25487084 kB) free!
[1585366815] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21215516 kB) free!
[1585367084] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2020/03/28 09:06:44 [error] 22134#0: *77935755 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585367204] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585367571] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.94, 6.04, 5.39
[1585367691] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.87, 6.80, 5.73
[1585367811] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.13, 7.57, 6.14
[1585368615] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1802740 kB) free!
[1585368908] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1585369028] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1585369815] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (16000316 kB) free!
[1585369900] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4757476 kB) free!
[1585370020] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4573432 kB) free!
[1585370140] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (4410752 kB) free!
[1585370242] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 32.1% (7911452 kB) free.
[1585370292] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 35.6% (17568440 kB) free.
[1585370415] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29328060 kB) free.
[1585370629] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186176 kB) free!
[1585370749] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182660 kB) free!
[1585370768] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1585370842] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182176 kB) free!
[1585370869] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184636 kB) free!
[1585370888] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585370892] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (970836 kB) free!
[1585370962] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185856 kB) free!
[1585371005] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1585371008] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1585371012] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (709332 kB) free!
[1585371015] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7449952 kB) free!
[1585371082] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173324 kB) free!
[1585371125] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1585371132] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (496512 kB) free!
[1585371135] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.4% (25530944 kB) free!
[1585371255] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20457692 kB) free!
[1585371608] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1585371892] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 60.6% (14864296 kB) free.
[1585371940] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2377580 kB) free!
[1585372325] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1585372445] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1585372455] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4135056 kB) free!
[1585372565] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1585373740] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (4400872 kB) free!
[1585373811] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1585374199] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (3870284 kB) free!
[1585374255] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18197904 kB) free!
[1585374319] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3661340 kB) free!
[1585374440] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (3447604 kB) free!
[1585375455] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2200512 kB) free!
[1585375565] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1585375640] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (1222516 kB) free!
[1585376228] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1585376348] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1585376655] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15311124 kB) free!
[1585376740] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1585376740] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2008928 kB) free!
[1585377440] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.3% (15552228 kB) free.
[1585377469] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 38.1% (9345216 kB) free.
[1585377508] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3894584 kB) free!
[1585377628] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3530768 kB) free!
[1585377748] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.3% (3046548 kB) free!
[1585377855] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1006560 kB) free!
[1585378348] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.1% (22738516 kB) free.
[1585379055] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13528296 kB) free!
[1585380255] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397548 kB) free!
[1585380892] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.9% (4627984 kB) free!
[1585381013] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.7% (4571304 kB) free!
[1585381133] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.5% (4524476 kB) free!
[1585381455] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13969448 kB) free!
[1585382655] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387068 kB) free!
[1585383855] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13695420 kB) free!
[1585384133] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2051040 kB) free!
[1585384424] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9454776 kB) free!
[1585384544] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9301240 kB) free!
[1585384565] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1585384665] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (9153808 kB) free!
[1585384670] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.1% (2732144 kB) free!
[1585384685] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1585384790] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (467980 kB) free!
[1585384805] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1585384910] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191356 kB) free!
[1585385055] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406656 kB) free!
[1585385309] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9716008 kB) free!
[1585385429] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9559820 kB) free!
[1585385454] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(480) < 103.255.146.164 [28/Mar/2020:14:15:30 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585385550] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (9453068 kB) free!
[1585385564] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1585385574] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585385684] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1585385740] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1585385804] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1585386855] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16771340 kB) free!
[1585387040] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4289952 kB) free!
[1585387160] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4056040 kB) free!
[1585387205] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1585387280] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.2% (3860224 kB) free!
[1585387455] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386496 kB) free!
[1585387805] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1585387880] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.1% (15511480 kB) free.
[1585387910] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 22.2% (5434452 kB) free.
[1585387925] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1585387942] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1585387974] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(133) < 92.42.104.166 [28/Mar/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585388045] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1585388062] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1585388094] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585388182] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1585388204] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1585388510] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (244828 kB) free!
[1585388630] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189616 kB) free!
[1585388655] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16917808 kB) free!
[1585388750] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1585388750] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203308 kB) free!
[1585389246] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1585389246] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1585389350] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 34.7% (8500280 kB) free.
[1585389382] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585389453] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 10.6% (2599600 kB) free!
[1585389465] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4601476 kB) free!
[1585389502] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1585389623] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1585389750] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4512012 kB) free!
[1585389855] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1610444 kB) free!
[1585390222] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585390342] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1585390446] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 361 processes
[1585390550] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187540 kB) free!
[1585390566] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 360 processes
[1585390604] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1585390654] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 42.7% (10453876 kB) free.
[1585390670] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179900 kB) free!
[1585390686] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 0 warn out of 362 processes
[1585390724] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1585390790] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193052 kB) free!
[1585390844] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1585391055] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19324196 kB) free!
[1585391655] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29883920 kB) free.
[1585392255] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4516576 kB) free!
[1585392294] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(452) < 103.255.146.164 [28/Mar/2020:16:14:53 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585392376] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (387072 kB) free!
[1585392414] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585392486] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1585392496] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16248928 kB) free!
[1585392590] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (3928560 kB) free!
[1585393086] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 5 crit, 0 warn out of 371 processes
[1585393096] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6356624 kB) free!
[1585393191] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196136 kB) free!
[1585393206] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 5 crit, 0 warn out of 373 processes
[1585393326] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 5 crit, 0 warn out of 369 processes
[1585393696] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17411544 kB) free!
[1585394296] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30165252 kB) free.
[1585394896] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3712372 kB) free!
[1585395016] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (375288 kB) free!
[1585395136] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31755836 kB) free.
[1585395644] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1585395644] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1585395736] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3079136 kB) free!
[1585395856] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402388 kB) free!
[1585395976] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32788100 kB) free.
[1585396576] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7219580 kB) free!
[1585396696] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3184272 kB) free!
[1585396816] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406528 kB) free!
[1585397444] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1585397480] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4336404 kB) free!
[1585397600] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4245356 kB) free!
[1585397720] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.5% (4145460 kB) free!
[1585397933] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 55.2% (13521924 kB) free.
[1585398017] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21779904 kB) free!
[1585398411] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.87, 3.57, 4.60
[1585398549] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1585398617] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31280636 kB) free.
[1585398669] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1585398921] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.0% (15489284 kB) free.
[1585399217] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (401364 kB) free!
[1585399244] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1585399337] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396752 kB) free!
[1585399457] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399108 kB) free!
[1585399869] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1585399926] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1585399989] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1585400211] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.72, 5.83, 5.14
[1585400331] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.25, 5.95, 5.27
[1585400391] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.4% (5005448 kB) free.
[1585400451] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.54, 6.01, 5.36
[1585400589] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 335 processes
[1585400657] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22454784 kB) free!
[1585400709] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1585400991] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4719492 kB) free!
[1585401111] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4660980 kB) free!
[1585401231] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (4607672 kB) free!
[1585401257] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398596 kB) free!
[1585402457] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28235556 kB) free.
[1585403057] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (593128 kB) free!
[1585403177] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (392392 kB) free!
[1585403297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35409924 kB) free.
[1585403451] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.68, 3.18, 4.32
[1585403897] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2775900 kB) free!
[1585404018] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395944 kB) free!
[1585404138] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2715456 kB) free!
[1585405338] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20651028 kB) free!
[1585405851] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.34, 6.10, 5.15
[1585405926] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1585405938] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36780552 kB) free.
[1585405971] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.21, 6.97, 5.57
[1585406091] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1585406091] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.30, 6.92, 5.72
[1585406538] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10164468 kB) free!
[1585406658] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6361036 kB) free!
[1585406778] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (11556784 kB) free!
[1585407291] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.84, 3.35, 4.37
[1585407978] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22718464 kB) free!
[1585408431] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2416856 kB) free!
[1585408578] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30271380 kB) free.
[1585409178] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7423712 kB) free!
[1585409298] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17364124 kB) free!
[1585409418] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (8718224 kB) free!
[1585409444] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1585409691] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.63, 6.99, 5.24
[1585409811] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.93, 7.71, 5.72
[1585409931] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.95, 7.81, 5.98
[1585409933] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.8% (3877808 kB) free!
[1585410053] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.2% (3469764 kB) free!
[1585410173] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 11.9% (2909628 kB) free!
[1585410773] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1870264 kB) free!
[1585411218] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14708336 kB) free!
[1585411818] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26921972 kB) free.
[1585412331] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.92, 3.69, 4.79
[1585412418] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3722184 kB) free!
[1585412538] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400308 kB) free!
[1585412658] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26874088 kB) free.
[1585413258] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3287044 kB) free!
[1585413378] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (373532 kB) free!
[1585413498] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26617428 kB) free.
[1585414098] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (688296 kB) free!
[1585414218] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28694168 kB) free.
[1585414432] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1585414818] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4115536 kB) free!
[1585414938] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (496536 kB) free!
[1585415058] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28814980 kB) free.
[1585415658] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5756232 kB) free!
[1585415778] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2175784 kB) free!
[1585415899] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27040456 kB) free.
[1585416498] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4537844 kB) free!
[1585416618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28664244 kB) free.
[1585417218] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7451572 kB) free!
[1585417338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3880124 kB) free!
[1585417458] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28624860 kB) free.
[1585418058] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6345856 kB) free!
[1585418178] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2888708 kB) free!
[1585418298] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26581988 kB) free.
[1585418898] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5120944 kB) free!
[1585419018] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28550792 kB) free.
[1585419618] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7239636 kB) free!
[1585419738] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3630076 kB) free!
[1585419858] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26573240 kB) free.