[1544208537] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544208657] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1544216069] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.23, 110.66, 89.95
[1544216189] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.47, 74.43, 79.19
[1544216309] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.07, 50.13, 69.71
[1544217029] SERVICE ALERT: qservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1544217509] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.90, 1.96, 19.88
[1544219309] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.87, 0.96, 3.69
[1544240861] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544240981] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544241057] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1544241177] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1544241284] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.0% (5440800 kB) free!
[1544241297] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1544241404] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8576820 kB) free!
[1544241525] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.8% (10726160 kB) free.
[1544241897] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544242018] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544242138] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1544242724] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.5% (5188700 kB) free!
[1544242781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544242844] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.7% (8701064 kB) free!
[1544242902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544242964] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.5% (10572168 kB) free.
[1544243022] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544243338] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544243458] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544243578] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1544247434] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 21.4% (3482860 kB) free.
[1544249909] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2.71, 100.62, 87.68
[1544250029] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.22, 67.76, 77.22
[1544250149] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.35, 45.76, 68.01
[1544250164] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.4% (6586620 kB) free!
[1544250178] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544250222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544250284] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10462540 kB) free.
[1544250298] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544250342] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544250898] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544250942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544251019] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544251062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544251139] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544251183] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544251349] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.92, 2.39, 19.75
[1544251739] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1544252741] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1544252861] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1544252981] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1544253149] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.39, 1.29, 3.95
[1544255381] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1544255684] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8299584 kB) free!
[1544255804] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10442532 kB) free.
[1544269717] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1544269949] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.48, 40.40, 27.75
[1544270069] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.32, 27.37, 24.51
[1544270189] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.95, 18.63, 21.65
[1544270204] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8591220 kB) free!
[1544270324] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10396716 kB) free.
[1544270382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544270503] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544270789] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 3.13, 11.69
[1544271517] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 300 processes
[1544271637] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 302 processes
[1544271757] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 303 processes
[1544271990] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.77, 0.81, 3.75
[1544274102] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544274157] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1544274222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544274343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1544274524] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.2% (6987276 kB) free!
[1544274540] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544274644] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10439216 kB) free.
[1544274660] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544276190] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 73.31, 229.64, 119.18
[1544276310] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 9.94, 153.69, 104.74
[1544276430] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1544276430] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.53, 102.90, 92.06
[1544276557] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1544276557] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1544277220] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.88, 8.19, 9.52
[1544277757] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1544277878] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1544277998] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1544278230] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 0.92, 13.68
[1544278420] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.99, 10.20, 10.20
[1544278540] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.41, 10.28, 10.23
[1544278660] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.46, 10.68, 10.37
[1544279430] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.83, 0.78, 4.29
[1544288798] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1544290230] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)