[1582396438] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2020/02/23 00:03:48 [error] 12255#0: *58173277 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1582396555] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17955168 kB) free!
[1582396558] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582396572] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180536 kB) free!
[1582396591] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179300 kB) free!
[1582396623] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186588 kB) free!
[1582396692] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186852 kB) free!
[1582396700] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.1% (29137064 kB) free.
[1582396711] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178692 kB) free!
[1582396743] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190476 kB) free!
[1582396812] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174308 kB) free!
[1582396831] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183056 kB) free!
[1582396863] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186184 kB) free!
[1582396920] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181376 kB) free!
[1582397041] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (181072 kB) free!
[1582397155] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28353080 kB) free.
[1582397161] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179456 kB) free!
[1582398292] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187788 kB) free!
[1582398352] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189032 kB) free!
[1582398355] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16704280 kB) free!
[1582398412] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (182800 kB) free!
[1582398472] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 67.3% (16505692 kB) free.
[1582398475] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12963568 kB) free!
[1582398595] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39212404 kB) free.
[1582398632] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2789696 kB) free!
[1582398692] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582398812] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582399195] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17305228 kB) free!
[1582399232] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191272 kB) free!
[1582399315] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13647824 kB) free!
[1582399435] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39202908 kB) free.
[1582399832] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1582399832] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.4% (10951824 kB) free.
[1582400035] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17585032 kB) free!
[1582400155] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13892880 kB) free!
[1582400275] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39161564 kB) free.
[1582400432] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1633020 kB) free!
[1582400552] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194216 kB) free!
[1582400672] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195768 kB) free!
[1582400875] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17482672 kB) free!
[1582400995] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13887368 kB) free!
[1582401115] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39142792 kB) free.
[1582401716] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16565160 kB) free!
[1582401836] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13105992 kB) free!
[1582401956] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38375136 kB) free.
[1582402556] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15634608 kB) free!
[1582402676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12312212 kB) free!
[1582402796] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35921408 kB) free.
[1582403396] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14732612 kB) free!
[1582403516] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39093700 kB) free.
[1582404116] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16930008 kB) free!
[1582404212] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1582404236] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13598256 kB) free!
[1582404332] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582404356] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37465584 kB) free.
[1582404559] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4801040 kB) free!
[1582404680] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225100 kB) free!
[1582404800] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246648 kB) free!
[1582404909] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2808384 kB) free!
[1582404956] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15952468 kB) free!
[1582405029] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (700888 kB) free!
[1582405076] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39073420 kB) free.
[1582405149] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 80.3% (39550092 kB) free.
[1582405676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11427132 kB) free!
[1582405796] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7387312 kB) free!
[1582405812] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.9% (10020296 kB) free.
[1582405916] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38575348 kB) free.
[1582406000] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 78.1% (38465540 kB) free.
[1582406517] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10898808 kB) free!
[1582406637] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6261468 kB) free!
[1582406757] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37393088 kB) free.
[1582407290] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(183) < 103.16.101.72 [23/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582407357] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10737888 kB) free!
[1582407358] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(237) < 2020/02/23 03:00:51 [error] 12249#0: *58210532 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1582407383] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3038) < 203.212.70.170 [23/Feb/2020:03:01:41 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582407412] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582407477] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32806332 kB) free.
[1582407478] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582407503] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582408077] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14179084 kB) free!
[1582408197] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10419876 kB) free!
[1582408317] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38550688 kB) free.
[1582408917] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14049936 kB) free!
[1582409037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40647268 kB) free.
[1582409637] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20190424 kB) free!
[1582409757] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37800752 kB) free.
[1582410357] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22001272 kB) free!
[1582410477] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18668316 kB) free!
[1582410597] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38105064 kB) free.
[1582411078] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2020/02/23 04:00:37 [error] 12244#0: *58222979 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582411198] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582411798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26419444 kB) free!
[1582411918] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582411918] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42758524 kB) free.
[1582412072] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582412518] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26106112 kB) free!
[1582412638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42776996 kB) free.
[1582413238] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26059900 kB) free!
[1582413358] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42781428 kB) free.
[1582413958] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25519148 kB) free!
[1582414078] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42417660 kB) free.
[1582414812] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.4% (2787964 kB) free!
[1582414932] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180476 kB) free!
[1582415052] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172036 kB) free!
[1582415278] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25273152 kB) free!
[1582415398] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41659528 kB) free.
[1582415463] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.4% (20203908 kB) free.
[1582415672] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.0% (11852820 kB) free.
[1582416035] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 964 processes
[1582416155] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 360 processes
[1582416598] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24972668 kB) free!
[1582416718] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41374916 kB) free.
[1582418398] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(20) < 2020/02/23 06:09:13 [error] 12243#0: *58250380 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1582418518] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582418804] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 87.2% (114999896 kB) free.
[1582418850] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39425620 kB) free.
[1582418895] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.0% (22055828 kB) free.
[1582418903] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(47) < 203.212.70.170 [23/Feb/2020:06:11:25 +0530] TCP 502 0 0 240.003 "10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 60.001, 60.001"
[1582418935] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.7% (14096412 kB) free.
[1582419023] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582419872] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (917496 kB) free!
[1582419992] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (221552 kB) free!
[1582420112] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204156 kB) free!
[1582421744] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.5% (21829352 kB) free.
[1582422118] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25774272 kB) free!
[1582422238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41923536 kB) free.
[1582423223] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(129) < 92.42.104.166 [23/Feb/2020:07:25:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582423343] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582425051] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1582425171] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1582425291] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1582425743] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [23/Feb/2020:08:03:53 +0530] TCP 502 0 0 127.016 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 7.015, 60.001, 60.000"
[1582425863] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582426319] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(60) < 2020/02/23 08:17:15 [error] 12246#0: *58304635 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582426439] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582429102] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183220 kB) free!
[1582429222] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182984 kB) free!
[1582429342] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174252 kB) free!
[1582429438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22448780 kB) free!
[1582429558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39456048 kB) free.
[1582430640] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(311) < 2020/02/23 09:29:30 [error] 12255#0: *58333015 upstream timed out (110: Connection timed out) while connecting to upstream, client: 182.18.144.246, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582430663] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(912) < 92.42.104.166 [23/Feb/2020:09:31:06 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582430758] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23044600 kB) free!
[1582430760] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582430783] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582430811] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(242) < 103.16.101.72 [23/Feb/2020:09:28:19 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582430878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40431924 kB) free.
[1582430931] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582431291] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1582431351] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (4054504 kB) free!
[1582431360] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/23 09:42:35 [error] 12257#0: *58310597 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/DlrLoad/msisdn"
[1582431411] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1582431472] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (4025652 kB) free!
[1582431478] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21275980 kB) free!
[1582431480] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582431592] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.1% (3981856 kB) free!
[1582431598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38430384 kB) free.
[1582432798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17457108 kB) free!
[1582432918] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38127428 kB) free.
[1582433050] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.6% (3835728 kB) free!
[1582433170] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.8% (3615972 kB) free!
[1582433291] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.1% (3466432 kB) free!
[1582433392] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1552408 kB) free!
[1582433519] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15623136 kB) free!
[1582433639] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35919124 kB) free.
[1582434491] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172888 kB) free!
[1582434839] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15560592 kB) free!
[1582434932] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1582434959] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36169132 kB) free.
[1582435053] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1582435140] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178756 kB) free!
[1582435260] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (193092 kB) free!
[1582435380] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (187308 kB) free!
[1582435559] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13943020 kB) free!
[1582435679] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34451868 kB) free.
[1582435680] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/23 10:49:12 [error] 12257#0: *58343315 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/DlrLoad/msisdn"
[1582435800] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582436279] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10969788 kB) free!
[1582436399] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30263900 kB) free.
[1582436400] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/23 11:09:42 [error] 12256#0: *58377496 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/DlrLoad/msisdn"
[1582436520] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582437599] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23777024 kB) free!
[1582437719] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43283104 kB) free.
[1582438319] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12334608 kB) free!
[1582438373] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 325 processes
[1582438439] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44268804 kB) free.
[1582438493] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 326 processes
[1582438613] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 325 processes
[1582439039] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18247024 kB) free!
[1582439159] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43370704 kB) free.
[1582439520] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2020/02/23 12:00:44 [error] 12254#0: *58403269 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582439640] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582439760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19341192 kB) free!
[1582439783] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 172.105.59.173 [23/Feb/2020:12:01:52 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1582439813] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1582439842] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21951740 kB) free.
[1582439880] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15118764 kB) free!
[1582439880] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 86.7% (21254840 kB) free.
[1582439903] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582439929] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22020500 kB) free.
[1582439961] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.0% (14303272 kB) free.
[1582440000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39880736 kB) free.
[1582440063] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188128 kB) free!
[1582440183] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175556 kB) free!
[1582440303] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1582440303] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187144 kB) free!
[1582440413] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1582440423] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1582440490] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190300 kB) free!
[1582440533] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1582440543] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1582440600] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14789088 kB) free!
[1582440610] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (183952 kB) free!
[1582440720] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39948688 kB) free.
[1582440730] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1582440730] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170008 kB) free!
[1582441061] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9617984 kB) free!
[1582441181] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9617276 kB) free!
[1582441301] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (8227928 kB) free!
[1582441320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (14990628 kB) free!
[1582441440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39854376 kB) free.
[1582442040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14548200 kB) free!
[1582442040] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/23 12:37:25 [error] 12255#0: *58419016 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/DlrLoad/msisdn"
[1582442160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39557716 kB) free.
[1582442161] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582442248] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172244 kB) free!
[1582442252] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582442368] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188924 kB) free!
[1582442372] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582442488] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168572 kB) free!
[1582442760] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13189180 kB) free!
[1582442880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38732952 kB) free.
[1582443480] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12724112 kB) free!
[1582443600] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.1% (23922268 kB) free!
[1582443720] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34529776 kB) free.
[1582444320] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9419060 kB) free!
[1582444441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35599040 kB) free.
[1582444901] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4657556 kB) free!
[1582445040] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11123236 kB) free!
[1582445160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36344800 kB) free.
[1582445760] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11779192 kB) free!
[1582445880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36013092 kB) free.
[1582446480] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12685644 kB) free!
[1582446600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35711508 kB) free.
[1582447200] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11771060 kB) free!
[1582447320] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35368480 kB) free.
[1582447729] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4382788 kB) free!
[1582447849] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (4321584 kB) free!
[1582447920] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12447436 kB) free!
[1582447970] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (4257696 kB) free!
[1582448040] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35125000 kB) free.
[1582448640] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12777828 kB) free!
[1582448760] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28714880 kB) free.
[1582448811] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1582448931] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1582449360] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13057672 kB) free!
[1582449480] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8980904 kB) free!
[1582449601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32414140 kB) free.
[1582450143] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582450164] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.8% (21524552 kB) free.
[1582450196] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.7% (21889740 kB) free.
[1582450200] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9546288 kB) free!
[1582450225] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.9% (39392016 kB) free.
[1582450264] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582450274] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.8% (21755044 kB) free.
[1582450316] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 57.2% (75410316 kB) free.
[1582450370] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2297756 kB) free!
[1582450703] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(164) < 185.255.8.58 [23/Feb/2020:15:00:24 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582450823] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582450863] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582450983] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582451103] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582451971] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186864 kB) free!
[1582452091] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181572 kB) free!
[1582452212] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192108 kB) free!
[1582452599] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.2% (2768692 kB) free!
[1582452719] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.2% (2770984 kB) free!
[1582452839] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2728752 kB) free!
[1582453131] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1582453251] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1582454039] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2399564 kB) free!
[1582454160] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(70) < 2020/02/23 16:01:02 [error] 12250#0: *58522042 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1582454280] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582454703] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1582454706] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.5% (21699388 kB) free.
[1582454785] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21871504 kB) free.
[1582454827] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.7% (21998292 kB) free.
[1582455533] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1582455653] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1582455682] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (1366292 kB) free!
[1582455773] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1582455802] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (1346872 kB) free!
[1582455922] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (1323552 kB) free!
[1582455989] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (660676 kB) free!
[1582456109] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (214940 kB) free!
[1582456229] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1582456229] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170488 kB) free!
[1582456680] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/23 16:44:32 [error] 12257#0: *58544986 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1582456800] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582461431] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4755224 kB) free!
[1582461551] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4733668 kB) free!
[1582461671] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (4710392 kB) free!
[1582462269] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.0% (21727208 kB) free.
[1582462317] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.1% (21596964 kB) free.
[1582462355] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21948284 kB) free.
[1582462400] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9749748 kB) free!
[1582462520] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9715408 kB) free!
[1582462640] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (9671880 kB) free!
[1582462831] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9673636 kB) free!
[1582462951] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9673648 kB) free!
[1582463071] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (9669152 kB) free!
[1582463110] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1582466372] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1582466492] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1582468269] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1582468321] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24504348 kB) free!
[1582468441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.0% (46136032 kB) free.
[1582468800] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2020/02/23 20:00:46 [error] 12245#0: *58637189 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582468824] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(14) < 203.212.70.170 [23/Feb/2020:20:02:53 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1582468920] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582468944] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582469641] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24638616 kB) free!
[1582469761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.2% (46335940 kB) free.
[1582470961] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25238356 kB) free!
[1582471081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45238508 kB) free.
[1582474892] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582475012] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1582476812] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1582476933] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1582481300] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (242416 kB) free!
[1582481420] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234536 kB) free!
[1582481540] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231784 kB) free!