[1584383429] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15630916 kB) free!
[1584383582] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1584383703] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1584384029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33758524 kB) free.
[1584384097] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199580 kB) free!
[1584384186] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185400 kB) free!
[1584384217] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190020 kB) free!
[1584384308] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174204 kB) free!
[1584384313] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186372 kB) free!
[1584384337] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197736 kB) free!
[1584384428] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169236 kB) free!
[1584384434] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166824 kB) free!
[1584384554] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186288 kB) free!
[1584384629] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10467144 kB) free!
[1584384652] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3946.49, 1107.71, 380.39
[1584384749] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6778260 kB) free!
[1584384772] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6211.36, 2913.23, 1130.21
[1584384869] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3202496 kB) free!
[1584384892] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3933.02, 3494.59, 1580.58
[1584384984] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 945.94, 2737.65, 1535.83
[1584385104] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 129.29, 1832.48, 1349.85
[1584385130] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 86.18, 1627.41, 1246.25
[1584385224] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.39, 1226.91, 1186.46
[1584385250] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.75, 1089.68, 1095.41
[1584385370] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.82, 729.87, 962.86
[1584385373] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (1537120 kB) free!
[1584385411] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(46) < 49.50.86.152 [17/Mar/2020:00:26:41 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584385433] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188828 kB) free!
[1584385472] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16254580 kB) free!
[1584385494] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (198892 kB) free!
[1584385532] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584385555] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (197524 kB) free!
[1584385824] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.77, 166.24, 622.65
[1584385970] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.88, 99.58, 505.48
[1584386073] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31124580 kB) free.
[1584386155] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 70.9% (17369024 kB) free.
[1584386354] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.4% (10456200 kB) free.
[1584386570] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.25, 15.33, 265.93
[1584386673] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8892752 kB) free!
[1584386692] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.41, 10.36, 229.57
[1584386793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5318164 kB) free!
[1584386913] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1790680 kB) free!
[1584386971] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1584387091] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 705 processes
[1584387211] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1584387554] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (2410676 kB) free!
[1584387674] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (724740 kB) free!
[1584387794] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174412 kB) free!
[1584387811] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584387902] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1584388022] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1584388113] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24082248 kB) free!
[1584388554] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (980844 kB) free!
[1584388675] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230188 kB) free!
[1584388713] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4623904 kB) free!
[1584388795] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235088 kB) free!
[1584389313] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15259928 kB) free!
[1584389827] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12834020 kB) free!
[1584389913] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27063876 kB) free.
[1584389948] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8701572 kB) free!
[1584390068] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (9665480 kB) free!
[1584390513] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6973828 kB) free!
[1584390633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3515760 kB) free!
[1584390753] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31137596 kB) free.
[1584390811] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584390931] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584390931] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1106) < 203.212.70.170 [17/Mar/2020:02:05:29 +0530] TCP 502 0 0 180.002 "10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.001, 60.000, 0.000"
[1584391052] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584391195] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 55.0% (27091252 kB) free.
[1584391353] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9496080 kB) free!
[1584391473] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6032420 kB) free!
[1584391531] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 707 processes
[1584391593] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33944488 kB) free.
[1584391645] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9868904 kB) free!
[1584391651] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 707 processes
[1584391765] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 11.0% (5412976 kB) free!
[1584391772] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 707 processes
[1584391885] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234492 kB) free!
[1584392193] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11664908 kB) free!
[1584392314] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8091276 kB) free!
[1584392434] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (12262268 kB) free!
[1584392737] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.4% (12103736 kB) free.
[1584393034] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13684668 kB) free!
[1584393085] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.5% (41624440 kB) free.
[1584393451] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(158) < 92.42.104.166 [17/Mar/2020:02:40:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584393572] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584394172] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584394234] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33938180 kB) free.
[1584394772] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584394834] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11335724 kB) free!
[1584394892] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584394954] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7719848 kB) free!
[1584395074] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32887332 kB) free.
[1584395492] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1584395613] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1584395674] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10115304 kB) free!
[1584395733] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584395794] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.1% (23813996 kB) free!
[1584395914] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30561864 kB) free.
[1584396514] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9833876 kB) free!
[1584396634] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33908352 kB) free.
[1584396933] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584397053] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584397234] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14373124 kB) free!
[1584397354] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (11003392 kB) free!
[1584397475] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33838620 kB) free.
[1584398074] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6881680 kB) free!
[1584398194] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (2980788 kB) free!
[1584398314] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33597320 kB) free.
[1584398822] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1584398914] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7295312 kB) free!
[1584398942] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1584399034] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3382512 kB) free!
[1584399154] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23611108 kB) free!
[1584399453] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584399573] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1584399693] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1584399754] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (8710620 kB) free!
[1584400074] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 366 processes
[1584400194] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1584400293] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584400354] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18657724 kB) free!
[1584400954] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27171616 kB) free.
[1584401494] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1584401614] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584401737] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168568 kB) free!
[1584401857] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178212 kB) free!
[1584401978] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168608 kB) free!
[1584402154] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16561148 kB) free!
[1584402214] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584402274] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27895696 kB) free.
[1584402334] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1584402454] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584402874] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18538712 kB) free!
[1584402994] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (15055388 kB) free!
[1584403028] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 78.1% (19154380 kB) free.
[1584403114] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35073816 kB) free.
[1584403394] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.2% (11157584 kB) free.
[1584403715] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16438580 kB) free!
[1584403835] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35763000 kB) free.
[1584404435] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17981596 kB) free!
[1584404555] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (14538144 kB) free!
[1584404675] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35319404 kB) free.
[1584405275] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16162528 kB) free!
[1584405395] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37761624 kB) free.
[1584405537] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 90.4% (119166912 kB) free.
[1584405995] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19095408 kB) free!
[1584406115] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37635424 kB) free.
[1584406327] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39325480 kB) free.
[1584406394] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.4% (20941736 kB) free.
[1584406715] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20254964 kB) free!
[1584406835] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16522172 kB) free!
[1584406955] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36343304 kB) free.
[1584406994] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196556 kB) free!
[1584407114] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166980 kB) free!
[1584407234] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170024 kB) free!
[1584407555] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17599436 kB) free!
[1584407675] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36684208 kB) free.
[1584408275] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18401888 kB) free!
[1584408395] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36812876 kB) free.
[1584408455] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584408575] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1584408695] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 708 processes
[1584408994] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.0% (20472848 kB) free.
[1584408996] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17853824 kB) free!
[1584409116] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36675284 kB) free.
[1584409295] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584409716] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18029020 kB) free!
[1584409836] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36755744 kB) free.
[1584410436] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17885448 kB) free!
[1584410495] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1584410556] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36667852 kB) free.
[1584410615] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584411156] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18036336 kB) free!
[1584411276] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36793736 kB) free.
[1584411876] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17510652 kB) free!
[1584411996] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35811756 kB) free.
[1584412415] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584412535] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1584412596] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16127112 kB) free!
[1584412716] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35449584 kB) free.
[1584413916] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18630420 kB) free!
[1584414036] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29945108 kB) free.
[1584414636] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18165936 kB) free!
[1584414757] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584414757] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28588012 kB) free.
[1584415356] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18015228 kB) free!
[1584415476] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14381824 kB) free!
[1584415596] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33419048 kB) free.
[1584416196] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10768064 kB) free!
[1584416316] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33553832 kB) free.
[1584416773] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584416798] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 23.29, 53.63, 45.56
[1584416903] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584416916] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9660252 kB) free!
[1584416923] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 21.98, 43.11, 42.74
[1584416978] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4900052 kB) free!
[1584417033] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584417036] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21399300 kB) free!
[1584417043] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 24.20, 36.58, 40.36
[1584417098] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.8% (4597144 kB) free!
[1584417156] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29343992 kB) free.
[1584417218] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.3% (4239724 kB) free!
[1584417335] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 710 processes
[1584417404] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584417455] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 711 processes
[1584417525] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584417575] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 711 processes
[1584417644] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 25.03, 26.26, 32.92
[1584417756] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5228988 kB) free!
[1584417876] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30508836 kB) free.
[1584417997] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.9% (4172348 kB) free!
[1584418117] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (3996640 kB) free!
[1584418238] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.5% (3815592 kB) free!
[1584418242] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38759 MB (75% inode=98%):
[1584418246] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 22.98, 22.58, 27.82
[1584418418] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (805904 kB) free!
[1584418476] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5653828 kB) free!
[1584418596] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1321580 kB) free!
[1584418716] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28721488 kB) free.
[1584419316] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4106640 kB) free!
[1584419437] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406500 kB) free!
[1584419438] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1989264 kB) free!
[1584419452] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584419455] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 416.80, 140.54, 65.47
[1584419557] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26718996 kB) free.
[1584419572] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38759 MB (75% inode=98%):
[1584419575] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 325.08, 183.64, 89.87
[1584419695] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1584419695] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.79, 130.47, 81.76
[1584419975] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584420157] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (836556 kB) free!
[1584420277] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (446468 kB) free!
[1584420397] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24688948 kB) free!
[1584420575] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 710 processes
[1584420695] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 705 processes
[1584420815] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 706 processes
[1584420997] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1456612 kB) free!
[1584421337] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.2% (6026552 kB) free!
[1584421457] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.5% (5173704 kB) free!
[1584421577] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (5171256 kB) free!
[1584422095] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 19.23, 18.35, 33.11
[1584422197] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (25051836 kB) free!
[1584422696] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 11.09, 14.02, 24.25
[1584422797] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1584422797] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28919724 kB) free.
[1584423274] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21703692 kB) free.
[1584423326] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.3% (20056496 kB) free.
[1584423397] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (405288 kB) free!
[1584423517] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (457592 kB) free!
[1584423637] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33234524 kB) free.
[1584423815] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584423878] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188572 kB) free!
[1584423977] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1790252 kB) free!
[1584423998] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187048 kB) free!
[1584424028] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176668 kB) free!
[1584424049] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584424068] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584424118] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1584424118] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192748 kB) free!
[1584424148] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188632 kB) free!
[1584424169] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584424188] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584424237] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8210028 kB) free!
[1584424268] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175640 kB) free!
[1584424289] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1584424309] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1584424357] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4024088 kB) free!
[1584424477] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388872 kB) free!
[1584424785] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3660188 kB) free!
[1584424905] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3358892 kB) free!
[1584425015] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 709 processes
[1584425025] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.8% (3089200 kB) free!
[1584425078] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16277488 kB) free!
[1584425135] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1584425255] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1584425597] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4394352 kB) free!
[1584425625] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1533576 kB) free!
[1584425678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30462796 kB) free.
[1584425717] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.3% (4245760 kB) free!
[1584425837] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (4030880 kB) free!
[1584425856] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1584426108] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584426228] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584426278] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6431476 kB) free!
[1584426337] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (4306308 kB) free!
[1584426348] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584426386] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4760228 kB) free!
[1584426398] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2405940 kB) free!
[1584426456] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 711 processes
[1584426457] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (4087176 kB) free!
[1584426506] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4282664 kB) free!
[1584426518] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 60.4% (14803532 kB) free.
[1584426518] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34324652 kB) free.
[1584426572] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(158) < 103.255.146.164 [17/Mar/2020:11:50:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584426576] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 716 processes
[1584426577] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (3848016 kB) free!
[1584426626] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (3809760 kB) free!
[1584426692] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(132) < 103.255.146.164 [17/Mar/2020:12:00:29 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584426696] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 718 processes
[1584426812] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584426826] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.8% (15230724 kB) free.
[1584427037] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (762360 kB) free!
[1584427118] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7323468 kB) free!
[1584427177] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1584427177] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1958488 kB) free!
[1584427226] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1186244 kB) free!
[1584427238] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3090780 kB) free!
[1584427358] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399220 kB) free!
[1584427496] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 28.14, 38.92, 38.96
[1584427616] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 26.74, 34.94, 37.51
[1584427736] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 27.38, 32.67, 36.39
[1584427827] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.4% (22808576 kB) free.
[1584427885] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9020940 kB) free!
[1584428005] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8641596 kB) free!
[1584428125] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.7% (8253356 kB) free!
[1584428336] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 17.97, 22.72, 29.59
[1584428558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28622216 kB) free.
[1584428613] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(126) < 103.255.146.164 [17/Mar/2020:12:29:18 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584428733] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584428873] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 27.92, 1183.81, 1038.58
[1584428993] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6.04, 792.98, 912.97
[1584429113] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.80, 531.38, 802.56
[1584429158] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4214244 kB) free!
[1584429278] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399620 kB) free!
[1584429398] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406244 kB) free!
[1584429605] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584429713] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.66, 73.49, 421.74
[1584429735] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584429865] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584429925] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3418144 kB) free!
[1584430118] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.2% (3978008 kB) free!
[1584430140] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1106.84, 575.24, 243.69
[1584430238] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.2% (3715056 kB) free!
[1584430270] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584430358] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.0% (3427456 kB) free!
[1584430393] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1182.31, 937.52, 473.86
[1584430467] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 71.6% (35264600 kB) free.
[1584430492] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2311.44, 546.12, 182.58
[1584430548] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1584430598] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24665332 kB) free!
[1584430612] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6166.82, 2604.81, 981.07
[1584430732] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6687.14, 3982.36, 1682.75
[1584430868] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 31.6% (7751168 kB) free.
[1584430958] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2182636 kB) free!
[1584430970] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 76.26, 1475.10, 1113.53
[1584431090] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.35, 987.80, 978.80
[1584431133] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.8% (21761244 kB) free.
[1584431134] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 203.212.70.170 [17/Mar/2020:13:10:17 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1584431169] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.4% (21316724 kB) free.
[1584431198] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1480460 kB) free!
[1584431211] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.19, 661.83, 860.43
[1584431211] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39216024 kB) free.
[1584431251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46668948 kB) free.
[1584431254] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584431299] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.3% (45482252 kB) free.
[1584431340] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.6% (20972148 kB) free.
[1584431811] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.48, 90.66, 451.90
[1584432376] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195824 kB) free!
[1584432455] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(72) < 103.255.146.164 [17/Mar/2020:13:31:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584432496] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174312 kB) free!
[1584432532] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.26, 13.94, 259.89
[1584432575] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584432616] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178132 kB) free!
[1584433015] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9087368 kB) free!
[1584433055] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11314872 kB) free!
[1584433135] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8598440 kB) free!
[1584433142] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 339 processes
[1584433176] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7341284 kB) free!
[1584433255] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (8599224 kB) free!
[1584433262] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1584433268] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4895516 kB) free!
[1584433296] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3438232 kB) free!
[1584433327] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (305104 kB) free!
[1584433383] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1584433388] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4807880 kB) free!
[1584433393] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.93, 10.85, 35.73
[1584433447] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (371496 kB) free!
[1584433508] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4753956 kB) free!
[1584433567] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (345384 kB) free!
[1584433775] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(86) < 103.255.146.164 [17/Mar/2020:13:50:27 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584433895] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584433993] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 15.86, 12.93, 24.77
[1584434496] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33904688 kB) free.
[1584434496] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584435067] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9213720 kB) free!
[1584435096] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(129) < 103.255.146.164 [17/Mar/2020:14:17:30 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584435096] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8097748 kB) free!
[1584435096] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 720 processes
[1584435187] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (8997964 kB) free!
[1584435216] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584435216] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3949480 kB) free!
[1584435217] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 720 processes
[1584435307] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (8825828 kB) free!
[1584435308] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1584435308] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1996544 kB) free!
[1584435336] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383312 kB) free!
[1584435337] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 720 processes
[1584435655] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4110872 kB) free!
[1584435773] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.9% (4867460 kB) free!
[1584435826] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (3865572 kB) free!
[1584435893] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4889976 kB) free!
[1584435936] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16935512 kB) free!
[1584435946] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (3673940 kB) free!
[1584436013] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4891220 kB) free!
[1584436066] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.6% (3509176 kB) free!
[1584436416] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(81) < 103.255.146.164 [17/Mar/2020:14:43:31 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584436536] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34699688 kB) free.
[1584436536] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584437136] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9732984 kB) free!
[1584437256] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5359116 kB) free!
[1584437266] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1584437266] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1531672 kB) free!
[1584437376] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1179320 kB) free!
[1584437866] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.9% (15236856 kB) free.
[1584437976] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17967148 kB) free!
[1584438493] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1584438576] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380044 kB) free!
[1584438613] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1584438907] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4248212 kB) free!
[1584439013] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 5.7% (1403188 kB) free!
[1584439176] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13795388 kB) free!
[1584439536] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(166) < 92.42.104.166 [17/Mar/2020:15:30:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584439657] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584439777] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34672468 kB) free.
[1584440136] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.1% (4673700 kB) free!
[1584440256] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.7% (4348656 kB) free!
[1584440347] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (4281892 kB) free!
[1584440376] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.3% (3991656 kB) free!
[1584440377] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8924072 kB) free!
[1584440389] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3146496 kB) free!
[1584440467] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (3999652 kB) free!
[1584440497] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4413100 kB) free!
[1584440509] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3713712 kB) free.
[1584440587] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.1% (3689260 kB) free!
[1584440617] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380080 kB) free!
[1584440748] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1584440856] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(54) < 103.255.146.164 [17/Mar/2020:15:49:42 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584440868] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 342 processes
[1584440976] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584440976] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1967568 kB) free!
[1584440988] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1584441109] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (2650112 kB) free!
[1584441187] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1706324 kB) free!
[1584441217] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20794244 kB) free!
[1584441229] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3312608 kB) free.
[1584441782] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584441817] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397840 kB) free!
[1584441902] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1584442177] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(602) < 103.13.97.241 [17/Mar/2020:16:18:59 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584442188] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1584442216] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584442297] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(120) < 203.212.70.170 [17/Mar/2020:16:20:08 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584442417] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18743456 kB) free!
[1584442417] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584442788] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1584442908] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1584443017] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401124 kB) free!
[1584443028] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 338 processes
[1584443137] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1584443553] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(49) < 2020/03/17 16:41:11 [error] 12242#0: *72199003 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:9726/147150, bytes from/to upstream:147150/9726
[1584443617] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16490660 kB) free!
[1584443628] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1584443628] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1584443673] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584444217] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (443692 kB) free!
[1584444218] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(89) < 103.255.146.164 [17/Mar/2020:16:43:54 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584444338] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584445393] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584445417] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35398884 kB) free.
[1584445508] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584446017] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10252628 kB) free!
[1584446137] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5850504 kB) free!
[1584446257] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1559716 kB) free!
[1584446628] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1584446748] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 335 processes
[1584446857] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20872032 kB) free!
[1584446868] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1584447458] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400808 kB) free!
[1584447468] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1584448058] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17087208 kB) free!
[1584448066] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.0% (3822452 kB) free!
[1584448186] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3663776 kB) free!
[1584448306] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.7% (15190384 kB) free.
[1584448477] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584448607] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584448611] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584448658] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420332 kB) free!
[1584448734] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584448737] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584448740] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2488 processes
[1584448855] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584449137] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584449338] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.3% (29702492 kB) free.
[1584449506] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584449858] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23490880 kB) free!
[1584450288] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584450289] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584450292] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584450408] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 143455 MB (63% inode=99%):
[1584450409] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1584450412] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2478 processes
[1584450458] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411500 kB) free!
[1584451058] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30043212 kB) free.
[1584451106] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9821268 kB) free!
[1584451226] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9679704 kB) free!
[1584451346] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.3% (9535796 kB) free!
[1584451658] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4112164 kB) free!
[1584451778] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (377584 kB) free!
[1584451898] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398132 kB) free!
[1584452137] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 718 processes
[1584452138] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(106) < 103.255.146.164 [17/Mar/2020:19:01:30 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584452258] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 716 processes
[1584452258] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584452376] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584452378] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 716 processes
[1584452498] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22327016 kB) free!
[1584452932] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 604.85, 139.84, 48.18
[1584453052] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 576.20, 410.14, 167.07
[1584453098] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394544 kB) free!
[1584453172] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 79.01, 274.96, 147.01
[1584453459] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(27) < 103.255.146.164 [17/Mar/2020:19:24:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584453579] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584453698] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20825104 kB) free!
[1584454298] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425308 kB) free!
[1584454429] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.7% (2229596 kB) free!
[1584454549] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2806612 kB) free!
[1584454669] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3610196 kB) free.
[1584454898] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20733616 kB) free!
[1584455269] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.2% (2147316 kB) free!
[1584455389] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.4% (2830960 kB) free!
[1584455499] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425180 kB) free!
[1584455509] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3437924 kB) free.
[1584455546] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4802296 kB) free!
[1584456099] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20736156 kB) free!
[1584456699] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400360 kB) free!
[1584456709] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2588652 kB) free!
[1584456743] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.4% (21073772 kB) free.
[1584456829] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 40.2% (6529580 kB) free.
[1584456907] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.2% (44460808 kB) free.
[1584457299] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18458932 kB) free!
[1584457899] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417252 kB) free!
[1584458499] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18344228 kB) free!
[1584458868] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584459099] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398300 kB) free!
[1584459699] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18141840 kB) free!
[1584459913] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 89.4% (21899836 kB) free.
[1584459972] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.3% (39559432 kB) free.
[1584460025] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.2% (21121136 kB) free.
[1584460060] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.2% (20873968 kB) free.
[1584460103] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.5% (14055896 kB) free.
[1584460197] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21794720 kB) free.
[1584460259] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.4% (45521180 kB) free.
[1584460299] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407100 kB) free!
[1584460899] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18406516 kB) free!
[1584460946] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (3842232 kB) free!
[1584461066] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196228 kB) free!
[1584461186] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193020 kB) free!
[1584461499] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391084 kB) free!
[1584462099] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17759108 kB) free!
[1584462379] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.8% (7804968 kB) free!
[1584462499] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (7806764 kB) free!
[1584462619] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (7806800 kB) free!
[1584462699] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414164 kB) free!
[1584463299] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16962040 kB) free!
[1584463900] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (464812 kB) free!
[1584464500] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14728736 kB) free!
[1584465100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38633112 kB) free.
[1584465700] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15558560 kB) free!
[1584465820] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11598524 kB) free!
[1584465940] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7567308 kB) free!
[1584466540] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26750060 kB) free.
[1584467140] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6364260 kB) free!
[1584467260] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2667864 kB) free!
[1584467380] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38369172 kB) free.
[1584467980] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15287064 kB) free!
[1584468019] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4806640 kB) free!
[1584468100] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11340096 kB) free!
[1584468220] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7507068 kB) free!
[1584468820] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24612632 kB) free!
[1584469420] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4680028 kB) free!