[1553279456] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 20.0% (3257884 kB) free!
[1553279582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553279702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553280292] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7298944 kB) free!
[1553280302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553280412] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (6009632 kB) free!
[1553280422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553280532] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4450112 kB) free!
[1553280542] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553281132] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.8% (8255228 kB) free!
[1553281577] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.89, 6.59, 5.79
[1553281697] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.69, 6.67, 5.92
[1553281732] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12320700 kB) free.
[1553281742] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553281817] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.78, 6.86, 6.09
[1553282332] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5435684 kB) free!
[1553282342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553282452] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3507116 kB) free!
[1553282462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553282572] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12326724 kB) free.
[1553282582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553283172] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6496764 kB) free!
[1553283182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553283292] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5240780 kB) free!
[1553283302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553283413] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12322900 kB) free.
[1553283422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553283944] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 325 processes
[1553284013] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7418468 kB) free!
[1553284023] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553284133] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6113148 kB) free!
[1553284143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553284253] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4586292 kB) free!
[1553284263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553284545] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 321 processes
[1553284853] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.9% (8305356 kB) free!
[1553285453] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12321140 kB) free.
[1553285463] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553285753] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186608 kB) free!
[1553285873] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186688 kB) free!
[1553285993] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188708 kB) free!
[1553286053] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5337248 kB) free!
[1553286063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553286173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12331456 kB) free.
[1553286183] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553286773] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7495284 kB) free!
[1553286783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553286784] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 26.6% (4329136 kB) free.
[1553286893] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6187812 kB) free!
[1553286903] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553287013] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4763484 kB) free!
[1553287024] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553287384] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (187436 kB) free!
[1553287504] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (651948 kB) free!
[1553287614] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.4% (8572420 kB) free!
[1553287624] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.0% (13322444 kB) free.
[1553287799] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 626 processes
[1553287919] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 627 processes
[1553288039] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 626 processes
[1553288214] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12327540 kB) free.
[1553288224] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553288639] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 313 processes
[1553288814] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5538732 kB) free!
[1553288824] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553288934] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (3634144 kB) free!
[1553288944] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553289054] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12319424 kB) free.
[1553289064] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553289617] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.99, 3.64, 4.59
[1553289654] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6505904 kB) free!
[1553289664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553289774] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5223872 kB) free!
[1553289784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553289894] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12327460 kB) free.
[1553289905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553290494] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7390256 kB) free!
[1553290504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553290614] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6087820 kB) free!
[1553290624] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553290734] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4561484 kB) free!
[1553290744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553291334] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.9% (8339008 kB) free!
[1553291935] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12326444 kB) free.
[1553291944] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1553292535] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5448076 kB) free!
[1553292544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553292655] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3502056 kB) free!
[1553292664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553292775] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12314816 kB) free.
[1553292785] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553293375] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6471304 kB) free!
[1553293384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553293495] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5204700 kB) free!
[1553293504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553293615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12316240 kB) free.
[1553293624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553294215] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7444744 kB) free!
[1553294224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553294335] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6155432 kB) free!
[1553294344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553294455] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4771884 kB) free!
[1553294464] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553295055] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.7% (8702884 kB) free!
[1553295656] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12309584 kB) free.
[1553295664] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553296256] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5649644 kB) free!
[1553296265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553296376] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (3885576 kB) free!
[1553296385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553296496] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12308600 kB) free.
[1553296505] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553297096] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6676532 kB) free!
[1553297105] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553297216] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5375332 kB) free!
[1553297225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553297336] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3379852 kB) free!
[1553297345] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553297936] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.5% (7613984 kB) free!
[1553297945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553298065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553298186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553298536] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12308724 kB) free.
[1553299136] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4447844 kB) free!
[1553299256] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12307328 kB) free.
[1553299386] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553299857] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6894184 kB) free!
[1553299977] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.4% (5619580 kB) free!
[1553299986] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553300097] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3780440 kB) free!
[1553300106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553300226] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553300697] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7816216 kB) free!
[1553300826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553300868] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47090 MB (20% inode=99%):
[1553300946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553300989] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47098 MB (20% inode=99%):
[1553301067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553301109] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57148 MB (25% inode=99%):
[1553301297] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12315944 kB) free.
[1553301666] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553301786] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553301897] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4577112 kB) free!
[1553301906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553302017] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12311948 kB) free.
[1553302617] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6987128 kB) free!
[1553302737] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.6% (5712744 kB) free!
[1553302857] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3915996 kB) free!
[1553303106] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553303457] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7935308 kB) free!
[1553303706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553303826] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553304057] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12313492 kB) free.
[1553304427] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553304547] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553304658] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4917564 kB) free!
[1553304667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553304778] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12321572 kB) free.
[1553305378] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7231596 kB) free!
[1553305498] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5956988 kB) free!
[1553305618] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4358160 kB) free!
[1553305867] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553306218] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.5% (8144860 kB) free!
[1553306467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553306587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553306707] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553306818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12321844 kB) free.
[1553307307] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553307418] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5230752 kB) free!
[1553307428] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553307538] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12318324 kB) free.
[1553307548] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553308138] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7424884 kB) free!
[1553308148] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553308258] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6109956 kB) free!
[1553308268] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553308378] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4666532 kB) free!
[1553308388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553308979] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.3% (8501032 kB) free!
[1553309579] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12304624 kB) free.
[1553309588] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553310179] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5347760 kB) free!
[1553310188] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553310299] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8201072 kB) free!
[1553310309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553310419] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12310304 kB) free.
[1553310908] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553311019] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6213012 kB) free!
[1553311028] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553311139] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4834068 kB) free!
[1553311148] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553311259] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12300836 kB) free.
[1553311859] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7141040 kB) free!
[1553311979] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.8% (5818552 kB) free!
[1553312099] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4179212 kB) free!
[1553312700] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5478716 kB) free!
[1553313548] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553314149] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553314269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553314389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553315100] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10035348 kB) free.
[1553315589] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553316189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553316309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553316900] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4134344 kB) free!
[1553316909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553317020] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12247024 kB) free.
[1553317030] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553317620] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3793284 kB) free!
[1553317629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553317740] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12167996 kB) free.
[1553317749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553317869] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553318341] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3166672 kB) free!
[1553318460] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11868692 kB) free.
[1553319069] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553319660] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4093632 kB) free!
[1553319669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553319780] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12215936 kB) free.
[1553319789] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553320380] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2914712 kB) free!
[1553320389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553320500] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11653764 kB) free.
[1553320510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553320630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553321700] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3532164 kB) free!
[1553321820] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12137312 kB) free.
[1553321830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553322420] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3060472 kB) free!
[1553322430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553322540] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11762860 kB) free.
[1553322550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553322670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553323141] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2526672 kB) free!
[1553323261] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11351736 kB) free.
[1553323871] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553324217] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8424120 kB) free!
[1553324337] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13160048 kB) free.
[1553324461] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3499656 kB) free!
[1553324471] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553324509] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46603 MB (20% inode=99%):
[1553324581] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12029412 kB) free.
[1553324591] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553324630] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56391 MB (25% inode=99%):
[1553324923] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.1% (3111156 kB) free!
[1553324937] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.4% (7110284 kB) free!
[1553325043] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (3031068 kB) free!
[1553325057] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.3% (11954188 kB) free.
[1553325163] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (2951572 kB) free!
[1553325181] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3034336 kB) free!
[1553325191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553325301] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11734176 kB) free.
[1553325311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553325431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553325901] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2712940 kB) free!
[1553326021] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1553326021] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11169268 kB) free.
[1553326257] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.4% (7574692 kB) free!
[1553326377] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.7% (12165944 kB) free.
[1553326624] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2882288 kB) free!
[1553326631] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553326744] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2810232 kB) free!
[1553326864] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.8% (2733320 kB) free!
[1553326977] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.9% (7343180 kB) free!
[1553327097] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.1% (11861020 kB) free.
[1553327221] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4023260 kB) free!
[1553327231] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553327341] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11920248 kB) free.
[1553327352] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553327697] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.8% (7280388 kB) free!
[1553327817] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11729456 kB) free.
[1553327942] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3974784 kB) free!
[1553327951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553328062] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11891076 kB) free.
[1553328071] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553328417] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.3% (6576056 kB) free!
[1553328537] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11457436 kB) free.
[1553328662] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4357076 kB) free!
[1553328671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553328763] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1380996 kB) free!
[1553328782] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11865640 kB) free.
[1553328791] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553329138] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.2% (6507676 kB) free!
[1553329258] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.4% (11045552 kB) free.
[1553329264] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1250080 kB) free!
[1553329382] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4455572 kB) free!
[1553329391] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553329502] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (2640648 kB) free!
[1553329511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553329622] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11472940 kB) free.
[1553329631] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553329858] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.2% (5993688 kB) free!
[1553329978] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10651768 kB) free.
[1553330222] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3472300 kB) free!
[1553330342] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11788640 kB) free.
[1553330578] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5762116 kB) free!
[1553330698] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10250476 kB) free.
[1553330942] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3800864 kB) free!
[1553331062] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11744356 kB) free.
[1553331298] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.0% (5405420 kB) free!
[1553331418] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9848292 kB) free!
[1553331538] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.1% (7911144 kB) free!
[1553331662] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3977080 kB) free!
[1553331783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11714800 kB) free.
[1553332138] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3818192 kB) free!
[1553332382] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4214860 kB) free!
[1553332502] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (2227616 kB) free!
[1553332622] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10921448 kB) free.
[1553332738] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.4% (5638228 kB) free!
[1553333222] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2342868 kB) free!
[1553333342] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11197040 kB) free.
[1553333938] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1553333938] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (1905160 kB) free!
[1553334538] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.3% (8527516 kB) free!
[1553334542] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3408176 kB) free!
[1553334662] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11539620 kB) free.
[1553335138] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4151216 kB) free!
[1553335262] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2782212 kB) free!
[1553335382] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11492644 kB) free.
[1553335641] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (4486776 kB) free!
[1553335738] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.6% (7200460 kB) free!
[1553335761] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (4354984 kB) free!
[1553335881] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4222488 kB) free!
[1553335982] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2289232 kB) free!
[1553336103] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11134960 kB) free.
[1553336702] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (2551704 kB) free!
[1553336822] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11212444 kB) free.
[1553337422] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3422252 kB) free!
[1553337542] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11385232 kB) free.
[1553337681] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2407708 kB) free!
[1553338138] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3054240 kB) free!
[1553338142] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3476380 kB) free!
[1553338262] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11327472 kB) free.
[1553338738] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.4% (6130828 kB) free!
[1553338862] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3747376 kB) free!
[1553338982] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (1836252 kB) free!
[1553339102] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10696248 kB) free.
[1553339231] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1553339338] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.4% (12011980 kB) free.
[1553339702] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3091180 kB) free!
[1553339822] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11289720 kB) free.
[1553339831] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1553339938] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.8% (7777560 kB) free!
[1553340059] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.1% (5000576 kB) free!
[1553340179] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.4% (11529360 kB) free.
[1553340402] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 19.7% (1445192 kB) free!
[1553340423] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3136476 kB) free!
[1553340431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553340522] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.6% (1440364 kB) free!
[1553340543] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11274896 kB) free.
[1553340552] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553340642] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 19.4% (1430064 kB) free!
[1553340779] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.7% (6260220 kB) free!
[1553340899] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11454704 kB) free.
[1553341143] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2856496 kB) free!
[1553341151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553341263] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11229672 kB) free.
[1553341271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553341391] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553341499] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8057952 kB) free!
[1553341619] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.5% (5158472 kB) free!
[1553341739] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11219816 kB) free.
[1553341863] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3784844 kB) free!
[1553341983] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11252700 kB) free.
[1553341991] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553342339] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.1% (6957340 kB) free!
[1553342459] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (4107368 kB) free!
[1553342579] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.5% (10588640 kB) free.
[1553342583] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3809384 kB) free!
[1553342591] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553342703] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11239876 kB) free.
[1553342711] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553343179] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.7% (6283092 kB) free!
[1553343299] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3682544 kB) free!
[1553343303] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3103212 kB) free!
[1553343311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553343420] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.5% (10098292 kB) free.
[1553343423] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11149168 kB) free.
[1553343432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553343552] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1553343552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553344019] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.4% (6096728 kB) free!
[1553344023] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2272996 kB) free!
[1553344139] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (3322208 kB) free!
[1553344143] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10502184 kB) free.
[1553344259] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.8% (9779676 kB) free!
[1553344752] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553345343] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3262516 kB) free!
[1553345352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1553345463] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11219276 kB) free.
[1553345472] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553346063] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2475892 kB) free!
[1553346072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553346183] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10896192 kB) free.
[1553346192] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553346312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553346659] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2953312 kB) free!
[1553346783] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (1850804 kB) free!
[1553346903] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10303776 kB) free.
[1553347259] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.5% (7619164 kB) free!
[1553347429] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45994 MB (20% inode=99%):
[1553347512] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553347549] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55784 MB (24% inode=99%):
[1553347859] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3214908 kB) free!
[1553348103] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3270020 kB) free!
[1553348112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553348223] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11157280 kB) free.
[1553348232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553348459] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.7% (7744912 kB) free!
[1553348823] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2766860 kB) free!
[1553348832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553348944] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10839940 kB) free.
[1553348952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553349059] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3084488 kB) free!
[1553349072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553349543] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2159608 kB) free!
[1553349659] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.5% (7617996 kB) free!
[1553349663] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10330892 kB) free.
[1553350260] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3479936 kB) free!
[1553350273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553350860] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.4% (8575572 kB) free!
[1553350863] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3174672 kB) free!
[1553350873] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553350983] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11155416 kB) free.
[1553350993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553351113] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553351460] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4416488 kB) free!
[1553351583] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2437212 kB) free!
[1553351703] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10605872 kB) free.
[1553352060] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11941700 kB) free.
[1553352313] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553352660] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7656428 kB) free!
[1553352780] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.1% (4998008 kB) free!
[1553352900] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.5% (11589368 kB) free.
[1553352903] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4108924 kB) free!
[1553352913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553353023] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (2259140 kB) free!
[1553353033] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553353144] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10538744 kB) free.
[1553353154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553353500] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7185360 kB) free!
[1553353620] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.2% (12431612 kB) free.
[1553353743] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (2337012 kB) free!
[1553353863] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10512776 kB) free.
[1553354463] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (2227076 kB) free!
[1553354583] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10417024 kB) free.
[1553354820] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6210128 kB) free!
[1553354940] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.2% (12421772 kB) free.
[1553355183] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (2319812 kB) free!
[1553355303] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10476752 kB) free.
[1553355540] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8338272 kB) free!
[1553355661] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.0% (5420228 kB) free!
[1553355781] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.5% (12052256 kB) free.
[1553355903] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2146924 kB) free!
[1553356023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10487948 kB) free.
[1553356381] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.2% (7470192 kB) free!
[1553356501] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4703156 kB) free!
[1553356621] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11441664 kB) free.
[1553357221] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.7% (7233464 kB) free!
[1553357223] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3561484 kB) free!
[1553357341] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (4375404 kB) free!
[1553357343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11129440 kB) free.
[1553357354] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553357461] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11101956 kB) free.
[1553357944] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3465768 kB) free!
[1553357954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553358061] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.2% (6973776 kB) free!
[1553358064] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11121364 kB) free.
[1553358075] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553358181] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.2% (12418528 kB) free.
[1553358664] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3717920 kB) free!
[1553358674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553358781] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8724812 kB) free!
[1553358784] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11127804 kB) free.
[1553358794] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553358901] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.0% (5928444 kB) free!
[1553359021] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.2% (12427108 kB) free.
[1553359384] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4647912 kB) free!
[1553359394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553359504] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (3009996 kB) free!
[1553359514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553359549] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44003 MB (19% inode=99%):
[1553359622] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7523548 kB) free!
[1553359624] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11124024 kB) free.
[1553359634] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553359669] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56295 MB (25% inode=99%):
[1553359742] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4852436 kB) free!
[1553359862] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11227740 kB) free.
[1553360224] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5032760 kB) free!
[1553360234] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553360344] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (3668944 kB) free!
[1553360354] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553360462] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.3% (6536808 kB) free!
[1553360464] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11120188 kB) free.
[1553360475] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553360582] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.3% (12456316 kB) free.
[1553361064] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (5997044 kB) free!
[1553361074] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553361182] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.9% (7818116 kB) free!
[1553361184] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4673948 kB) free!
[1553361194] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553361302] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.4% (5147352 kB) free!
[1553361305] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (2987284 kB) free!
[1553361314] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553361422] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11438328 kB) free.
[1553361905] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6830352 kB) free!
[1553362022] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.9% (6828344 kB) free!
[1553362142] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12508648 kB) free.
[1553362505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11111888 kB) free.
[1553362514] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553362742] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7894748 kB) free!
[1553362862] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12528740 kB) free.
[1553363105] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3788024 kB) free!
[1553363114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553363225] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11118664 kB) free.
[1553363234] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553363462] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8747736 kB) free!
[1553363583] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.0% (5915380 kB) free!
[1553363703] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.1% (11857600 kB) free.
[1553363825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6324180 kB) free!
[1553363835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553363945] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.1% (5000492 kB) free!
[1553363955] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553364065] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1553364065] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3604820 kB) free!
[1553364075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553364303] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.4% (7105632 kB) free!
[1553364423] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.5% (12571460 kB) free.
[1553364665] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7600928 kB) free!
[1553365023] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7863284 kB) free!
[1553365069] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44278 MB (19% inode=99%):
[1553365143] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.5% (12590904 kB) free.
[1553365189] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56283 MB (25% inode=99%):
[1553365265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11111532 kB) free.
[1553365275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553365743] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9027536 kB) free!