[1551465422] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8759536 kB) free!
[1551465542] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7370732 kB) free!
[1551465662] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13262724 kB) free.
[1551465793] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 19.3% (1420104 kB) free!
[1551465913] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 18.9% (1388376 kB) free!
[1551465914] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3071424 kB) free!
[1551466033] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 18.7% (1371744 kB) free!
[1551466034] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (2974224 kB) free!
[1551466154] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.8% (2894608 kB) free!
[1551467580] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36980 MB (17% inode=99%):
[1551467700] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44659 MB (21% inode=99%):
[1551468900] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41090 MB (19% inode=99%):
[1551469020] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 45195 MB (21% inode=99%):
[1551470129] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49828 MB (22% inode=99%):
[1551470354] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1625196 kB) free!
[1551470462] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8368744 kB) free!
[1551470582] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6945596 kB) free!
[1551470656] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (1473704 kB) free!
[1551470702] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13067512 kB) free.
[1551470731] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551470776] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 20.0% (1473844 kB) free!
[1551470896] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 20.0% (1473464 kB) free!
[1551471302] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8202528 kB) free!
[1551471329] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44452 MB (19% inode=99%):
[1551471332] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551471422] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6634452 kB) free!
[1551471449] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50116 MB (22% inode=99%):
[1551471452] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551471542] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13071684 kB) free.
[1551472049] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46020 MB (20% inode=99%):
[1551472052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551472142] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8364744 kB) free!
[1551472169] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41924 MB (18% inode=99%):
[1551472172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551472262] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7039412 kB) free!
[1551472289] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41924 MB (18% inode=99%):
[1551472292] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551472382] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13061460 kB) free.
[1551472620] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43243 MB (20% inode=99%):
[1551472740] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41195 MB (19% inode=99%):
[1551472860] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 45305 MB (21% inode=99%):
[1551472982] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8257616 kB) free!
[1551473102] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6718880 kB) free!
[1551473223] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13061064 kB) free.
[1551473489] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48067 MB (21% inode=99%):
[1551473822] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7786136 kB) free!
[1551473942] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13065752 kB) free.
[1551474092] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551474542] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8512608 kB) free!
[1551474662] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7139760 kB) free!
[1551474689] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551474692] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551474782] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13065624 kB) free.
[1551474809] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551474813] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551475382] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7941968 kB) free!
[1551475409] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551475412] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551475502] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13061336 kB) free.
[1551475529] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551475532] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551476102] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8682660 kB) free!
[1551476129] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551476132] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551476222] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7392868 kB) free!
[1551476249] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551476252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551476342] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13065056 kB) free.
[1551476369] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551476372] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551476460] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42882 MB (20% inode=99%):
[1551476580] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40868 MB (19% inode=99%):
[1551476700] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42579 MB (20% inode=99%):
[1551476942] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8139392 kB) free!
[1551476969] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551476972] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551477063] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6480208 kB) free!
[1551477089] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50114 MB (22% inode=99%):
[1551477092] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551477183] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13065656 kB) free.
[1551477689] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551477692] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551477783] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7631028 kB) free!
[1551477809] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41410 MB (18% inode=99%):
[1551477813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551477903] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13059816 kB) free.
[1551477929] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49602 MB (22% inode=99%):
[1551477933] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551478503] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8337516 kB) free!
[1551478529] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551478533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551478623] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6832784 kB) free!
[1551478649] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50114 MB (22% inode=99%):
[1551478653] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551478743] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13064508 kB) free.
[1551479249] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46019 MB (20% inode=99%):
[1551479253] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551479343] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7816436 kB) free!
[1551479369] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551479373] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551479463] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13058492 kB) free.
[1551479489] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551479493] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551480063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8553796 kB) free!
[1551480089] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44387 MB (19% inode=99%):
[1551480093] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551480183] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7220044 kB) free!
[1551480209] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40291 MB (17% inode=99%):
[1551480214] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551480304] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13059632 kB) free.
[1551480330] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551480334] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551480903] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7985892 kB) free!
[1551480929] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551480934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551481023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13063772 kB) free.
[1551481049] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551481054] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551481500] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46731 MB (22% inode=99%):
[1551481623] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8675400 kB) free!
[1551481649] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551481654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551481743] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7381840 kB) free!
[1551481769] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551481774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551481863] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13059308 kB) free.
[1551481889] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551481894] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551482463] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8153352 kB) free!
[1551482489] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44835 MB (19% inode=99%):
[1551482494] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551482583] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6506572 kB) free!
[1551482609] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551482614] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551482703] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13056560 kB) free.
[1551483209] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551483215] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551483303] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7590108 kB) free!
[1551483329] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551483335] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551483423] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13061744 kB) free.
[1551483449] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551483455] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551483900] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42736 MB (20% inode=99%):
[1551484020] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46846 MB (22% inode=99%):
[1551484024] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8357284 kB) free!
[1551484049] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44067 MB (19% inode=99%):
[1551484055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551484144] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6834732 kB) free!
[1551484169] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50114 MB (22% inode=99%):
[1551484175] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551484264] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13057964 kB) free.
[1551484769] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46018 MB (20% inode=99%):
[1551484775] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551484864] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7749688 kB) free!
[1551484889] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551484895] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551484984] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13062244 kB) free.
[1551485009] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50114 MB (22% inode=99%):
[1551485015] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551485584] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8479280 kB) free!
[1551485609] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41410 MB (18% inode=99%):
[1551485616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551485704] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7031808 kB) free!
[1551485729] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41410 MB (18% inode=99%):
[1551485736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551485824] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13058972 kB) free.
[1551485849] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551485856] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551486424] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7800968 kB) free!
[1551486449] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551486456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551486544] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13061500 kB) free.
[1551486569] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50114 MB (22% inode=99%):
[1551486576] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551487144] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8477788 kB) free!
[1551487170] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551487176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551487265] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7046812 kB) free!
[1551487290] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551487296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551487385] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13060272 kB) free.
[1551487410] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1551487410] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551487416] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1551487416] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551487620] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42812 MB (20% inode=99%):
[1551487740] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46960 MB (22% inode=99%):
[1551487777] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1551487898] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1551487985] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7803048 kB) free!
[1551488010] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551488017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551488105] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13052616 kB) free.
[1551488130] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50114 MB (22% inode=99%):
[1551488137] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551488705] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8470500 kB) free!
[1551488730] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1551488737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551488825] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7085428 kB) free!
[1551488850] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551488857] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551488933] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 530 processes
[1551488945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13059352 kB) free.
[1551488970] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50114 MB (22% inode=99%):
[1551488977] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551488993] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 532 processes
[1551489545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7865380 kB) free!
[1551489570] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44098 MB (19% inode=99%):
[1551489577] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551489665] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13056016 kB) free.
[1551489690] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50115 MB (22% inode=99%):
[1551489697] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551490140] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42974 MB (20% inode=99%):
[1551490260] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47084 MB (22% inode=99%):
[1551490265] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8557796 kB) free!
[1551490290] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551490297] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551490385] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7190980 kB) free!
[1551490410] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551490418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551490505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13052964 kB) free.
[1551490530] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50114 MB (22% inode=99%):
[1551490538] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551491106] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7792400 kB) free!
[1551491130] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551491138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551491226] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13054176 kB) free.
[1551491250] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50114 MB (22% inode=99%):
[1551491258] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551491826] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8445888 kB) free!
[1551491850] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41922 MB (18% inode=99%):
[1551491858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551491946] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7033228 kB) free!
[1551491970] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41921 MB (18% inode=99%):
[1551491978] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551492066] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13042720 kB) free.
[1551492090] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50113 MB (22% inode=99%):
[1551492098] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1551492660] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43050 MB (20% inode=99%):
[1551492666] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2054036 kB) free!
[1551492690] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41921 MB (18% inode=99%):
[1551492698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551492780] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47154 MB (22% inode=99%):
[1551492786] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5285708 kB) free!
[1551492810] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49601 MB (22% inode=99%):
[1551492818] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551492906] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.4% (14965264 kB) free.
[1551493380] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42931 MB (20% inode=99%):
[1551493410] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41921 MB (18% inode=99%):
[1551493419] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551493500] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47023 MB (22% inode=99%):
[1551493506] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9554380 kB) free!
[1551493531] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41921 MB (18% inode=99%):
[1551493539] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551493626] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.4% (14962408 kB) free.
[1551493651] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50113 MB (22% inode=99%):
[1551493659] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551494251] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41409 MB (18% inode=99%):
[1551494259] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551494371] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41405 MB (18% inode=99%):
[1551494379] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551494491] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49599 MB (22% inode=99%):
[1551494979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551495091] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41407 MB (18% inode=99%):
[1551495099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551495211] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49597 MB (22% inode=99%):
[1551495219] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551495300] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42933 MB (20% inode=99%):
[1551495420] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44729 MB (21% inode=99%):
[1551495811] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41919 MB (18% inode=99%):
[1551495820] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551495931] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41919 MB (18% inode=99%):
[1551495940] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551496051] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50111 MB (22% inode=99%):
[1551496060] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551496626] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8972296 kB) free!
[1551496651] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41393 MB (18% inode=99%):
[1551496660] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551496746] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.0% (14787628 kB) free.
[1551496771] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49075 MB (21% inode=99%):
[1551496780] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551496900] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551497864] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (698456 kB) free!
[1551498571] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45439 MB (20% inode=99%):
[1551498691] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45439 MB (20% inode=99%):
[1551498811] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41343 MB (18% inode=99%):
[1551499633] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (686720 kB) free!
[1551500555] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (1283632 kB) free!
[1551500676] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (1272224 kB) free!
[1551500796] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1261728 kB) free!
[1551500896] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (668456 kB) free!
[1551501812] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49991 MB (22% inode=99%):
[1551503346] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9465196 kB) free!
[1551503466] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.6% (14611064 kB) free.
[1551503500] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551504101] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551504172] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24437 MB (10% inode=99%):
[1551504221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551504341] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551504773] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24795 MB (11% inode=99%):
[1551504812] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45823 MB (20% inode=99%):
[1551504932] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45695 MB (20% inode=99%):
[1551505052] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41599 MB (18% inode=99%):
[1551505373] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24435 MB (10% inode=99%):
[1551507420] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42533 MB (20% inode=99%):
[1551507452] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1551507452] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49269 MB (21% inode=99%):
[1551507460] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 333 processes
[1551507540] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40485 MB (19% inode=99%):
[1551507562] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.53, 7.69, 5.46
[1551507593] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 19.5% (1430944 kB) free!
[1551507660] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49370 MB (23% inode=99%):
[1551507682] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.48, 5.27, 4.85
[1551507713] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 19.2% (1413096 kB) free!
[1551507833] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 18.9% (1391628 kB) free!
[1551508060] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 317 processes
[1551508542] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1551508542] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551509142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551509252] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46955 MB (20% inode=99%):
[1551509262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551509372] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45422 MB (20% inode=99%):
[1551509382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551509466] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9434972 kB) free!
[1551509492] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41326 MB (18% inode=99%):
[1551509586] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8053548 kB) free!
[1551509706] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.9% (14234744 kB) free.
[1551509927] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 366 processes
[1551510047] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1551510242] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.73, 5.25, 5.03
[1551510362] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.06, 5.25, 5.06
[1551510482] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.15, 5.26, 5.08
[1551511082] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.54, 4.63, 4.85
[1551511872] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.7% (22243220 kB) free.
[1551511893] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48694 MB (21% inode=99%):
[1551512033] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.3% (7012020 kB) free.
[1551512233] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 94.8% (6971744 kB) free.
[1551512296] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 94.3% (6933888 kB) free.
[1551512354] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.7% (14581476 kB) free.
[1551512982] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551513660] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40587 MB (19% inode=99%):
[1551513693] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47153 MB (20% inode=99%):
[1551513780] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 36491 MB (17% inode=99%):
[1551513813] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45105 MB (20% inode=99%):
[1551513900] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43639 MB (20% inode=99%):
[1551513907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9302448 kB) free!
[1551513933] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43793 MB (19% inode=99%):
[1551514027] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7898708 kB) free!
[1551514147] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.7% (14128872 kB) free.
[1551514182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551514302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551514422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551514500] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46506 MB (22% inode=99%):
[1551515081] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2587596 kB) free!
[1551515201] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.8% (2412240 kB) free!
[1551515282] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.95, 5.88, 5.35
[1551515321] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.9% (2265272 kB) free!
[1551515402] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.81, 5.88, 5.42
[1551515522] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.98, 5.88, 5.46
[1551515947] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9753008 kB) free!
[1551516067] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7677724 kB) free!
[1551516187] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (5998616 kB) free!
[1551516333] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49247 MB (21% inode=99%):
[1551516722] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551516722] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.33, 4.63, 4.93
[1551517121] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1260364 kB) free!
[1551517387] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4370184 kB) free!
[1551517501] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40196 MB (19% inode=99%):
[1551517621] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47327 MB (22% inode=99%):
[1551518022] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551518587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12039116 kB) free.
[1551518733] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47181 MB (20% inode=99%):
[1551518821] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36013 MB (17% inode=99%):
[1551518853] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45133 MB (20% inode=99%):
[1551518941] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43167 MB (20% inode=99%):
[1551518973] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45133 MB (20% inode=99%):
[1551519061] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 40095 MB (19% inode=99%):
[1551519222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551519342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551519462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551520261] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551520261] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 43984 MB (21% inode=99%):
[1551520861] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 35923 MB (17% inode=99%):
[1551520981] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43083 MB (20% inode=99%):
[1551521101] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 40011 MB (19% inode=99%):
[1551521587] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9688500 kB) free!
[1551521707] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8083012 kB) free!
[1551521827] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551521827] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.2% (6503332 kB) free!
[1551522301] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44059 MB (21% inode=99%):
[1551522901] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39953 MB (19% inode=99%):
[1551523021] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44041 MB (21% inode=99%):
[1551523284] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189820 kB) free!
[1551523303] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 104.63, 65.32, 29.02
[1551523404] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195308 kB) free!
[1551523423] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 105.25, 78.55, 38.27
[1551523525] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1285320 kB) free!
[1551523543] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 59.32, 75.92, 42.56
[1551523621] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 35841 MB (17% inode=99%):
[1551523721] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 63.1% (10252204 kB) free.
[1551523741] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44012 MB (21% inode=99%):
[1551523773] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48683 MB (21% inode=99%):
[1551523789] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54655 MB (24% inode=99%):
[1551524227] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4929736 kB) free!
[1551524341] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38198 MB (18% inode=99%):
[1551524461] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43498 MB (20% inode=99%):
[1551524488] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4908120 kB) free!
[1551524581] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41258 MB (19% inode=99%):
[1551524608] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4870352 kB) free!
[1551524728] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4851476 kB) free!
[1551524743] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.31, 2.73, 12.67
[1551525181] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 43960 MB (21% inode=99%):
[1551525723] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.65, 5.44, 5.02
[1551525781] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 35465 MB (16% inode=99%):
[1551525843] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.85, 5.26, 5.01
[1551525901] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41865 MB (20% inode=99%):
[1551525943] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.82, 1.76, 4.66
[1551525963] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.59, 5.00, 4.94
[1551526021] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 39689 MB (19% inode=99%):
[1551526062] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551526154] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (2888740 kB) free!
[1551526274] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (2708216 kB) free!
[1551526394] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.6% (2532572 kB) free!
[1551526608] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4786900 kB) free!
[1551526627] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11418380 kB) free.
[1551526728] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4719364 kB) free!
[1551526773] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47135 MB (20% inode=99%):
[1551526848] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4663196 kB) free!
[1551526893] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45151 MB (20% inode=99%):
[1551526994] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1424384 kB) free!
[1551527013] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43103 MB (19% inode=99%):
[1551527221] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46311 MB (22% inode=99%):
[1551527262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551527382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551527502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551527513] SERVICE ALERT: sig4-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 274 processes
[1551527633] SERVICE ALERT: sig4-vm1;CPU_Procs;OK;HARD;1;CPU OK: 271 processes
[1551527763] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1551527822] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 34628 MB (16% inode=99%):
[1551527942] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41747 MB (20% inode=99%):
[1551528062] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 38675 MB (18% inode=99%):
[1551529027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9718056 kB) free!
[1551529147] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6960000 kB) free!
[1551529267] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5186816 kB) free!
[1551529305] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8970040 kB) free!
[1551529425] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.3% (14914108 kB) free.
[1551529789] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47154 MB (20% inode=99%):
[1551529909] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49587 MB (22% inode=99%):
[1551530025] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8739216 kB) free!
[1551530128] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2304644 kB) free!
[1551530145] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.3% (14910164 kB) free.
[1551530467] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4803004 kB) free!
[1551530745] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9094884 kB) free!
[1551530763] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.63, 5.33, 5.14
[1551530865] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.2% (15381684 kB) free.
[1551530883] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.56, 5.08, 5.07
[1551531003] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.39, 4.78, 4.96
[1551531465] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9472532 kB) free!
[1551531586] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.9% (15243052 kB) free.
[1551531813] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48110 MB (21% inode=99%):
[1551532267] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5103208 kB) free!
[1551532413] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40430 MB (17% inode=99%):
[1551532533] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48109 MB (21% inode=99%):
[1551532823] SERVICE ALERT: sig5-vm1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1551532853] HOST ALERT: sig5-vm1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1551532917] HOST ALERT: sig5-vm1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1551532943] SERVICE ALERT: sig5-vm1;CPU_Procs;OK;HARD;1;CPU OK: 273 processes
[1551533385] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8880236 kB) free!
[1551533505] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.1% (14833080 kB) free.
[1551533734] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46571 MB (20% inode=99%):
[1551533854] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44523 MB (19% inode=99%):
[1551533974] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44523 MB (19% inode=99%):
[1551534055] SERVICE ALERT: sig6-vm1;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig6-vm1' was out of bounds)
[1551534105] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9335660 kB) free!
[1551534175] SERVICE ALERT: sig6-vm1;CPU_Procs;OK;HARD;1;CPU OK: 271 processes
[1551534225] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.8% (14707068 kB) free.
[1551534825] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9607292 kB) free!
[1551534945] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.6% (14590440 kB) free.
[1551536745] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9044336 kB) free!
[1551536866] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.4% (6115296 kB) free!
[1551536986] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12529972 kB) free.
[1551537586] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.8% (6792408 kB) free!
[1551537706] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.5% (13055172 kB) free.
[1551538302] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1551538306] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.8% (7810076 kB) free!
[1551538426] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.2% (13921192 kB) free.
[1551538616] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 187.97, 44.82, 15.32
[1551538660] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 316 processes
[1551538737] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 500.75, 211.82, 80.09
[1551538857] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 179.59, 216.23, 99.91
[1551539026] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9104480 kB) free!
[1551539125] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (2749776 kB) free!
[1551539143] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4.41, 102.95, 86.24
[1551539146] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.2% (13891936 kB) free.
[1551539263] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.08, 69.41, 75.97
[1551539383] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.38, 46.85, 66.91
[1551539462] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551539502] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551539746] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9804268 kB) free!
[1551539866] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.3% (7032628 kB) free!
[1551539986] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.1% (12853836 kB) free.
[1551540103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551540223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551540343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551540574] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48589 MB (21% inode=99%):
[1551540583] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.44, 2.18, 19.38
[1551540586] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8387012 kB) free!
[1551540657] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.72, 1.19, 14.95
[1551540707] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.2% (13885612 kB) free.
[1551540943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1551541063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551541183] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551541267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11481020 kB) free.
[1551541774] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40492 MB (18% inode=99%):
[1551541783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551541857] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.73, 0.76, 4.66
[1551541867] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5046696 kB) free!
[1551541894] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48684 MB (21% inode=99%):
[1551541903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551541987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11478488 kB) free.
[1551542125] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188720 kB) free!
[1551542383] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.45, 1.66, 4.26
[1551542573] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 31443 MB (13% inode=99%):
[1551542594] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 71.0% (11538888 kB) free.
[1551543048] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 23.7% (5845900 kB) free.
[1551543094] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44045 MB (19% inode=99%):
[1551543103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551543187] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6064940 kB) free!
[1551543214] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39949 MB (17% inode=99%):
[1551543223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551543307] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11474456 kB) free.
[1551543325] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3324880 kB) free!
[1551543328] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 29.2% (7165360 kB) free.
[1551543334] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48140 MB (21% inode=99%):
[1551543343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551543632] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 336 processes
[1551543706] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7518976 kB) free!
[1551543826] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.2% (13399404 kB) free.
[1551543943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551544063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551544183] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551544426] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8522360 kB) free!
[1551544546] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.0% (13810556 kB) free.
[1551544660] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 318 processes
[1551544832] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 336 processes
[1551546307] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7002380 kB) free!
[1551546334] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44107 MB (19% inode=99%):
[1551546346] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7187216 kB) free!
[1551546427] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.6% (5722084 kB) free!
[1551546454] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42571 MB (18% inode=99%):
[1551546466] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.0% (12808664 kB) free.
[1551546547] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11591116 kB) free.
[1551546574] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48651 MB (21% inode=99%):
[1551546583] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551547066] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.0% (6896620 kB) free!
[1551547147] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6900544 kB) free!
[1551547174] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40459 MB (18% inode=99%):
[1551547183] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551547187] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.3% (12486488 kB) free.
[1551547267] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5373496 kB) free!
[1551547294] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40451 MB (17% inode=99%):
[1551547303] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551547387] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11594492 kB) free.
[1551547414] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1551547414] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48641 MB (21% inode=99%):
[1551547423] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551547786] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7650736 kB) free!
[1551547906] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.1% (13851724 kB) free.
[1551547987] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6821548 kB) free!
[1551548014] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40451 MB (17% inode=99%):
[1551548023] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551548107] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.8% (5314736 kB) free!
[1551548134] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40451 MB (17% inode=99%):
[1551548143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551548228] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1551548228] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11591528 kB) free.
[1551548254] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48642 MB (21% inode=99%):
[1551548263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551548506] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9256320 kB) free!
[1551548626] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.2% (13873004 kB) free.
[1551548827] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6851136 kB) free!
[1551548854] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40450 MB (17% inode=99%):
[1551548864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551548947] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.0% (5411064 kB) free!
[1551548974] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40450 MB (17% inode=99%):
[1551548984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551549067] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11592492 kB) free.
[1551549094] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48642 MB (21% inode=99%):
[1551549104] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551549226] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9313940 kB) free!
[1551549346] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.2% (13889984 kB) free.
[1551549667] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6716960 kB) free!
[1551549694] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40450 MB (17% inode=99%):
[1551549704] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551549787] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5124620 kB) free!
[1551549814] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48642 MB (21% inode=99%):
[1551549824] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551549907] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11580132 kB) free.
[1551550414] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44738 MB (19% inode=99%):
[1551550424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551550507] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6855748 kB) free!
[1551550534] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42722 MB (19% inode=99%):
[1551550544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551550627] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.5% (5665056 kB) free!
[1551550655] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 38594 MB (17% inode=99%):
[1551550664] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1551550664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551550747] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11592388 kB) free.
[1551551347] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6835716 kB) free!