[1565893989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6692248 kB) free!
[1565894010] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565894110] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10830792 kB) free.
[1565894130] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565894163] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (1791028 kB) free!
[1565894285] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184856 kB) free!
[1565894405] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188736 kB) free!
[1565894570] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565894709] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6287384 kB) free!
[1565894730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565894829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10838816 kB) free.
[1565894850] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565895429] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7005248 kB) free!
[1565895450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565895549] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10840604 kB) free.
[1565895570] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565896149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7284192 kB) free!
[1565896170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565896269] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.1% (5487756 kB) free!
[1565896291] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565896389] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10839904 kB) free.
[1565896766] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.73, 0.57, 11.61
[1565896890] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565896989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5897960 kB) free!
[1565897010] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565897109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10828336 kB) free.
[1565897130] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565897402] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 579 processes
[1565897462] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 574 processes
[1565897524] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 577 processes
[1565897584] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 572 processes
[1565897709] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6351372 kB) free!
[1565897730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565897829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10839848 kB) free.
[1565897850] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565897966] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.38, 3.50
[1565898005] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1565898005] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.1% (9393804 kB) free.
[1565898429] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6757688 kB) free!
[1565898450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565898550] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10841032 kB) free.
[1565898570] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565899149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7472724 kB) free!
[1565899170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565899269] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6302376 kB) free!
[1565899291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1565899389] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10836768 kB) free.
[1565899411] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1565899411] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565899989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6012636 kB) free!
[1565900011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565900109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10827804 kB) free.
[1565900131] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565900709] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5440636 kB) free!
[1565900829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10289156 kB) free.
[1565901331] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565901451] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565901572] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565902172] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565902205] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182988 kB) free!
[1565902325] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197664 kB) free!
[1565902446] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197644 kB) free!
[1565902452] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217056 kB) free!
[1565902572] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (229604 kB) free!
[1565902629] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7438008 kB) free!
[1565902692] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215684 kB) free!
[1565902701] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1565902749] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6137496 kB) free!
[1565902772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565902821] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565902870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10825272 kB) free.
[1565902892] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565903421] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565903469] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7747392 kB) free!
[1565903492] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565903541] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1565903589] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6736308 kB) free!
[1565903612] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565903709] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10821952 kB) free.
[1565903732] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565903892] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.9% (41838856 kB) free.
[1565904309] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5972796 kB) free!
[1565904332] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565904429] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10687624 kB) free.
[1565904452] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565904573] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1565904573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565904981] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 45.6% (11169676 kB) free.
[1565905046] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 23.3% (5720332 kB) free.
[1565905629] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5645992 kB) free!
[1565905749] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10481280 kB) free.
[1565906373] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565906949] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7581948 kB) free!
[1565906973] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565907069] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6461168 kB) free!
[1565907093] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565907141] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1565907190] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10821040 kB) free.
[1565907213] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565907261] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565907789] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7173532 kB) free!
[1565907813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565907909] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5522280 kB) free!
[1565907934] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565908029] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10820248 kB) free.
[1565908533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565908629] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6446196 kB) free!
[1565908653] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565908749] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10821800 kB) free.
[1565908773] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565909349] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6975048 kB) free!
[1565909373] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565909469] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5577568 kB) free!
[1565909493] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565909589] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10821016 kB) free.
[1565910093] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565910189] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6084396 kB) free!
[1565910213] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565910246] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1565910309] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10813120 kB) free.
[1565910334] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565910446] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.1% (3955308 kB) free!
[1565910566] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.8% (3630708 kB) free!
[1565910566] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4817.48, 4149.44, 2266.16
[1565910686] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 25.0% (6137144 kB) free.
[1565910686] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 706.92, 2831.06, 2014.30
[1565910806] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 95.82, 1894.48, 1770.12
[1565910861] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1565910909] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7201484 kB) free!
[1565910933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565910981] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1565911030] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (6031888 kB) free!
[1565911053] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1565911150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10822968 kB) free.
[1565911173] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1565911581] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565911702] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1565911750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7060868 kB) free!
[1565911773] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565911870] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.5% (5690856 kB) free!
[1565911893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565911990] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10823040 kB) free.
[1565912013] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565912590] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6687396 kB) free!
[1565912613] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565912710] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10816084 kB) free.
[1565912733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565913310] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7703564 kB) free!
[1565913334] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565913430] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6662216 kB) free!
[1565913454] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565913550] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10822168 kB) free.
[1565913574] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565914150] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7561956 kB) free!
[1565914174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565914270] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6392864 kB) free!
[1565914294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565914391] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10819436 kB) free.
[1565914414] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565914990] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7172332 kB) free!
[1565915006] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 0.56, 19.72
[1565915014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565915110] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5891460 kB) free!
[1565915134] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565915230] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10819500 kB) free.
[1565915254] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565915830] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7131016 kB) free!
[1565915855] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565915950] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.8% (5826360 kB) free!
[1565915975] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565916070] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10812212 kB) free.
[1565916095] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565916670] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7001268 kB) free!
[1565916695] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565916790] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.6% (5697428 kB) free!
[1565916807] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.52, 3.28
[1565916815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565916846] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 65.0% (16031040 kB) free.
[1565916910] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10812776 kB) free.
[1565916935] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565917510] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7034408 kB) free!
[1565917535] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565917630] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.6% (5696404 kB) free!
[1565917655] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565917751] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10820700 kB) free.
[1565917775] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565918350] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6804508 kB) free!
[1565918376] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565918470] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10815264 kB) free.
[1565918496] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565918584] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 544 processes
[1565918644] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 542 processes
[1565918705] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 540 processes
[1565919070] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7502240 kB) free!
[1565919096] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565919190] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6334060 kB) free!
[1565919216] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565919310] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10814612 kB) free.
[1565919336] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565919501] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565919621] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1565919910] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7322728 kB) free!
[1565919936] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565920030] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6100132 kB) free!
[1565920056] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565920150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10819244 kB) free.
[1565920177] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1565920177] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565920750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6936992 kB) free!
[1565920776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565920870] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5513648 kB) free!
[1565920896] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565920990] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10816888 kB) free.
[1565921496] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565921591] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6822508 kB) free!
[1565921616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565921711] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10708668 kB) free.
[1565921736] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565922311] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7524204 kB) free!
[1565922336] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565922431] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6229500 kB) free!
[1565922456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565922551] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10819364 kB) free.
[1565922576] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565923151] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7669676 kB) free!
[1565923177] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1565923271] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6585720 kB) free!
[1565923297] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565923391] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10817172 kB) free.
[1565923417] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565923970] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565923991] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7926968 kB) free!
[1565924017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565924090] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1565924111] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6897904 kB) free!
[1565924137] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565924210] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1565924231] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.2% (5523864 kB) free!
[1565924257] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565924486] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4880628 kB) free!
[1565924606] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4859524 kB) free!
[1565924726] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2016624 kB) free!
[1565924781] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3942824 kB) free!
[1565924831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10813556 kB) free.
[1565924901] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3908252 kB) free!
[1565925021] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167772 kB) free!
[1565925457] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565925578] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565925698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565925846] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.9% (3185768 kB) free!
[1565925966] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.9% (3174004 kB) free!
[1565926086] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (3158044 kB) free!
[1565926221] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565926341] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565926610] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565927231] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6848232 kB) free!
[1565927351] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10811472 kB) free.
[1565927498] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565927951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6919504 kB) free!
[1565928071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10801388 kB) free.
[1565928671] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7175420 kB) free!
[1565928698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565928791] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5523356 kB) free!
[1565928818] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565928911] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10798516 kB) free.
[1565929087] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191076 kB) free!
[1565929418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565929511] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5327548 kB) free!
[1565929539] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565929631] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10121892 kB) free.
[1565930138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565930212] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565930258] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565930332] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1565930378] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565930452] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1565931046] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 362 processes
[1565931166] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 358 processes
[1565931286] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 356 processes
[1565931750] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.54, 7.61, 5.77
[1565931870] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.79, 6.87, 5.74
[1565931990] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.91, 6.41, 5.72
[1565932178] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1565932178] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565932487] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1565932590] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.43, 4.60, 4.91
[1565932631] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5820144 kB) free!
[1565932751] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10800516 kB) free.
[1565933087] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 357 processes
[1565933190] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.00, 7.11, 6.20
[1565933207] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 359 processes
[1565933310] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.38, 6.72, 6.18
[1565933327] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1565933351] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6605052 kB) free!
[1565933378] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565933430] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.44, 5.88, 5.94
[1565933471] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10790628 kB) free.
[1565933498] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565934030] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1565934030] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.97, 3.14, 4.59
[1565934035] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.11, 5.99, 5.23
[1565934071] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6935104 kB) free!
[1565934098] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565934155] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.96, 5.74, 5.25
[1565934191] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10796376 kB) free.
[1565934218] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565934276] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.71, 6.03, 5.40
[1565934791] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7068940 kB) free!
[1565934818] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565934911] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10786260 kB) free.
[1565934938] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565936111] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5794348 kB) free!
[1565936139] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565936231] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10786888 kB) free.
[1565936259] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565936327] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 360 processes
[1565936430] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.34, 5.89, 5.23
[1565936448] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1565936550] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.22, 5.30, 5.10
[1565936568] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1565936670] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.44, 5.14, 5.07
[1565936704] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 629 processes
[1565936764] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 626 processes
[1565936824] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 627 processes
[1565936866] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 11.2% (2738308 kB) free!
[1565936885] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 623 processes
[1565936926] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 11.2% (2737320 kB) free!
[1565936986] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 11.2% (2738696 kB) free!
[1565937046] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 11.1% (2732136 kB) free!
[1565937168] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1565937168] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1565937184] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 634 processes
[1565937244] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 636 processes
[1565937270] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.63, 4.10, 4.61
[1565937304] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 638 processes
[1565937346] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (181980 kB) free!
[1565937364] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 0 warn out of 639 processes
[1565937394] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1565937454] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1565937459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565937514] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 63 users currently logged in
[1565937574] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1565937579] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565937665] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 1 warn out of 645 processes
[1565937699] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565937946] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.3% (4973412 kB) free.
[1565937965] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 647 processes
[1565938031] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8728504 kB) free!
[1565938151] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7378664 kB) free!
[1565938271] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5557504 kB) free!
[1565938299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565938871] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10238796 kB) free.
[1565938899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565939019] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565939140] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565939471] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8979700 kB) free!
[1565939591] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7700348 kB) free!
[1565939674] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1565939711] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6268048 kB) free!
[1565940274] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1565940276] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.95, 3.62, 4.62
[1565940312] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1565940312] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 34.5% (16986736 kB) free.
[1565940438] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.7% (23360912 kB) free.
[1565940452] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.2% (23098948 kB) free.
[1565940472] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.4% (36672312 kB) free.
[1565940493] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.4% (22884684 kB) free.
[1565940511] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 32.0% (5194400 kB) free.
[1565940646] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 452 processes
[1565940874] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1565940874] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1565941081] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4732184 kB) free!
[1565941111] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (179132 kB) free!
[1565941201] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (1078648 kB) free!
[1565941231] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173360 kB) free!
[1565941237] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4913828 kB) free!
[1565941321] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217280 kB) free!
[1565941352] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156492 kB) free!
[1565941357] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 20.0% (4897440 kB) free!
[1565941368] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 368 processes
[1565941478] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.9% (4884168 kB) free!
[1565941488] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 368 processes
[1565941512] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9388224 kB) free!
[1565941608] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 369 processes
[1565941632] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8057288 kB) free!
[1565941752] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11880392 kB) free.
[1565942076] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.16, 6.07, 5.10
[1565942196] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.14, 5.40, 4.95
[1565942208] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1565942352] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8188212 kB) free!
[1565942374] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1565942472] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11873328 kB) free.
[1565942740] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565942797] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.21, 6.95, 5.88
[1565942917] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.02, 6.42, 5.81
[1565943037] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.87, 6.19, 5.80
[1565943072] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8234816 kB) free!
[1565943192] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11879988 kB) free.
[1565943646] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4515388 kB) free!
[1565943706] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4521640 kB) free!
[1565943766] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 18.4% (4516844 kB) free!
[1565943792] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8318344 kB) free!
[1565943826] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 18.4% (4517552 kB) free!
[1565943913] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6412984 kB) free!
[1565944033] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11876044 kB) free.
[1565944540] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565944633] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6572600 kB) free!
[1565944660] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565944753] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11878380 kB) free.
[1565945260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565945353] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6611292 kB) free!
[1565945380] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565945473] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11863744 kB) free.
[1565945974] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1565945980] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565946073] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6715736 kB) free!
[1565946100] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565946193] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11870264 kB) free.
[1565946700] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565946793] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6804760 kB) free!
[1565946821] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565946913] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11859864 kB) free.
[1565947199] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565947420] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565947474] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1565947513] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7038844 kB) free!
[1565947540] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565947633] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11869036 kB) free.
[1565947660] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565947695] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.71, 5.39, 5.07
[1565947794] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565947815] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.97, 5.56, 5.17
[1565947935] SERVICE ALERT: webservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.61, 5.60, 5.23
[1565948234] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7192928 kB) free!
[1565948260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565948354] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11869124 kB) free.
[1565948380] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565948671] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565948954] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7318908 kB) free!
[1565948980] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565949074] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11859600 kB) free.
[1565949100] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565949674] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7484040 kB) free!
[1565949701] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565949794] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11865856 kB) free.
[1565949821] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565950126] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 8.1% (1983972 kB) free!
[1565950341] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 399 processes
[1565950394] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7479188 kB) free!
[1565950421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565950461] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1565950514] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11852336 kB) free.
[1565950541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565951026] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 20.0% (4901648 kB) free!
[1565951074] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1565951114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7463960 kB) free!
[1565951141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565951208] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 2 warn out of 375 processes
[1565951234] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11861200 kB) free.
[1565951261] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565951328] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 2 warn out of 372 processes
[1565951448] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 377 processes
[1565951671] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.18, 7.05, 6.02
[1565951674] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1565951791] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.70, 6.91, 6.09
[1565951834] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7409884 kB) free!
[1565951861] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565951893] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.6% (37734204 kB) free.
[1565951911] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.85, 6.27, 5.97
[1565951934] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 85.9% (13967416 kB) free.
[1565951954] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11862676 kB) free.
[1565951981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565952048] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 376 processes
[1565952102] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565952274] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1565952503] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (1315172 kB) free!
[1565952511] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.39, 3.78, 4.75
[1565952545] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (185240 kB) free!
[1565952623] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (251372 kB) free!
[1565952648] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1565952665] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (168480 kB) free!
[1565952735] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.47, 4.48, 4.88
[1565952743] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230796 kB) free!
[1565952785] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178752 kB) free!
[1565952862] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 397 processes
[1565952982] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565953034] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.33, 5.61, 5.13
[1565953154] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.34, 5.81, 5.26
[1565953194] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1565953199] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1565953274] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.55, 5.65, 5.26
[1565953474] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1565953726] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 32.6% (8000956 kB) free.
[1565953754] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565953774] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1565953902] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565954354] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7736168 kB) free!
[1565954399] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565954448] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 379 processes
[1565954474] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11850064 kB) free.
[1565954535] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.84, 5.27, 5.09
[1565954568] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 375 processes
[1565954655] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.96, 5.19, 5.09
[1565954688] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 375 processes
[1565954775] SERVICE ALERT: webservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.32, 5.28, 5.14
[1565954912] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.87, 6.47, 5.35
[1565955032] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.37, 6.75, 5.58
[1565955074] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7682984 kB) free!
[1565955102] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565955152] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.81, 6.76, 5.73
[1565955194] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11854176 kB) free.
[1565955222] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565955274] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1565955398] SERVICE ALERT: sigqscluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565955461] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182320 kB) free!
[1565955519] SERVICE ALERT: sigqscluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1565955574] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1565955581] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180952 kB) free!
[1565955639] SERVICE ALERT: sigqscluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1565955665] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 653 processes
[1565955674] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.15, 4.32, 4.76
[1565955683] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.3% (22370008 kB) free.
[1565955725] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 652 processes
[1565955734] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.2% (40031840 kB) free.
[1565955774] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.7% (14262056 kB) free.
[1565955786] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 651 processes
[1565955794] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7919436 kB) free!
[1565955822] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1565955826] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (594828 kB) free!
[1565955886] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (594516 kB) free!
[1565955914] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11841724 kB) free.
[1565955942] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565955946] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 2.4% (594396 kB) free!
[1565956006] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 2.5% (601160 kB) free!
[1565956342] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8950808 kB) free!
[1565956352] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1565956352] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.85, 3.55, 4.44
[1565956384] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (150496 kB) free!
[1565956462] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (3226920 kB) free!
[1565956488] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1565956504] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (154672 kB) free!
[1565956514] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7928012 kB) free!
[1565956543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565956575] SERVICE FLAPPING ALERT: webservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1565956575] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.35, 4.37, 4.68
[1565956582] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1565956582] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229184 kB) free!
[1565956624] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1565956624] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (193784 kB) free!
[1565956634] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11851848 kB) free.
[1565956663] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565956884] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1565957004] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1565957078] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2406476 kB) free!
[1565957124] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1565957234] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8055440 kB) free!
[1565957263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565957355] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11853016 kB) free.
[1565957383] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565957552] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.78, 6.08, 5.07
[1565957672] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.47, 6.25, 5.26
[1565957674] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1565957775] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.35, 5.38, 5.13
[1565957792] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.95, 6.40, 5.43
[1565957895] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.46, 5.11, 5.07
[1565957954] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8245260 kB) free!
[1565957983] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565958015] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.77, 4.69, 4.92
[1565958074] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11843996 kB) free.
[1565958103] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565958674] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8387044 kB) free!
[1565958703] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565958794] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6534928 kB) free!
[1565958823] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565958914] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11837628 kB) free.
[1565958942] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 602.43, 335.28, 138.14
[1565959062] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 612.82, 427.63, 195.88
[1565959077] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 588.59, 380.01, 168.16
[1565959182] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 614.38, 489.49, 246.65
[1565959197] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 591.46, 450.13, 219.55
[1565959274] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.45, 5.70, 5.42
[1565959317] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 501.95, 477.83, 258.45
[1565959394] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.78, 5.37, 5.33
[1565959424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565959514] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6724312 kB) free!
[1565959514] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.55, 5.38, 5.34
[1565959544] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565959634] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11844884 kB) free.
[1565959774] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 54 users currently logged in
[1565959891] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.5% (3789572 kB) free!
[1565959906] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 29.0% (7103740 kB) free.
[1565960011] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3485016 kB) free!
[1565960074] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1565960088] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 371 processes
[1565960125] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1565960131] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1927068 kB) free!
[1565960134] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1565960144] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565960194] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1565960208] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1565960234] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6776336 kB) free!
[1565960254] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1565960264] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565960354] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11849320 kB) free.
[1565960792] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.40, 4.29, 4.75
[1565960864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565960892] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9102264 kB) free!
[1565960954] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6959368 kB) free!
[1565960984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565961012] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8879316 kB) free!
[1565961074] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11830404 kB) free.
[1565961104] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565961132] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (8671780 kB) free!
[1565961154] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1565961325] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1565961408] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 4 crit, 2 warn out of 402 processes
[1565961445] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;SOFT;2;USERS CRITICAL - 13 users currently logged in
[1565961454] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1565961514] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1565961528] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 4 crit, 2 warn out of 395 processes
[1565961565] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1565961574] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1565961634] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1565961649] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 2 warn out of 398 processes
[1565961674] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7016320 kB) free!
[1565961704] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565961795] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11836948 kB) free.
[1565961825] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565961993] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.06, 7.03, 5.87
[1565962113] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.16, 7.10, 6.04
[1565962165] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1565962165] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1565962182] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 0.73, 13.00
[1565962199] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1565962233] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.30, 7.14, 6.17
[1565962239] SERVICE ALERT: sigqscluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565962317] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.64, 2.84, 13.03
[1565962394] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7513372 kB) free!
[1565962424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565962514] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11828024 kB) free.
[1565962544] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565962837] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.31, 4.85, 4.79
[1565963114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8013992 kB) free!
[1565963144] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565963234] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11837840 kB) free.
[1565963264] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565963382] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.55, 0.57, 3.99
[1565963449] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 374 processes
[1565963834] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8227624 kB) free!
[1565963864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565963954] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11836416 kB) free.
[1565963984] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565964117] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.32, 2.50, 4.10
[1565964133] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4206384 kB) free!
[1565964554] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8305248 kB) free!
[1565964584] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565964633] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.73, 4.43, 4.93
[1565964674] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6420264 kB) free!
[1565964705] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565964794] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11834072 kB) free.
[1565964995] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.1% (21601796 kB) free.
[1565965047] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.6% (22452768 kB) free.
[1565965081] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 56.3% (27736040 kB) free.
[1565965113] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 40.3% (19855680 kB) free.
[1565965155] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.9% (14280272 kB) free.
[1565965249] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 385 processes
[1565965304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565965369] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 6 crit, 0 warn out of 392 processes
[1565965394] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6539600 kB) free!
[1565965424] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565965490] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 3 warn out of 392 processes
[1565965514] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11822404 kB) free.
[1565965684] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.5% (5665328 kB) free!
[1565965763] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151216 kB) free!
[1565965804] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247908 kB) free!
[1565965833] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.40, 5.63, 5.04
[1565965883] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (183404 kB) free!
[1565965925] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (245424 kB) free!
[1565965953] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.55, 5.97, 5.23
[1565966004] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157608 kB) free!
[1565966024] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565966073] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.46, 5.92, 5.31
[1565966114] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.20, 4.06, 4.59
[1565966115] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6881780 kB) free!
[1565966144] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565966235] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11830572 kB) free.
[1565966264] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1565966264] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565966365] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1565966434] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1565966673] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.22, 4.46, 4.83
[1565966690] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1565966835] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6994892 kB) free!
[1565966864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565966955] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11833868 kB) free.
[1565966984] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565967015] SERVICE FLAPPING ALERT: webservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1565967555] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7216176 kB) free!
[1565967585] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565967675] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11823012 kB) free.
[1565967705] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565968252] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565968275] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7161940 kB) free!
[1565968305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565968395] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11831140 kB) free.
[1565968425] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565968995] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7514496 kB) free!
[1565969025] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565969115] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11820320 kB) free.
[1565969145] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565969715] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7750820 kB) free!
[1565969745] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565969835] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11829560 kB) free.
[1565969865] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565970436] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7857860 kB) free!
[1565970465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565970556] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11832600 kB) free.
[1565970586] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565971085] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 602 processes
[1565971146] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 605 processes
[1565971156] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7955244 kB) free!
[1565971185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565971276] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11825048 kB) free.
[1565971302] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.3% (14518476 kB) free.
[1565971305] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565971347] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 53.9% (26546456 kB) free.
[1565971765] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1565971794] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565971876] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7987980 kB) free!
[1565971905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565971914] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172200 kB) free!
[1565971953] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (7472152 kB) free!
[1565971996] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11827196 kB) free.
[1565972025] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565972034] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (162876 kB) free!
[1565972073] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234300 kB) free!
[1565972154] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161616 kB) free!
[1565972193] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248056 kB) free!
[1565972596] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8010640 kB) free!
[1565972625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565972716] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11818496 kB) free.
[1565972745] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565973316] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7979840 kB) free!
[1565973345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565973436] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11827156 kB) free.
[1565973466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565974036] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8063948 kB) free!
[1565974065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565974156] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11828796 kB) free.
[1565974185] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565974702] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4833836 kB) free!
[1565974757] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8134660 kB) free!
[1565974785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565974822] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4823780 kB) free!
[1565974877] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11820072 kB) free.
[1565974905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565974942] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4803716 kB) free!
[1565975477] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8173904 kB) free!
[1565975505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565975597] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11827864 kB) free.
[1565975625] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565976197] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8433492 kB) free!
[1565976225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565976317] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6817448 kB) free!
[1565976345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565976437] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11824648 kB) free.
[1565976466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565977037] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6955652 kB) free!
[1565977065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565977157] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11828024 kB) free.
[1565977185] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1565977490] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565977757] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7173344 kB) free!
[1565977785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565977877] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11819556 kB) free.
[1565977905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565978073] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565978365] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1565978477] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7313548 kB) free!
[1565978505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565978598] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11825324 kB) free.
[1565978625] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565979142] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.9% (5124172 kB) free.
[1565979197] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7910360 kB) free!
[1565979225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565979317] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11818564 kB) free.
[1565979346] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565979917] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8242048 kB) free!
[1565979945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565980037] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6541032 kB) free!
[1565980065] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565980157] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11825292 kB) free.