[1563647544] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (14687680 kB) free.
[1563647584] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563647626] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7123056 kB) free!
[1563647704] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563647746] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11417788 kB) free.
[1563648009] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196020 kB) free!
[1563648129] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198608 kB) free!
[1563648144] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (11502632 kB) free!
[1563648249] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204100 kB) free!
[1563648264] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (7782232 kB) free!
[1563648289] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173660 kB) free!
[1563648304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563648346] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7319844 kB) free!
[1563648384] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (4439348 kB) free!
[1563648409] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184768 kB) free!
[1563648424] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563648466] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11421472 kB) free.
[1563648529] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177160 kB) free!
[1563648984] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (14946264 kB) free.
[1563649024] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563649066] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7442732 kB) free!
[1563649144] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563649186] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11425328 kB) free.
[1563649584] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10863448 kB) free!
[1563649704] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (7501048 kB) free!
[1563649745] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563649786] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7203036 kB) free!
[1563649824] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (3924808 kB) free!
[1563649865] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563649906] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11421208 kB) free.
[1563650424] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (15160400 kB) free.
[1563650465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563650506] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6926924 kB) free!
[1563650585] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563650626] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11425000 kB) free.
[1563651024] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (12783608 kB) free!
[1563651144] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (9220004 kB) free!
[1563651185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563651226] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6614192 kB) free!
[1563651264] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (5918392 kB) free!
[1563651302] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (2258284 kB) free!
[1563651305] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563651346] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11414684 kB) free.
[1563651362] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (2252988 kB) free!
[1563651422] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 9.2% (2255700 kB) free!
[1563651482] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.2% (2257756 kB) free!
[1563651864] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (15157992 kB) free.
[1563651905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563651947] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6329436 kB) free!
[1563652025] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563652067] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11423408 kB) free.
[1563652367] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 360 processes
[1563652465] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13058832 kB) free!
[1563652487] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1563652585] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (9396336 kB) free!
[1563652626] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563652705] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (5474756 kB) free!
[1563652746] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563653087] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 359 processes
[1563653207] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1563653267] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7976452 kB) free!
[1563653305] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (15150292 kB) free.
[1563653346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563653387] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11423720 kB) free.
[1563653466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563653905] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (11185056 kB) free!
[1563653987] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7940996 kB) free!
[1563654025] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (7404152 kB) free!
[1563654066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563654107] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11420244 kB) free.
[1563654145] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (3588120 kB) free!
[1563654186] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563654707] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7996056 kB) free!
[1563654745] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (12863340 kB) free!
[1563654786] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563654827] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11423840 kB) free.
[1563654906] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563655427] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8110764 kB) free!
[1563655507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563655547] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6419580 kB) free!
[1563655627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563655668] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11420376 kB) free.
[1563655945] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (4534112 kB) free!
[1563656227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563656347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563656545] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (14621752 kB) free.
[1563656867] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8004152 kB) free!
[1563656947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563656987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11321932 kB) free.
[1563657067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563657145] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (10000432 kB) free!
[1563657265] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (6142196 kB) free!
[1563657385] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (15162332 kB) free.
[1563657587] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7895524 kB) free!
[1563657667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563657707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11422624 kB) free.
[1563657787] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563657985] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12386788 kB) free!
[1563658105] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (9008784 kB) free!
[1563658226] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (5487536 kB) free!
[1563658307] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7926156 kB) free!
[1563658388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563658427] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11419888 kB) free.
[1563658508] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563658826] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (15126296 kB) free.
[1563659027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7997828 kB) free!
[1563659108] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563659147] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11415732 kB) free.
[1563659228] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563659426] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12689548 kB) free!
[1563659546] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (9196112 kB) free!
[1563659666] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (5634676 kB) free!
[1563659747] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7805020 kB) free!
[1563659828] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563659867] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11421436 kB) free.
[1563659948] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563660266] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (15114004 kB) free.
[1563660467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7726264 kB) free!
[1563660548] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563660588] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11423608 kB) free.
[1563660668] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563660866] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (12445436 kB) free!
[1563660986] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (8762068 kB) free!
[1563661106] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (5097788 kB) free!
[1563661187] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7658124 kB) free!
[1563661269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563661307] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11422300 kB) free.
[1563661389] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563661706] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (14767552 kB) free.
[1563661907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7595208 kB) free!
[1563661989] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563662027] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11422476 kB) free.
[1563662109] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563662306] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10884156 kB) free!
[1563662426] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (7567380 kB) free!
[1563662546] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (4006824 kB) free!
[1563662627] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7456084 kB) free!
[1563662747] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11415616 kB) free.
[1563662890] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3126768 kB) free!
[1563663011] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (3122412 kB) free!
[1563663131] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (3128460 kB) free!
[1563663146] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (13759396 kB) free.
[1563663309] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563663347] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7274408 kB) free!
[1563663429] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563663467] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11423304 kB) free.
[1563663746] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (9435064 kB) free!
[1563663867] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (5980480 kB) free!
[1563663987] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (14985372 kB) free.
[1563664030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563664067] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7181320 kB) free!
[1563664150] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563664187] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11420888 kB) free.
[1563664587] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11866016 kB) free!
[1563664707] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (8091148 kB) free!
[1563664750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563664788] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6927820 kB) free!
[1563664827] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (4290332 kB) free!
[1563664870] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563664908] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11422020 kB) free.
[1563665427] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (12799344 kB) free!
[1563665470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563665508] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6860008 kB) free!
[1563665590] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563665628] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11420212 kB) free.
[1563666131] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1510476 kB) free!
[1563666190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563666228] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6595436 kB) free!
[1563666310] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563666348] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11131220 kB) free.
[1563666627] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (5822248 kB) free!
[1563666910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563666948] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9591600 kB) free!
[1563667031] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563667068] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8200248 kB) free!
[1563667151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563667188] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.2% (6486412 kB) free!
[1563667788] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10713048 kB) free.
[1563668388] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9140040 kB) free!
[1563668508] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7743064 kB) free!
[1563668629] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11379160 kB) free.
[1563669627] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (15372196 kB) free.
[1563669828] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8716176 kB) free!
[1563669948] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7262888 kB) free!
[1563670068] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11353608 kB) free.
[1563670827] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12682288 kB) free!
[1563670860] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1228.88, 747.36, 318.64
[1563670947] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (9288488 kB) free!
[1563670980] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1239.52, 910.75, 430.50
[1563671068] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (6008116 kB) free!
[1563671100] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 385.63, 760.77, 438.66
[1563671318] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2613368 kB) free!
[1563671668] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (13293472 kB) free.
[1563671951] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563672868] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (11785140 kB) free!
[1563672988] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (8176148 kB) free!
[1563673068] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7696624 kB) free!
[1563673108] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (4563640 kB) free!
[1563673118] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2466800 kB) free!
[1563673151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563673188] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11345728 kB) free.
[1563673271] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563673708] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (14851936 kB) free.
[1563673788] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7894500 kB) free!
[1563673871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563673909] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6159068 kB) free!
[1563673991] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563674029] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11343988 kB) free.
[1563674100] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.08, 0.69, 17.83
[1563674308] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (11320824 kB) free!
[1563674428] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (7946100 kB) free!
[1563674548] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (4577264 kB) free!
[1563674592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563674712] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563675148] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (14016840 kB) free.
[1563675229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7582016 kB) free!
[1563675312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563675349] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11335224 kB) free.
[1563675432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563675552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563675748] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (10458664 kB) free!
[1563675868] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (7144176 kB) free!
[1563675900] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.49, 0.55, 2.97
[1563675988] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (3957680 kB) free!
[1563676621] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175156 kB) free!
[1563676742] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (148168 kB) free!
[1563676752] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1563676752] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563676862] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166320 kB) free!
[1563677149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7854856 kB) free!
[1563677188] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (15055548 kB) free.
[1563677269] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11330364 kB) free.
[1563677869] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7920204 kB) free!
[1563677952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563677989] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7354816 kB) free!
[1563678073] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563678109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11329648 kB) free.
[1563678388] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (12990300 kB) free!
[1563678508] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.3% (10057828 kB) free!
[1563678628] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (7004500 kB) free!
[1563678672] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563678710] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6173488 kB) free!
[1563678792] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563678830] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11328192 kB) free.
[1563679228] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (4964936 kB) free!
[1563679392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563679430] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6329168 kB) free!
[1563679512] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563679550] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11319756 kB) free.
[1563680112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563680150] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6680688 kB) free!
[1563680232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563680270] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11325684 kB) free.
[1563680428] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (14647944 kB) free.
[1563680807] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 360 processes
[1563680832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563680870] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6759920 kB) free!
[1563680927] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1563680953] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563680990] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11328488 kB) free.
[1563681028] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (12064120 kB) free!
[1563681148] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (8793532 kB) free!
[1563681268] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (5508132 kB) free!
[1563681552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563681590] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7042296 kB) free!
[1563681672] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563681710] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11326480 kB) free.
[1563682171] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 127.12, 73.93, 30.68
[1563682272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563682291] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 19.39, 50.38, 27.30
[1563682310] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7334012 kB) free!
[1563682392] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563682411] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.45, 34.84, 24.41
[1563682430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11329696 kB) free.
[1563682468] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (14124332 kB) free.
[1563682992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563683011] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.66, 7.12, 14.15
[1563683031] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7262240 kB) free!
[1563683068] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11224160 kB) free!
[1563683112] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563683148] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244996 kB) free!
[1563683151] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11315588 kB) free.
[1563683188] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (7869476 kB) free!
[1563683268] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232484 kB) free!
[1563683308] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (4295704 kB) free!
[1563683388] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235204 kB) free!
[1563683908] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (14103408 kB) free.
[1563684509] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (9781008 kB) free!
[1563684629] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (6139572 kB) free!
[1563684749] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (2507912 kB) free!
[1563684811] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.15, 3.35, 4.81
[1563685349] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (8631948 kB) free!
[1563685512] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563685551] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9620864 kB) free!
[1563685632] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563685671] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8276620 kB) free!
[1563685752] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563685791] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.5% (6649468 kB) free!
[1563686391] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10701320 kB) free.
[1563686549] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (4246152 kB) free!
[1563686991] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9289848 kB) free!
[1563687111] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (8003564 kB) free!
[1563687125] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.08, 6.07, 5.01
[1563687232] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6268576 kB) free!
[1563687245] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.49, 5.78, 5.03
[1563687365] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.82, 5.53, 5.04
[1563687749] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (13601032 kB) free.
[1563687832] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563687832] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10903268 kB) free.
[1563688349] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (11051672 kB) free!
[1563688432] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9731028 kB) free!
[1563688469] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (6867888 kB) free!
[1563688552] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.1% (8433132 kB) free!
[1563688566] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.71, 4.70, 4.84
[1563688589] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2771924 kB) free!
[1563688672] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7055100 kB) free!
[1563688752] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563689166] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.75, 5.34, 5.13
[1563689189] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (12686640 kB) free!
[1563689272] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11121644 kB) free.
[1563689286] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.60, 5.81, 5.33
[1563689352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563689406] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.35, 6.12, 5.51
[1563689472] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563689592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563689789] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (13236204 kB) free.
[1563690395] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (9181188 kB) free!
[1563690515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (5054656 kB) free!
[1563690635] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (1165508 kB) free!
[1563691072] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6697452 kB) free!
[1563691192] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11034368 kB) free.
[1563691207] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1563691207] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.23, 4.24, 4.86
[1563691235] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (14571184 kB) free.
[1563691682] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 45.3% (11108260 kB) free.
[1563691792] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9484892 kB) free!
[1563691807] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.42, 5.88, 5.28
[1563691835] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (11042044 kB) free!
[1563691912] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8069980 kB) free!
[1563691927] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.33, 5.56, 5.23
[1563691955] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (6761528 kB) free!
[1563692033] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6364376 kB) free!
[1563692048] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.44, 5.59, 5.29
[1563692075] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (2518036 kB) free!
[1563692592] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563692633] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11280936 kB) free.
[1563692675] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (15034748 kB) free.
[1563692744] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.0% (14794820 kB) free.
[1563692819] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.1% (21244740 kB) free.
[1563692876] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.3% (21395264 kB) free.
[1563692929] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39754364 kB) free.
[1563693015] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.3% (13701488 kB) free.
[1563693079] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 81.1% (19869240 kB) free.
[1563693192] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563693233] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6174856 kB) free!
[1563693312] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563693353] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11272264 kB) free.
[1563693875] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (8909636 kB) free!
[1563693912] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563693995] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (4216432 kB) free!
[1563694032] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563694115] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (267788 kB) free!
[1563694553] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7580088 kB) free!
[1563694632] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563694673] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11280644 kB) free.
[1563694715] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (15145716 kB) free.
[1563694752] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563694872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563695048] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.89, 3.71, 4.66
[1563695315] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (11360128 kB) free!
[1563695435] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (6925676 kB) free!
[1563695556] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2712796 kB) free!
[1563695752] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (2442540 kB) free!
[1563695872] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.1% (2446652 kB) free!
[1563695992] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (2444416 kB) free!
[1563696156] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (14831560 kB) free.
[1563696592] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (188284 kB) free!
[1563696756] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13103736 kB) free!
[1563696876] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (8492200 kB) free!
[1563696900] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 13082.96, 5419.75, 2027.70
[1563696996] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (4126584 kB) free!
[1563697020] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14291.47, 8417.55, 3537.54
[1563697140] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14454.04, 10423.28, 4864.24
[1563697272] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563697448] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.33, 5.39, 5.19
[1563697568] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.91, 4.87, 5.03
[1563697673] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6979912 kB) free!
[1563697688] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.97, 4.64, 4.93
[1563697794] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11275684 kB) free.
[1563698196] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (14240828 kB) free.
[1563698393] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7421676 kB) free!
[1563698513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11274060 kB) free.
[1563698796] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10942668 kB) free!
[1563698916] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (6948844 kB) free!
[1563699036] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (3059548 kB) free!
[1563699072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563699113] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7064120 kB) free!
[1563699192] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563699233] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11271464 kB) free.
[1563699636] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (14681796 kB) free.
[1563699793] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563699833] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6866236 kB) free!
[1563699913] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563699953] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11265304 kB) free.
[1563700236] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12254744 kB) free!
[1563700356] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (8224448 kB) free!
[1563700476] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (4383908 kB) free!
[1563700513] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563700553] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6672520 kB) free!
[1563700633] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563700673] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11270840 kB) free.
[1563701076] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (10160916 kB) free!
[1563701233] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563701273] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6384236 kB) free!
[1563701327] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 362 processes
[1563701353] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563701393] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11271660 kB) free.
[1563701447] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1563701649] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 359 processes
[1563701677] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (13501444 kB) free.
[1563701769] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 360 processes
[1563701827] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1571420 kB) free!
[1563701889] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1563701947] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (338556 kB) free!
[1563702067] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (314964 kB) free!
[1563702277] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (9285504 kB) free!
[1563702397] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (5188724 kB) free!
[1563702489] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1563702517] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (1375904 kB) free!
[1563702553] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563702594] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9724892 kB) free!
[1563702663] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 627.39, 130.44, 42.58
[1563702674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563702714] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8318412 kB) free!
[1563702783] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 84.96, 87.38, 37.45
[1563702794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563702834] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6755848 kB) free!
[1563702903] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.59, 58.53, 32.94
[1563702952] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.8% (22157376 kB) free.
[1563703000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.5% (26630328 kB) free.
[1563703064] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.5% (14865404 kB) free.
[1563703140] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.70, 0.62, 19.70
[1563703503] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.33, 8.24, 17.49
[1563703743] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9599080 kB) free!
[1563703863] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9050292 kB) free!
[1563703983] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (8502428 kB) free!
[1563704157] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197924 kB) free!
[1563704277] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166416 kB) free!
[1563704288] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1563704398] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199212 kB) free!
[1563704634] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10686484 kB) free.
[1563704808] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12660936 kB) free!
[1563704928] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (8755752 kB) free!
[1563705048] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (20649064 kB) free.
[1563705194] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563705303] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.39, 2.84
[1563706223] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (3045040 kB) free!
[1563706248] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13102024 kB) free!
[1563706343] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (2990612 kB) free!
[1563706368] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (9409340 kB) free!
[1563706394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563706434] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6333160 kB) free!
[1563706463] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (2928304 kB) free!
[1563706488] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (20472408 kB) free.
[1563706515] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563706554] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11254112 kB) free.
[1563707114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563707234] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563707270] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.7% (1897384 kB) free!
[1563707390] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (1019280 kB) free!
[1563707510] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (218660 kB) free!
[1563707583] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (2433204 kB) free!
[1563707755] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7693408 kB) free!
[1563707834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563707875] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11263488 kB) free.
[1563707954] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563708289] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (9524352 kB) free!
[1563708409] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (20309416 kB) free.
[1563708475] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7527384 kB) free!
[1563708554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563708595] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11261628 kB) free.
[1563708674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563708794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563709247] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 361 processes
[1563709367] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1563709463] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1599284 kB) free!
[1563709609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12848312 kB) free!
[1563709729] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.9% (9123928 kB) free!
[1563709849] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (20188732 kB) free.
[1563709994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563710327] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1563710327] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 69.0% (17033876 kB) free.
[1563710369] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 67.9% (33473064 kB) free.
[1563710395] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6146320 kB) free!
[1563710431] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.9% (13805012 kB) free.
[1563710515] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11258012 kB) free.
[1563710523] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.3% (14842048 kB) free.
[1563710927] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (1362768 kB) free!
[1563711047] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (512212 kB) free!
[1563711049] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12674216 kB) free!
[1563711115] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6366336 kB) free!
[1563711167] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169704 kB) free!
[1563711169] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.5% (8893504 kB) free!
[1563711235] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11257944 kB) free.
[1563711289] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (19743704 kB) free.
[1563711794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563711835] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6493764 kB) free!
[1563711914] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563711955] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11248216 kB) free.
[1563712489] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (11535908 kB) free!
[1563712514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563712556] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6232864 kB) free!
[1563712609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (7464180 kB) free!
[1563712634] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563712676] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11255144 kB) free.
[1563712729] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (18535584 kB) free.
[1563713234] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563713276] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6077328 kB) free!
[1563713354] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563713396] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11258704 kB) free.
[1563713930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (10741216 kB) free!
[1563713954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563714050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (19962876 kB) free.
[1563714075] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563714596] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7510308 kB) free!
[1563714674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563714716] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11254420 kB) free.
[1563714794] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563715140] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.62, 1.45, 3.35
[1563715250] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563715316] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7442012 kB) free!
[1563715394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563715436] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11255048 kB) free.
[1563715514] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563715850] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (8664940 kB) free!
[1563715970] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (18967100 kB) free.
[1563716036] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7212628 kB) free!
[1563716156] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563716156] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11244112 kB) free.
[1563716714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563716757] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7078732 kB) free!
[1563716834] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563716877] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11250600 kB) free.
[1563717170] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (10732276 kB) free!
[1563717209] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.8% (21181064 kB) free.
[1563717290] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (19779996 kB) free.
[1563717435] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563717477] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7410624 kB) free!
[1563717555] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563717597] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11254424 kB) free.
[1563717675] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563718197] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7546224 kB) free!
[1563718275] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563718317] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11253056 kB) free.
[1563718327] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.1% (1807112 kB) free!
[1563718395] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563718447] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (1403224 kB) free!
[1563718490] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12527552 kB) free!
[1563718567] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1176248 kB) free!
[1563718610] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (9024560 kB) free!
[1563718730] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (19332188 kB) free.
[1563718917] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7442596 kB) free!
[1563718995] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563719037] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11255052 kB) free.
[1563719115] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563719637] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7253752 kB) free!
[1563719757] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11243420 kB) free.
[1563719930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (9911080 kB) free!
[1563720050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (19691932 kB) free.
[1563720316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563720357] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7443380 kB) free!
[1563720436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563720477] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11253172 kB) free.
[1563720556] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563721078] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7363296 kB) free!
[1563721198] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11253112 kB) free.
[1563721250] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (11179520 kB) free!
[1563721370] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (19669200 kB) free.
[1563721756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563721798] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7158628 kB) free!
[1563721876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563721918] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11251520 kB) free.
[1563722020] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563722476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563722518] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7186956 kB) free!
[1563722570] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11231536 kB) free!
[1563722596] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563722596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563722638] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11249264 kB) free.
[1563722690] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (19667740 kB) free.
[1563722940] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1006.30, 337.03, 120.03
[1563723060] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1255.49, 653.78, 262.41
[1563723180] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1288.93, 865.65, 387.51
[1563723197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563723219] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.4% (3799332 kB) free!
[1563723238] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7156628 kB) free!
[1563723317] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563723339] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 24.0% (5920608 kB) free.
[1563723358] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11242116 kB) free.
[1563723381] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1257 processes
[1563723501] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1256 processes
[1563723621] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1222 processes
[1563723890] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12202860 kB) free!
[1563723917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563723958] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7305612 kB) free!
[1563724010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (19655276 kB) free.
[1563724037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563724078] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11247648 kB) free.
[1563724157] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563724222] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 357 processes
[1563724757] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563724858] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.8% (14756756 kB) free.
[1563725810] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (9897604 kB) free!
[1563725879] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7255044 kB) free!
[1563725930] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (18585664 kB) free.
[1563725999] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11250580 kB) free.
[1563726558] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563726599] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7128852 kB) free!
[1563726678] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563726719] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11244592 kB) free.
[1563726780] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 0.49, 13.16
[1563727278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563727319] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7056288 kB) free!
[1563727398] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563727439] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11247508 kB) free.
[1563727731] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12587648 kB) free!
[1563727851] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (19722824 kB) free.
[1563727967] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 361 processes
[1563727980] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 0.28, 3.86
[1563727998] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563728039] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6745752 kB) free!
[1563728087] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1563728118] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563728159] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11234956 kB) free.
[1563728451] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (12970608 kB) free!
[1563728571] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (19741712 kB) free.
[1563728687] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 359 processes
[1563728718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563728759] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6322136 kB) free!
[1563728807] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1563728838] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563728879] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11243116 kB) free.
[1563729171] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (12966136 kB) free!
[1563729291] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (19761604 kB) free.
[1563729439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563729479] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6450408 kB) free!
[1563729559] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563729600] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11243984 kB) free.
[1563729891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12694392 kB) free!
[1563730011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (19775784 kB) free.
[1563730159] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563730199] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6693256 kB) free!
[1563730279] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1563730319] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11244700 kB) free.
[1563730611] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12576036 kB) free!
[1563730731] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (19786252 kB) free.
[1563730879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563730919] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7052928 kB) free!
[1563730999] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563731039] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11244472 kB) free.
[1563731599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563731639] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7268836 kB) free!
[1563731719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563731759] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11240256 kB) free.
[1563732211] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1335.31, 804.57, 342.93
[1563732320] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563732331] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1347.73, 984.93, 465.00
[1563732359] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7081568 kB) free!
[1563732440] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563732451] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1349.38, 1105.61, 572.27
[1563732479] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11242500 kB) free.
[1563732591] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9348456 kB) free!
[1563732672] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.5% (1711004 kB) free!
[1563732711] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9348180 kB) free!
[1563732792] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.5% (1713840 kB) free!
[1563732831] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (9346660 kB) free!
[1563732912] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (1710260 kB) free!
[1563733040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563733079] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7241996 kB) free!
[1563733131] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12227880 kB) free!
[1563733160] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563733199] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11245480 kB) free.
[1563733251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (19319216 kB) free.
[1563733501] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 4862 MB (9% inode=97%):
[1563733760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563733799] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7377956 kB) free!