[1546626877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546626997] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546627117] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546627405] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 189.98, 121.14, 48.00
[1546627525] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 25.93, 81.20, 42.23
[1546627645] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.02, 54.53, 37.18
[1546628245] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 7.78, 19.74
[1546630045] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.69, 3.44
[1546634845] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15.95, 73.13, 40.63
[1546634917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546634965] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.67, 49.15, 35.79
[1546635038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546635085] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1546635085] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.71, 33.01, 31.49
[1546635158] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546635685] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.58, 4.96, 16.80
[1546636109] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546637485] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.68, 3.00
[1546642285] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.17, 40.16, 36.79
[1546642405] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.62, 27.06, 32.40
[1546642526] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.80, 18.34, 28.55
[1546643126] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.94, 3.26, 15.39
[1546644326] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.72, 0.76, 4.77
[1546649726] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.26, 35.86, 35.45
[1546649846] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.70, 24.21, 31.23
[1546649966] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.71, 16.44, 27.53
[1546650566] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 2.82, 14.77
[1546651766] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.86, 0.89, 4.65
[1546657166] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.17, 43.42, 37.66
[1546657286] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.63, 29.27, 33.17
[1546657406] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.92, 21.80, 29.95
[1546658006] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.33, 3.42, 15.98
[1546659206] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.62, 0.73, 4.87
[1546664009] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1546664129] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1546664249] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1546664255] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1546664375] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1546664495] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1546664849] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1546664980] SERVICE ALERT: sig5;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546664987] SERVICE ALERT: sig5;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546665010] HOST ALERT: sig5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1546665100] SERVICE ALERT: sig5;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1546665104] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1546665107] SERVICE ALERT: sig5;Total Processes;OK;HARD;1;PROCS OK: 304 processes
[1546665111] SERVICE ALERT: sig3;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546665141] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1546665142] SERVICE ALERT: sig3;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig3' was out of bounds)
[1546665142] SERVICE ALERT: sig3;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig3' was out of bounds)
[1546665205] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1546665231] SERVICE ALERT: sig3;Memory;OK;HARD;1;OK - 94.4% (15346236 kB) free.
[1546665262] SERVICE ALERT: sig3;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37946 MB (74% inode=98%):
[1546665262] SERVICE ALERT: sig3;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1546665521] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 288 processes
[1546665641] SERVICE ALERT: sig5;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1546665806] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 68.40, 19.23, 7.06
[1546665926] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.60, 12.98, 6.24
[1546666046] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.76, 8.88, 5.56
[1546666646] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.57, 1.74, 3.21
[1546672958] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546673078] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546673516] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546673636] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546673756] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546673846] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.08, 35.40, 32.25
[1546673966] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.66, 23.85, 28.40
[1546674087] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.64, 16.18, 25.04
[1546674687] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.78, 2.70, 13.39
[1546675887] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.66, 0.61, 4.10
[1546677695] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1546677815] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1546677878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546677935] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1546677998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1546678567] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1546678627] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1546678687] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1546678747] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 41 users currently logged in
[1546679647] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 38 users currently logged in
[1546679756] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 13 users currently logged in
[1546680998] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546681118] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546681159] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 923 processes
[1546681238] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1546681279] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 355 processes
[1546681287] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.45, 38.67, 33.31
[1546681407] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.16, 26.19, 29.39
[1546681527] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.62, 17.71, 25.89
[1546682127] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 2.97, 13.92
[1546683327] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.72, 0.65, 4.29
[1546684152] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (639124 kB) free!
[1546684273] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (407904 kB) free!
[1546684394] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163552 kB) free!
[1546688127] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 25.33, 79.10, 41.08
[1546688247] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.45, 52.95, 36.11
[1546688367] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.30, 35.75, 31.85
[1546688967] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.87, 5.43, 17.07
[1546690709] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 359 processes
[1546690767] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.43, 0.60, 2.99
[1546691135] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1546691309] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 5 warn out of 361 processes
[1546692956] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546693076] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546693196] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546695567] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 6.52, 61.04, 38.42
[1546695687] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.15, 40.97, 33.81
[1546695807] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.62, 27.60, 29.78
[1546696196] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546696407] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.52, 4.72, 16.37
[1546697607] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.53, 0.69, 4.96
[1546698638] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546698758] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1546702109] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 362 processes
[1546703607] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.17, 39.53, 33.23
[1546703727] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.28, 26.49, 29.22
[1546703847] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.21, 17.78, 25.70
[1546703996] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546704116] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546704236] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546704447] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 2.56, 13.56
[1546705647] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.14, 0.23, 3.87
[1546706636] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546709246] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1546709366] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes