[1549737174] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.7% (2616480 kB) free!
[1549737227] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10303808 kB) free.
[1549737295] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (1977112 kB) free!
[1549737415] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1285240 kB) free!
[1549737429] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549737549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549737669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549737827] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7561508 kB) free!
[1549737914] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.2% (7017916 kB) free!
[1549737947] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9664728 kB) free!
[1549737977] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 390 processes
[1549738067] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.7% (8730420 kB) free!
[1549738270] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549738439] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 15.2% (2468464 kB) free!
[1549738497] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.87, 4.59, 4.81
[1549738514] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3330980 kB) free!
[1549738559] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 13.6% (2215192 kB) free!
[1549738577] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 396 processes
[1549738608] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45230 MB (20% inode=99%):
[1549738679] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 12.1% (1967948 kB) free!
[1549738728] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44178 MB (19% inode=99%):
[1549738848] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44131 MB (19% inode=99%):
[1549738870] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1549738990] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549739110] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1549739279] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (739100 kB) free!
[1549739596] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.15, 4.59, 5.02
[1549739697] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.89, 5.86, 5.45
[1549739710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549739714] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.6% (5246596 kB) free!
[1549739716] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.07, 3.73, 4.65
[1549739777] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1549739777] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1549739817] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.54, 5.77, 5.47
[1549739830] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549739937] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.91, 5.81, 5.52
[1549739950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549740226] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46485 MB (20% inode=99%):
[1549740314] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (290100 kB) free!
[1549740346] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46485 MB (20% inode=99%):
[1549740466] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46485 MB (20% inode=99%):
[1549741015] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 39.6% (9700724 kB) free.
[1549741066] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47895 MB (21% inode=99%):
[1549741233] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1549741248] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56712 MB (25% inode=99%):
[1549741293] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 51 users currently logged in
[1549741353] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 51 users currently logged in
[1549741413] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 51 users currently logged in
[1549741666] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43351 MB (19% inode=99%):
[1549741787] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43159 MB (19% inode=99%):
[1549741907] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43150 MB (19% inode=99%):
[1549742013] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1549742114] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.4% (12046292 kB) free.
[1549743914] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9216832 kB) free!
[1549744034] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.9% (5362892 kB) free!
[1549744137] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.87, 4.86, 4.93
[1549744154] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.4% (7598300 kB) free!
[1549744737] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.82, 5.25, 5.07
[1549744857] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.82, 5.11, 5.04
[1549744977] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.82, 5.00, 5.01
[1549745267] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4407836 kB) free!
[1549745350] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549745577] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.30, 4.38, 4.70
[1549745778] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 394 processes
[1549745815] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (4161724 kB) free!
[1549745867] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1549745867] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6054132 kB) free!
[1549745898] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 396 processes
[1549745935] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.9% (3172492 kB) free!
[1549745951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549746018] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 398 processes
[1549746055] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1984764 kB) free!
[1549746071] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549746191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549746554] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1760828 kB) free!
[1549746648] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38528 MB (17% inode=99%):
[1549746768] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 38535 MB (17% inode=99%):
[1549746888] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 38535 MB (17% inode=99%):
[1549748383] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 21.6% (3508352 kB) free.
[1549748591] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549748867] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3822416 kB) free!
[1549749191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549749312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1549749432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549750067] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5357272 kB) free!
[1549750218] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 388 processes
[1549750632] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549750818] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 389 processes
[1549751232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549751267] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2791120 kB) free!
[1549751352] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549751688] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24090 MB (10% inode=99%):
[1549751952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549752072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549752193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549753067] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5317708 kB) free!
[1549753667] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3585700 kB) free!
[1549756067] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.2% (5506712 kB) free!
[1549756154] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1549756977] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549757267] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3976676 kB) free!
[1549757867] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5052928 kB) free!
[1549759067] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3247372 kB) free!
[1549759393] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549762218] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549763577] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.96, 5.88, 5.48
[1549763697] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.36, 5.67, 5.45
[1549763817] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.24, 6.00, 5.60
[1549765154] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11711592 kB) free.
[1549765393] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1549765754] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8899584 kB) free!
[1549765874] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10518980 kB) free.
[1549766867] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5022420 kB) free!
[1549767074] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9643248 kB) free!
[1549767193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549767194] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10566400 kB) free.
[1549767313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549767433] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549767794] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8788304 kB) free!
[1549767915] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10567168 kB) free.
[1549768033] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549768067] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3420272 kB) free!
[1549768153] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549768667] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5085744 kB) free!
[1549768754] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549768874] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549769474] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549769594] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549769714] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7647160 kB) free!
[1549769834] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10278048 kB) free.
[1549769868] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3805960 kB) free!
[1549770119] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 310 processes
[1549770194] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549770239] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 308 processes
[1549770314] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549770359] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 307 processes
[1549770417] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.90, 4.94, 5.00
[1549770914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549771017] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.01, 5.12, 5.10
[1549771034] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.9% (6343460 kB) free!
[1549771034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549771137] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.76, 5.19, 5.11
[1549771154] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9022132 kB) free!
[1549771155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549771183] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3252072 kB) free!
[1549771257] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.65, 5.64, 5.28
[1549771274] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10471884 kB) free.
[1549771303] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.9% (3239252 kB) free!
[1549771423] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.9% (3226616 kB) free!
[1549771874] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.8% (6301596 kB) free!
[1549771994] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8895156 kB) free!
[1549772114] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10388348 kB) free.
[1549772355] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549772715] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.9% (7838376 kB) free!
[1549772835] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8851112 kB) free!
[1549772955] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10260372 kB) free.
[1549772955] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549773075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549773195] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549773555] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6196588 kB) free!
[1549773675] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8695388 kB) free!
[1549773795] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10174244 kB) free.
[1549774257] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1549774257] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.99, 4.88, 4.94
[1549774857] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.29, 5.35, 5.10
[1549774977] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.73, 5.21, 5.08
[1549774995] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1549774995] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549775097] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.51, 5.22, 5.10
[1549775596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549775716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549775759] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1549775836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549776195] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6223188 kB) free!
[1549776315] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8689852 kB) free!
[1549776359] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 312 processes
[1549776435] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9902816 kB) free.
[1549776479] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1549776656] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (2596356 kB) free!
[1549776916] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.64, 17.11, 11.94
[1549777036] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.7% (6271572 kB) free!
[1549777036] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.88, 12.06, 10.72
[1549777079] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 313 processes
[1549777156] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8811868 kB) free!
[1549777156] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.81, 8.66, 9.63
[1549777199] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 314 processes
[1549777256] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1852384 kB) free!
[1549777276] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9866316 kB) free!
[1549777319] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 316 processes
[1549777713] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 53 users currently logged in
[1549777773] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 53 users currently logged in
[1549777833] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 53 users currently logged in
[1549777893] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 53 users currently logged in
[1549778236] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549778476] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9903040 kB) free.
[1549778519] SERVICE FLAPPING ALERT: webservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1549778519] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1549778836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549778956] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.45, 2.29, 3.94
[1549778957] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549779120] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 314 processes
[1549779240] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 314 processes
[1549779360] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 313 processes
[1549779556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549779676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549779676] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9463388 kB) free!
[1549779796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549779796] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.0% (6428784 kB) free!
[1549779916] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 74.3% (36608652 kB) free.
[1549780016] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 52.8% (26026560 kB) free.
[1549780042] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.7% (23205544 kB) free.
[1549780105] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.1% (15292920 kB) free.
[1549780176] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.1% (14962692 kB) free.
[1549780400] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27129 MB (12% inode=99%):
[1549780524] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 92.6% (15049096 kB) free.
[1549781136] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (1503656 kB) free!
[1549781243] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 92.6% (15050364 kB) free.
[1549781507] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 66686 MB (29% inode=99%):
[1549781596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549781603] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1549781724] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1549781844] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1549782196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549782317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549782437] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549782444] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1549782483] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 601 processes
[1549782544] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 602 processes
[1549782604] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 603 processes
[1549782664] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 0 warn out of 603 processes
[1549782964] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 596 processes
[1549783239] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1549783560] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 317 processes
[1549783839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549783960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549784559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549784679] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549784760] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1549785024] SERVICE ALERT: sig2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3074132 kB) free!
[1549785144] SERVICE ALERT: sig2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3022372 kB) free!
[1549785264] SERVICE ALERT: sig2;Memory;WARNING;HARD;3;WARNING - 18.3% (2973556 kB) free!
[1549785279] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549785297] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.45, 4.80, 4.95
[1549785399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549785423] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1549785516] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 91.1% (14806480 kB) free.
[1549785519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549785619] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 406 processes
[1549785739] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 404 processes
[1549785859] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1549786119] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549786239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549786360] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549786593] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 61 users currently logged in
[1549786712] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 14.1% (2292268 kB) free!
[1549786832] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 12.7% (2056896 kB) free!
[1549786952] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 11.4% (1857712 kB) free!
[1549786959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549786998] SERVICE ALERT: sig4;Memory;WARNING;SOFT;1;WARNING - 19.9% (3238016 kB) free!
[1549787079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549787118] SERVICE ALERT: sig4;Memory;WARNING;SOFT;2;WARNING - 19.5% (3162564 kB) free!
[1549787199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549787238] SERVICE ALERT: sig4;Memory;WARNING;HARD;3;WARNING - 18.8% (3060316 kB) free!
[1549787275] SERVICE ALERT: sig3;Memory;WARNING;SOFT;1;WARNING - 18.0% (2926804 kB) free!
[1549787395] SERVICE ALERT: sig3;Memory;WARNING;SOFT;2;WARNING - 17.4% (2831360 kB) free!
[1549787466] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 87.0% (14135952 kB) free.
[1549787493] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 60 users currently logged in
[1549787516] SERVICE ALERT: sig3;Memory;WARNING;HARD;3;WARNING - 16.8% (2723392 kB) free!
[1549787552] SERVICE ALERT: sig3;Memory;OK;HARD;1;OK - 90.6% (14730724 kB) free.
[1549787618] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 90.7% (14741228 kB) free.
[1549787659] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 410 processes
[1549787756] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1549787799] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549787876] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1549787956] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.88, 5.19, 5.78
[1549787982] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (1482004 kB) free!
[1549787996] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1549788074] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 17.9% (2906180 kB) free!
[1549788076] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.56, 4.03, 5.29
[1549788102] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (600660 kB) free!
[1549788117] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1549788197] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.72, 3.25, 4.85
[1549788200] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;2;(Return code of 255 for service 'Memory' on host 'sig1' was out of bounds)
[1549788203] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1549788223] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (150384 kB) free!
[1549788238] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1549788267] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1549788320] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.7% (15384344 kB) free.
[1549788358] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1549788399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549788447] HOST ALERT: sig2;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1549788507] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1549788511] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1549788519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549788627] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1549788639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549788747] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1549788917] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9298904 kB) free!
[1549789037] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.7% (12163812 kB) free.
[1549789229] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23085 MB (10% inode=99%):
[1549789369] SERVICE ALERT: sig3;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1549789372] HOST ALERT: sig3;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.60)
[1549789380] SERVICE ALERT: sig3;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig3' was out of bounds)
[1549789380] SERVICE ALERT: sig3;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig3' was out of bounds)
[1549789436] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1549789459] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 399 processes
[1549789489] SERVICE ALERT: sig3;Memory;OK;HARD;1;OK - 94.6% (15383644 kB) free.
[1549789500] SERVICE ALERT: sig3;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 200351 MB (89% inode=99%):
[1549789500] SERVICE ALERT: sig3;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1549789579] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 396 processes
[1549789637] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.9% (7345280 kB) free!
[1549789700] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1549789757] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9438724 kB) free!
[1549789796] SERVICE ALERT: sig1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 311 processes
[1549789839] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549789877] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.9% (7837448 kB) free!
[1549790299] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 396 processes
[1549790419] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1549790439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549790477] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4474204 kB) free!
[1549790539] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 399 processes
[1549790559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549790679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549790769] SERVICE ALERT: sig2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1549790769] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'sig2' was out of bounds)
[1549790769] SERVICE ALERT: sig2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig2' was out of bounds)
[1549790769] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2' was out of bounds)
[1549790769] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1549790769] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig2' was out of bounds)
[1549790772] SERVICE ALERT: sig2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig2' was out of bounds)
[1549790889] SERVICE ALERT: sig2;Sensors Status;OK;HARD;1;SENSORS OK
[1549790890] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38723 MB (75% inode=98%):
[1549790890] SERVICE ALERT: sig2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1549790890] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1549790890] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 327 processes
[1549790893] SERVICE ALERT: sig2;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.13, 0.05
[1549790893] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1549791077] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.9% (7840564 kB) free!
[1549791223] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 32.3% (5252088 kB) free.
[1549791279] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549791297] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1549791358] SERVICE ALERT: sig2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 313 processes
[1549791369] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 315 processes
[1549791448] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3911108 kB) free!
[1549791489] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 312 processes
[1549791568] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (3707208 kB) free!
[1549791610] SERVICE ALERT: sig3;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 312 processes
[1549791624] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 298 processes
[1549791678] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1549791678] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4223788 kB) free!
[1549791688] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (3497456 kB) free!
[1549791744] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 299 processes
[1549791824] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (890156 kB) free!
[1549791864] SERVICE ALERT: sig4;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 298 processes
[1549791879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549791944] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169236 kB) free!
[1549791999] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549792064] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171312 kB) free!
[1549792120] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549792288] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2391372 kB) free!
[1549792464] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1549792558] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1549792810] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1549792939] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1549792939] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1549792996] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.84, 9.74, 5.69
[1549793099] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 661 processes
[1549793116] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.94, 11.45, 6.82
[1549793219] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 661 processes
[1549793236] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 11.08, 11.68, 7.49
[1549793339] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 405 processes
[1549793396] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1549793478] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.5% (6180960 kB) free!
[1549793920] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549794520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549794640] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1549794678] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2893372 kB) free!
[1549794760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549795064] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1549795064] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 28.8% (4673380 kB) free.
[1549795560] SERVICE FLAPPING ALERT: webservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1549795888] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 24.9% (6105164 kB) free.
[1549796193] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 61 users currently logged in
[1549796355] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1549796475] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1549796478] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.8% (5800960 kB) free!
[1549796595] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1549797078] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4842784 kB) free!
[1549797088] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (4256504 kB) free!
[1549797208] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (4112264 kB) free!
[1549797328] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1549797328] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (3823404 kB) free!
[1549797795] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1549798036] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.98, 3.45, 4.87
[1549798278] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 68.9% (33945740 kB) free.
[1549798528] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1665556 kB) free!
[1549798636] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.78, 4.09, 5.06
[1549798756] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.98, 4.89, 5.29
[1549798877] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.49, 3.97, 4.90
[1549798947] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1549799114] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.0% (23043048 kB) free.
[1549799539] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 406 processes
[1549799659] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 406 processes
[1549799780] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 406 processes
[1549800380] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1549800620] SERVICE ALERT: sig2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1549800623] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1549800743] SERVICE ALERT: sig2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig2' was out of bounds)
[1549800961] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig2' was out of bounds)
[1549801092] SERVICE ALERT: sig2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig2' was out of bounds)
[1549801092] SERVICE ALERT: sig2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig2' was out of bounds)
[1549801092] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2' was out of bounds)
[1549801092] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig2' was out of bounds)
[1549801092] SERVICE ALERT: sig2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig2' was out of bounds)
[1549801095] SERVICE ALERT: sig2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig2' was out of bounds)
[1549801125] SERVICE ALERT: sig2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig2' was out of bounds)
[1549801188] HOST ALERT: sig2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.150)
[1549801293] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 60 users currently logged in
[1549802193] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 61 users currently logged in
[1549802364] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27165 MB (12% inode=99%):
[1549804064] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1549806160] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549807198] SERVICE ALERT: smscapp1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1549807318] SERVICE ALERT: smscapp1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1549807593] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 58 users currently logged in
[1549808715] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177104 kB) free!
[1549808764] HOST ALERT: sig2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1549808835] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194008 kB) free!
[1549808881] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1549808955] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181720 kB) free!
[1549809012] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38723 MB (75% inode=98%):
[1549809012] SERVICE ALERT: sig2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1549809012] SERVICE ALERT: sig2;Sensors Status;OK;HARD;1;SENSORS OK
[1549809012] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 324 processes
[1549809012] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1549809015] SERVICE ALERT: sig2;CPU Load;OK;HARD;1;OK - load average: 0.01, 0.07, 0.05
[1549809045] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 94.5% (15355300 kB) free.
[1549809078] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1549809140] SERVICE ALERT: sig2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 192689 MB (85% inode=99%):
[1549810361] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549811780] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549811800] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.1% (23070408 kB) free.
[1549814806] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194880 kB) free!
[1549814926] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195428 kB) free!
[1549815046] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188536 kB) free!
[1549816361] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549816481] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549817913] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 83.7% (20527772 kB) free.
[1549817964] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24042 MB (10% inode=99%):
[1549818513] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189252 kB) free!
[1549818633] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190580 kB) free!
[1549818753] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189292 kB) free!
[1549819353] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 23.7% (5802888 kB) free.
[1549819954] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197068 kB) free!
[1549820074] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188340 kB) free!
[1549820194] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187820 kB) free!
[1549821994] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (4057468 kB) free!