[1573756290] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (9577708 kB) free!
[1573756601] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/11/14 23:57:59 [error] 24130#0: *6468794 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573756722] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573756890] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19559540 kB) free!
[1573757035] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179116 kB) free!
[1573757143] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185512 kB) free!
[1573757157] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187548 kB) free!
[1573757177] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202492 kB) free!
[1573757263] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183812 kB) free!
[1573757277] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175132 kB) free!
[1573757297] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189164 kB) free!
[1573757383] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188668 kB) free!
[1573757418] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171832 kB) free!
[1573757490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28862136 kB) free.
[1573758090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12355956 kB) free!
[1573758210] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (16717192 kB) free!
[1573758330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29861800 kB) free.
[1573758449] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191832 kB) free!
[1573758509] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 60.9% (14929108 kB) free.
[1573758578] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/15 00:35:02 [error] 23644#0: *2997346 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573758698] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/11/15 00:41:28 [error] 23642#0: *2998667 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573758700] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (186888 kB) free!
[1573758819] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573758930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13636464 kB) free!
[1573759050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10586428 kB) free!
[1573759171] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31612452 kB) free.
[1573759770] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15708364 kB) free!
[1573759890] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12626680 kB) free!
[1573760010] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (9643776 kB) free!
[1573760072] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9826128 kB) free!
[1573760192] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9486776 kB) free!
[1573760312] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (8928224 kB) free!
[1573760610] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16992136 kB) free!
[1573761065] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.9% (27596892 kB) free.
[1573761512] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2473152 kB) free!
[1573761810] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30803196 kB) free.
[1573762265] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.8% (26061292 kB) free!
[1573762385] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.9% (27563668 kB) free.
[1573762410] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13699708 kB) free!
[1573762530] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10275272 kB) free!
[1573762650] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30478260 kB) free.
[1573762985] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26364144 kB) free!
[1573763105] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.9% (27547124 kB) free.
[1573763250] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12847076 kB) free!
[1573763370] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31593580 kB) free.
[1573763970] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16317348 kB) free!
[1573764090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13051972 kB) free!
[1573764210] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (9975692 kB) free!
[1573764810] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17847972 kB) free!
[1573765185] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239004 kB) free!
[1573765305] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240108 kB) free!
[1573765426] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 88.7% (43721632 kB) free.
[1573765504] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217352 kB) free!
[1573765624] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233628 kB) free!
[1573765744] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224068 kB) free!
[1573766010] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12533680 kB) free!
[1573766344] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 94.0% (46315896 kB) free.
[1573766610] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16827644 kB) free!
[1573766678] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1573766798] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1573767628] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(196) < 49.50.86.152 [15/Nov/2019:03:10:27 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1573767666] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(219) < 180.179.244.32 [15/Nov/2019:03:10:14 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1573767749] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 49.50.86.152 [15/Nov/2019:03:10:28 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1573767786] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(36) < 203.212.70.170 [15/Nov/2019:03:11:22 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1573767810] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30921128 kB) free.
[1573767819] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(264) < 2019/11/15 03:11:22 [error] 23644#0: *3029997 upstream timed out (110: Connection timed out) while connecting to upstream, client: 121.241.242.109, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573767869] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573767907] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573767939] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573768034] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 62.2% (15337912 kB) free.
[1573768410] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14113404 kB) free!
[1573768530] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10814648 kB) free!
[1573768650] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30695212 kB) free.
[1573769250] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13064784 kB) free!
[1573769370] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31495396 kB) free.
[1573769970] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14805532 kB) free!
[1573770090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11281700 kB) free!
[1573770210] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30630788 kB) free.
[1573770810] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12596676 kB) free!
[1573770930] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31496092 kB) free.
[1573771530] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13861804 kB) free!
[1573771651] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31496592 kB) free.
[1573772250] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15254500 kB) free!
[1573772370] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11707468 kB) free!
[1573772490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30801856 kB) free.
[1573773090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12469456 kB) free!
[1573773210] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31505872 kB) free.
[1573773810] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14541504 kB) free!
[1573773930] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11240184 kB) free!
[1573774050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30882684 kB) free.
[1573774650] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14426976 kB) free!
[1573774770] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11116768 kB) free!
[1573774890] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30813396 kB) free.
[1573775490] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13847780 kB) free!
[1573775610] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31251280 kB) free.
[1573776210] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16047820 kB) free!
[1573776330] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12776980 kB) free!
[1573776451] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31459884 kB) free.
[1573777050] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15375028 kB) free!
[1573777170] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12200544 kB) free!
[1573777290] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31455812 kB) free.
[1573777890] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15372908 kB) free!
[1573778010] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12337044 kB) free!
[1573778130] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12960684 kB) free!
[1573778235] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177688 kB) free!
[1573778355] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184468 kB) free!
[1573778475] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192852 kB) free!
[1573778730] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18542228 kB) free!
[1573779330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28445724 kB) free.
[1573779477] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 81.9% (20079996 kB) free.
[1573779584] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 53.1% (13009016 kB) free.
[1573779620] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.9% (21308016 kB) free.
[1573779930] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12361640 kB) free!
[1573780050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9094020 kB) free!
[1573780170] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30539412 kB) free.
[1573780770] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13725532 kB) free!
[1573780890] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10316608 kB) free!
[1573781010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31063036 kB) free.
[1573781610] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13327292 kB) free!
[1573781731] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9699800 kB) free!
[1573781851] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30689244 kB) free.
[1573782451] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12883484 kB) free!
[1573782571] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9131912 kB) free!
[1573782691] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30362816 kB) free.
[1573783291] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12564532 kB) free!
[1573783411] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8696924 kB) free!
[1573783531] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4851148 kB) free!
[1573783734] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4185272 kB) free!
[1573783854] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4044916 kB) free!
[1573783974] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.4% (3915464 kB) free!
[1573784731] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20620396 kB) free!
[1573785331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28907916 kB) free.
[1573785774] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1433336 kB) free!
[1573785931] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6760600 kB) free!
[1573786051] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2584012 kB) free!
[1573786171] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27601832 kB) free.
[1573786771] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7440236 kB) free!
[1573786892] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3525856 kB) free!
[1573787012] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27734960 kB) free.
[1573787612] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8165316 kB) free!
[1573787732] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4729908 kB) free!
[1573787852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27340924 kB) free.
[1573788452] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10381552 kB) free!
[1573788572] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6889184 kB) free!
[1573788692] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3071080 kB) free!
[1573788774] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 28.4% (6387440 kB) free.
[1573789892] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17472920 kB) free!
[1573790492] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27176064 kB) free.
[1573791092] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5543120 kB) free!
[1573791212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1583904 kB) free!
[1573791332] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28312216 kB) free.
[1573791932] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5826004 kB) free!
[1573792052] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1462180 kB) free!
[1573792173] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411664 kB) free!
[1573792974] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (3782328 kB) free!
[1573793094] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3574160 kB) free!
[1573793214] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1573793214] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.7% (3297920 kB) free!
[1573793373] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22197664 kB) free!
[1573793386] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4626168 kB) free!
[1573793506] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (4460796 kB) free!
[1573793626] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (4291188 kB) free!
[1573793814] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1622528 kB) free!
[1573793973] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400904 kB) free!
[1573795077] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (4165072 kB) free!
[1573795173] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22864060 kB) free!
[1573795198] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3992832 kB) free!
[1573795318] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (3805292 kB) free!
[1573795426] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2107800 kB) free!
[1573795773] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1214480 kB) free!
[1573795918] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 25.2% (6176364 kB) free.
[1573796973] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24777216 kB) free!
[1573797118] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (4081080 kB) free!
[1573797238] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3900444 kB) free!
[1573797358] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (3724908 kB) free!
[1573797573] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2352568 kB) free!
[1573798173] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13383948 kB) free!
[1573798558] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1573798558] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1903420 kB) free!
[1573798773] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26480668 kB) free.
[1573799373] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4523016 kB) free!
[1573799493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30621904 kB) free.
[1573799815] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 30.0% (6738064 kB) free.
[1573800093] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10568256 kB) free!
[1573800214] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6210160 kB) free!
[1573800334] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2074824 kB) free!
[1573801169] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.2% (19802312 kB) free.
[1573801232] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.4% (22406952 kB) free.
[1573801320] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.8% (57772788 kB) free.
[1573801369] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.0% (37946828 kB) free.
[1573801426] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.3% (15753616 kB) free.
[1573801609] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.1% (14483944 kB) free.
[1573801773] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181328 kB) free!
[1573801893] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175232 kB) free!
[1573801975] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9093676 kB) free!
[1573802013] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178372 kB) free!
[1573802095] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9090156 kB) free!
[1573802215] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (9088716 kB) free!
[1573803020] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4700392 kB) free!
[1573803124] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17518992 kB) free!
[1573803140] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4355284 kB) free!
[1573803244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28629876 kB) free.
[1573803260] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (4225484 kB) free!
[1573803544] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8851480 kB) free!
[1573803664] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8560908 kB) free!
[1573803784] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8269920 kB) free!
[1573803845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4318256 kB) free!
[1573803965] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401440 kB) free!
[1573804015] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (3916012 kB) free!
[1573804085] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32923940 kB) free.
[1573804135] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.2% (3646928 kB) free!
[1573804255] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.0% (3365564 kB) free!
[1573804685] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10749276 kB) free!
[1573804805] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6451440 kB) free!
[1573804855] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1898588 kB) free!
[1573804925] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2430028 kB) free!
[1573805060] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2410992 kB) free!
[1573805215] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3313292 kB) free!
[1573805525] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14550864 kB) free!
[1573805584] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4599388 kB) free!
[1573806067] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.3844.610> (10.147.212.158:50720 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573806125] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27924392 kB) free.
[1573806187] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573806415] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1573806415] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.2% (40001992 kB) free.
[1573806463] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4612244 kB) free!
[1573806479] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1573806479] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2202304 kB) free!
[1573806493] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.9% (22523416 kB) free.
[1573806508] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.4% (47027144 kB) free.
[1573806523] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 29.4% (6610756 kB) free.
[1573806590] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.4% (22552572 kB) free.
[1573806725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5923572 kB) free!
[1573806845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1845464 kB) free!
[1573806965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31028508 kB) free.
[1573807463] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4653760 kB) free!
[1573807565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8625612 kB) free!
[1573807583] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4535084 kB) free!
[1573807685] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4652984 kB) free!
[1573807703] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4424052 kB) free!
[1573807793] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165256 kB) free!
[1573807805] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (626840 kB) free!
[1573807913] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181900 kB) free!
[1573808033] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164120 kB) free!
[1573808051] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1573808171] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1573808342] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4175404 kB) free!
[1573808462] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (3995048 kB) free!
[1573808582] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.9% (3799464 kB) free!
[1573809006] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25814600 kB) free!
[1573809606] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3521268 kB) free!
[1573809783] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1875736 kB) free!
[1573810103] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1868132 kB) free!
[1573810206] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16134972 kB) free!
[1573810232] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573810383] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 30.0% (6746348 kB) free.
[1573810678] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.3% (21060120 kB) free.
[1573810689] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1573810767] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1573810767] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 62.4% (15293144 kB) free.
[1573810806] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32350256 kB) free.
[1573810809] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 993.98, 445.04, 171.47
[1573810809] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1573810929] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1057.56, 651.03, 280.13
[1573810929] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 400 processes
[1573810998] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1573811049] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1066.42, 788.89, 375.62
[1573811118] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 381 processes
[1573811238] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 436 processes
[1573811281] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188660 kB) free!
[1573811321] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(513) < 2019/11/15 15:18:04 [error] 24132#0: *6638060 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573811402] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179752 kB) free!
[1573811406] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9160420 kB) free!
[1573811441] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573811522] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189796 kB) free!
[1573811526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (5070372 kB) free!
[1573811530] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1573811646] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1201224 kB) free!
[1573811649] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.02, 253.77, 339.71
[1573812246] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13866188 kB) free!
[1573812939] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/15 15:44:51 [error] 23643#0: *3354347 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /favicon.ico HTTP/1.1", host: "148qservercluster1.com:8091", referrer: "http://148qservercluster1.com:8091/api/index.html"
[1573812971] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1573813059] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573813091] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1573813446] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1491876 kB) free!
[1573813998] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4798388 kB) free!
[1573814046] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16249184 kB) free!
[1573814119] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.5% (4779968 kB) free!
[1573814239] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.4% (4750252 kB) free!
[1573814646] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29872480 kB) free.
[1573814815] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9412388 kB) free!
[1573814935] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9412200 kB) free!
[1573815032] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4714872 kB) free!
[1573815055] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (9412136 kB) free!
[1573815152] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4634216 kB) free!
[1573815183] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4434376 kB) free!
[1573815246] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5543220 kB) free!
[1573815272] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4515716 kB) free!
[1573815303] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4232816 kB) free!
[1573815366] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1183380 kB) free!
[1573815423] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.8% (3991184 kB) free!
[1573815486] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34633384 kB) free.
[1573816087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11235300 kB) free!
[1573816207] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6679584 kB) free!
[1573816255] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250320 kB) free!
[1573816327] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1907056 kB) free!
[1573816623] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1998592 kB) free!
[1573816771] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 232 processes
[1573816841] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/15 16:45:17 [error] 24131#0: *6654847 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573816891] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 229 processes
[1573816927] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15243300 kB) free!
[1573816961] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573817527] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28267116 kB) free.
[1573817562] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/15 17:00:24 [error] 24129#0: *6657720 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573817625] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9603032 kB) free!
[1573817672] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2330076 kB) free!
[1573817682] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573817746] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9517824 kB) free!
[1573817859] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/15 17:03:53 [error] 23638#0: *3418474 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573817866] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (9430164 kB) free!
[1573817979] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573818127] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5197092 kB) free!
[1573818247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1034080 kB) free!
[1573818367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32452372 kB) free.
[1573818707] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21659620 kB) free.
[1573818761] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.9% (22760372 kB) free.
[1573818846] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 29.8% (6699384 kB) free.
[1573818885] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.2% (40003376 kB) free.
[1573818967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7075000 kB) free!
[1573819087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2444204 kB) free!
[1573819207] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381040 kB) free!
[1573819679] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(3) < Channel error on connection <0.28398.7496> (10.147.212.191:40976 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1573819800] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573819891] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 235 processes
[1573819913] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193288 kB) free!
[1573820011] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 236 processes
[1573820033] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198532 kB) free!
[1573820131] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 240 processes
[1573820153] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191696 kB) free!
[1573820375] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4827076 kB) free!
[1573820400] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.765.7418> (10.147.212.191:40990 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1573820407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26459700 kB) free.
[1573820495] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4721536 kB) free!
[1573820522] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573820615] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4583124 kB) free!
[1573820650] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.0% (3378412 kB) free!
[1573820693] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216008 kB) free!
[1573820770] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3063044 kB) free!
[1573820813] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247976 kB) free!
[1573820890] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.3% (2751276 kB) free!
[1573820933] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227312 kB) free!
[1573821007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1722628 kB) free!
[1573821121] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < Channel error on connection <0.31665.7466> (10.147.212.191:41038 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1573821127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29757448 kB) free.
[1573821243] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573821490] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 29.8% (6685620 kB) free.
[1573821727] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13182260 kB) free!
[1573821842] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.7921.7522> (10.147.212.191:41010 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1573821848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9105400 kB) free!
[1573821964] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(1) < Channel error on connection <0.28216.7512> (10.147.212.191:41040 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1573821968] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5148344 kB) free!
[1573822085] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573822138] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 79.1% (19527396 kB) free.
[1573822225] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 50.1% (66099016 kB) free.
[1573822266] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.2% (40008916 kB) free.
[1573822304] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 65.7% (16095016 kB) free.
[1573822358] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 92.4% (22642640 kB) free.
[1573822443] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 95.7% (47164952 kB) free.
[1573822777] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183196 kB) free!
[1573822897] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170792 kB) free!
[1573823017] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189624 kB) free!
[1573823518] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179776 kB) free!
[1573823561] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196896 kB) free!
[1573823638] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181780 kB) free!
[1573823681] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192864 kB) free!
[1573823758] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173116 kB) free!
[1573823804] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1573823804] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199048 kB) free!
[1573823900] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (209680 kB) free!
[1573824021] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198276 kB) free!
[1573824141] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187448 kB) free!
[1573824415] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (186480 kB) free!
[1573824535] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (184816 kB) free!
[1573824580] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/15 18:49:58 [error] 23646#0: *3513322 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/DlrLoad/msisdn"
[1573824635] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25656256 kB) free!
[1573824655] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (184264 kB) free!
[1573824700] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573824755] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21400276 kB) free!
[1573824875] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17484456 kB) free!
[1573825475] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29931352 kB) free.
[1573825482] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/15 19:10:21 [error] 24132#0: *6682504 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573825531] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 227 processes
[1573825602] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573826473] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9138112 kB) free!
[1573826593] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8547032 kB) free!
[1573826675] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19825908 kB) free!
[1573826713] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (8546148 kB) free!
[1573826795] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (48048752 kB) free.
[1573826890] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3978268 kB) free!
[1573827010] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (3871648 kB) free!
[1573827130] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.6% (3735328 kB) free!
[1573827238] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 70.3% (17345624 kB) free.
[1573827267] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.2% (22853416 kB) free.
[1573827296] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 65.9% (16141164 kB) free.
[1573827328] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.5% (22921408 kB) free.
[1573827362] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.1% (14477036 kB) free.
[1573827398] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1573827402] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/15 19:45:13 [error] 24133#0: *6689104 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573827518] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1573827523] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573827845] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194148 kB) free!
[1573827965] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187796 kB) free!
[1573828085] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178964 kB) free!
[1573828595] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14869320 kB) free!
[1573828715] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27173688 kB) free.
[1573828930] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1817108 kB) free!
[1573829315] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5968376 kB) free!
[1573829322] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/15 20:15:20 [error] 24130#0: *6694806 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573829435] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1899664 kB) free!
[1573829442] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573829555] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30513892 kB) free.
[1573830042] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/15 20:30:10 [error] 24132#0: *6697575 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573830155] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8930492 kB) free!
[1573830162] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573830275] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4969776 kB) free!
[1573830395] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1446876 kB) free!
[1573831595] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24191500 kB) free!
[1573831930] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 29.5% (6613812 kB) free.
[1573832196] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2293368 kB) free!
[1573832316] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9281424 kB) free!
[1573832436] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9174048 kB) free!
[1573832556] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9069108 kB) free!
[1573832762] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (159968 kB) free!
[1573832796] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16246112 kB) free!
[1573832882] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (157776 kB) free!
[1573833002] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (147860 kB) free!
[1573833101] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/15 21:14:27 [error] 23639#0: *3595003 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573833221] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573833313] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1573833396] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27530400 kB) free.
[1573833821] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/15 21:30:25 [error] 23639#0: *3603658 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:14001, upstream: "10.147.212.38:5001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573833871] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 103.255.146.164 [15/Nov/2019:21:30:25 +0530] TCP 502 0 0 0.001 "10.147.212.37:5001, 10.147.212.49:5001, 10.147.212.38:5001" "0, 0, 0" "0, 0, 0" "0.001, 0.000, 0.000"
[1573833913] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4750048 kB) free!
[1573833941] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573833992] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573833996] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7311696 kB) free!
[1573834116] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3598148 kB) free!
[1573834236] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30333224 kB) free.
[1573834475] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1573834506] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1573834836] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10439452 kB) free!
[1573834956] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6796960 kB) free!
[1573835076] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29651956 kB) free.
[1573835676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12730252 kB) free!
[1573835796] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8811276 kB) free!
[1573835916] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5090312 kB) free!
[1573836288] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.4% (21328192 kB) free.
[1573836319] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.6% (35298404 kB) free.
[1573836346] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.4% (14523472 kB) free.
[1573836378] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40172124 kB) free.
[1573837431] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 297 processes
[1573837494] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.5% (3823660 kB) free!
[1573837551] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1573837614] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192536 kB) free!
[1573837734] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197744 kB) free!
[1573838531] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4344564 kB) free!
[1573838651] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4238392 kB) free!
[1573838771] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.5% (4145224 kB) free!
[1573839756] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4790596 kB) free!
[1573839905] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24436320 kB) free!
[1573840025] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.8% (26132760 kB) free!
[1573840145] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.2% (21416492 kB) free!
[1573842341] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(49) < 2019/11/15 23:55:36 [error] 23648#0: *3675228 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573842384] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25523932 kB) free!
[1573842461] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(231) < 2019/11/15 23:57:39 [error] 23644#0: *3676284 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573842504] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.0% (48818952 kB) free.
[1573842581] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found