[1576693860] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26346136 kB) free!
[1576693980] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.7% (23305060 kB) free!
[1576694100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42900668 kB) free.
[1576694211] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576694331] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576694415] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176208 kB) free!
[1576694466] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187180 kB) free!
[1576694535] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (213024 kB) free!
[1576694586] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175640 kB) free!
[1576694655] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (339240 kB) free!
[1576694655] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204436 kB) free!
[1576694700] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25949372 kB) free!
[1576694706] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186832 kB) free!
[1576694746] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180196 kB) free!
[1576694775] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (335172 kB) free!
[1576694821] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.6% (23153960 kB) free!
[1576694866] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172692 kB) free!
[1576694895] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (329844 kB) free!
[1576694941] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42794424 kB) free.
[1576694986] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165992 kB) free!
[1576695541] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25720496 kB) free!
[1576695661] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44908224 kB) free.
[1576697915] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9264460 kB) free!
[1576698035] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8914920 kB) free!
[1576698155] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (8623924 kB) free!
[1576698661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25334332 kB) free!
[1576698755] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4058904 kB) free!
[1576698781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44268276 kB) free.
[1576701781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25766964 kB) free!
[1576701901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43419948 kB) free.
[1576702731] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1576702746] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240328 kB) free!
[1576702851] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1576702866] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (218164 kB) free!
[1576702986] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240484 kB) free!
[1576703451] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576703572] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576703586] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.1% (43892740 kB) free.
[1576703656] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.7% (10530408 kB) free.
[1576705227] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/12/19 03:10:27 [error] 28867#0: *23164556 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1576705348] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2019/12/19 03:10:29 [error] 28875#0: *23164673 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1576705468] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576705469] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576705501] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26158408 kB) free!
[1576705621] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43044820 kB) free.
[1576706221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26236308 kB) free!
[1576706341] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43532364 kB) free.
[1576710771] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576710891] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576711141] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23973964 kB) free!
[1576711261] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44570208 kB) free.
[1576715656] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (150916 kB) free!
[1576715776] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (783192 kB) free!
[1576715897] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189720 kB) free!
[1576716306] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.9% (21050312 kB) free.
[1576716495] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.0% (21072588 kB) free.
[1576716500] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1576716586] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.1% (21844636 kB) free.
[1576720491] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576720611] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576727068] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2019/12/19 09:06:12 [error] 28874#0: *23249201 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576727189] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576731555] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.16247.1608> (10.147.212.158:40958 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1576731675] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1576732103] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 1584 processes
[1576732224] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 1572 processes
[1576734178] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.1% (22477516 kB) free.
[1576734221] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.8% (38827204 kB) free.
[1576735700] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1576735820] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1576735940] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1576736669] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2019/12/19 11:54:02 [error] 20759#0: *18301456 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576736789] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2019/12/19 11:55:45 [error] 20756#0: *18301728 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576736909] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2019/12/19 11:56:36 [error] 20756#0: *18301866 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576737787] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.5% (3325088 kB) free!
[1576737907] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.2% (3745212 kB) free!
[1576738027] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (3689136 kB) free!
[1576738262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26189384 kB) free!
[1576738382] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.0% (48722628 kB) free.
[1576738628] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2042112 kB) free!
[1576738982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25101472 kB) free!
[1576739102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.0% (47441896 kB) free.
[1576739437] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1576739437] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22327168 kB) free.
[1576739500] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576739505] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/19 12:41:41 [error] 20756#0: *18309079 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576739513] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576739702] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23865432 kB) free!
[1576739822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45728868 kB) free.
[1576740113] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2019/12/19 12:51:36 [error] 20758#0: *18310653 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576740233] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576740422] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22105068 kB) free!
[1576740542] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44605880 kB) free.
[1576740833] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2019/12/19 13:01:36 [error] 20756#0: *18312261 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576740907] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4605512 kB) free!
[1576740953] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(5) < 2019/12/19 13:05:46 [error] 20753#0: *18312927 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576741027] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4438456 kB) free!
[1576741073] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576741142] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20644584 kB) free!
[1576741147] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (4269736 kB) free!
[1576741262] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40119268 kB) free.
[1576741674] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2019/12/19 13:15:45 [error] 20757#0: *18314507 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576741794] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/19 13:19:02 [error] 20756#0: *18315026 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576741914] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2019/12/19 13:20:46 [error] 20756#0: *18315295 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576742295] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4785200 kB) free!
[1576742415] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4628944 kB) free!
[1576742535] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4507188 kB) free!
[1576742670] SERVICE ALERT: webservercluster1;Total Processes;OK;HARD;1;PROCS OK: 515 processes
[1576742947] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1839084 kB) free!
[1576744335] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2203240 kB) free!
[1576745386] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4355332 kB) free!
[1576745507] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4247892 kB) free!
[1576745627] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (4127132 kB) free!
[1576747588] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(68) < 2019/12/19 14:53:18 [error] 28870#0: *23390066 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.27.86.30, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1576747708] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576748462] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21900268 kB) free!
[1576748582] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (47615072 kB) free.
[1576748627] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2089032 kB) free!
[1576748908] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2019/12/19 15:13:50 [error] 28866#0: *23400472 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:462/17, bytes from/to upstream:17/462
[1576749029] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576749182] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23263888 kB) free!
[1576749186] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9404260 kB) free!
[1576749302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34769660 kB) free.
[1576749306] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9239636 kB) free!
[1576749426] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9081876 kB) free!
[1576749623] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1576749628] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/19 15:26:58 [error] 28860#0: *23407415 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:153/17, bytes from/to upstream:17/153
[1576749649] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.3% (3274184 kB) free!
[1576749743] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1576749748] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576749769] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (2475756 kB) free!
[1576749838] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 75.5% (18493688 kB) free.
[1576749881] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576749881] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22861088 kB) free.
[1576749902] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24806204 kB) free!
[1576749921] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 75.5% (18494216 kB) free.
[1576749959] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.6% (20250460 kB) free.
[1576750022] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.0% (21032496 kB) free!
[1576750142] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (46062388 kB) free.
[1576750348] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/19 15:39:01 [error] 28871#0: *23413807 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1576750468] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576750742] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21300192 kB) free!
[1576750863] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.4% (45310016 kB) free.
[1576751462] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23598372 kB) free!
[1576751582] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19698448 kB) free!
[1576751628] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.8% (5839028 kB) free!
[1576751668] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/19 16:04:09 [error] 28871#0: *23427020 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:233/17, bytes from/to upstream:17/233
[1576751687] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176156 kB) free!
[1576751702] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44554204 kB) free.
[1576751748] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (2371968 kB) free!
[1576751789] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576751807] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (204068 kB) free!
[1576751869] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215528 kB) free!
[1576751928] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181148 kB) free!
[1576752302] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20661216 kB) free!
[1576752422] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17110104 kB) free!
[1576752542] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42037252 kB) free.
[1576753026] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4904796 kB) free!
[1576753142] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18435672 kB) free!
[1576753262] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42858260 kB) free.
[1576753355] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(102) < 203.212.70.170 [19/Dec/2019:16:29:01 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1576753475] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576753862] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18045124 kB) free!
[1576753982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14252928 kB) free!
[1576754102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41442112 kB) free.
[1576754702] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16588024 kB) free!
[1576754822] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12794944 kB) free!
[1576754942] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37235192 kB) free.
[1576755143] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 337 processes
[1576755263] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1576755543] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13235108 kB) free!
[1576755663] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.9% (52661608 kB) free.
[1576756407] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9356020 kB) free!
[1576756527] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (9225420 kB) free!
[1576756647] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (9102992 kB) free!
[1576760050] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (2428464 kB) free!
[1576760125] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167832 kB) free!
[1576760170] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (2298608 kB) free!
[1576760211] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576760245] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166356 kB) free!
[1576760288] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.5% (20854168 kB) free.
[1576760290] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2179424 kB) free!
[1576760331] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1576760333] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.0% (22554868 kB) free.
[1576760365] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183824 kB) free!
[1576760426] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 55.9% (27531724 kB) free.
[1576760463] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21756008 kB) free!
[1576760466] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.2% (14894600 kB) free.
[1576760504] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.2% (21121352 kB) free.
[1576760583] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (47573640 kB) free.
[1576760764] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188156 kB) free!
[1576760884] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187528 kB) free!
[1576760898] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174052 kB) free!
[1576760938] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182404 kB) free!
[1576761004] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175880 kB) free!
[1576761018] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175464 kB) free!
[1576761059] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169352 kB) free!
[1576761104] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184272 kB) free!
[1576761138] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177836 kB) free!
[1576761179] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177040 kB) free!
[1576761183] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25439436 kB) free!
[1576761225] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190944 kB) free!
[1576761303] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21373728 kB) free!
[1576761345] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191216 kB) free!
[1576761423] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46779504 kB) free.
[1576761531] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576761651] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1576762023] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16165180 kB) free!
[1576762047] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4865720 kB) free!
[1576762144] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12380336 kB) free!
[1576762264] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9446516 kB) free!
[1576762864] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18033320 kB) free!
[1576763464] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29167400 kB) free.
[1576763990] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22661924 kB) free.
[1576764027] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1576764027] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.8% (22758024 kB) free.
[1576764064] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8695592 kB) free!
[1576764068] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 38.3% (18887076 kB) free.
[1576764109] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 49.8% (65659948 kB) free.
[1576764154] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1576764154] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.2% (22843456 kB) free.
[1576764207] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.4% (23132556 kB) free.
[1576769558] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4537368 kB) free!
[1576769616] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180160 kB) free!
[1576769678] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191672 kB) free!
[1576769736] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169624 kB) free!
[1576769798] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165364 kB) free!
[1576769857] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194792 kB) free!
[1576770001] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189280 kB) free!
[1576770037] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179248 kB) free!
[1576770121] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175396 kB) free!
[1576770157] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192512 kB) free!
[1576770241] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183212 kB) free!
[1576770277] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186748 kB) free!
[1576771280] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1576771280] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.3% (22529432 kB) free.
[1576771310] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.2% (22599864 kB) free.
[1576771343] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.7% (23201000 kB) free.
[1576771380] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.6% (22942252 kB) free.
[1576777887] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183888 kB) free!
[1576777916] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169280 kB) free!
[1576777950] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182828 kB) free!
[1576777986] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181104 kB) free!
[1576778007] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194260 kB) free!
[1576778036] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188520 kB) free!
[1576778071] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168048 kB) free!
[1576778106] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192736 kB) free!
[1576778127] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191872 kB) free!
[1576778156] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181036 kB) free!
[1576778191] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164044 kB) free!
[1576778226] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174640 kB) free!