[1551119442] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551120042] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551120163] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551120762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551120882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551121482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551121602] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551122202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551122322] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551122922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551123043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551123642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551123762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551123882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551124482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551124602] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551125202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551125322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551125443] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551126042] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551126162] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551126762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551126882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551127482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551127602] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551128202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551128322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551128443] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551129042] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551129162] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551129762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551129882] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1551129882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551130482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551130602] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551131202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551131323] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551131443] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551132043] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551132163] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551132763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551132883] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551133483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551133603] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551134204] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551134324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551134444] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551135044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551135164] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551135764] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551135884] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551136484] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551136605] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551137204] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551137324] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551137924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551138044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551138644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551138764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551139364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551139485] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551139605] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551140205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551140325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551140925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551141045] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551141645] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551141765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551141885] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551142486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551142606] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551143206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551143326] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551143410] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4812404 kB) free!
[1551143470] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4816132 kB) free!
[1551143530] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.7% (4817768 kB) free!
[1551143590] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.7% (4818104 kB) free!
[1551143926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551144046] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551144646] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551144766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551144887] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551145486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551145606] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551146206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551146326] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551146926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551147046] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551147646] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551147766] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1551148367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551148487] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551149087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551149207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551149327] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551149927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551150047] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551150647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551150768] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551151367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551151487] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551152087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551152207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551152327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551155327] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551155927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551156047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551156167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551157967] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1551157968] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551158366] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1567460 kB) free!
[1551158564] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 330 processes
[1551158567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551158687] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551158807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551159649] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.1% (2612232 kB) free!
[1551159769] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.2% (2471864 kB) free!
[1551159889] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.4% (2339760 kB) free!
[1551160007] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551160022] SERVICE ALERT: webserver2;Root Partition;OK;HARD;1;DISK OK - free space: / 14448 MB (28.23% inode=99%):
[1551160071] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.8% (22739564 kB) free.
[1551160090] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 37.9% (9280528 kB) free.
[1551160122] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.1% (14641648 kB) free.
[1551160192] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 93.8% (15240356 kB) free.
[1551160364] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 331 processes
[1551160607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1551160727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551160847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 307 processes
[1551161623] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.45, 12.87, 13.07
[1551161730] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 14.5% (2350044 kB) free!
[1551161743] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.41, 9.08, 11.65
[1551161850] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 13.7% (2222980 kB) free!
[1551161863] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.51, 6.54, 10.41
[1551161971] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.9% (2098964 kB) free!
[1551162254] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 337 processes
[1551163063] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.41, 1.48, 3.87
[1551163171] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 21.6% (3516300 kB) free.
[1551163454] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 325 processes
[1551164216] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (34.4% change >= 20.0% threshold)
[1551164371] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3248512 kB) free!
[1551164491] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3221308 kB) free!
[1551164611] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.8% (3213384 kB) free!
[1551164663] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (159512 kB) free!
[1551164737] SERVICE ALERT: sig2-vm1;Sensors Status;UNKNOWN;SOFT;1;SENSORS UNKNOWN - command not found (did you install lmsensors?)
[1551164783] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (153196 kB) free!
[1551164857] SERVICE ALERT: sig2-vm1;Sensors Status;UNKNOWN;SOFT;2;SENSORS UNKNOWN - command not found (did you install lmsensors?)
[1551164904] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 95.2% (7003800 kB) free.
[1551164977] SERVICE ALERT: sig2-vm1;Sensors Status;WARNING;HARD;3;WARNING - sensors returned state 1
[1551165565] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 280 processes
[1551165647] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551165685] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 280 processes
[1551165805] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 280 processes
[1551166106] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (123528 kB) free!
[1551166226] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (117848 kB) free!
[1551166248] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551166346] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (132356 kB) free!
[1551166368] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551166406] SERVICE ALERT: sig2-vm1;CPU_Procs;OK;HARD;1;CPU OK: 279 processes
[1551166488] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551168994] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1551169060] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1551169060] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.5% (14864540 kB) free.
[1551169183] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1551169183] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1551169245] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 95.4% (7014052 kB) free.
[1551169634] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (149972 kB) free!
[1551169754] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (148080 kB) free!
[1551169874] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (146924 kB) free!
[1551171206] SERVICE ALERT: sig2-vm1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 286 processes
[1551171288] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551171326] SERVICE ALERT: sig2-vm1;CPU_Procs;OK;HARD;1;CPU OK: 281 processes
[1551171888] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551172009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551172129] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551172184] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 91.8% (6747804 kB) free.
[1551172200] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178944 kB) free!
[1551172320] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (175460 kB) free!
[1551172408] SERVICE ALERT: sig2-vm2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1551172440] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176148 kB) free!
[1551172528] SERVICE ALERT: sig2-vm2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1551172648] SERVICE ALERT: sig2-vm2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1551173329] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551173929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551174049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551174169] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551174620] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1551174740] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1551175664] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1551175678] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (481516 kB) free!
[1551175730] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (148984 kB) free!
[1551175798] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (492792 kB) free!
[1551175850] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176660 kB) free!
[1551175918] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (490440 kB) free!
[1551175970] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172492 kB) free!
[1551176114] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (30.8% change >= 20.0% threshold)
[1551176169] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.5% (23156764 kB) free.
[1551176208] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.4% (15337840 kB) free.
[1551176249] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.4% (14049188 kB) free.
[1551176550] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (127456 kB) free!
[1551176670] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (126620 kB) free!
[1551176790] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (125760 kB) free!
[1551176856] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.2% (2310408 kB) free!
[1551176976] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.2% (2302908 kB) free!
[1551177096] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.1% (2293856 kB) free!
[1551177148] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 93.3% (6864128 kB) free.
[1551177459] SERVICE ALERT: sig3-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1551177579] SERVICE ALERT: sig3-vm1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1551177696] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1551177696] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (631556 kB) free!
[1551177699] SERVICE ALERT: sig3-vm1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1551177754] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (122760 kB) free!
[1551177869] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 93.5% (6874544 kB) free.
[1551177956] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.4% (14043476 kB) free.
[1551178299] SERVICE ALERT: sig3-vm1;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1551179569] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551180025] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 336 processes
[1551180132] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1551180132] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (26.1% change >= 20.0% threshold)
[1551180169] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551180289] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551180409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551180571] SERVICE ALERT: sig3-vm2;Memory;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1551180691] SERVICE ALERT: sig3-vm2;Memory;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1551180811] SERVICE ALERT: sig3-vm2;Memory;UNKNOWN;HARD;3;NRPE: Unable to read output
[1551181683] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (29.5% change >= 20.0% threshold)
[1551181683] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1551181874] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (176036 kB) free!
[1551182075] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (124648 kB) free!
[1551182195] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (220864 kB) free!
[1551182258] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.2% (6999936 kB) free.
[1551182809] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551183013] SERVICE ALERT: sig4-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 280 processes
[1551183133] SERVICE ALERT: sig4-vm1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 280 processes
[1551183253] SERVICE ALERT: sig4-vm1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 280 processes
[1551183409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551183462] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (119812 kB) free!
[1551183529] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551183582] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (126276 kB) free!
[1551183649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551183702] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (117480 kB) free!
[1551183854] SERVICE ALERT: sig4-vm1;CPU_Procs;OK;HARD;1;CPU OK: 276 processes
[1551184481] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9464376 kB) free!
[1551184601] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.3% (10996416 kB) free.
[1551185051] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (30.3% change >= 20.0% threshold)
[1551185051] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1551185062] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (30.3% change >= 20.0% threshold)
[1551185062] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1551185125] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.4% (7019048 kB) free.
[1551185175] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.9% (6908360 kB) free.
[1551185201] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9833380 kB) free!
[1551185263] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (156948 kB) free!
[1551185321] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11413532 kB) free.
[1551185383] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (156840 kB) free!
[1551185503] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (157096 kB) free!
[1551186521] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9582372 kB) free!
[1551186641] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.9% (11260912 kB) free.
[1551188163] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551188449] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551189050] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551189170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551189290] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551189890] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551190490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551190610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551190730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551190841] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9480600 kB) free!
[1551190961] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11254996 kB) free.
[1551191930] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551192025] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 334 processes
[1551193130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551193250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551193361] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9711112 kB) free!
[1551193370] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551193481] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11438656 kB) free.
[1551194791] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 314 processes
[1551194911] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 316 processes
[1551195031] SERVICE ALERT: sig1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 314 processes
[1551195881] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9278008 kB) free!
[1551196001] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11432520 kB) free.
[1551196831] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1551196970] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551197570] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551197690] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551198290] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551198410] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551199010] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551199130] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551199731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551199851] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551199971] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551200571] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551200691] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551200801] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9543432 kB) free!
[1551200921] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.6% (11153832 kB) free.
[1551201291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551201411] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551201521] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9619008 kB) free!
[1551201641] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.6% (11143556 kB) free.
[1551202011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551202132] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551202241] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9781076 kB) free!
[1551202361] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11332520 kB) free.
[1551202731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551202851] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551202971] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551203571] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551203691] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551204291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551204411] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551205011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551205132] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551205731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes