[1579026732] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.7% (30896788 kB) free.
[1579026767] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22848544 kB) free!
[1579026807] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579026887] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19332972 kB) free!
[1579026927] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579027007] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16023892 kB) free!
[1579027471] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171188 kB) free!
[1579028127] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579028207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34867400 kB) free.
[1579028247] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579028523] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (1645460 kB) free!
[1579028585] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163884 kB) free!
[1579028645] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (175444 kB) free!
[1579028705] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (174104 kB) free!
[1579028856] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1579028976] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579029005] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 66.8% (16383576 kB) free.
[1579029096] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579029408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22457744 kB) free!
[1579029528] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (18644176 kB) free!
[1579029648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45495876 kB) free.
[1579029696] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1579029817] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1579030248] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23180556 kB) free!
[1579030329] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1579030368] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.6% (19191736 kB) free!
[1579030449] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/01/15 01:01:28 [error] 21677#0: *24896931 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.155:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579030488] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15678496 kB) free!
[1579030569] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579031048] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579031496] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579031616] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579031688] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33372784 kB) free.
[1579031822] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579031942] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 403 processes
[1579032062] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 402 processes
[1579032216] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579032336] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579032456] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579032662] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1579032888] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22185444 kB) free!
[1579033008] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18689928 kB) free!
[1579033057] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1579033128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45065640 kB) free.
[1579033177] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1579033728] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22926920 kB) free!
[1579033777] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579033848] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.6% (19276904 kB) free!
[1579033897] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579033968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (45016968 kB) free.
[1579034017] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1579034568] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22800340 kB) free!
[1579034617] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579034688] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19333364 kB) free!
[1579034737] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1579034808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44932248 kB) free.
[1579034857] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1579035259] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8781740 kB) free!
[1579035379] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (4239516 kB) free!
[1579035408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22853152 kB) free!
[1579035458] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579035499] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (934100 kB) free!
[1579035528] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19482752 kB) free!
[1579035568] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (2214636 kB) free!
[1579035578] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579035648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44894420 kB) free.
[1579035688] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236924 kB) free!
[1579035808] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222344 kB) free!
[1579036099] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 85.7% (42257112 kB) free.
[1579036100] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.65, 5.67, 5.20
[1579036220] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.38, 5.25, 5.10
[1579036248] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22737204 kB) free!
[1579036340] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.08, 4.86, 4.98
[1579036369] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19400772 kB) free!
[1579036408] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.7% (42705268 kB) free.
[1579036471] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.3% (10371668 kB) free.
[1579036489] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44848260 kB) free.
[1579037089] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22592860 kB) free!
[1579037209] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.6% (19268616 kB) free!
[1579037329] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43720200 kB) free.
[1579037378] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579037470] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2199) < 182.18.144.246 [15/Jan/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579037498] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1579037590] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(120) < 49.50.64.122 [15/Jan/2020:03:01:49 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1579037710] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579037770] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2811) < 2020/01/15 03:02:42 [error] 21675#0: *24907391 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579037872] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(209) < 103.16.101.72 [15/Jan/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1579037890] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579037929] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21740984 kB) free!
[1579037944] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(260) < 2020/01/15 03:01:32 [error] 18356#0: *39142992 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579037993] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579038049] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18485436 kB) free!
[1579038064] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579038169] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42073680 kB) free.
[1579038665] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 03:17:06 [error] 18370#0: *39147135 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1579038769] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21544220 kB) free!
[1579038785] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579038889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44701844 kB) free.
[1579039385] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 03:31:35 [error] 18360#0: *39150776 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579039489] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23731968 kB) free!
[1579039505] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579039609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.3% (20151044 kB) free!
[1579039729] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44587740 kB) free.
[1579040182] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.1% (13319948 kB) free!
[1579040329] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22585632 kB) free!
[1579040449] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.6% (19265184 kB) free!
[1579040570] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41710268 kB) free.
[1579040783] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12605876 kB) free!
[1579041169] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14111540 kB) free!
[1579041289] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (43993988 kB) free.
[1579041889] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17584928 kB) free!
[1579041983] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (13672152 kB) free!
[1579042009] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13832004 kB) free!
[1579042129] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43936696 kB) free.
[1579042298] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579042729] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17260452 kB) free!
[1579042849] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12235024 kB) free!
[1579042969] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41678216 kB) free.
[1579043569] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16802880 kB) free!
[1579043689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43815576 kB) free.
[1579043705] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 04:41:36 [error] 18370#0: *39170717 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1579043825] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579044289] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19748964 kB) free!
[1579044409] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15458260 kB) free!
[1579044529] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42100680 kB) free.
[1579044871] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194120 kB) free!
[1579044991] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171084 kB) free!
[1579045112] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194316 kB) free!
[1579045129] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19892824 kB) free!
[1579045250] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45731044 kB) free.
[1579045625] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 05:09:06 [error] 18370#0: *39177298 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579045745] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579045849] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579046290] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 05:22:13 [error] 21682#0: *24913549 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:47/0, bytes from/to upstream:0/47
[1579046410] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579046529] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 27.5% (6779100 kB) free.
[1579047129] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188940 kB) free!
[1579047250] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191552 kB) free!
[1579047370] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194988 kB) free!
[1579047545] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 05:48:05 [error] 18362#0: *39186423 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579047665] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/15 05:50:07 [error] 18362#0: *39186876 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579047698] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579047785] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579047818] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579047938] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579048385] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 06:01:04 [error] 18366#0: *39189518 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1579048505] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579048538] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579048583] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1579048583] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12464920 kB) free!
[1579048849] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24662080 kB) free!
[1579048969] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.8% (47179764 kB) free.
[1579049569] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24509552 kB) free!
[1579049689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46679808 kB) free.
[1579049705] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 06:20:36 [error] 18356#0: *39194142 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579049783] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.7% (31204920 kB) free.
[1579049825] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579050289] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24037088 kB) free!
[1579050383] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.5% (17783432 kB) free!
[1579050409] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (46055276 kB) free.
[1579050425] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 06:32:06 [error] 18367#0: *39197456 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579050503] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 38.3% (50445836 kB) free.
[1579050545] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579050766] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.7% (21888612 kB) free.
[1579050818] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.2% (38519292 kB) free.
[1579050870] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.7% (22715264 kB) free.
[1579050904] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.9% (22517596 kB) free.
[1579051009] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24193048 kB) free!
[1579051129] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45536788 kB) free.
[1579051145] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 06:43:06 [error] 18366#0: *39201088 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579051265] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579051369] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196648 kB) free!
[1579051442] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244212 kB) free!
[1579051489] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187396 kB) free!
[1579051506] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169000 kB) free!
[1579051562] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216212 kB) free!
[1579051609] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1579051609] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193740 kB) free!
[1579051626] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182168 kB) free!
[1579051683] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220296 kB) free!
[1579051729] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (26003852 kB) free!
[1579051746] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189976 kB) free!
[1579051850] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45120460 kB) free.
[1579051865] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 06:54:06 [error] 18359#0: *39204805 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579051986] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/15 07:02:04 [error] 18357#0: *39207535 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579052106] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579052449] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24065500 kB) free!
[1579052569] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44694000 kB) free.
[1579052706] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 07:14:05 [error] 18370#0: *39211523 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579052826] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579052957] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4892828 kB) free!
[1579053077] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.9% (4883716 kB) free!
[1579053169] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25606144 kB) free!
[1579053197] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.9% (4881912 kB) free!
[1579053289] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44431512 kB) free.
[1579053426] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 07:23:35 [error] 18369#0: *39214688 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579053546] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579053889] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24887408 kB) free!
[1579054009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44095544 kB) free.
[1579054609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24752296 kB) free!
[1579054729] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43824844 kB) free.
[1579054746] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 07:48:05 [error] 18363#0: *39222947 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579054867] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579055329] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24976236 kB) free!
[1579055449] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43582484 kB) free.
[1579055466] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 08:00:06 [error] 18366#0: *39227033 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579055586] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579055704] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.7% (22387376 kB) free.
[1579055752] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.3% (40057728 kB) free.
[1579055783] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.9% (23260856 kB) free.
[1579056049] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22272472 kB) free!
[1579056169] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42966668 kB) free.
[1579056186] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 08:07:05 [error] 18360#0: *39229273 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579056306] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579056307] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194564 kB) free!
[1579056358] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217664 kB) free!
[1579056427] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186684 kB) free!
[1579056478] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (246900 kB) free!
[1579056547] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194332 kB) free!
[1579056599] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1579056599] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224876 kB) free!
[1579056769] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22985036 kB) free!
[1579056889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42556652 kB) free.
[1579056906] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 08:16:37 [error] 18366#0: *39232497 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:122/17, bytes from/to upstream:17/122
[1579057026] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579057490] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22563784 kB) free!
[1579057610] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42137224 kB) free.
[1579057627] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 08:37:05 [error] 18368#0: *39239410 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579057747] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/15 08:38:06 [error] 18367#0: *39239738 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579057867] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579058210] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (21996236 kB) free!
[1579058330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41709856 kB) free.
[1579058930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22719676 kB) free!
[1579059050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29273892 kB) free.
[1579059067] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 09:00:35 [error] 18366#0: *39247164 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579059187] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/15 09:01:36 [error] 18369#0: *39247520 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579059307] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579059650] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20239200 kB) free!
[1579059733] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.1% (22476364 kB) free.
[1579059770] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16056716 kB) free!
[1579059785] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.7% (38788296 kB) free.
[1579059890] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39133924 kB) free.
[1579059907] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 09:10:36 [error] 18370#0: *39250729 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579059939] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579059986] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195792 kB) free!
[1579060027] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190708 kB) free!
[1579060028] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579060059] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579060106] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192168 kB) free!
[1579060148] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190380 kB) free!
[1579060179] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1579060226] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193020 kB) free!
[1579060235] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579060268] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203428 kB) free!
[1579060336] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174256 kB) free!
[1579060356] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1579060388] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214568 kB) free!
[1579060456] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177732 kB) free!
[1579060490] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17133124 kB) free!
[1579060508] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247652 kB) free!
[1579060576] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173612 kB) free!
[1579060610] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39440320 kB) free.
[1579060627] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 09:19:36 [error] 18368#0: *39253734 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579060628] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216916 kB) free!
[1579060703] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579060747] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/15 09:27:35 [error] 18365#0: *39256413 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579060779] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1579060867] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579061210] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16854520 kB) free!
[1579061330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37976536 kB) free.
[1579061467] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 09:37:05 [error] 18355#0: *39259588 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579061587] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579061930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15899140 kB) free!
[1579061979] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1579062050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33633632 kB) free.
[1579062068] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 80.5% (19737932 kB) free.
[1579062099] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579062219] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1579062651] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16539728 kB) free!
[1579062771] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12654280 kB) free!
[1579062787] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 09:55:06 [error] 18366#0: *39265771 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579062891] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34686996 kB) free.
[1579062908] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/15 10:03:35 [error] 18364#0: *39268591 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579063028] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579063491] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12845424 kB) free!
[1579063611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35667716 kB) free.
[1579063628] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 10:14:36 [error] 18362#0: *39272261 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579063748] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579064211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14210840 kB) free!
[1579064331] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10295368 kB) free!
[1579064348] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 10:21:35 [error] 18369#0: *39274636 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579064451] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33880976 kB) free.
[1579064468] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/01/15 10:30:37 [error] 18366#0: *39277718 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579064489] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.7% (22389764 kB) free.
[1579064536] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39761672 kB) free.
[1579064571] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.2% (22835192 kB) free.
[1579064588] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579065051] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11473732 kB) free!
[1579065097] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196812 kB) free!
[1579065119] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579065138] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (253520 kB) free!
[1579065171] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7588012 kB) free!
[1579065217] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185752 kB) free!
[1579065239] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579065258] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (250316 kB) free!
[1579065291] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31867780 kB) free.
[1579065337] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177132 kB) free!
[1579065359] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579065378] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (261164 kB) free!
[1579065788] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/15 10:53:06 [error] 18364#0: *39285472 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579065819] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1579065819] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579065891] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10428936 kB) free!
[1579065908] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579066011] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6416284 kB) free!
[1579066131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30258036 kB) free.
[1579066508] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 11:02:36 [error] 18366#0: *39289161 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579066628] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579066731] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8348148 kB) free!
[1579066852] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4377728 kB) free!
[1579066972] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28340972 kB) free.
[1579067228] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 11:11:36 [error] 18365#0: *39292324 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579067338] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3293836 kB) free!
[1579067348] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579067411] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/15 11:20:06 [error] 21682#0: *24928360 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:7985/87, bytes from/to upstream:87/7985
[1579067458] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2888440 kB) free!
[1579067531] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/01/15 11:22:06 [error] 21682#0: *24928459 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:4169/121, bytes from/to upstream:121/4169
[1579067572] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (381816 kB) free!
[1579067578] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.9% (2444272 kB) free!
[1579067651] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579067692] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (375888 kB) free!
[1579067812] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385192 kB) free!
[1579068178] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (296520 kB) free!
[1579068251] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/15 11:29:58 [error] 21679#0: *24928813 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:8939/176, bytes from/to upstream:176/8939
[1579068278] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (4420016 kB) free!
[1579068371] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/15 11:35:37 [error] 21680#0: *24929077 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:8409/87, bytes from/to upstream:87/8409
[1579068398] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (4176664 kB) free!
[1579068412] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13689812 kB) free!
[1579068491] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579068518] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.1% (3946356 kB) free!
[1579068549] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 11:35:36 [error] 18364#0: *39300882 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579068669] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/15 11:39:34 [error] 18370#0: *39302205 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579068732] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8124580 kB) free!
[1579068772] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174012 kB) free!
[1579068789] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579068852] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7599664 kB) free!
[1579068892] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179040 kB) free!
[1579068972] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.2% (7014956 kB) free!
[1579069012] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164872 kB) free!
[1579069012] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31029776 kB) free.
[1579069091] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/15 11:47:59 [error] 21675#0: *24929628 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:4487/51, bytes from/to upstream:51/4487
[1579069139] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3636400 kB) free!
[1579069212] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579069259] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 12.8% (3161940 kB) free!
[1579069379] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.8% (2657780 kB) free!
[1579069389] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 11:46:05 [error] 18366#0: *39304499 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579069397] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 10.0% (2461236 kB) free!
[1579069509] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579069572] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3835100 kB) free!
[1579069593] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1579069612] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8844716 kB) free!
[1579069713] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1579069718] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1994456 kB) free!
[1579069732] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5106860 kB) free!
[1579069833] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1579069852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33584056 kB) free.
[1579069978] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.1% (16852344 kB) free.
[1579069979] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (483684 kB) free!
[1579070109] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 12:01:08 [error] 18363#0: *39309383 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579070173] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.7% (30392380 kB) free.
[1579070229] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579070411] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 12:05:48 [error] 21683#0: *24930434 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:2172/87, bytes from/to upstream:87/2172
[1579070433] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1579070452] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9857828 kB) free!
[1579070531] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579070572] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5700876 kB) free!
[1579070608] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9218124 kB) free!
[1579070692] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1579070692] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (1953432 kB) free!
[1579070728] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8933952 kB) free!
[1579070830] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 12:17:06 [error] 18370#0: *39314531 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579070846] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 91.0% (22313268 kB) free.
[1579070849] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8642888 kB) free!
[1579070878] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22088292 kB) free.
[1579070921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 48.6% (64072760 kB) free.
[1579070950] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579071024] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.2% (36588952 kB) free.
[1579071058] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.8% (22249136 kB) free.
[1579071096] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1579071096] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.2% (22854300 kB) free.
[1579071131] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 12:20:47 [error] 21682#0: *24931242 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:5017/176, bytes from/to upstream:176/5017
[1579071179] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.4% (23066872 kB) free.
[1579071251] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579071359] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579071488] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193724 kB) free!
[1579071608] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188192 kB) free!
[1579071627] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232360 kB) free!
[1579071633] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 707 processes
[1579071728] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183640 kB) free!
[1579071747] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243052 kB) free!
[1579071753] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1579071867] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (253576 kB) free!
[1579071873] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1579072276] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) <                  {amqp_error,frame_error,
[1579072397] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1579072452] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(734) < 2020/01/15 12:38:03 [error] 21669#0: *24933479 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579072473] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1579072473] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1579072510] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(722) < 103.255.146.164 [15/Jan/2020:12:38:03 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579072572] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579072630] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579072649] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3722972 kB) free!
[1579072728] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26107020 kB) free!
[1579072848] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20593368 kB) free!
[1579072868] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.5% (4522812 kB) free!
[1579072969] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15653696 kB) free!
[1579072988] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.8% (4356308 kB) free!
[1579073108] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.2% (4204764 kB) free!
[1579073159] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579073172] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(838) < 2020/01/15 12:49:12 [error] 21676#0: *24935634 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579073230] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(838) < 103.255.146.164 [15/Jan/2020:12:49:12 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579073279] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579073292] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579073351] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579073399] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579073569] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28984892 kB) free.
[1579073892] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 13:02:46 [error] 21681#0: *24936266 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:5627/104, bytes from/to upstream:104/5627
[1579073926] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21990692 kB) free.
[1579073974] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.5% (39181884 kB) free.
[1579074000] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1579074000] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579074012] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579074029] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.5% (46063012 kB) free.
[1579074169] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5430716 kB) free!
[1579074219] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1579074289] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1485796 kB) free!
[1579074409] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31890680 kB) free.
[1579074530] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188532 kB) free!
[1579074580] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238416 kB) free!
[1579074650] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194040 kB) free!
[1579074699] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186084 kB) free!
[1579074700] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226336 kB) free!
[1579074770] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188020 kB) free!
[1579074820] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169448 kB) free!
[1579074820] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218500 kB) free!
[1579074908] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2069240 kB) free!
[1579074940] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173296 kB) free!
[1579075009] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8221560 kB) free!
[1579075129] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4045140 kB) free!
[1579075249] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378888 kB) free!
[1579076019] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579076139] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1579076449] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29260844 kB) free.
[1579076739] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1579076859] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1579076979] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1579077049] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5867460 kB) free!
[1579077169] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (2002952 kB) free!
[1579077289] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (470392 kB) free!
[1579077612] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 14:08:50 [error] 21681#0: *24939400 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3851/121, bytes from/to upstream:121/3851
[1579077732] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579078489] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28183804 kB) free.
[1579078508] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 47.1% (11536020 kB) free.
[1579078551] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 55.1% (13602080 kB) free.
[1579078590] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 29.2% (14370540 kB) free.
[1579078621] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 70.3% (17219856 kB) free.
[1579078779] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579078800] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579078920] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579078933] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 14:30:15 [error] 21677#0: *24940361 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:8423/87, bytes from/to upstream:87/8423
[1579078979] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (3994732 kB) free!
[1579079040] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579079053] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579079090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5170732 kB) free!
[1579079099] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.8% (3782764 kB) free!
[1579079156] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175540 kB) free!
[1579079198] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231568 kB) free!
[1579079210] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1283032 kB) free!
[1579079219] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.7% (3517668 kB) free!
[1579079228] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193872 kB) free!
[1579079276] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195384 kB) free!
[1579079318] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240032 kB) free!
[1579079330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34396312 kB) free.
[1579079348] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192792 kB) free!
[1579079380] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579079387] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1579079396] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186248 kB) free!
[1579079438] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (249324 kB) free!
[1579079468] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190856 kB) free!
[1579079500] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1579079507] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1579079930] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9404628 kB) free!
[1579080050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5169200 kB) free!
[1579080069] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (4421088 kB) free!
[1579080100] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1579080170] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1161468 kB) free!
[1579080189] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4211608 kB) free!
[1579080220] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579080309] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.3% (4003268 kB) free!
[1579080419] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1064504 kB) free!
[1579080499] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9665064 kB) free!
[1579080619] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9635068 kB) free!
[1579080739] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (9477692 kB) free!
[1579080770] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14978028 kB) free!
[1579080820] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1579080853] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 15:01:03 [error] 21683#0: *24941736 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:9787/87, bytes from/to upstream:87/9787
[1579080940] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1579080973] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579081019] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1579081019] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.1% (16862720 kB) free.
[1579081060] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579081370] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31600660 kB) free.
[1579081509] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2136288 kB) free!
[1579081660] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1579081780] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1579081970] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7458868 kB) free!
[1579082090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3316108 kB) free!
[1579082210] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395356 kB) free!
[1579082380] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1579082501] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1579082555] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579082621] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1579082675] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1579082718] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171172 kB) free!
[1579082810] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13277692 kB) free!
[1579082838] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196356 kB) free!
[1579082958] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180184 kB) free!
[1579083373] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 15:38:30 [error] 21682#0: *24943383 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:11452/87, bytes from/to upstream:87/11452
[1579083410] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31353116 kB) free.
[1579083493] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579083821] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 355 processes
[1579083873] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579083941] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579084010] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7199152 kB) free!
[1579084061] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1579084131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3059476 kB) free!
[1579084251] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13392680 kB) free!
[1579084758] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1579084758] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 80.0% (19614268 kB) free.
[1579084851] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11902032 kB) free!
[1579085261] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1579085261] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 355 processes
[1579085451] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30078228 kB) free.
[1579085539] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4317904 kB) free!
[1579085861] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1579086051] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5344312 kB) free!
[1579086171] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1182416 kB) free!
[1579086240] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1579086291] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405188 kB) free!
[1579086493] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 16:35:30 [error] 21676#0: *24945959 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:7258/87, bytes from/to upstream:87/7258
[1579086613] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579087061] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1579087181] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1579087301] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1579087473] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1579087491] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20862908 kB) free!
[1579087593] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1579087648] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.8% (22162328 kB) free.
[1579087688] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 27.3% (13447904 kB) free.
[1579087713] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 705 processes
[1579087722] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 41.3% (20377068 kB) free.
[1579087748] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 61.1% (14972480 kB) free.
[1579087751] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(772) < 103.255.146.164 [15/Jan/2020:16:58:19 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579087782] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.8% (21524148 kB) free.
[1579087813] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(784) < 2020/01/15 16:58:19 [error] 21673#0: *24948493 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579087871] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579087933] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579088091] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382104 kB) free!
[1579088258] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175188 kB) free!
[1579088328] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232952 kB) free!
[1579088378] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182708 kB) free!
[1579088387] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194536 kB) free!
[1579088448] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232864 kB) free!
[1579088499] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172336 kB) free!
[1579088501] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1579088507] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189644 kB) free!
[1579088533] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(801) < 2020/01/15 17:11:54 [error] 21677#0: *24950712 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579088568] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242676 kB) free!
[1579088621] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1579088627] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167960 kB) free!
[1579088654] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(735) < 2020/01/15 17:13:08 [error] 21683#0: *24952266 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579088774] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(4046) < 2020/01/15 17:16:14 [error] 21680#0: *24960591 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579088913] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1579089071] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(25577) < 103.255.146.164 [15/Jan/2020:17:21:11 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579089192] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(7689) < 103.255.146.164 [15/Jan/2020:17:23:11 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579089312] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(7536) < 103.255.146.164 [15/Jan/2020:17:25:12 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579089513] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1579089633] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1579089753] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1579089891] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15846296 kB) free!
[1579090421] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1579090491] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33290456 kB) free.
[1579090512] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579090541] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579090575] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579091091] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8262620 kB) free!
[1579091175] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 17:53:32 [error] 21677#0: *25067412 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3533/104, bytes from/to upstream:104/3533
[1579091212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4099884 kB) free!
[1579091219] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4490764 kB) free!
[1579091295] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579091332] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398376 kB) free!
[1579091339] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4314324 kB) free!
[1579091459] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.3% (16903948 kB) free.
[1579091741] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1579091861] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579091895] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 18:02:19 [error] 21683#0: *25067827 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:24500/104, bytes from/to upstream:104/24500
[1579091932] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13252120 kB) free!
[1579092016] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579092461] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579092532] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34059368 kB) free.
[1579092581] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579092615] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 18:17:16 [error] 21677#0: *25068484 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:5394/104, bytes from/to upstream:104/5394
[1579092659] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579092701] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1579092735] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579093132] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8717860 kB) free!
[1579093252] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4430136 kB) free!
[1579093302] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 359 processes
[1579093353] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1579093372] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394596 kB) free!
[1579093422] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579093473] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1579093542] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1579093972] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13975856 kB) free!
[1579094142] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579094535] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 18:47:05 [error] 21674#0: *25069877 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3485/104, bytes from/to upstream:104/3485
[1579094572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35806992 kB) free.
[1579094655] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579094742] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 355 processes
[1579094862] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579094958] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579094982] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1579095172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6796608 kB) free!
[1579095292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2080888 kB) free!
[1579095412] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403796 kB) free!
[1579095638] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9096304 kB) free!
[1579095758] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8936036 kB) free!
[1579095878] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.8% (8776000 kB) free!
[1579096012] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18645208 kB) free!
[1579096456] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/15 19:17:21 [error] 21677#0: *25071308 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3081/87, bytes from/to upstream:87/3081
[1579096576] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579096612] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29174584 kB) free.
[1579096750] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.6% (19012664 kB) free.
[1579096795] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 57.5% (14190264 kB) free.
[1579096831] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (5537496 kB) free!
[1579096835] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 55.1% (27172232 kB) free.
[1579097212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11503564 kB) free!
[1579097333] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6892612 kB) free!
[1579097356] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188664 kB) free!
[1579097398] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195068 kB) free!
[1579097438] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235712 kB) free!
[1579097453] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2732596 kB) free!
[1579097476] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164564 kB) free!
[1579097507] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1579097518] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189312 kB) free!
[1579097558] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244668 kB) free!
[1579097596] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197068 kB) free!
[1579097627] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1579097638] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186152 kB) free!
[1579097678] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232412 kB) free!
[1579097747] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1579097846] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 0.8% (189768 kB) free!
[1579097966] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 0.8% (194464 kB) free!
[1579097982] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 355 processes
[1579098053] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19967524 kB) free!
[1579098087] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.6% (22203560 kB) free.
[1579098102] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1579098347] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1579098653] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405404 kB) free!
[1579098702] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 359 processes
[1579098822] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1579098912] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(715) < 103.255.146.164 [15/Jan/2020:20:01:23 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579098942] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 352 processes
[1579098976] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(726) < 2020/01/15 20:01:23 [error] 21683#0: *25074783 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579099033] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579099096] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579099253] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16987396 kB) free!
[1579099853] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39473508 kB) free.
[1579100078] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4539212 kB) free!
[1579100453] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13888744 kB) free!
[1579100573] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8919264 kB) free!
[1579100693] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4270548 kB) free!
[1579100959] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4614452 kB) free!
[1579101079] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (4498028 kB) free!
[1579101199] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (4371920 kB) free!
[1579101293] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19633680 kB) free!
[1579101343] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1579101496] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/15 20:39:59 [error] 21676#0: *25076626 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:2831/87, bytes from/to upstream:87/2831
[1579101617] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579101893] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (454500 kB) free!
[1579101943] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579102063] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1579102183] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579102493] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17816372 kB) free!
[1579102816] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/15 21:01:06 [error] 21683#0: *25077596 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:8279/87, bytes from/to upstream:87/8279
[1579102936] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579103080] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 41.9% (20641368 kB) free.
[1579103093] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36812440 kB) free.
[1579103118] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 58.7% (14473420 kB) free.
[1579103147] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 32.0% (15774692 kB) free.
[1579103186] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.1% (22328148 kB) free.
[1579103688] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238792 kB) free!
[1579103693] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13233168 kB) free!
[1579103727] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187008 kB) free!
[1579103808] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233676 kB) free!
[1579103813] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9060016 kB) free!
[1579103847] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194660 kB) free!
[1579103928] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219256 kB) free!
[1579103933] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (5242636 kB) free!
[1579103967] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178744 kB) free!
[1579104363] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.9% (20693112 kB) free.
[1579104386] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (1264740 kB) free!
[1579104395] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36680228 kB) free.
[1579104421] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 45.3% (22342908 kB) free.
[1579104455] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.7% (22219316 kB) free.
[1579104800] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2454916 kB) free!
[1579104966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195400 kB) free!
[1579105028] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (249320 kB) free!
[1579105059] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (1288956 kB) free!
[1579105087] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195700 kB) free!
[1579105148] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238104 kB) free!
[1579105180] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (1282040 kB) free!
[1579105207] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193400 kB) free!
[1579105268] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233472 kB) free!
[1579105300] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1265180 kB) free!
[1579106798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21363036 kB) free!
[1579106918] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17318884 kB) free!
[1579107038] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13499368 kB) free!
[1579107583] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1579107638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27609144 kB) free.
[1579108184] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 353 processes
[1579108304] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1579108839] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19397264 kB) free!
[1579108904] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579108959] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15706544 kB) free!
[1579109024] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579109079] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (12198888 kB) free!
[1579109144] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579109679] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25837016 kB) free!
[1579110279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39747432 kB) free.
[1579110879] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18440168 kB) free!
[1579110944] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579110999] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14815552 kB) free!
[1579111119] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.8% (45870656 kB) free.
[1579111544] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579111573] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3794056 kB) free!
[1579111665] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1579111693] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237632 kB) free!
[1579111719] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21533092 kB) free!
[1579111785] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1579111813] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232008 kB) free!
[1579111839] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17723384 kB) free!
[1579111959] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14178144 kB) free!
[1579112303] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4436912 kB) free!
[1579112423] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2886976 kB) free!
[1579112543] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3476016 kB) free!
[1579112559] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27057596 kB) free.