[1557253901] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557254021] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557254141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1557256154] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557258726] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 270.00, 150.88, 63.01
[1557258846] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 274.66, 192.09, 88.76
[1557258966] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 275.87, 219.81, 111.45
[1557259541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1557260141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557260261] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557260862] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557260982] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557261367] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.57, 1.71, 12.60
[1557262182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557262302] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557262567] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.40, 1.64, 4.68
[1557262902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557262903] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 298.24, 186.75, 80.76
[1557263022] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557263023] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 276.05, 219.36, 105.78
[1557263143] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 37.30, 146.81, 92.96
[1557263375] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557263622] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557263743] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1557264942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557264943] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.64, 0.82, 13.68
[1557265062] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557265662] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557265782] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557266143] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.33, 3.99
[1557266382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557266502] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557267103] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1557267702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557267822] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557268422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557268542] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557269142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1557269262] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1557271062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557271182] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557271782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557271902] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557273102] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557273222] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1557273823] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557273943] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1557274543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557274663] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557275263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557275383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557275503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557277754] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 258.59, 111.47, 42.43
[1557277874] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 278.48, 167.78, 71.44
[1557277994] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 281.26, 205.48, 96.94
[1557278296] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.9% (20572700 kB) free.
[1557279794] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.82, 1.08, 14.86
[1557280994] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.48, 4.45
[1557281503] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1557282103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557282223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557282343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557283567] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 259.64, 100.30, 37.82
[1557283687] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 294.46, 166.37, 69.61
[1557283807] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 298.99, 210.49, 97.51
[1557286207] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.68, 1.88, 14.62
[1557287422] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3072832 kB) free!
[1557287542] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (2998400 kB) free!
[1557287662] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (2926372 kB) free!
[1557288007] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.37, 1.49, 3.49
[1557288343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1557288943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557289064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557289184] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1557289184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557290062] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1229496 kB) free!
[1557291558] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57798 MB (25% inode=99%):
[1557291584] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557292184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557292284] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.0% (14472876 kB) free.
[1557292304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557292424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557294895] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3170560 kB) free!
[1557295016] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (3164880 kB) free!
[1557295136] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (3165096 kB) free!
[1557295543] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 221.53, 81.44, 29.83
[1557295664] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 257.76, 141.63, 58.14
[1557295784] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 262.50, 181.86, 83.01
[1557295884] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 587 processes
[1557295945] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 584 processes
[1557296403] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170960 kB) free!
[1557296481] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (1235172 kB) free!
[1557296491] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (373876 kB) free!
[1557296523] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165688 kB) free!
[1557296601] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (1201956 kB) free!
[1557296611] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184160 kB) free!
[1557296643] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172492 kB) free!
[1557296721] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1175528 kB) free!
[1557296731] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179084 kB) free!
[1557297224] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557297584] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.30, 1.03, 14.50
[1557297824] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557297945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557298065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557298344] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 592 processes
[1557298404] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 589 processes
[1557298784] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.27, 0.37, 4.32
[1557299027] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.1% (22975956 kB) free.
[1557299091] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.3% (22142560 kB) free.
[1557299153] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.5% (22671472 kB) free.
[1557300200] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.33, 5.07, 5.01
[1557300320] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.88, 5.36, 5.12
[1557300440] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.01, 5.89, 5.34
[1557301065] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557301293] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.7% (2559104 kB) free!
[1557301413] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.5% (2363648 kB) free!
[1557301534] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.3% (2163524 kB) free!
[1557301665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557301785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557301906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557302134] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1180880 kB) free!
[1557302944] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1557303172] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 367 processes
[1557303292] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 381 processes
[1557303440] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.39, 4.77, 4.94
[1557303607] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 459.79, 358.73, 178.04
[1557303641] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5554 MB (10% inode=97%):
[1557303727] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 459.66, 392.19, 212.21
[1557303847] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 87.85, 284.20, 195.45
[1557304241] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7007 MB (13% inode=98%):
[1557304306] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1557304906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557305026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557305147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557306247] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.62, 2.04, 16.63
[1557306794] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 113.92, 184.48, 103.44
[1557306914] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.66, 123.58, 90.95
[1557307034] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.28, 82.80, 79.97
[1557307040] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.02, 5.16, 5.04
[1557307160] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.95, 5.45, 5.16
[1557307280] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.89, 5.20, 5.10
[1557307880] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1557307880] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.72, 4.13, 4.62
[1557308047] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.57, 1.71, 3.84
[1557308834] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 0.71, 11.90
[1557309947] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1557310034] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.62, 0.60, 3.82
[1557310209] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1557310548] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1557310668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557310758] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47047 MB (20% inode=99%):
[1557310775] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2971268 kB) free!
[1557310788] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557310878] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46662 MB (20% inode=99%):
[1557310880] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.87, 5.73, 5.17
[1557310895] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2859792 kB) free!
[1557310998] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46487 MB (20% inode=99%):
[1557311000] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.87, 5.03, 4.98
[1557311015] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.0% (2763764 kB) free!
[1557311388] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557311612] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 315 processes
[1557311733] SERVICE ALERT: db1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 309 processes
[1557311853] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 308 processes
[1557311988] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557311997] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.3% (14190056 kB) free.
[1557312053] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.0% (14623864 kB) free.
[1557312102] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.3% (14196052 kB) free.
[1557312108] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557312190] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8842256 kB) free!
[1557312228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557312310] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8842904 kB) free!
[1557312430] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.9% (8843800 kB) free!
[1557312453] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 307 processes
[1557312544] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1557312809] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1557312929] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1557313049] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1557313053] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1557313447] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 286.22, 194.32, 88.35
[1557313567] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 287.47, 225.19, 112.51
[1557313687] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1557313687] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 287.06, 245.73, 133.71
[1557314344] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1557314634] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173068 kB) free!
[1557314691] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173436 kB) free!
[1557314754] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172160 kB) free!
[1557314765] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196708 kB) free!
[1557314811] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183408 kB) free!
[1557314874] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172112 kB) free!
[1557314886] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169996 kB) free!
[1557314931] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185096 kB) free!
[1557315006] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164180 kB) free!
[1557315584] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 316.52, 194.55, 83.57
[1557315704] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 193.93, 205.61, 102.11
[1557315824] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 27.35, 138.07, 89.91
[1557316087] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.21, 2.24, 12.30
[1557316744] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1557317077] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 97.4% (47996948 kB) free.
[1557317112] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.6% (23336360 kB) free.
[1557317154] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.5% (22909592 kB) free.
[1557317217] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.6% (22948136 kB) free.
[1557317287] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.12, 2.01, 4.76
[1557317344] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1557317344] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1557317624] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.17, 0.61, 13.20
[1557317628] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557317849] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 343 processes
[1557317969] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 343 processes
[1557318062] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (147916 kB) free!
[1557318089] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 343 processes
[1557318111] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (179620 kB) free!
[1557318182] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (154812 kB) free!
[1557318229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557318232] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180168 kB) free!
[1557318302] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174716 kB) free!
[1557318349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557318352] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181088 kB) free!
[1557318469] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557318689] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1557318824] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.28, 3.85
[1557319069] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557319401] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557319521] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169020 kB) free!
[1557319641] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178864 kB) free!
[1557319669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557319761] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199028 kB) free!
[1557319789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1557319909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557319998] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 84513 MB (37% inode=99%):
[1557320102] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 54.3% (8830888 kB) free.
[1557320221] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.8% (3866200 kB) free!
[1557320341] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.9% (3904428 kB) free!
[1557320461] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (3883580 kB) free!
[1557320509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557320887] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 42.54, 10.40, 4.61
[1557320945] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8987828 kB) free!
[1557321007] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.73, 13.63, 6.81
[1557321065] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.9% (14246996 kB) free.
[1557321110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557321127] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.11, 9.62, 6.17
[1557321230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557321350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1557321727] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.46, 3.20, 4.22
[1557321950] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1557322144] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1557322162] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1557322162] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.1% (3477108 kB) free!
[1557322261] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1557322261] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1198636 kB) free!
[1557322454] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.3% (23509544 kB) free.
[1557322550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557322560] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.5% (22923468 kB) free.
[1557322623] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 84.8% (13777692 kB) free.
[1557322670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557322790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557322865] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9677352 kB) free!
[1557322985] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.1% (14850688 kB) free.
[1557323390] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557323664] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194784 kB) free!
[1557323784] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197032 kB) free!
[1557323904] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196780 kB) free!
[1557323991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557324111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557324231] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557324785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8754728 kB) free!
[1557324831] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557324905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.4% (13998636 kB) free.
[1557325431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557325551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557325671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557326271] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557326871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557326992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557327112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557327434] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 90.91, 23.09, 8.06
[1557327554] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 14.95, 17.91, 8.11
[1557327674] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.35, 12.13, 7.18
[1557327902] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (3224368 kB) free!
[1557328022] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3130072 kB) free!
[1557328142] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (3038036 kB) free!
[1557328274] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.23, 2.27, 4.09
[1557329105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8711240 kB) free!
[1557329225] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.3% (13967584 kB) free.
[1557329344] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1557329474] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4.88, 82.17, 62.12
[1557329527] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2.69, 45.35, 40.89
[1557329594] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.93, 55.11, 54.63
[1557329647] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.81, 30.92, 36.14
[1557329714] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.50, 37.08, 48.09
[1557329716] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 96.2% (23738176 kB) free.
[1557329735] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1557329735] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.2% (14660620 kB) free.
[1557329767] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.73, 21.28, 31.98
[1557330112] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557330367] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.12, 4.48, 17.64
[1557330713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557330833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557330914] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1557330914] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 1.39, 13.72
[1557330953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557331025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8477792 kB) free!
[1557331145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.0% (13774720 kB) free.
[1557331553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557332153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557332167] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.58, 1.72, 4.06
[1557332273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557332393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557332714] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.58, 0.57, 3.47
[1557333364] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557335121] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1557335393] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557335993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557336113] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557336233] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557336833] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557337433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557337554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557337674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557337745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8914272 kB) free!
[1557337865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.3% (13948584 kB) free.
[1557338274] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557338745] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1557338874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557338994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557339114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557339665] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8908852 kB) free!
[1557339714] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557339785] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.2% (13909116 kB) free.