[1550428667] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3516660 kB) free!
[1550428787] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9598996 kB) free!
[1550428907] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.8% (6296672 kB) free!
[1550430707] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10358928 kB) free.
[1550431017] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47348 MB (21% inode=99%):
[1550431308] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4696916 kB) free!
[1550431428] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10350488 kB) free.
[1550431617] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45300 MB (20% inode=99%):
[1550431737] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43248 MB (19% inode=99%):
[1550431857] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43252 MB (19% inode=99%):
[1550431937] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550432457] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47305 MB (21% inode=99%):
[1550432628] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.3% (5075052 kB) free!
[1550432748] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10332476 kB) free.
[1550433137] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550433257] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550433377] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550433657] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46279 MB (20% inode=99%):
[1550433777] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45257 MB (20% inode=99%):
[1550433897] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1550433897] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43207 MB (19% inode=99%):
[1550433948] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.6% (5226208 kB) free!
[1550434068] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10317712 kB) free.
[1550434577] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550434668] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.0% (5404652 kB) free!
[1550434788] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10381288 kB) free.
[1550435988] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4490676 kB) free!
[1550436108] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.3% (9995420 kB) free.
[1550436377] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550436497] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550437124] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1550437308] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4423984 kB) free!
[1550437428] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9975792 kB) free.
[1550438628] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4278856 kB) free!
[1550438748] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9775228 kB) free!
[1550438868] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.9% (6869940 kB) free!
[1550439498] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550439618] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550440095] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550440158] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550440215] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550440278] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1550440335] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550440398] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1550440668] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10374580 kB) free.
[1550441535] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550441598] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550441868] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.3% (5091260 kB) free!
[1550441988] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10188252 kB) free.
[1550442297] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47220 MB (21% inode=99%):
[1550442618] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550442738] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550443188] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4938868 kB) free!
[1550443308] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10142104 kB) free.
[1550445108] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3942484 kB) free!
[1550445228] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9076412 kB) free!
[1550445348] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.7% (5788860 kB) free!
[1550445738] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1550446338] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550446458] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550448297] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1550448349] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9900552 kB) free.
[1550449549] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4699272 kB) free!
[1550449669] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10191672 kB) free.
[1550450058] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550450097] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43221 MB (19% inode=99%):
[1550450178] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550450217] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42325 MB (18% inode=99%):
[1550450298] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550450337] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45267 MB (20% inode=99%):
[1550451469] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4757596 kB) free!
[1550451499] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550451589] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9832040 kB) free!
[1550451709] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.7% (7226332 kB) free!
[1550452099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550452219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550452339] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550454109] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.5% (10122104 kB) free.
[1550457109] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.5% (5643784 kB) free!
[1550457229] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.1% (10897792 kB) free.
[1550460229] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.5% (5669016 kB) free!
[1550460349] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10864808 kB) free.
[1550461184] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (2993368 kB) free!
[1550461304] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (2945568 kB) free!
[1550461424] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.8% (2896252 kB) free!
[1550462149] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.1% (5473964 kB) free!
[1550462269] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.7% (10694080 kB) free.
[1550462390] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.4% (22790972 kB) free.
[1550462443] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.1% (23069160 kB) free.
[1550462522] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.2% (14989520 kB) free.
[1550462570] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 95.0% (23273572 kB) free.
[1550463135] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550463255] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550463375] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550463469] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4754960 kB) free!
[1550463589] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.5% (10114772 kB) free.
[1550463739] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550464189] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1550464339] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550464459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550464580] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550464789] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3858052 kB) free!
[1550464909] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9202152 kB) free!
[1550464945] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 54.8% (26991820 kB) free.
[1550465175] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550465653] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 530 processes
[1550465714] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 530 processes
[1550465775] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550465780] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550465895] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550466015] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550466170] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4845132 kB) free!
[1550466290] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4837328 kB) free!
[1550466380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550466410] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (4829292 kB) free!
[1550466500] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550466600] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.4% (22654384 kB) free.
[1550466798] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550466918] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550467038] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550467100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550467220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550467340] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550467721] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550467815] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 7 warn out of 345 processes
[1550467841] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550467940] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550467961] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550468061] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550468181] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1550468181] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550468415] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 343 processes
[1550469581] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 27.62, 26.76, 12.43
[1550469701] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.91, 18.35, 11.08
[1550469821] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.76, 12.70, 9.90
[1550470354] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (8998412 kB) free!
[1550470474] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10431516 kB) free.
[1550470736] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7800628 kB) free!
[1550470795] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4390364 kB) free!
[1550470856] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11565500 kB) free.
[1550470915] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (4327264 kB) free!
[1550471021] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.46, 1.71, 3.82
[1550471035] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4277492 kB) free!
[1550471075] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.4% (6592708 kB) free!
[1550471181] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550471195] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (3297016 kB) free!
[1550471315] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217796 kB) free!
[1550471456] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8629540 kB) free!
[1550471576] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12218948 kB) free.
[1550471635] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1691428 kB) free!
[1550471781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550471902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550472022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550472176] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8538416 kB) free!
[1550472296] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11869164 kB) free.
[1550472754] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 52.2% (25748220 kB) free.
[1550472807] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1550472807] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.7% (22866628 kB) free.
[1550472896] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9016528 kB) free!
[1550472992] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 44.1% (21714904 kB) free.
[1550473357] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.9% (7842492 kB) free!
[1550473477] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11238352 kB) free.
[1550474078] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8502428 kB) free!
[1550474198] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.2% (12427664 kB) free.
[1550474527] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (165700 kB) free!
[1550474648] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (157148 kB) free!
[1550474768] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174064 kB) free!
[1550474798] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220644 kB) free!
[1550474838] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1550474918] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4688548 kB) free!
[1550475022] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550475039] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (1047360 kB) free!
[1550475047] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195580 kB) free!
[1550475167] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178988 kB) free!
[1550475287] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165860 kB) free!
[1550475622] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550475742] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550475863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550476512] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 581 processes
[1550476517] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1550476573] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 579 processes
[1550476637] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1550476757] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1550477196] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9560384 kB) free!
[1550477316] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7665120 kB) free!
[1550477357] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1550477436] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11902852 kB) free.
[1550478036] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7339752 kB) free!
[1550478052] SERVICE ALERT: sig2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3198384 kB) free!
[1550478156] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10898264 kB) free.
[1550478173] SERVICE ALERT: sig2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3159056 kB) free!
[1550478293] SERVICE ALERT: sig2;Memory;WARNING;HARD;3;WARNING - 19.3% (3129796 kB) free!
[1550478558] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.4% (23140932 kB) free.
[1550478559] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24005 MB (10% inode=99%):
[1550478620] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1550478620] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 50.7% (24982780 kB) free.
[1550478756] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7999240 kB) free!
[1550478760] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.1% (14976740 kB) free.
[1550478855] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 44.0% (21679840 kB) free.
[1550478893] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 93.8% (15239936 kB) free.
[1550478934] SERVICE ALERT: sig4;Memory;WARNING;SOFT;1;WARNING - 19.6% (3177984 kB) free!
[1550479054] SERVICE ALERT: sig4;Memory;WARNING;SOFT;2;WARNING - 19.5% (3161540 kB) free!
[1550479174] SERVICE ALERT: sig4;Memory;WARNING;HARD;3;WARNING - 19.1% (3107188 kB) free!
[1550479227] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5763956 kB) free!
[1550479347] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (590764 kB) free!
[1550479467] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (221756 kB) free!
[1550479653] SERVICE ALERT: sig3;Memory;WARNING;SOFT;1;WARNING - 19.8% (3215344 kB) free!
[1550479773] SERVICE ALERT: sig3;Memory;WARNING;SOFT;2;WARNING - 19.5% (3167776 kB) free!
[1550479894] SERVICE ALERT: sig3;Memory;WARNING;HARD;3;WARNING - 19.3% (3134896 kB) free!
[1550480615] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (1353108 kB) free!
[1550480634] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 48.4% (23864688 kB) free.
[1550480689] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22818272 kB) free.
[1550480788] SERVICE ALERT: sig3;Memory;OK;HARD;1;OK - 93.5% (15196624 kB) free.
[1550480857] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 93.8% (15240156 kB) free.
[1550481234] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (218732 kB) free!
[1550481264] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550481354] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (2103780 kB) free!
[1550481474] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234344 kB) free!
[1550481864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550481896] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1550481984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550482074] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.5% (5155276 kB) free!
[1550482104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550482676] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237360 kB) free!
[1550483069] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8875092 kB) free!
[1550483189] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12513200 kB) free.
[1550483789] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9626556 kB) free!
[1550483796] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 52.2% (25729600 kB) free.
[1550483884] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 43.9% (21622520 kB) free.
[1550484296] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (4267040 kB) free!
[1550484409] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2425328 kB) free!
[1550484416] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.4% (4040356 kB) free!
[1550484529] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4778336 kB) free!
[1550484536] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (3781196 kB) free!
[1550484559] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45875 MB (20% inode=99%):
[1550484649] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230220 kB) free!
[1550484671] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 47.0% (23147912 kB) free.
[1550484714] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.2% (22756000 kB) free.
[1550485258] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3004352 kB) free!
[1550485276] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (292936 kB) free!
[1550485366] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.7% (2717352 kB) free!
[1550485378] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2844564 kB) free!
[1550485396] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (251472 kB) free!
[1550485486] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.9% (2418656 kB) free!
[1550485499] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.6% (2702196 kB) free!
[1550485516] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.5% (6649512 kB) free!
[1550485606] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.1% (2122272 kB) free!
[1550485961] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550486099] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 21.4% (3478432 kB) free.
[1550486116] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (1503872 kB) free!
[1550486206] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (991908 kB) free!
[1550486631] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 47.2% (23252456 kB) free.
[1550486699] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3037268 kB) free!
[1550486819] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2878104 kB) free!
[1550486929] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1550486929] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.8% (15077964 kB) free.
[1550486939] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.7% (2708704 kB) free!
[1550487118] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4917024 kB) free!
[1550487123] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (56.6% change >= 20.0% threshold)
[1550487123] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1550487154] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1550487154] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 88.8% (14432064 kB) free.
[1550487189] SERVICE ALERT: db1;Redis Memory Usage;CRITICAL;SOFT;1;NRPE: Unable to read output
[1550487230] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (56.6% change >= 20.0% threshold)
[1550487231] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1550487231] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1550487241] SERVICE ALERT: db1;Redis Processes;CRITICAL;SOFT;1;NRPE: Unable to read output
[1550487241] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.9% (22686220 kB) free.
[1550487273] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215236 kB) free!
[1550487289] SERVICE ALERT: db2;Redis Memory Usage;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1550487309] SERVICE ALERT: db1;Redis Memory Usage;CRITICAL;SOFT;2;NRPE: Unable to read output
[1550487348] SERVICE ALERT: db2;Redis Processes;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1550487361] SERVICE ALERT: db1;Redis Processes;CRITICAL;SOFT;2;NRPE: Unable to read output
[1550487393] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 50.0% (24634580 kB) free.
[1550487409] SERVICE ALERT: db2;Redis Memory Usage;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1550487429] SERVICE ALERT: db1;Redis Memory Usage;CRITICAL;HARD;3;NRPE: Unable to read output
[1550487468] SERVICE ALERT: db2;Redis Processes;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1550487481] SERVICE ALERT: db1;Redis Processes;CRITICAL;HARD;3;NRPE: Unable to read output
[1550487504] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1550487504] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550487529] SERVICE ALERT: db2;Redis Memory Usage;UNKNOWN;HARD;3;NRPE: Unable to read output
[1550487588] SERVICE ALERT: db2;Redis Processes;UNKNOWN;HARD;3;NRPE: Unable to read output
[1550487754] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 52 users currently logged in
[1550487777] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (55.2% change >= 20.0% threshold)
[1550487777] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1550487777] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1550487777] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1550487804] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3229004 kB) free!
[1550487814] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 52 users currently logged in
[1550487874] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1550487925] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 19.6% (3181472 kB) free!
[1550487934] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 52 users currently logged in
[1550487982] HOST ALERT: sig2;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1550487993] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244500 kB) free!
[1550488045] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 19.1% (3109968 kB) free!
[1550488082] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;HARD;1;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1550488104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550488109] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1550488114] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213976 kB) free!
[1550488198] SERVICE ALERT: sig3;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig3' was out of bounds)
[1550488225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550488234] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (1477824 kB) free!
[1550488320] SERVICE ALERT: sig3;CPU Load;OK;HARD;1;OK - load average: 0.47, 0.21, 0.08
[1550488334] SERVICE ALERT: sig3;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 323 processes
[1550488345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550488454] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (1315532 kB) free!
[1550488455] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 324 processes
[1550488482] SERVICE ALERT: sig3;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550488495] SERVICE ALERT: sig3;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550488497] SERVICE ALERT: sig3;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550488511] SERVICE ALERT: sig3;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig3' was out of bounds)
[1550488511] SERVICE ALERT: sig3;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'sig3' was out of bounds)
[1550488512] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1550488574] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (1178132 kB) free!
[1550488576] SERVICE ALERT: sig3;CPU_Procs;CRITICAL;HARD;2;(Return code of 255 for service 'CPU_Procs' on host 'sig3' was out of bounds)
[1550488600] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1550488603] SERVICE ALERT: sig3;Sensors Status;OK;HARD;1;SENSORS OK
[1550488615] SERVICE ALERT: sig3;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1550488617] SERVICE ALERT: sig3;HDD_Root status;OK;HARD;1;DISK OK - free space: / 32960 MB (64% inode=98%):
[1550488631] SERVICE ALERT: sig3;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 184869 MB (82% inode=99%):
[1550488632] SERVICE ALERT: sig3;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550488665] SERVICE ALERT: sig4;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'sig4' was out of bounds)
[1550488694] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (1004016 kB) free!
[1550488695] SERVICE ALERT: sig4;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550488696] SERVICE ALERT: sig4;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550488702] SERVICE ALERT: sig4;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550488725] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1550488785] SERVICE ALERT: sig4;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4' was out of bounds)
[1550488789] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1550488802] SERVICE ALERT: sig2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 313 processes
[1550488816] SERVICE ALERT: sig4;Sensors Status;OK;HARD;1;SENSORS OK
[1550488817] SERVICE ALERT: sig4;Total Processes;OK;HARD;1;PROCS OK: 314 processes
[1550488823] SERVICE ALERT: sig4;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550488975] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4545968 kB) free!
[1550489095] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4432860 kB) free!
[1550489188] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 3 warn out of 301 processes
[1550489215] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4299568 kB) free!
[1550489308] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 301 processes
[1550489357] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1550489386] SERVICE FLAPPING ALERT: sig4;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1550489386] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 93.0% (15121584 kB) free.
[1550489402] SERVICE ALERT: sig2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 316 processes
[1550489428] SERVICE ALERT: sig4;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 299 processes
[1550489434] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 43.0% (21192188 kB) free.
[1550489454] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1550489454] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.0% (22694296 kB) free.
[1550489470] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 93.6% (15210844 kB) free.
[1550489477] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1550489522] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 2 warn out of 316 processes
[1550489533] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.8% (23241032 kB) free.
[1550489642] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 2 warn out of 314 processes
[1550489762] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 2 warn out of 314 processes
[1550489839] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1550489882] SERVICE ALERT: sig2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 313 processes
[1550490003] SERVICE ALERT: sig2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 315 processes
[1550490034] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1550490123] SERVICE ALERT: sig2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 315 processes
[1550490375] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 309 processes
[1550490634] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (226980 kB) free!
[1550490754] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (4730628 kB) free!
[1550490874] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (558812 kB) free!
[1550490975] SERVICE ALERT: sig3;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 311 processes
[1550491039] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1550491277] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1550491397] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1550491916] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1550491997] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1550492036] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1550492053] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 327 processes
[1550492117] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1550492193] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig1' was out of bounds)
[1550492239] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1550492313] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550492338] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (2951688 kB) free!
[1550492415] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 339 processes
[1550492433] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550492458] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (2722020 kB) free!
[1550492545] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550492578] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (195884 kB) free!
[1550492967] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (210868 kB) free!
[1550493005] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3232988 kB) free!
[1550493034] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 55 users currently logged in
[1550493087] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (216836 kB) free!
[1550493094] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 55 users currently logged in
[1550493125] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 37.2% (6038616 kB) free.
[1550493146] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550493154] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 55 users currently logged in
[1550493207] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (209200 kB) free!
[1550493214] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 55 users currently logged in
[1550493266] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550493303] SERVICE ALERT: sig2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550493306] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1550493376] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1550493386] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550493424] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 94.8% (15411376 kB) free.
[1550493428] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1550493596] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 51.6% (25433272 kB) free.
[1550493634] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.6% (23199304 kB) free.
[1550493680] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.9% (15103628 kB) free.
[1550493729] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550493976] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 315 processes
[1550494096] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 316 processes
[1550494196] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227972 kB) free!
[1550494216] SERVICE FLAPPING ALERT: sig3;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1550494216] SERVICE ALERT: sig3;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 318 processes
[1550494228] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550494317] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (562796 kB) free!
[1550494348] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 45.4% (22395356 kB) free.
[1550494436] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 318 processes
[1550494556] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 320 processes
[1550494676] SERVICE ALERT: sig1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 319 processes
[1550494759] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51373 MB (22% inode=99%):
[1550494816] SERVICE ALERT: sig3;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 316 processes
[1550495558] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.1% (6442756 kB) free!
[1550495614] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 46 users currently logged in
[1550495678] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.6% (15559060 kB) free.
[1550495786] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550496015] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550496278] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8427456 kB) free!
[1550496386] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550496398] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4684280 kB) free!
[1550496507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550496518] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.5% (14515436 kB) free.
[1550496615] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550496627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550496735] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550496855] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550497118] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8460484 kB) free!
[1550497227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550497238] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4757532 kB) free!
[1550497272] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 578 processes
[1550497332] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 578 processes
[1550497347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550497358] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.3% (13963964 kB) free.
[1550497363] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550497392] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 577 processes
[1550497452] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 0 warn out of 578 processes
[1550497467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550497752] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 572 processes
[1550497958] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.0% (5433408 kB) free!
[1550498078] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (1665072 kB) free!
[1550498198] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11431264 kB) free.
[1550498359] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45753 MB (20% inode=99%):
[1550498479] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45757 MB (20% inode=99%):
[1550498599] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45332 MB (20% inode=99%):
[1550498667] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550498798] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4158296 kB) free!
[1550498918] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (361672 kB) free!
[1550499038] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.9% (10811136 kB) free.
[1550499267] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550499387] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550499508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550499638] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4900304 kB) free!
[1550499659] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1550499758] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (1240964 kB) free!
[1550499878] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11074624 kB) free.
[1550500039] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550500455] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550500479] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4923844 kB) free!
[1550500599] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (1258396 kB) free!
[1550500719] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10824780 kB) free.
[1550500785] SERVICE FLAPPING ALERT: sig4;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550501319] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.6% (5232288 kB) free!
[1550501439] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (1579656 kB) free!
[1550501523] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1550501559] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.5% (10576008 kB) free.
[1550501876] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1550502159] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4607564 kB) free!
[1550502279] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11752404 kB) free.
[1550502879] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1550503086] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550503101] SERVICE ALERT: sig1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550503110] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'sig1' was out of bounds)
[1550503110] SERVICE ALERT: sig1;Sensors Status;CRITICAL;SOFT;1;(Return code of 255 for service 'Sensors Status' on host 'sig1' was out of bounds)
[1550503116] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1550503165] SERVICE ALERT: sig1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig1' was out of bounds)
[1550503180] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1550503206] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1550503221] SERVICE ALERT: sig1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37144 MB (72% inode=98%):
[1550503230] SERVICE ALERT: sig1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 147178 MB (65% inode=99%):
[1550503230] SERVICE ALERT: sig1;Sensors Status;OK;HARD;1;SENSORS OK
[1550503286] SERVICE ALERT: sig1;Total Processes;OK;HARD;1;PROCS OK: 321 processes
[1550503289] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1550503326] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1550503336] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550503456] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550503576] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550504308] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550504679] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (277972 kB) free!
[1550504800] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216052 kB) free!
[1550504884] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1049716 kB) free!
[1550504909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 299 processes
[1550504920] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217004 kB) free!
[1550505004] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (212072 kB) free!
[1550505029] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 1 warn out of 305 processes
[1550505038] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186784 kB) free!
[1550505125] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234632 kB) free!
[1550505149] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 302 processes
[1550505158] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192556 kB) free!
[1550505278] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (205760 kB) free!
[1550505520] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 35.1% (17291332 kB) free.
[1550506120] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (298568 kB) free!
[1550506216] SERVICE FLAPPING ALERT: sig3;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550506240] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 35.0% (17266368 kB) free.
[1550507525] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.2% (5529612 kB) free!
[1550508149] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550508659] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.52, 5.54, 5.04
[1550508779] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.64, 5.53, 5.09
[1550508899] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.73, 5.61, 5.18
[1550509847] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1550509911] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1550510416] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1550510776] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550513525] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1293524 kB) free!
[1550513549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550513670] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes