[1550341876] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1550341900] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9322608 kB) free!
[1550341996] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1550342020] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.4% (15457668 kB) free.
[1550342096] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46310 MB (20% inode=99%):
[1550342216] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44260 MB (19% inode=99%):
[1550342336] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47779 MB (21% inode=99%):
[1550342620] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9210268 kB) free!
[1550342647] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550342740] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.3% (15448392 kB) free.
[1550342936] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44251 MB (19% inode=99%):
[1550343056] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48349 MB (21% inode=99%):
[1550343247] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550343340] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7743148 kB) free!
[1550343367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550343461] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.8% (14692028 kB) free.
[1550343487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550343656] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44489 MB (19% inode=99%):
[1550343776] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48586 MB (21% inode=99%):
[1550344376] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44488 MB (19% inode=99%):
[1550344496] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43974 MB (19% inode=99%):
[1550344616] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47562 MB (21% inode=99%):
[1550344660] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9281048 kB) free!
[1550344780] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.3% (15441496 kB) free.
[1550345216] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44488 MB (19% inode=99%):
[1550345287] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550345336] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48586 MB (21% inode=99%):
[1550345380] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8310936 kB) free!
[1550345500] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.9% (15216928 kB) free.
[1550345887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550345936] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44490 MB (19% inode=99%):
[1550346007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550346056] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48586 MB (21% inode=99%):
[1550346100] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.8% (7279276 kB) free!
[1550346127] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550346220] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.8% (14208104 kB) free.
[1550346727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550346847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550346967] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550347420] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8627056 kB) free!
[1550347540] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.9% (15231152 kB) free.
[1550347568] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550347856] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44486 MB (19% inode=99%):
[1550347977] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48583 MB (21% inode=99%):
[1550348140] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8060744 kB) free!
[1550348260] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.6% (14584416 kB) free.
[1550349460] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8145016 kB) free!
[1550349580] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.8% (14681612 kB) free.
[1550350568] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550350688] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550350808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550351380] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.2% (5018288 kB) free!
[1550351409] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550351500] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11324804 kB) free.
[1550352700] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.4% (5118412 kB) free!
[1550352820] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9355380 kB) free!
[1550352940] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.9% (6373556 kB) free!
[1550354141] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.3% (10017816 kB) free.
[1550355341] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.9% (5360856 kB) free!
[1550355461] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10849616 kB) free.
[1550356661] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.2% (5521340 kB) free!
[1550356781] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11227368 kB) free.
[1550357409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550357529] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550359181] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8093088 kB) free!
[1550359301] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.9% (12283488 kB) free.
[1550361101] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7404052 kB) free!
[1550361221] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11199816 kB) free.
[1550362329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550362449] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550363049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550363169] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550363289] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550363576] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44434 MB (19% inode=99%):
[1550363621] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.0% (6397320 kB) free!
[1550363696] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44434 MB (19% inode=99%):
[1550363741] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11657164 kB) free.
[1550363816] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48017 MB (21% inode=99%):
[1550364416] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44432 MB (19% inode=99%):
[1550364536] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48528 MB (21% inode=99%):
[1550365736] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44143 MB (19% inode=99%):
[1550365784] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1608056 kB) free!
[1550365856] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47470 MB (21% inode=99%):
[1550366456] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44398 MB (19% inode=99%):
[1550366576] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48496 MB (21% inode=99%):
[1550367176] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44355 MB (19% inode=99%):
[1550367196] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1550367296] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43843 MB (19% inode=99%):
[1550367316] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1550367416] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47429 MB (21% inode=99%):
[1550367942] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.4% (6128840 kB) free!
[1550368016] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44351 MB (19% inode=99%):
[1550368062] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.4% (11026744 kB) free.
[1550368089] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550368136] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47935 MB (21% inode=99%):
[1550368689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550368736] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44349 MB (19% inode=99%):
[1550368810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550368857] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48446 MB (21% inode=99%):
[1550368930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550369456] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44348 MB (19% inode=99%):
[1550369576] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48446 MB (21% inode=99%):
[1550370176] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44346 MB (19% inode=99%):
[1550370296] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43834 MB (19% inode=99%):
[1550370375] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550370416] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47420 MB (21% inode=99%):
[1550370640] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1550370678] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550370730] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550371016] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44337 MB (19% inode=99%):
[1550371136] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48434 MB (21% inode=99%):
[1550371330] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550371450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550371570] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550371736] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44336 MB (19% inode=99%):
[1550371856] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48433 MB (21% inode=99%):
[1550372262] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9344444 kB) free!
[1550372382] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10543640 kB) free.
[1550372456] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44335 MB (19% inode=99%):
[1550372576] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43823 MB (19% inode=99%):
[1550372696] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47406 MB (21% inode=99%):
[1550373296] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44334 MB (19% inode=99%):
[1550373370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550373416] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48335 MB (21% inode=99%):
[1550373970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550374016] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44228 MB (19% inode=99%):
[1550374090] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550374136] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48326 MB (21% inode=99%):
[1550374210] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550374736] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44228 MB (19% inode=99%):
[1550374856] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43651 MB (19% inode=99%):
[1550374976] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47293 MB (21% inode=99%):
[1550375576] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44217 MB (19% inode=99%):
[1550375697] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48311 MB (21% inode=99%):
[1550376296] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44216 MB (19% inode=99%):
[1550376416] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43706 MB (19% inode=99%):
[1550376536] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47159 MB (20% inode=99%):
[1550377736] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47280 MB (21% inode=99%):
[1550378336] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43950 MB (19% inode=99%):
[1550378410] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550378456] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47661 MB (21% inode=99%):
[1550378982] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3992080 kB) free!
[1550379010] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550379056] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44110 MB (19% inode=99%):
[1550379102] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9068308 kB) free!
[1550379130] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550379159] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (206996 kB) free!
[1550379176] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48206 MB (21% inode=99%):
[1550379222] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.7% (5769996 kB) free!
[1550379404] SERVICE ALERT: sig2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3177528 kB) free!
[1550379524] SERVICE ALERT: sig2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3160928 kB) free!
[1550379644] SERVICE ALERT: sig2;Memory;WARNING;HARD;3;WARNING - 19.4% (3145036 kB) free!
[1550379730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550379776] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44186 MB (19% inode=99%):
[1550379850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550379896] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48281 MB (21% inode=99%):
[1550379970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550380496] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44167 MB (19% inode=99%):
[1550380575] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550380616] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44175 MB (19% inode=99%):
[1550380695] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550380736] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47757 MB (21% inode=99%):
[1550380815] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550380894] SERVICE ALERT: sig4;Memory;WARNING;SOFT;1;WARNING - 19.7% (3203020 kB) free!
[1550381014] SERVICE ALERT: sig4;Memory;WARNING;SOFT;2;WARNING - 19.6% (3187464 kB) free!
[1550381022] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3891888 kB) free!
[1550381134] SERVICE ALERT: sig4;Memory;WARNING;HARD;3;WARNING - 19.5% (3170240 kB) free!
[1550381336] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44111 MB (19% inode=99%):
[1550381456] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48207 MB (21% inode=99%):
[1550381623] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.7% (5258724 kB) free!
[1550382056] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44100 MB (19% inode=99%):
[1550382176] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48194 MB (21% inode=99%):
[1550382615] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550382776] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44090 MB (19% inode=99%):
[1550382896] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44098 MB (19% inode=99%):
[1550382970] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550383016] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47060 MB (20% inode=99%):
[1550383541] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.80, 8.94, 6.50
[1550383570] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550383661] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.57, 6.48, 5.90
[1550383690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550383781] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.28, 4.76, 5.34
[1550383810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550384023] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1550384023] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2584696 kB) free!
[1550384381] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.13, 1.67, 3.37
[1550385213] SERVICE ALERT: sig3;Memory;WARNING;SOFT;1;WARNING - 19.9% (3229516 kB) free!
[1550385223] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.1% (5474440 kB) free!
[1550385333] SERVICE ALERT: sig3;Memory;WARNING;SOFT;2;WARNING - 19.8% (3210688 kB) free!
[1550385453] SERVICE ALERT: sig3;Memory;WARNING;HARD;3;WARNING - 19.6% (3189232 kB) free!
[1550385615] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550385735] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550385855] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550386016] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1550386016] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48040 MB (21% inode=99%):
[1550386616] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45862 MB (20% inode=99%):
[1550386736] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43934 MB (19% inode=99%):
[1550386810] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550386856] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47518 MB (21% inode=99%):
[1550386878] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550386998] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550387023] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (1938464 kB) free!
[1550387118] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550387410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550387456] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43932 MB (19% inode=99%):
[1550387530] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550387576] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48027 MB (21% inode=99%):
[1550388130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550388176] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43926 MB (19% inode=99%):
[1550388223] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.7% (7247504 kB) free!
[1550388250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550388296] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48008 MB (21% inode=99%):
[1550388316] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1550388370] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550388436] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1550388556] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1550388824] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2462856 kB) free!
[1550388896] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45968 MB (20% inode=99%):
[1550389016] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43920 MB (19% inode=99%):
[1550389136] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47500 MB (21% inode=99%):
[1550389424] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5140244 kB) free!
[1550389736] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43917 MB (19% inode=99%):
[1550389856] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48013 MB (21% inode=99%):
[1550390119] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46954 MB (20% inode=99%):
[1550390239] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46937 MB (20% inode=99%):
[1550390359] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46808 MB (20% inode=99%):
[1550390456] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43882 MB (19% inode=99%):
[1550390493] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4776156 kB) free!
[1550390576] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47747 MB (21% inode=99%):
[1550390613] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4736580 kB) free!
[1550390733] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4709576 kB) free!
[1550391176] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45698 MB (20% inode=99%):
[1550391296] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43650 MB (19% inode=99%):
[1550391370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550391416] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47232 MB (21% inode=99%):
[1550391824] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4046380 kB) free!
[1550391970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550392017] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43649 MB (19% inode=99%):
[1550392090] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550392137] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47743 MB (21% inode=99%):
[1550392210] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550392244] SERVICE ALERT: sig2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1603000 kB) free!
[1550392424] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.5% (6174224 kB) free!
[1550392737] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43613 MB (19% inode=99%):
[1550392857] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47708 MB (21% inode=99%):
[1550392993] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 553 processes
[1550393053] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 554 processes
[1550393253] SERVICE ALERT: sig3;Memory;OK;HARD;1;OK - 58.2% (9455304 kB) free.
[1550393444] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 55.3% (8996060 kB) free.
[1550393457] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45684 MB (20% inode=99%):
[1550393577] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43635 MB (19% inode=99%):
[1550393624] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (1465328 kB) free!
[1550393697] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47218 MB (21% inode=99%):
[1550393734] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 56.2% (9130172 kB) free.
[1550394010] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550394297] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43633 MB (19% inode=99%):
[1550394417] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47729 MB (21% inode=99%):
[1550394610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550394730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550394824] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.7% (6754564 kB) free!
[1550394850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550395017] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43751 MB (19% inode=99%):
[1550395137] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47847 MB (21% inode=99%):
[1550395737] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45798 MB (20% inode=99%):
[1550395857] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43750 MB (19% inode=99%):
[1550395977] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47334 MB (21% inode=99%):
[1550396024] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4193316 kB) free!
[1550396577] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43744 MB (19% inode=99%):
[1550396625] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.3% (7030820 kB) free!
[1550396650] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550396697] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47844 MB (21% inode=99%):
[1550396718] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550397251] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550397297] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43747 MB (19% inode=99%):
[1550397371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550397417] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1550397417] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47843 MB (21% inode=99%):
[1550397491] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550398017] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43730 MB (19% inode=99%):
[1550398137] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43353 MB (19% inode=99%):
[1550398257] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46809 MB (20% inode=99%):
[1550398533] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2386048 kB) free!
[1550399025] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2151440 kB) free!
[1550399625] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.3% (5075420 kB) free!
[1550400057] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47448 MB (21% inode=99%):
[1550400255] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550400491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550400657] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45392 MB (20% inode=99%):
[1550400777] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43342 MB (19% inode=99%):
[1550400825] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2976656 kB) free!
[1550400855] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550400897] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47439 MB (21% inode=99%):
[1550400975] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550401092] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550401095] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550401212] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550401332] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550401425] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.2% (5513924 kB) free!
[1550401497] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43334 MB (19% inode=99%):
[1550401617] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42830 MB (19% inode=99%):
[1550401737] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46412 MB (20% inode=99%):
[1550403225] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3069428 kB) free!
[1550403537] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1550403537] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47274 MB (21% inode=99%):
[1550403825] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.3% (5549432 kB) free!
[1550404137] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43170 MB (19% inode=99%):
[1550404257] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43178 MB (19% inode=99%):
[1550404332] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550404377] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46759 MB (20% inode=99%):
[1550404933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550405049] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4912980 kB) free!
[1550405053] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550405169] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4904228 kB) free!
[1550405289] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (4895716 kB) free!
[1550405625] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3138032 kB) free!
[1550405653] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550405773] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550405893] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550406225] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.1% (5452608 kB) free!
[1550406493] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550406613] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550406733] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550408025] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3888328 kB) free!
[1550408440] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550408560] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550408625] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.4% (6595236 kB) free!
[1550408680] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550409280] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550409777] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47279 MB (21% inode=99%):
[1550409825] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4626920 kB) free!
[1550410377] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43303 MB (19% inode=99%):
[1550410425] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.9% (6834756 kB) free!
[1550410480] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550410497] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42350 MB (18% inode=99%):
[1550410518] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 9 users currently logged in
[1550410600] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550410617] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46318 MB (20% inode=99%):
[1550410638] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 9 users currently logged in
[1550410720] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550410758] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1550410933] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550411533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550411626] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2384448 kB) free!
[1550411654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550411774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550411817] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3153044 kB) free!
[1550411937] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3126208 kB) free!
[1550412057] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.1% (3098572 kB) free!
[1550412226] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.3% (5095240 kB) free!
[1550414026] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2524004 kB) free!
[1550414626] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5102896 kB) free!
[1550416426] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (2420664 kB) free!
[1550416575] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550417175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550417295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550417415] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550417626] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.2% (7488792 kB) free!
[1550418615] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550418827] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2572924 kB) free!
[1550419215] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550419335] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550419427] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.0% (5404496 kB) free!
[1550419455] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550419756] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1550419876] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1550420056] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550420176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550420296] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550420627] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2452976 kB) free!
[1550420896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550421016] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550421136] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550421227] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.3% (5552180 kB) free!
[1550421584] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 42.9% (6969300 kB) free.
[1550421677] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1550421797] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1550421884] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 369 processes
[1550421917] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1550422004] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 367 processes
[1550422017] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550422124] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 367 processes
[1550422336] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550422937] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550423027] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3058208 kB) free!
[1550423057] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550423627] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.9% (5389016 kB) free!
[1550423657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550423777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550423897] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550424497] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550424617] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550424737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550424827] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (2587700 kB) free!
[1550425095] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550425121] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1550425158] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550425258] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169792 kB) free!
[1550425427] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.2% (7486828 kB) free!
[1550426627] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10430756 kB) free.
[1550427227] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.9% (5384484 kB) free!
[1550427347] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10409044 kB) free.
[1550427658] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.4% (14857480 kB) free.
[1550427947] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4600076 kB) free!
[1550428067] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10390268 kB) free.