[1578335409] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3518828 kB) free!
[1578335471] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 54.4% (26801780 kB) free.
[1578337425] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184760 kB) free!
[1578337486] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197052 kB) free!
[1578337547] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (164460 kB) free!
[1578337607] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 60.5% (14838088 kB) free.
[1578338724] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26272864 kB) free!
[1578338844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (44039768 kB) free.
[1578340254] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(670) < 103.255.146.164 [07/Jan/2020:01:16:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578340304] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(682) < 2020/01/07 01:16:40 [error] 21677#0: *23093483 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:41/0, bytes from/to upstream:0/41
[1578340374] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578340424] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578341817] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1578341937] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1578342537] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1578342657] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1578342757] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.2% (10355772 kB) free.
[1578343886] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8489012 kB) free!
[1578344006] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.5% (5688328 kB) free!
[1578344126] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (5678876 kB) free!
[1578344294] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3432304 kB) free!
[1578344414] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228204 kB) free!
[1578344535] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214744 kB) free!
[1578344726] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 83.1% (40961044 kB) free.
[1578345135] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.8% (43257532 kB) free.
[1578346857] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(220) < 2020/01/07 03:10:22 [error] 18365#0: *34149595 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578346974] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2830) < 172.105.59.173 [07/Jan/2020:03:11:42 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1578346978] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2020/01/07 03:11:22 [error] 18357#0: *34149618 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578347025] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3255) < 2020/01/07 03:13:17 [error] 21678#0: *23120883 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:115/0, bytes from/to upstream:0/115
[1578347094] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578347098] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578347145] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578348657] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1578348777] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1578357157] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187980 kB) free!
[1578357277] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (160696 kB) free!
[1578357398] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176848 kB) free!
[1578357541] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21990520 kB) free.
[1578357609] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 60.0% (79052148 kB) free.
[1578357662] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.4% (37665000 kB) free.
[1578357703] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 80.1% (19621680 kB) free.
[1578357734] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.2% (22599160 kB) free.
[1578358000] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1578358101] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1578358120] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1578358147] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194288 kB) free!
[1578358196] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180128 kB) free!
[1578358221] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1578358240] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1578358266] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (246912 kB) free!
[1578358267] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189444 kB) free!
[1578358316] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179064 kB) free!
[1578358337] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196692 kB) free!
[1578358341] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578358386] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (264908 kB) free!
[1578358387] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188152 kB) free!
[1578358434] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.5% (21437368 kB) free.
[1578358457] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197208 kB) free!
[1578358481] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.1% (20991516 kB) free.
[1578358506] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236488 kB) free!
[1578358527] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.6% (37756720 kB) free.
[1578358568] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.5% (20709904 kB) free.
[1578358596] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1578358716] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1578358836] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1578359037] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188164 kB) free!
[1578359088] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181296 kB) free!
[1578359138] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220660 kB) free!
[1578359157] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187344 kB) free!
[1578359176] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190284 kB) free!
[1578359208] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188836 kB) free!
[1578359258] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (263648 kB) free!
[1578359277] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190752 kB) free!
[1578359296] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171072 kB) free!
[1578359329] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1578359329] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190304 kB) free!
[1578359378] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1578359378] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250432 kB) free!
[1578359416] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196680 kB) free!
[1578359504] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.1% (22222180 kB) free.
[1578359506] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4384456 kB) free!
[1578359539] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.1% (38488168 kB) free.
[1578359626] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (4146632 kB) free!
[1578359746] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (3657432 kB) free!
[1578360036] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1578360080] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.4% (21911596 kB) free.
[1578360107] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191720 kB) free!
[1578360127] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.1% (22581480 kB) free.
[1578360141] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1578360146] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (262020 kB) free!
[1578360161] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180232 kB) free!
[1578360163] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188464 kB) free!
[1578360197] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21976428 kB) free.
[1578360240] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.1% (38496536 kB) free.
[1578360346] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1578360346] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189840 kB) free!
[1578360804] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192152 kB) free!
[1578360847] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219548 kB) free!
[1578360924] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193760 kB) free!
[1578360955] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22090012 kB) free.
[1578360967] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232200 kB) free!
[1578360988] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.3% (38586040 kB) free.
[1578361026] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.4% (22412376 kB) free.
[1578361240] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578361558] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194280 kB) free!
[1578361597] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236104 kB) free!
[1578361678] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187856 kB) free!
[1578361717] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (261536 kB) free!
[1578361798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197244 kB) free!
[1578361837] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227732 kB) free!
[1578364036] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197868 kB) free!
[1578364156] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185324 kB) free!
[1578364276] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184208 kB) free!
[1578364627] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21862464 kB) free.
[1578364672] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.5% (37687196 kB) free.
[1578364712] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.4% (22398388 kB) free.
[1578365237] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189268 kB) free!
[1578365277] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (264788 kB) free!
[1578365357] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189448 kB) free!
[1578365397] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (253396 kB) free!
[1578365477] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177688 kB) free!
[1578365517] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (281252 kB) free!
[1578368478] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22076056 kB) free.
[1578368518] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 52.1% (25677264 kB) free.
[1578368544] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 773.87, 291.48, 107.27
[1578368664] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 106.18, 195.73, 94.52
[1578368694] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(788) < 103.255.146.164 [07/Jan/2020:09:12:24 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578368745] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(807) < 2020/01/07 09:15:06 [error] 21683#0: *23188411 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:51/17, bytes from/to upstream:17/51
[1578368814] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578368865] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578369087] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197356 kB) free!
[1578369126] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (248808 kB) free!
[1578369207] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193536 kB) free!
[1578369246] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234120 kB) free!
[1578369327] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170144 kB) free!
[1578369366] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (247004 kB) free!
[1578372039] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.1% (38971964 kB) free.
[1578372087] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22071036 kB) free.
[1578372695] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193968 kB) free!
[1578372815] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166288 kB) free!
[1578372935] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181176 kB) free!
[1578373239] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240540 kB) free!
[1578373359] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234260 kB) free!
[1578373479] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238404 kB) free!
[1578373841] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578373961] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578374081] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578374265] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(81) < 2020/01/07 10:47:45 [error] 21679#0: *23205244 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:42306/87, bytes from/to upstream:87/42306
[1578374385] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(83) < 2020/01/07 10:49:44 [error] 21677#0: *23205775 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:34577/104, bytes from/to upstream:104/34577
[1578374506] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(86) < 2020/01/07 10:51:41 [error] 21683#0: *23206297 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:15993/123, bytes from/to upstream:123/15993
[1578374582] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(3) < Channel error on connection <0.22890.4107> (10.147.212.158:46800 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1578374703] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1578375414] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(66) < 103.209.99.7 [07/Jan/2020:11:05:32 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1578375485] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22067616 kB) free.
[1578375511] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578375534] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578375541] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.1% (38484720 kB) free.
[1578376088] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191012 kB) free!
[1578376145] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (242664 kB) free!
[1578376208] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189200 kB) free!
[1578376265] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (262268 kB) free!
[1578376328] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196176 kB) free!
[1578376386] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233580 kB) free!
[1578376499] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/07 11:16:38 [error] 18367#0: *34329153 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1578376619] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578376906] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578377219] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/07 11:36:32 [error] 18370#0: *34339805 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1578377334] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(42) < 103.209.99.7 [07/Jan/2020:11:34:45 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1578377339] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578377454] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578377506] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(72) < 2020/01/07 11:34:45 [error] 21681#0: *23214783 no live upstreams while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "backend_13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578377626] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578378514] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1578378634] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 331 processes
[1578378738] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4516868 kB) free!
[1578378754] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1578378858] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4397768 kB) free!
[1578378978] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (4267248 kB) free!
[1578379204] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.9% (22190368 kB) free.
[1578379242] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.7% (38271424 kB) free.
[1578379807] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186580 kB) free!
[1578379846] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (222080 kB) free!
[1578379928] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192620 kB) free!
[1578379966] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228212 kB) free!
[1578380048] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194080 kB) free!
[1578380086] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (245116 kB) free!
[1578380778] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (215608 kB) free!
[1578382664] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22074928 kB) free.
[1578382708] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.6% (38220560 kB) free.
[1578382744] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.1% (22567800 kB) free.
[1578383026] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/07 13:06:38 [error] 21679#0: *23230719 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:47/0, bytes from/to upstream:0/47
[1578383146] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578383268] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165452 kB) free!
[1578383315] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (246160 kB) free!
[1578383388] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193236 kB) free!
[1578383435] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227592 kB) free!
[1578383508] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187528 kB) free!
[1578383555] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234068 kB) free!
[1578385148] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.1% (2467108 kB) free!
[1578385236] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1578385268] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (2369108 kB) free!
[1578385356] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1578385388] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1578385388] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2291820 kB) free!
[1578387088] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.8% (4845844 kB) free!
[1578387208] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.3% (4721236 kB) free!
[1578387281] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1578387328] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.8% (4619492 kB) free!
[1578387346] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(947) < 2020/01/07 14:22:59 [error] 21673#0: *23246092 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578387466] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578387654] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(934) < 103.255.146.164 [07/Jan/2020:14:22:59 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578387775] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578388335] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9681112 kB) free!
[1578388455] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9543448 kB) free!
[1578388575] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.1% (9409288 kB) free!
[1578388666] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/07 14:45:58 [error] 21677#0: *23250036 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:10943/104, bytes from/to upstream:104/10943
[1578388711] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4821504 kB) free!
[1578388786] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578388831] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4731988 kB) free!
[1578388951] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4619184 kB) free!
[1578388954] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1578389074] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 338 processes
[1578389194] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 332 processes
[1578389794] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1578390928] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1865664 kB) free!
[1578391186] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/07 15:27:57 [error] 21681#0: *23257275 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:2831/70, bytes from/to upstream:70/2831
[1578391306] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578391351] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1544604 kB) free!
[1578391575] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4306524 kB) free!
[1578392081] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1578394306] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/07 16:20:38 [error] 21679#0: *23266301 recv() failed (104: Connection reset by peer) while proxying connection, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:8461/140, bytes from/to upstream:140/8461
[1578394308] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578394355] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578394388] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1578394426] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578394644] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20283944 kB) free!
[1578394764] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41099448 kB) free.
[1578395965] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21649632 kB) free!
[1578396085] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41921972 kB) free.
[1578396174] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(681) < 183.82.105.122 [07/Jan/2020:16:44:06 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578396226] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(694) < 2020/01/07 16:44:06 [error] 21681#0: *23271908 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578396295] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578396346] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578396685] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20267788 kB) free!
[1578396777] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1578396805] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41096536 kB) free.
[1578396898] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1578397405] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19662216 kB) free!
[1578397525] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40378116 kB) free.
[1578398125] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18748728 kB) free!
[1578398146] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/07 17:16:00 [error] 21679#0: *23277715 recv() failed (104: Connection reset by peer) while proxying connection, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:20015/104, bytes from/to upstream:104/20015
[1578398245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39761704 kB) free.
[1578398267] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578398694] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(665) < 183.82.105.122 [07/Jan/2020:17:31:23 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578398726] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9637696 kB) free!
[1578398815] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578398845] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16955556 kB) free!
[1578398847] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9543936 kB) free!
[1578398866] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(681) < 2020/01/07 17:36:12 [error] 21679#0: *23282683 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.154:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578398965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39045368 kB) free.
[1578398967] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (9476884 kB) free!
[1578398986] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2020/01/07 17:38:27 [error] 21677#0: *23283099 connect() failed (111: Connection refused) while connecting to upstream, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "10.147.212.143:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578399106] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578399565] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16828420 kB) free!
[1578399685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37522092 kB) free.
[1578400285] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14836784 kB) free!
[1578400306] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2020/01/07 18:01:45 [error] 21676#0: *23287329 recv() failed (104: Connection reset by peer) while proxying connection, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:40577/104, bytes from/to upstream:104/40577
[1578400405] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.7% (47122076 kB) free.
[1578400426] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(80) < 2020/01/07 18:03:46 [error] 21679#0: *23287834 recv() failed (104: Connection reset by peer) while proxying connection, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:21095/104, bytes from/to upstream:104/21095
[1578400547] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(83) < 2020/01/07 18:05:43 [error] 21672#0: *23288344 recv() failed (104: Connection reset by peer) while proxying connection, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:10835/121, bytes from/to upstream:121/10835
[1578401215] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(161) < 92.42.104.166 [07/Jan/2020:18:11:29 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1578401335] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(141) < 92.42.104.166 [07/Jan/2020:18:18:35 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1578401455] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578401605] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18786184 kB) free!
[1578401725] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13659932 kB) free!
[1578401845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42031600 kB) free.
[1578401939] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/07 18:26:09 [error] 18355#0: *34579948 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: rq-1-1gbps.server, request: "GET /favicon.ico HTTP/1.1", host: "rq-1-1gbps.server:8091"
[1578402059] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578402347] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578402445] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13358384 kB) free!
[1578402565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44706852 kB) free.
[1578402610] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 21 users currently logged in
[1578402730] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 21 users currently logged in
[1578402850] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 21 users currently logged in
[1578403165] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15537072 kB) free!
[1578403259] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(39) < 2020/01/07 18:41:41 [error] 18368#0: *34461027 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:95645/177158, bytes from/to upstream:177158/95645
[1578403285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10966792 kB) free!
[1578403379] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578403405] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43929380 kB) free.
[1578403450] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 20 users currently logged in
[1578404005] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12729924 kB) free!
[1578404125] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8775720 kB) free!
[1578404245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40757392 kB) free.
[1578404845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12070892 kB) free!
[1578404965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42524848 kB) free.
[1578405566] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12623756 kB) free!
[1578405686] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8659704 kB) free!
[1578405806] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41113872 kB) free.
[1578406167] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4741500 kB) free!
[1578406406] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12534292 kB) free!
[1578406526] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41580036 kB) free.
[1578407126] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14142872 kB) free!
[1578407246] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10437660 kB) free!
[1578407366] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40696636 kB) free.
[1578407455] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(625) < 183.82.105.122 [07/Jan/2020:19:56:48 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578407557] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578407575] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578407677] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1578407747] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(625) < 2020/01/07 19:56:48 [error] 21672#0: *23310454 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578407868] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578407966] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11700572 kB) free!
[1578408086] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41516124 kB) free.
[1578408277] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578408397] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578408517] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578408686] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12875872 kB) free!
[1578408806] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9389312 kB) free!
[1578408926] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39681748 kB) free.
[1578409117] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1578409526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10360884 kB) free!
[1578409646] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40929188 kB) free.
[1578409717] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578409838] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578409958] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578410247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11688760 kB) free!
[1578410367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41232740 kB) free.
[1578410579] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2020/01/07 20:51:45 [error] 18361#0: *34591945 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:4804793/1383316, bytes from/to upstream:1383316/4804793
[1578410699] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578410967] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18363692 kB) free!
[1578411087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41059296 kB) free.
[1578411158] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1578411158] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1578411687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19053532 kB) free!
[1578411758] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578411807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40922732 kB) free.
[1578411878] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1578411998] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1578412407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15815128 kB) free!
[1578412527] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15592936 kB) free!
[1578412598] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1578412647] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37397016 kB) free.
[1578413199] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578413247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12896684 kB) free!
[1578413319] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578413367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38261656 kB) free.
[1578413439] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1578413697] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1578413818] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1578413967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11563120 kB) free!
[1578414039] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578414087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38792824 kB) free.
[1578414159] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578414279] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1578414687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12634408 kB) free!
[1578414807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39396208 kB) free.
[1578415407] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (12985100 kB) free!
[1578415528] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39022052 kB) free.
[1578416128] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12354592 kB) free!
[1578416248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38623868 kB) free.
[1578416848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12591368 kB) free!
[1578416968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38227160 kB) free.
[1578417279] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1578417399] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578417467] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(727) < 2020/01/07 22:47:33 [error] 21683#0: *23340919 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578417519] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1578417568] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12763408 kB) free!
[1578417587] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578417688] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33484200 kB) free.
[1578417775] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(714) < 183.82.105.122 [07/Jan/2020:22:47:33 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578417895] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578418288] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12687804 kB) free!
[1578418408] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37455964 kB) free.
[1578418719] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578418839] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578418960] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578419008] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12457892 kB) free!
[1578419128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37112268 kB) free.
[1578419728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12591736 kB) free!
[1578419848] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36824956 kB) free.
[1578420448] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15311032 kB) free!
[1578420568] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36687912 kB) free.
[1578420760] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1578421017] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1578421137] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1578421169] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15238016 kB) free!
[1578421270] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2407528 kB) free!
[1578421289] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36656712 kB) free.
[1578421360] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578421390] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (212844 kB) free!
[1578421510] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218572 kB) free!