[1583778677] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27321124 kB) free.
[1583779033] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192828 kB) free!
[1583779064] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177832 kB) free!
[1583779154] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182572 kB) free!
[1583779184] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174456 kB) free!
[1583779274] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177948 kB) free!
[1583779277] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5475108 kB) free!
[1583779296] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201444 kB) free!
[1583779304] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185876 kB) free!
[1583779397] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36377548 kB) free.
[1583779416] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202512 kB) free!
[1583779537] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198880 kB) free!
[1583779997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12006776 kB) free!
[1583780117] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8084612 kB) free!
[1583780237] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4265108 kB) free!
[1583780730] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192896 kB) free!
[1583780791] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180764 kB) free!
[1583780837] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14328224 kB) free!
[1583780851] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (184140 kB) free!
[1583780911] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (165348 kB) free!
[1583781211] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.9% (17127536 kB) free.
[1583781674] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.8% (11297188 kB) free.
[1583782037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36342312 kB) free.
[1583782042] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(130) < 92.42.104.166 [10/Mar/2020:00:55:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583782163] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583782504] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4396556 kB) free!
[1583782637] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13074712 kB) free!
[1583782757] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9234584 kB) free!
[1583782874] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1853708 kB) free!
[1583782877] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36321436 kB) free.
[1583782994] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (261052 kB) free!
[1583783114] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181524 kB) free!
[1583783477] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14129424 kB) free!
[1583783597] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10300952 kB) free!
[1583783715] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229912 kB) free!
[1583783717] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36342116 kB) free.
[1583783835] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238968 kB) free!
[1583783955] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235584 kB) free!
[1583784318] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14920612 kB) free!
[1583784438] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11131884 kB) free!
[1583784558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36358828 kB) free.
[1583784754] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1456204 kB) free!
[1583784874] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1316772 kB) free!
[1583784994] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1536528 kB) free!
[1583785158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14450980 kB) free!
[1583785278] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10806508 kB) free!
[1583785398] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36083236 kB) free.
[1583785942] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2020/03/10 02:01:56 [error] 12243#0: *67900465 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1583785998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13705876 kB) free!
[1583786062] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(49) < 2020/03/10 02:03:35 [error] 12242#0: *67900962 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583786118] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10162844 kB) free!
[1583786182] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583786238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33705240 kB) free.
[1583786355] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 52.0% (25640876 kB) free.
[1583786362] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(154) < 92.42.104.166 [10/Mar/2020:02:01:02 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583786482] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583786838] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13142792 kB) free!
[1583786958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36373640 kB) free.
[1583787120] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235536 kB) free!
[1583787338] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 45.0% (11024924 kB) free.
[1583787558] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15504156 kB) free!
[1583787678] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11849328 kB) free!
[1583787798] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35125992 kB) free.
[1583788321] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.4% (42099100 kB) free.
[1583788398] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13946888 kB) free!
[1583788518] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36302468 kB) free.
[1583789119] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15252592 kB) free!
[1583789239] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11641416 kB) free!
[1583789359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33595424 kB) free.
[1583789482] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3718) < 103.16.101.72 [10/Mar/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1583789504] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (2534952 kB) free!
[1583789602] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(22) < 172.105.59.173 [10/Mar/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583789722] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583789727] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(252) < 103.16.101.72 [10/Mar/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583789782] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(301) < 2020/03/10 03:00:22 [error] 12254#0: *67915593 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583789848] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583789902] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583789959] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13285016 kB) free!
[1583790079] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35791532 kB) free.
[1583790679] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13915032 kB) free!
[1583790799] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35671944 kB) free.
[1583791304] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2432788 kB) free!
[1583791399] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14593972 kB) free!
[1583791519] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36250212 kB) free.
[1583792119] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15638188 kB) free!
[1583792239] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36253568 kB) free.
[1583792839] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16317196 kB) free!
[1583792959] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34664452 kB) free.
[1583793419] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1583793539] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1583793559] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14348884 kB) free!
[1583793679] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30949504 kB) free.
[1583794880] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11211008 kB) free!
[1583795000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35618604 kB) free.
[1583795600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11346472 kB) free!
[1583795720] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583795720] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35591544 kB) free.
[1583796320] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12637516 kB) free!
[1583796440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35610980 kB) free.
[1583796939] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170556 kB) free!
[1583797040] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12947128 kB) free!
[1583797059] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179968 kB) free!
[1583797160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36610112 kB) free.
[1583797179] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185196 kB) free!
[1583797760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17731188 kB) free!
[1583797880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36338368 kB) free.
[1583797904] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.9% (19096688 kB) free.
[1583798114] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.2% (10652852 kB) free.
[1583798480] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17532332 kB) free!
[1583798600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35884892 kB) free.
[1583799800] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20077060 kB) free!
[1583799920] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36287692 kB) free.
[1583800920] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 89.7% (118237032 kB) free.
[1583800960] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.4% (39152920 kB) free.
[1583800997] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.1% (20847092 kB) free.
[1583801720] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18685372 kB) free!
[1583801723] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(185) < 203.212.70.170 [10/Mar/2020:06:21:08 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583801840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36735256 kB) free.
[1583801843] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583801902] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2020/03/10 06:21:16 [error] 12254#0: *67964145 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1583802022] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583802314] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195308 kB) free!
[1583802434] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198780 kB) free!
[1583802554] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (209444 kB) free!
[1583803040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19388472 kB) free!
[1583803160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36238196 kB) free.
[1583804126] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21949632 kB) free.
[1583804360] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18580956 kB) free!
[1583804480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36300724 kB) free.
[1583806280] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19690560 kB) free!
[1583806400] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36033292 kB) free.
[1583808201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17908216 kB) free!
[1583808321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34898816 kB) free.
[1583810121] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18010200 kB) free!
[1583810241] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34457956 kB) free.
[1583811441] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14036076 kB) free!
[1583811561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33355400 kB) free.
[1583812761] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11876792 kB) free!
[1583812881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32356540 kB) free.
[1583813481] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9946868 kB) free!
[1583813601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31262328 kB) free.
[1583814201] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8340996 kB) free!
[1583814322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29340288 kB) free.
[1583815521] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6652348 kB) free!
[1583815533] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4651280 kB) free!
[1583815641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2269304 kB) free!
[1583815653] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4514652 kB) free!
[1583815761] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20769712 kB) free!
[1583815773] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4378000 kB) free!
[1583816361] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383364 kB) free!
[1583816373] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1927424 kB) free!
[1583816961] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15102140 kB) free!
[1583817404] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.2% (4211764 kB) free!
[1583817525] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.3% (4006332 kB) free!
[1583817645] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.5% (3800748 kB) free!
[1583818585] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1583818705] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 705 processes
[1583818761] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1583818761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27417720 kB) free.
[1583818825] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1583818845] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (1080920 kB) free!
[1583818964] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (7669712 kB) free!
[1583819084] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.5% (7632572 kB) free!
[1583819204] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (8464716 kB) free!
[1583819361] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2891320 kB) free!
[1583819425] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1583819459] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (3872216 kB) free!
[1583819481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31910532 kB) free.
[1583819579] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3602856 kB) free!
[1583819699] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3337408 kB) free!
[1583820081] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7648384 kB) free!
[1583820201] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33215356 kB) free.
[1583820299] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2114704 kB) free!
[1583820802] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8301664 kB) free!
[1583820922] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32146568 kB) free.
[1583821522] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11165472 kB) free!
[1583821605] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4137212 kB) free!
[1583821642] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7071376 kB) free!
[1583821762] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31039500 kB) free.
[1583822099] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.3% (15345364 kB) free.
[1583822362] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5169100 kB) free!
[1583822482] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30714232 kB) free.
[1583822504] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172944 kB) free!
[1583822624] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202368 kB) free!
[1583822743] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583822744] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199892 kB) free!
[1583822842] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583822850] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 85.6% (20971060 kB) free.
[1583822863] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583822902] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.9% (20699940 kB) free.
[1583822933] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.4% (39127852 kB) free.
[1583822939] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1583822962] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1583822984] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583822998] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 69.7% (17091544 kB) free.
[1583823059] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1583823082] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8592760 kB) free!
[1583823082] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1583823202] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4236724 kB) free!
[1583823322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28799000 kB) free.
[1583823395] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (238304 kB) free!
[1583823515] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (314536 kB) free!
[1583823635] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (319780 kB) free!
[1583823922] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3530988 kB) free!
[1583824042] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29383468 kB) free.
[1583824102] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166000 kB) free!
[1583824204] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (213420 kB) free!
[1583824222] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182676 kB) free!
[1583824259] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1583824324] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186860 kB) free!
[1583824342] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1583824342] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172140 kB) free!
[1583824379] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1583824444] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177492 kB) free!
[1583824500] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1583824642] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (4993560 kB) free!
[1583824762] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (565724 kB) free!
[1583824882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26832960 kB) free.
[1583825004] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4425164 kB) free!
[1583825125] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (4300676 kB) free!
[1583825245] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (4190740 kB) free!
[1583825483] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1629948 kB) free!
[1583825603] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29443620 kB) free.
[1583826203] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4821332 kB) free!
[1583826323] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384712 kB) free!
[1583826443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26799376 kB) free.
[1583827043] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1509700 kB) free!
[1583827045] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1601132 kB) free!
[1583827144] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8393676 kB) free!
[1583827163] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29627316 kB) free.
[1583827184] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1583827264] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (7908984 kB) free!
[1583827384] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.5% (7616332 kB) free!
[1583827444] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1583827444] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 57.2% (14019096 kB) free.
[1583827763] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4158344 kB) free!
[1583827883] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (415216 kB) free!
[1583828003] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27375032 kB) free.
[1583828603] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1873852 kB) free!
[1583828723] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30183972 kB) free.
[1583829323] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6023436 kB) free!
[1583829443] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2089516 kB) free!
[1583829563] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28508764 kB) free.
[1583830164] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5652184 kB) free!
[1583830284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29369172 kB) free.
[1583830321] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9580776 kB) free!
[1583830441] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9514164 kB) free!
[1583830561] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.0% (20484760 kB) free.
[1583830561] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.1% (9431520 kB) free!
[1583830598] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.7% (20518504 kB) free.
[1583830625] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.1% (14161696 kB) free.
[1583830884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5387716 kB) free!
[1583830984] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (508316 kB) free!
[1583831004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28949968 kB) free.
[1583831604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6374132 kB) free!
[1583831724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2115696 kB) free!
[1583831844] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583831844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28061792 kB) free.
[1583832443] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(170) < 92.42.104.166 [10/Mar/2020:14:50:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583832444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3978668 kB) free!
[1583832563] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583832564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28060660 kB) free.
[1583833163] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(149) < 92.42.104.166 [10/Mar/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583833164] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2139032 kB) free!
[1583833283] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583833284] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (442700 kB) free!
[1583833404] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31384416 kB) free.
[1583834004] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7539920 kB) free!
[1583834109] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.0% (38942332 kB) free.
[1583834124] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2909720 kB) free!
[1583834147] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 50.6% (66732096 kB) free.
[1583834763] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183252 kB) free!
[1583834883] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184820 kB) free!
[1583835003] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190788 kB) free!
[1583835026] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1583835146] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1583835266] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1583835313] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9576636 kB) free!
[1583835433] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9467060 kB) free!
[1583835553] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1583835553] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (9370344 kB) free!
[1583835866] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1583835961] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4701832 kB) free!
[1583836444] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1583836466] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1583836586] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1583836706] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1583837154] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23782352 kB) free!
[1583837274] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19522288 kB) free!
[1583837306] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583837306] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1583837394] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46637168 kB) free.
[1583837907] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1583837994] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22438004 kB) free!
[1583838027] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 707 processes
[1583838114] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.6% (46948696 kB) free.
[1583838147] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1583838553] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3997824 kB) free!
[1583838714] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23730048 kB) free!
[1583838747] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1583838834] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19339320 kB) free!
[1583838954] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46292760 kB) free.
[1583839555] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21076656 kB) free!
[1583839675] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (16641856 kB) free!
[1583839795] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44829148 kB) free.
[1583840395] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19221108 kB) free!
[1583840515] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14881828 kB) free!
[1583840635] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28089188 kB) free.
[1583841235] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3487888 kB) free!
[1583841355] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31390104 kB) free.
[1583841955] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5870512 kB) free!
[1583842052] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.0% (4653364 kB) free!
[1583842075] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1613424 kB) free!
[1583842172] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.6% (4554056 kB) free!
[1583842195] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29078840 kB) free.
[1583842204] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1583842293] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.2% (4455924 kB) free!
[1583842795] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4284092 kB) free!
[1583842915] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30486560 kB) free.
[1583843428] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9738004 kB) free!
[1583843515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6300276 kB) free!
[1583843548] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9658372 kB) free!
[1583843635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1857536 kB) free!
[1583843668] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9579312 kB) free!
[1583843756] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30066956 kB) free.
[1583844355] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4442264 kB) free!
[1583844437] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.8% (20672472 kB) free.
[1583844475] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384660 kB) free!
[1583844493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.5% (53409512 kB) free.
[1583844535] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.2% (39028916 kB) free.
[1583844577] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (45221868 kB) free.
[1583844682] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583844693] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194136 kB) free!
[1583844802] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1583844849] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1583844914] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 13.4% (3289524 kB) free!
[1583844969] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1583845034] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 9.1% (2221296 kB) free!
[1583845089] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1583845154] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 1.2% (295012 kB) free!
[1583845402] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583845523] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1583845643] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181036 kB) free!
[1583845643] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583845689] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1583845763] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174160 kB) free!
[1583845883] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171892 kB) free!
[1583846099] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1583846219] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1583846243] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1583846289] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1583846345] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.2% (5026376 kB) free!
[1583846409] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1583846465] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.2% (5026556 kB) free!
[1583846493] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 69.1% (16929080 kB) free.
[1583846529] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1583846585] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (4969648 kB) free!
[1583846644] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (3168424 kB) free!
[1583846764] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.7% (3104596 kB) free!
[1583846884] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (3043976 kB) free!
[1583846903] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6735528 kB) free!
[1583847023] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1482400 kB) free!
[1583847130] SERVICE FLAPPING ALERT: secondary-server;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583847130] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1583847144] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39510508 kB) free.
[1583847185] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4804024 kB) free!
[1583847743] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6761264 kB) free!
[1583847863] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (937164 kB) free!
[1583847983] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42911964 kB) free.
[1583848084] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2342560 kB) free!
[1583848330] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1583848450] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1583848570] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1583848583] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9112772 kB) free!
[1583848604] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4717588 kB) free!
[1583848703] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3864436 kB) free!
[1583848724] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4633672 kB) free!
[1583848754] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 63.0% (15442788 kB) free.
[1583848823] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381816 kB) free!
[1583848845] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (4541352 kB) free!
[1583849170] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1583850023] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25752292 kB) free!
[1583850147] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583850623] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2318456 kB) free!
[1583850868] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4754304 kB) free!
[1583851823] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1583851823] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30074936 kB) free.
[1583852445] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2341832 kB) free!
[1583853023] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14842396 kB) free!
[1583853143] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41657056 kB) free.
[1583853744] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18400960 kB) free!
[1583853864] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (14488588 kB) free!
[1583853984] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41517264 kB) free.
[1583854584] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17553012 kB) free!
[1583854704] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (14043584 kB) free!
[1583854824] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41365688 kB) free.
[1583855424] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17195260 kB) free!
[1583855544] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13552908 kB) free!
[1583855664] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39042936 kB) free.
[1583856264] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16603320 kB) free!
[1583856384] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41240000 kB) free.
[1583856984] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18942232 kB) free!
[1583857104] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15347180 kB) free!
[1583857224] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40036164 kB) free.
[1583857824] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16558380 kB) free!
[1583857945] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41023916 kB) free.
[1583858544] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17010616 kB) free!
[1583858585] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583858664] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41113332 kB) free.
[1583859264] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18427828 kB) free!
[1583859384] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14588348 kB) free!
[1583859504] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39883756 kB) free.
[1583860104] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15074892 kB) free!
[1583860224] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39670488 kB) free.
[1583860570] SERVICE FLAPPING ALERT: secondary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583860824] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16648320 kB) free!
[1583860944] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40775740 kB) free.
[1583861544] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18406400 kB) free!
[1583861664] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28670024 kB) free.
[1583862083] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(164) < 92.42.104.166 [10/Mar/2020:23:05:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583862204] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583862264] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20576844 kB) free!
[1583862384] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16910720 kB) free!
[1583862504] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39914108 kB) free.
[1583863104] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18280172 kB) free!
[1583863225] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40681352 kB) free.
[1583863824] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18782436 kB) free!
[1583863944] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40676568 kB) free.
[1583864544] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19316724 kB) free!
[1583864664] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40691796 kB) free.
[1583864870] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 367 processes
[1583864990] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes