[1548873316] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548873436] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548873556] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548875996] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1548876556] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548880222] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548880343] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548880942] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548881062] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548881662] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548881782] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548883796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548883916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548884036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548884606] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9659388 kB) free!
[1548884636] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548884726] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12468972 kB) free.
[1548885382] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548885502] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548886102] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548886222] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548888022] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548888143] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548889942] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548890062] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548891236] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548891357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548891477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548891862] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548891983] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548892077] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1548892077] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548892582] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548892702] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548894326] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9574488 kB) free!
[1548894447] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12464948 kB) free.
[1548894502] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548894622] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548895222] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548895342] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548897142] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548897262] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548898462] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548898582] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1548899277] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548899397] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548899518] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548900982] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548901102] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548901702] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548901822] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548902422] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548902542] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548902662] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548903262] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548903383] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1548903503] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548903718] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548904646] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9696064 kB) free!
[1548904703] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548904766] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12406164 kB) free.
[1548905303] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 331 processes
[1548905423] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 331 processes
[1548905517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548905543] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1548905637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548905757] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548906143] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1548906143] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 333 processes
[1548906357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548906477] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548906744] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1548907078] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548907198] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548907318] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548907344] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 333 processes
[1548907767] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9317092 kB) free!
[1548907887] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12344808 kB) free.
[1548908210] SERVICE ALERT: db2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5514 MB (10% inode=98%):
[1548908518] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548908544] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548908638] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548908758] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548908810] SERVICE ALERT: db2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7100 MB (13% inode=98%):
[1548909144] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548909264] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548909384] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548909959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548909984] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 332 processes
[1548910079] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548910104] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 334 processes
[1548910224] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548910679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548910799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548910919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548911424] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548912119] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1548912624] SERVICE ALERT: db1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1548912719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1548912744] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 333 processes
[1548912840] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548912864] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 330 processes
[1548913439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548913464] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548913518] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 25.4% (4124972 kB) free.
[1548913559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548913679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548914110] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548914230] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548914350] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548914665] SERVICE ALERT: db1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1548914785] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548915263] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1548915318] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3050212 kB) free!
[1548915385] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548915438] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2981204 kB) free!
[1548915505] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548915558] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.9% (2909528 kB) free!
[1548915725] SERVICE ALERT: secondary-server;Nginx Status;OK;HARD;1;NGINX OK -  0.087 sec. response time, Active: 0 (Writing: 0 Reading: 0 Waiting: 0) ReqPerSec: 0.000 ConnPerSec: 0.000 ReqPerConn: 0.000
[1548916150] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548916705] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548916825] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548917350] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548917425] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548917470] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548917545] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548917590] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548917957] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548918077] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548918126] SERVICE ALERT: secondary-server;Nginx Status;CRITICAL;SOFT;1;NGINX CRITICAL - 500 Can't connect to 10.147.212.153:80 (Connection refused)
[1548918146] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548918197] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548918246] SERVICE ALERT: secondary-server;Nginx Status;OK;HARD;1;NGINX OK -  0.083 sec. response time, Active: 0 (Writing: 0 Reading: 0 Waiting: 0) ReqPerSec: 0.000 ConnPerSec: 0.000 ReqPerConn: 0.000
[1548918266] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548918479] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1548918479] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1548918565] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (3827516 kB) free!
[1548918866] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548918986] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548919079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548919199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548919319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548919397] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548919997] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548920117] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548920365] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1864640 kB) free!
[1548920487] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9510044 kB) free!
[1548920519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1548920607] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12075832 kB) free.
[1548920692] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.8% (3225484 kB) free!
[1548920966] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (3671560 kB) free!
[1548921120] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548921207] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9835000 kB) free!
[1548921240] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1548921317] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548921327] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12288572 kB) free.
[1548921437] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548921557] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548921840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1548921960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1548922157] SERVICE FLAPPING ALERT: webservercluster2;Current Users;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1548922157] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548922560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548922586] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548922680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548922706] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1548922800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548922990] SERVICE FLAPPING ALERT: db2;Current Users;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1548922990] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548923306] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1548923426] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548924048] SERVICE ALERT: primary-server;Nginx Status;CRITICAL;SOFT;1;NGINX CRITICAL - 500 Can't connect to 10.147.212.141:80 (Connection refused)
[1548924168] SERVICE ALERT: primary-server;Nginx Status;OK;HARD;1;NGINX OK -  0.089 sec. response time, Active: 0 (Writing: 0 Reading: 0 Waiting: 0) ReqPerSec: 0.000 ConnPerSec: 0.000 ReqPerConn: 0.000
[1548924927] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9569936 kB) free!
[1548925047] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12052320 kB) free.
[1548925159] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1612420 kB) free!
[1548925826] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548925947] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548925947] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548927000] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1548927146] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548927266] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1548927600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548927720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548927841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548928167] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2199248 kB) free!
[1548928466] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548928492] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (1478540 kB) free!
[1548928586] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548928610] SERVICE ALERT: db2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5503 MB (10% inode=98%):
[1548929590] SERVICE FLAPPING ALERT: db2;Current Users;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1548929641] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548929957] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548930077] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548930197] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548930241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548930361] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548930961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548931082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548932186] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548932281] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548932306] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548932401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548932521] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548933071] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47175 MB (20% inode=99%):
[1548933191] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47158 MB (20% inode=99%):
[1548933311] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47094 MB (20% inode=99%):
[1548933721] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548934106] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548934227] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548934321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548934347] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548934442] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548934647] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9668648 kB) free!
[1548934768] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11852252 kB) free.
[1548935041] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548935161] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548935367] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9603820 kB) free!
[1548935487] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11776712 kB) free.
[1548935761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548935881] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1548936087] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9701968 kB) free!
[1548936197] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548936207] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11705900 kB) free.
[1548936481] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548936601] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548936807] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9341784 kB) free!
[1548936927] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11629160 kB) free.
[1548937201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548937322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548937442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548937527] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8965164 kB) free!
[1548937647] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11562176 kB) free.
[1548938247] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8894564 kB) free!
[1548938367] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11385660 kB) free.
[1548938642] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1548938967] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8672872 kB) free!
[1548939087] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10747100 kB) free.
[1548939242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548939362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548939482] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1548940010] SERVICE ALERT: db2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 5948 MB (11% inode=98%):
[1548940082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548940203] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1548940802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548940922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548941042] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1548941488] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8761412 kB) free!
[1548941608] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7708780 kB) free!
[1548941728] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10453280 kB) free.
[1548941810] SERVICE ALERT: db2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 3983 MB (7% inode=97%):
[1548942198] SERVICE FLAPPING ALERT: webservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1548942242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548942328] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9174504 kB) free!
[1548942362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548942448] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8125212 kB) free!
[1548942482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548942568] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6910384 kB) free!
[1548943768] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10068560 kB) free.
[1548944882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1548946082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548946202] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1548947968] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7295340 kB) free!
[1548948002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548948088] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 20.0% (9850788 kB) free!
[1548948122] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548948208] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.6% (8669604 kB) free!
[1548948242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548949408] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1548949408] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10028456 kB) free.
[1548949442] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548951747] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548951867] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548951987] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548952442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548952562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548952588] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548952682] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548952997] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548953008] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9516744 kB) free!
[1548953117] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548953128] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7921184 kB) free!
[1548953237] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548953248] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10001020 kB) free.
[1548955588] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548955682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548955708] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548955802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548955828] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548955922] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548956249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9544308 kB) free!
[1548956369] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7937156 kB) free!
[1548956428] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1548956428] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548956489] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (9989980 kB) free.
[1548957628] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548957748] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548958348] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548958468] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548958588] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548958923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548959043] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548959163] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes