[1546454839] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 6.53, 75.92, 50.01
[1546454959] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.63, 51.07, 44.04
[1546455079] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1546455079] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.75, 34.42, 38.80
[1546455324] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.1% (2289596 kB) free!
[1546455484] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1546456084] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1546456204] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1546456279] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.26, 1.69, 11.41
[1546456324] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1546456556] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546456676] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546456766] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1546456886] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1546457479] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.82, 1.05, 3.89
[1546457553] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546457673] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546457876] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546457996] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546458116] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546462916] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 11 users currently logged in
[1546467079] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8.94, 64.29, 39.15
[1546467199] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.92, 43.29, 34.50
[1546467319] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.95, 29.22, 30.41
[1546467919] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.62, 4.54, 16.28
[1546469119] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.69, 4.95
[1546474519] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.17, 37.52, 33.15
[1546474639] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.07, 25.37, 29.22
[1546474760] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.58, 17.15, 25.74
[1546475360] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.59, 2.97, 13.87
[1546476560] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.80, 0.75, 4.31
[1546481360] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 110.43, 23.45, 8.13
[1546481480] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 162.25, 129.87, 53.47
[1546481600] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.64, 87.21, 47.10
[1546482800] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.81, 2.29, 13.47
[1546484000] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.85, 0.91, 4.31
[1546493309] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 363 processes
[1546493609] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8078548 kB) free!
[1546493673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546493729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10162512 kB) free.
[1546493793] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546493909] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 370 processes
[1546494800] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1546495916] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546496000] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 175.99, 145.14, 60.03
[1546496036] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546496120] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 24.54, 97.38, 52.84
[1546496156] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 17 users currently logged in
[1546496241] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.26, 65.51, 46.56
[1546496729] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7817504 kB) free!
[1546496849] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10160532 kB) free.
[1546497356] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 11 users currently logged in
[1546497441] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 1.88, 13.33
[1546498641] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.63, 0.80, 4.30
[1546500356] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546500476] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546500596] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546502996] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1546502996] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546503596] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546503716] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546503836] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546504041] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 669.64, 172.85, 58.69
[1546504161] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 107.30, 131.44, 58.18
[1546504281] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1546504281] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.20, 88.25, 51.24
[1546504593] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546504650] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7314404 kB) free!
[1546504713] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1546504770] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10156248 kB) free.
[1546505481] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.83, 2.48, 14.77
[1546506681] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.91, 0.89, 4.70
[1546508036] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546511371] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9712276 kB) free!
[1546511481] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 605.95, 152.05, 51.51
[1546511491] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8278804 kB) free!
[1546511601] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 98.32, 117.40, 51.83
[1546511611] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9938520 kB) free.
[1546511721] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.08, 78.89, 45.67
[1546512513] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546512633] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546512753] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1546512836] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 17 users currently logged in
[1546512921] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.81, 2.19, 13.13
[1546512956] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 17 users currently logged in
[1546513076] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 17 users currently logged in
[1546514121] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.78, 0.88, 4.26
[1546514486] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1546514606] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1546516411] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9696944 kB) free!
[1546516531] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.7% (8719256 kB) free!
[1546516652] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7390332 kB) free!
[1546517852] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9915520 kB) free.
[1546518153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546518273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546518393] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1546519406] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1546519524] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1468876 kB) free!
[1546519526] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1546520124] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 11.3% (1828924 kB) free!
[1546520276] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1546521921] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.03, 11.73, 10.18
[1546522041] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.28, 10.46, 9.93
[1546522161] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.25, 7.26, 8.83
[1546522761] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.30, 1.71, 4.99
[1546523124] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1592628 kB) free!
[1546523361] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.89, 4.76, 5.01
[1546523481] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.76, 3.42, 4.50
[1546523852] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9437248 kB) free!
[1546523876] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546523972] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7961388 kB) free!
[1546524092] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9900628 kB) free.
[1546524394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546524514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546524634] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546525881] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.54, 4.94, 3.25
[1546526001] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.15, 3.52, 2.93
[1546526726] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1546526846] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1546527692] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9733432 kB) free!
[1546527812] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8386372 kB) free!
[1546527932] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.9% (9812012 kB) free!
[1546528234] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546528355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546528475] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546529075] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546532676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546532796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546532916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546533516] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1546533516] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546534401] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546537716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546537836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546537956] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes