[1570041920] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (205608 kB) free!
[1570042041] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191708 kB) free!
[1570042161] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177496 kB) free!
[1570045437] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 449 processes
[1570045926] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.3% (5075636 kB) free!
[1570046046] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4185584 kB) free!
[1570046166] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3314168 kB) free!
[1570046641] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.02, 14.95, 7.71
[1570046761] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.38, 10.20, 6.85
[1570046881] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 6.98, 6.07
[1570046893] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.12, 7.18, 6.80
[1570046932] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.99, 6.81, 6.11
[1570047013] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.64, 6.39, 6.55
[1570047052] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.70, 4.80, 5.46
[1570047133] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.29, 5.60, 6.24
[1570047172] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.93, 3.55, 4.92
[1570047482] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.84, 1.59, 3.56
[1570047667] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4148668 kB) free!
[1570047733] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.56, 3.43, 4.76
[1570048867] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1452452 kB) free!
[1570049828] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (258976 kB) free!
[1570049882] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 5.50, 35.57, 22.79
[1570049948] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240436 kB) free!
[1570050002] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.03, 23.90, 20.07
[1570050068] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227748 kB) free!
[1570050122] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.32, 16.11, 17.68
[1570050668] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.3% (45006636 kB) free.
[1570051542] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 297 processes
[1570051662] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1570051922] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1570051922] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.12, 0.85, 3.09
[1570053066] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 75.2% (18426916 kB) free.
[1570053117] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 50.1% (12291164 kB) free.
[1570058168] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1570058288] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1570062122] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1570063901] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 431 processes
[1570063962] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 428 processes
[1570064517] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170732 kB) free!
[1570064637] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199220 kB) free!
[1570064758] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174060 kB) free!
[1570065066] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200344 kB) free!
[1570065187] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175008 kB) free!
[1570065307] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185032 kB) free!
[1570066661] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 421 processes
[1570066721] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 426 processes
[1570066781] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 439 processes
[1570067052] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 369 processes
[1570067172] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1570073637] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 572 processes
[1570073757] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 632 processes
[1570073878] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 631 processes
[1570077657] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 95.1% (23303776 kB) free.
[1570077709] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.6% (23579540 kB) free.
[1570077776] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.5% (22928712 kB) free.
[1570077948] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 59.1% (77893620 kB) free.
[1570078014] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39807804 kB) free.
[1570078080] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.5% (23396948 kB) free.
[1570078225] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 93.2% (22845960 kB) free.
[1570078342] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 95.3% (46957552 kB) free.
[1570078822] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (207196 kB) free!
[1570078942] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (209876 kB) free!
[1570079062] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (209384 kB) free!
[1570079288] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1570079409] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1570079517] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188916 kB) free!
[1570079637] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194156 kB) free!
[1570079757] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172664 kB) free!
[1570080836] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.45, 9.18, 6.42
[1570080954] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1570080956] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.34, 10.32, 7.18
[1570081075] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1570081076] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 14.60, 11.81, 8.10
[1570081195] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1570081384] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (1454516 kB) free!
[1570081504] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (204416 kB) free!
[1570081624] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201504 kB) free!
[1570082372] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.8% (15082412 kB) free.
[1570082411] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.0% (21725996 kB) free.
[1570082454] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.6% (22689108 kB) free.
[1570082533] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.47, 8.75, 6.26
[1570082653] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.36, 7.06, 5.94
[1570082773] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.50, 5.56, 5.53
[1570082825] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215048 kB) free!
[1570082883] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186444 kB) free!
[1570082945] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234368 kB) free!
[1570083003] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186564 kB) free!
[1570083016] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183444 kB) free!
[1570083064] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191908 kB) free!
[1570083066] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (245356 kB) free!
[1570083069] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179896 kB) free!
[1570083125] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196216 kB) free!
[1570083136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192064 kB) free!
[1570083137] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 81.82, 23.89, 9.10
[1570083184] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165852 kB) free!
[1570083189] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186204 kB) free!
[1570083257] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1570083257] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173960 kB) free!
[1570083257] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 133.97, 63.05, 25.25
[1570083304] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177744 kB) free!
[1570083309] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1570083309] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193060 kB) free!
[1570083347] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1570083373] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.88, 4.07, 4.48
[1570083378] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 141.23, 89.26, 39.43
[1570083467] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1570083490] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1570083540] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.5% (22681944 kB) free.
[1570083587] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1570083610] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1570083624] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.9% (17977144 kB) free.
[1570083706] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.6% (22700544 kB) free.
[1570083730] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1570083777] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.8% (23240944 kB) free.
[1570083847] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.2% (38033168 kB) free.
[1570083973] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.53, 12.00, 8.97
[1570084093] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.24, 11.32, 9.08
[1570084213] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.25, 10.38, 9.00
[1570084235] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190280 kB) free!
[1570084317] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174704 kB) free!
[1570084356] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177076 kB) free!
[1570084391] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187304 kB) free!
[1570084437] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191896 kB) free!
[1570084455] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236396 kB) free!
[1570084476] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175252 kB) free!
[1570084511] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176680 kB) free!
[1570084558] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170972 kB) free!
[1570084575] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231744 kB) free!
[1570084631] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1570084631] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193976 kB) free!
[1570084695] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1570084695] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (249124 kB) free!
[1570084746] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195468 kB) free!
[1570084866] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186952 kB) free!
[1570084987] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1570084987] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191340 kB) free!
[1570085178] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.14, 3.00, 15.33
[1570085530] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1570085995] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1570086495] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 42.2% (20805640 kB) free.
[1570086613] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1570086613] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.72, 3.30, 4.39
[1570086978] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.46, 1.28, 3.36
[1570087095] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.3% (5091072 kB) free!
[1570087215] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (2068660 kB) free!
[1570087336] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250768 kB) free!
[1570088536] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 28.0% (13809012 kB) free.
[1570088876] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 95.7% (23454240 kB) free.
[1570088916] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.2% (23344320 kB) free.
[1570088968] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.7% (22975768 kB) free.
[1570088995] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21198476 kB) free!
[1570089007] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.6% (22362096 kB) free.
[1570089098] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 65.4% (86202296 kB) free.
[1570089136] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (251100 kB) free!
[1570089256] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244120 kB) free!
[1570089376] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246724 kB) free!
[1570089613] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.79, 7.10, 5.46
[1570089613] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165544 kB) free!
[1570089733] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.81, 7.36, 5.76
[1570089734] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179800 kB) free!
[1570089853] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.94, 7.95, 6.17
[1570089854] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181144 kB) free!
[1570091224] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.2% (20536844 kB) free.
[1570091275] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.0% (39893260 kB) free.
[1570091835] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184476 kB) free!
[1570091955] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177940 kB) free!
[1570092068] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9467836 kB) free!
[1570092075] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166060 kB) free!
[1570092188] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9194540 kB) free!
[1570092254] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.76, 3.81, 4.34
[1570092308] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (8944580 kB) free!
[1570092594] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 11.4% (2686292 kB) free!
[1570092714] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (2307688 kB) free!
[1570092834] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1962240 kB) free!
[1570093365] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.0% (23185684 kB) free.
[1570093407] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.7% (23026492 kB) free.
[1570094285] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (689188 kB) free!
[1570094405] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230464 kB) free!
[1570094525] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241204 kB) free!
[1570094575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177764 kB) free!
[1570094696] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191332 kB) free!
[1570094708] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4698068 kB) free!
[1570094816] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163576 kB) free!
[1570094923] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1570094976] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1570095483] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1570096325] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 28.4% (14015752 kB) free.
[1570096616] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 60.1% (14840396 kB) free.
[1570097816] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.9% (4171376 kB) free!
[1570097936] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194296 kB) free!
[1570098056] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184628 kB) free!
[1570098125] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229976 kB) free!
[1570098245] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (254336 kB) free!
[1570098365] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233140 kB) free!
[1570099075] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1570099195] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1570100376] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4764456 kB) free!
[1570100496] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4640520 kB) free!
[1570100530] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1570100616] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4515932 kB) free!
[1570100650] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1570100770] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1570101277] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.1% (22975412 kB) free.
[1570101337] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 56.5% (27824336 kB) free.
[1570101372] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 50.7% (24997812 kB) free.
[1570101855] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1570101884] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196508 kB) free!
[1570102004] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185476 kB) free!
[1570102124] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165488 kB) free!
[1570102545] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227100 kB) free!
[1570102665] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (250388 kB) free!
[1570102724] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4392056 kB) free!
[1570102785] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250332 kB) free!
[1570102844] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (4260860 kB) free!
[1570102964] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4135076 kB) free!
[1570103016] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1971996 kB) free!
[1570104706] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24744648 kB) free!
[1570104764] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2230084 kB) free!
[1570104826] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.8% (45867956 kB) free.
[1570104854] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.08, 5.04, 5.09
[1570104974] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.53, 4.81, 4.99
[1570105342] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9798420 kB) free!
[1570105414] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3576636 kB) free!
[1570105426] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23040552 kB) free!
[1570105462] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9671348 kB) free!
[1570105534] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.0% (3293900 kB) free!
[1570105546] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43920072 kB) free.
[1570105574] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.38, 5.22, 5.09
[1570105582] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9545296 kB) free!
[1570105654] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.3% (2905648 kB) free!
[1570105694] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.97, 5.04, 5.06
[1570105814] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.86, 4.26, 4.78
[1570106256] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.9% (23070620 kB) free.
[1570106746] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (21986928 kB) free!
[1570106866] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 42.4% (55847520 kB) free.
[1570108481] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 542 processes
[1570108541] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 531 processes
[1570109182] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4528260 kB) free!
[1570110008] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1570110128] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1570110970] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1570111079] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.2% (22257780 kB) free.
[1570111114] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.2% (22595408 kB) free.
[1570111153] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 40.6% (20024484 kB) free.
[1570111185] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 55.8% (27481508 kB) free.
[1570111221] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.0% (20587264 kB) free.
[1570114087] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177088 kB) free!
[1570114207] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (596060 kB) free!
[1570114327] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (494400 kB) free!
[1570115396] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1570116476] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.24, 2.43, 4.98
[1570116525] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175508 kB) free!
[1570116625] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185832 kB) free!
[1570116645] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169076 kB) free!
[1570116746] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177916 kB) free!
[1570116765] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190376 kB) free!
[1570116866] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190924 kB) free!
[1570117373] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 82.88, 60.04, 28.66
[1570117493] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 82.50, 67.43, 35.18
[1570117613] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 24.07, 54.24, 34.56
[1570118213] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.83, 7.86, 18.42
[1570120719] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.5% (23083984 kB) free.
[1570120754] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.2% (23328076 kB) free.
[1570120788] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.5% (23164676 kB) free.
[1570120796] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (228552 kB) free!
[1570120916] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226636 kB) free!
[1570121036] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243232 kB) free!
[1570121213] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.03, 1.94, 4.48
[1570121325] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168688 kB) free!
[1570121445] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178272 kB) free!
[1570121528] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1570121565] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199564 kB) free!
[1570121648] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes