[1569612168] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 22.78, 19.20, 8.40
[1569612288] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.75, 13.10, 7.47
[1569612357] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.95, 12.84, 8.33
[1569612408] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.21, 9.08, 6.69
[1569612477] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.15, 8.94, 7.45
[1569612490] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.00, 6.63, 6.48
[1569612597] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 51.80, 22.46, 12.29
[1569612610] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 132.93, 50.83, 22.57
[1569612730] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.66, 34.58, 20.04
[1569613430] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 449 processes
[1569613608] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.25, 24.53, 25.95
[1569614291] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 356 processes
[1569614411] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 414 processes
[1569614468] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.1% (3971616 kB) free!
[1569614563] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 366 processes
[1569614633] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 779 processes
[1569614753] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 779 processes
[1569614864] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1070 processes
[1569614873] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 334 processes
[1569614984] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1069 processes
[1569615069] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185628 kB) free!
[1569615104] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 412 processes
[1569615163] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 2 warn out of 348 processes
[1569615408] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 5.92, 18.00
[1569616197] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.11, 2.08, 12.07
[1569616330] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.77, 3.46, 14.10
[1569617208] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1569617208] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.85, 0.96, 3.36
[1569617397] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.54, 3.74
[1569617530] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.55, 1.18, 4.61
[1569619154] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.6% (47094824 kB) free.
[1569620896] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 50.2% (12296880 kB) free.
[1569621073] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 47.5% (11630796 kB) free.
[1569624720] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569624840] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1569626040] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569626160] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1569626747] SERVICE ALERT: webservercluster1;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24032 MB (10% inode=99%):
[1569626946] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (2765032 kB) free!
[1569627066] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (2765792 kB) free!
[1569627186] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (2765296 kB) free!
[1569627408] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1569630360] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569630480] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569632308] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 537 processes
[1569632368] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 536 processes
[1569632428] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 529 processes
[1569632880] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569633000] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1569634228] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 535 processes
[1569634269] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.6% (12990436 kB) free.
[1569634288] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 529 processes
[1569635300] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171060 kB) free!
[1569635421] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194196 kB) free!
[1569635469] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (221564 kB) free!
[1569635473] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175752 kB) free!
[1569635543] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193356 kB) free!
[1569635589] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189276 kB) free!
[1569635594] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171248 kB) free!
[1569635709] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186820 kB) free!
[1569635714] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177744 kB) free!
[1569639600] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1569639720] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569640320] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1569640440] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569641031] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 559 processes
[1569641151] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 560 processes
[1569641272] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 559 processes
[1569645352] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 14.66, 10.77, 7.32
[1569645472] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.18, 9.85, 7.39
[1569645592] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.30, 10.34, 7.91
[1569647745] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.2% (23498724 kB) free.
[1569647809] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.5% (22917348 kB) free.
[1569647867] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39535540 kB) free.
[1569647916] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.1% (22322344 kB) free.
[1569649126] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.7% (2626392 kB) free!
[1569649246] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.1% (2470880 kB) free!
[1569649366] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2277628 kB) free!
[1569649420] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1569649540] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1569649889] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 565 processes
[1569649949] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 568 processes
[1569650156] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193088 kB) free!
[1569650273] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (213968 kB) free!
[1569650276] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163336 kB) free!
[1569650393] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (223980 kB) free!
[1569650396] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187600 kB) free!
[1569650513] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232580 kB) free!
[1569651555] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7198244 kB) free!
[1569651675] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 13.5% (6666968 kB) free!
[1569651795] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.4% (6125668 kB) free!
[1569652395] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3615096 kB) free!
[1569652672] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 633 processes
[1569655192] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 27.13, 23.61, 21.13
[1569655614] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4656996 kB) free!
[1569655734] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4439648 kB) free!
[1569655854] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4246000 kB) free!
[1569656392] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 11.10, 13.01, 17.02
[1569657054] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1703068 kB) free!
[1569658009] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.42, 8.76, 6.61
[1569658129] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.91, 6.19, 5.93
[1569658249] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.75, 4.39, 5.31
[1569658849] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.93, 1.41, 3.25
[1569661099] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4175480 kB) free!
[1569661219] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3751044 kB) free!
[1569661339] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3333916 kB) free!
[1569661931] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 108.17, 54.85, 23.24
[1569661939] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1591080 kB) free!
[1569662051] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 112.44, 74.15, 34.15
[1569662171] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 113.00, 87.04, 43.72
[1569663139] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.6% (22993124 kB) free.
[1569663371] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.60, 7.38, 18.03
[1569664192] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.59, 3.57, 4.93
[1569665171] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.01, 2.04, 4.74
[1569665392] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.71, 4.33, 5.29
[1569665512] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.83, 3.86, 5.00
[1569666712] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.24, 8.20, 6.69
[1569666832] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.06, 7.31, 6.54
[1569666952] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.65, 6.61, 6.37
[1569669952] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.61, 2.76, 4.72
[1569669971] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.28, 7.48, 5.12
[1569670091] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.72, 7.41, 5.40
[1569670211] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1569670211] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.17, 5.85, 5.08
[1569670811] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.88, 1.81, 3.35
[1569674040] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 406 processes
[1569674152] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.14, 8.54, 6.26
[1569674160] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1569674272] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.91, 8.21, 6.41
[1569674392] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1569674392] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.31, 7.01, 6.20
[1569675844] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.3% (23016756 kB) free.
[1569675962] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 58.2% (14260240 kB) free.
[1569675991] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 43.0% (21177876 kB) free.
[1569676020] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 33.1% (16333200 kB) free.
[1569676045] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 63.8% (15642300 kB) free.
[1569677055] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180408 kB) free!
[1569677175] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190016 kB) free!
[1569677194] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (256840 kB) free!
[1569677296] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194324 kB) free!
[1569677314] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (218968 kB) free!
[1569677434] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219044 kB) free!
[1569682192] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1569682212] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569682792] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.09, 3.28, 4.71
[1569683392] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.69, 5.02, 5.24
[1569683512] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.03, 4.22, 4.91
[1569684042] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22001776 kB) free.
[1569684074] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 45.1% (22249928 kB) free.
[1569684248] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45720 MB (20% inode=99%):
[1569684368] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45481 MB (20% inode=99%):
[1569684488] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45482 MB (20% inode=99%):
[1569684961] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1569685081] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569685088] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47377 MB (21% inode=99%):
[1569685688] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46520 MB (20% inode=99%):
[1569685808] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46117 MB (20% inode=99%):
[1569685846] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4716004 kB) free!
[1569685928] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45874 MB (20% inode=99%):
[1569685966] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4671324 kB) free!
[1569686086] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4614680 kB) free!
[1569686686] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1569686686] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1757608 kB) free!
[1569687509] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.7% (22875048 kB) free.
[1569688717] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185680 kB) free!
[1569688837] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172604 kB) free!
[1569688957] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196612 kB) free!
[1569689223] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216592 kB) free!
[1569689343] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 61.4% (30246236 kB) free.
[1569689528] SERVICE FLAPPING ALERT: smscapp1;HDD_Home status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1569689528] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 91083 MB (40% inode=99%):
[1569692586] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171352 kB) free!
[1569693224] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167780 kB) free!
[1569693284] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171084 kB) free!
[1569693344] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (170412 kB) free!
[1569693404] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (163556 kB) free!
[1569694349] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 569 processes
[1569694409] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 573 processes
[1569694469] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 1 warn out of 573 processes
[1569694530] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 1 warn out of 569 processes