[1596479965] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596479989] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596480004] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (217872 kB) free!
[1596480020] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188708 kB) free!
[1596480086] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596480113] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 9216 open files (0% of max 2424156)
[1596480124] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182560 kB) free!
[1596480140] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191876 kB) free!
[1596480244] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202636 kB) free!
[1596480261] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193148 kB) free!
[1596480529] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596480654] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.58% system=4.78% iowait=0.00% idle=93.64%
[1596481337] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.84, 48.91, 45.27
[1596481444] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.6% (21134896 kB) free.
[1596481461] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.9% (20799356 kB) free.
[1596481937] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 37.74, 43.11, 43.98
[1596482537] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1596482537] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.76, 51.36, 47.22
[1596483137] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 45.59, 43.26, 44.65
[1596483381] SERVICE ALERT: webserver2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596483434] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192388 kB) free!
[1596483446] SERVICE ALERT: webserver2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=4.54% system=1.66% iowait=0.03% idle=93.78%
[1596483495] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202088 kB) free!
[1596483555] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 72.4% (17735760 kB) free.
[1596484517] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.3% (3022972 kB) free!
[1596484637] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170036 kB) free!
[1596484693] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (386748 kB) free!
[1596484757] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196180 kB) free!
[1596484813] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (445548 kB) free!
[1596484920] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220612 kB) free!
[1596484933] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (449956 kB) free!
[1596485040] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241956 kB) free!
[1596485160] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239204 kB) free!
[1596485357] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 21.0% (5140064 kB) free.
[1596486135] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16160.64, 12526.25, 6158.45
[1596486255] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16168.54, 13731.82, 7372.33
[1596486375] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16169.53, 14538.46, 8439.00
[1596486961] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 37.9% (18681152 kB) free.
[1596487157] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174984 kB) free!
[1596487277] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174016 kB) free!
[1596487397] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203460 kB) free!
[1596487444] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191156 kB) free!
[1596487461] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181468 kB) free!
[1596487564] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21923880 kB) free.
[1596487581] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 74.4% (18233924 kB) free.
[1596488597] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1596488597] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 62.8% (15387908 kB) free.
[1596488649] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 65.8% (32443892 kB) free.
[1596489197] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (859044 kB) free!
[1596489249] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 78.7% (38781960 kB) free.
[1596489318] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.9% (23010144 kB) free.
[1596489375] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.30, 3.26, 608.47
[1596489975] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.11, 1.41, 319.37
[1596490575] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 103.255.146.164 [04/Aug/2020:03:00:07 +0530] TCP 502 0 0 1.002 "10.147.212.142:3000" "0" "0" "1.002"
[1596490695] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596492137] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 54.04, 47.27, 45.07
[1596493337] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 28.49, 34.73, 41.80
[1596498737] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.78, 51.35, 47.84
[1596498917] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179124 kB) free!
[1596499037] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180168 kB) free!
[1596499157] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178752 kB) free!
[1596499757] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.0% (23028144 kB) free.
[1596499937] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 43.60, 39.99, 42.76
[1596501137] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.20, 45.17, 45.05
[1596502456] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (209584 kB) free!
[1596504733] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 43.7% (57643800 kB) free.
[1596505694] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1504 processes
[1596505815] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1501 processes
[1596505935] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1500 processes
[1596511157] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596511937] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1596512175] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16163.74, 11428.13, 5267.41
[1596512295] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16194.85, 13006.93, 6592.94
[1596512415] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16198.95, 14063.33, 7757.75
[1596512639] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 613 processes
[1596512759] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 615 processes
[1596513049] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596513064] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596513073] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596513085] SERVICE ALERT: pservercluster2;Messages Log Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596513087] SERVICE ALERT: pservercluster2;Current Users;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596513094] SERVICE ALERT: pservercluster2;CPU_Procs;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596513119] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513179] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513194] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596513203] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513207] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596513215] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513224] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513249] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513309] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513323] SERVICE ALERT: pservercluster2;HDD_Home status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596513324] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513327] SERVICE ALERT: pservercluster2;Current Users;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596513345] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596513354] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513355] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596513369] SERVICE ALERT: pservercluster2;Open-Files;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596513391] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513485] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513521] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513584] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596513589] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513594] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513615] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513629] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513651] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513714] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596513719] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513724] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513759] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513844] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596513849] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513854] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513889] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596513927] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596513929] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596513979] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596514127] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 320 processes
[1596514247] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1596514902] SERVICE ALERT: pservercluster1;zombie_procs;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596514915] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515022] SERVICE ALERT: pservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596515035] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1596515057] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;SOFT;1;(1) < Aug  4 09:51:50 PSERVERCLUSTER1 xinetd[1593]: nrpe: fork failed: Cannot allocate memory (errno = 12)
[1596515150] SERVICE ALERT: pservercluster1;CPU-STATISTICS;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596515161] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515169] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515182] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515192] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515265] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515270] SERVICE ALERT: pservercluster1;CPU-STATISTICS;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596515286] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515294] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515295] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 103.255.146.164 [04/Aug/2020:09:52:13 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1596515307] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515317] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515367] SERVICE ALERT: pservercluster1;Current Users;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596515375] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596515385] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1596515390] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596515411] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515415] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596515416] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.00, 3.56, 621.58
[1596515419] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515442] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515492] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515495] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596515505] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596515594] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596515617] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515622] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515625] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515640] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515671] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515719] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515745] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515747] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515765] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515796] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515844] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515867] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515870] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515890] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596515921] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596516015] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 103.255.146.164 [04/Aug/2020:10:01:50 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1596516016] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.31, 1.89, 326.54
[1596516136] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596517300] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596517420] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596517545] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596517935] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(150) < 182.18.184.244 [04/Aug/2020:10:41:30 +0530] TCP 502 0 0 0.000 "10.147.212.49:5500" "0" "0" "0.000"
[1596518055] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596519218] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596519244] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38287 MB (74% inode=98%):
[1596519249] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 328 processes
[1596519254] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1596519255] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 103.27.86.30 [04/Aug/2020:11:04:10 +0530] TCP 502 0 0 63.010 "10.147.212.143:13001, 10.147.212.142:13001, 10.147.212.155:13001, 10.147.212.154:13001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 3.008, 60.001, 0.000"
[1596519289] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 1 process with STATE = Z
[1596519314] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=2.30% system=6.28% iowait=0.00% idle=91.43%
[1596519324] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596519327] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596519329] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 140176 MB (62% inode=99%):
[1596519354] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1596519375] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596519379] SERVICE ALERT: pservercluster2;Open-Files;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596519444] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 317 processes
[1596519467] SERVICE ALERT: pservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596519470] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.77, 1.16, 1.28
[1596519490] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1596519521] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 195824 MB (88% inode=99%):
[1596519595] SERVICE ALERT: pservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.39% system=4.00% iowait=0.00% idle=94.61%
[1596519612] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1596519615] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.43, 1.50, 1.59
[1596519619] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 27.7% (6831760 kB) free.
[1596519642] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39511 MB (77% inode=99%):
[1596519651] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 57.1% (13989364 kB) free.
[1596519945] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596519959] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1596519975] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(22) < 103.255.146.164 [04/Aug/2020:11:14:08 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1596519979] SERVICE FLAPPING ALERT: pservercluster2;Open-Files;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1596519979] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 11648 open files (0% of max 2424156)
[1596520080] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1596520090] SERVICE ALERT: pservercluster1;Sensors Status;WARNING;SOFT;1;NRPE: Unable to read output
[1596520095] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596520200] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1596520210] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1596520215] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14912.67, 7320.49, 2878.62
[1596520305] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 9952 open files (0% of max 2440681)
[1596520335] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596520450] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520456] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16667.44, 12100.62, 5839.98
[1596520570] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 661 processes
[1596520667] SERVICE ALERT: pservercluster1;zombie_procs;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596520675] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520690] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 422 processes
[1596520726] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520792] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520795] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520800] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520812] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520815] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520817] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16015.63, 11972.03, 5739.63
[1596520824] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520846] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520847] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520910] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520912] SERVICE ALERT: pservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596520915] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15807.72, 14625.09, 8946.27
[1596520926] SERVICE ALERT: pservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=0.95% system=9.86% iowait=0.00% idle=89.19%
[1596520932] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1596520935] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1596520937] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16029.32, 13315.22, 6987.52
[1596520944] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 67.7% (16698148 kB) free.
[1596520967] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596520971] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596521030] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 7584 open files (0% of max 2440681)
[1596521057] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1596521057] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16031.25, 14214.00, 8084.10
[1596521087] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39511 MB (77% inode=99%):
[1596521146] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596521295] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 103.255.146.164 [04/Aug/2020:11:33:48 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1596521415] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596521571] SERVICE FLAPPING ALERT: pservercluster1;HDD_Home status;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1596521571] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 194653 MB (88% inode=99%):
[1596521907] SERVICE ALERT: pservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596523456] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1596523456] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.38, 3.50, 531.00
[1596523457] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.89, 10.85, 892.73
[1596524057] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.88, 2.85, 468.55
[1596524447] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 323 processes
[1596524567] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1596527515] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.98, 1.05, 239.80
[1596528484] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596528600] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 421 processes
[1596528604] SERVICE ALERT: sservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596529367] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 319 processes
[1596529487] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1596531379] SERVICE FLAPPING ALERT: pservercluster2;Open-Files;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596532456] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1596532971] SERVICE FLAPPING ALERT: pservercluster1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596535215] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(40) < 182.18.184.244 [04/Aug/2020:15:24:43 +0530] TCP 502 0 0 0.000 "10.147.212.49:5501" "0" "0" "0.000"
[1596535335] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596535457] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596536610] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.8% (38360628 kB) free.
[1596536771] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(4) < Channel error on connection <0.1128.232> (10.147.212.49:60852 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1596536892] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1596536917] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.5% (14717164 kB) free.
[1596537254] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596537256] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596537294] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537314] SERVICE ALERT: pservercluster2;CPU-STATISTICS;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596537324] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596537329] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537332] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537346] SERVICE ALERT: pservercluster2;Messages Log Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596537359] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537374] SERVICE ALERT: pservercluster2;Sensors Status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596537379] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537381] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537419] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537434] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1596537449] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537450] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537457] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537466] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1596537479] SERVICE ALERT: pservercluster2;CPU_Procs;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596537494] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537495] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537499] SERVICE ALERT: pservercluster2;Open-Files;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596537501] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596537544] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537554] SERVICE ALERT: pservercluster2;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Unable to read output
[1596537569] SERVICE ALERT: pservercluster2;HDD_Home status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596537575] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537582] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537591] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537604] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537619] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537620] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537656] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537745] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537781] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537849] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537906] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596537974] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596538099] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596538106] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596538159] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596538174] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596538256] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.8% (6483032 kB) free!
[1596538376] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.6% (6418940 kB) free!
[1596538496] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.4% (6362724 kB) free!
[1596538506] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 69.5% (17029220 kB) free.
[1596538695] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1596538699] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38284 MB (74% inode=98%):
[1596538706] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.65, 32.09, 23.54
[1596538745] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 1 process with STATE = Z
[1596538765] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=2.40% system=6.57% iowait=0.00% idle=91.02%
[1596538774] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 125328 MB (55% inode=99%):
[1596538775] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596538782] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596538804] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1596538819] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 10368 open files (0% of max 2424156)
[1596538945] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 333 processes
[1596539058] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15782.95, 8442.74, 3420.14
[1596539178] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16137.77, 11007.10, 4968.84
[1596539298] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16185.86, 12723.11, 6329.80
[1596539392] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596539906] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16199.39, 13063.86, 6652.28
[1596540026] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16203.95, 14103.01, 7810.47
[1596540135] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(39) < 182.18.184.244 [04/Aug/2020:16:46:51 +0530] TCP 502 0 0 0.001 "10.147.212.49:5501" "0" "0" "0.001"
[1596540146] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1596540146] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16203.90, 14798.19, 8828.19
[1596540255] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596540634] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 75.9% (24837080 kB) free.
[1596540856] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(19) < 182.18.184.244 [04/Aug/2020:16:58:18 +0530] TCP 502 0 0 0.001 "10.147.212.49:5500" "0" "0" "0.001"
[1596540976] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596541547] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 96.1% (23558772 kB) free.
[1596541666] SERVICE ALERT: sig1-vm2;Memory;OK;HARD;1;OK - 97.3% (15891620 kB) free.
[1596542298] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.35, 5.55, 743.11
[1596542898] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.19, 1.95, 390.02
[1596543146] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.98, 2.23, 503.20
[1596543976] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2480) < 182.18.184.244 [04/Aug/2020:17:56:12 +0530] TCP 502 0 0 0.000 "10.147.212.49:5501" "0" "0" "0.000"
[1596544096] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(958) < 182.18.184.244 [04/Aug/2020:17:58:13 +0530] TCP 502 0 0 0.000 "10.147.212.49:5501" "0" "0" "0.000"
[1596544217] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(960) < 182.18.184.244 [04/Aug/2020:18:00:14 +0530] TCP 502 0 0 0.001 "10.147.212.49:5501" "0" "0" "0.001"
[1596544449] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9540472 kB) free!
[1596544569] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9467884 kB) free!
[1596544689] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.1% (9394236 kB) free!
[1596546017] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596547817] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 103.255.146.164 [04/Aug/2020:18:57:46 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1596547937] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596551649] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9745544 kB) free!
[1596551769] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9714680 kB) free!
[1596551889] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (9689580 kB) free!
[1596552146] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1596553840] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 13.5% (4410356 kB) free!
[1596553960] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 13.4% (4396556 kB) free!
[1596554080] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 13.4% (4376872 kB) free!
[1596555280] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 45.1% (14763668 kB) free.
[1596555489] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4807116 kB) free!
[1596556098] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16122.12, 10562.54, 4664.20
[1596556155] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 721 processes
[1596556218] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16200.93, 12434.32, 6065.44
[1596556276] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 729 processes
[1596556338] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16205.40, 13682.33, 7295.00
[1596556396] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 746 processes
[1596556996] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 326 processes
[1596557547] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15879.64, 9040.10, 3747.12
[1596557672] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596557712] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596557797] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596557837] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596557957] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 73.6% (18043316 kB) free.
[1596557987] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596557997] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596558009] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596558024] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596558112] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596558122] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596558134] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596558145] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596558149] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596558232] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596558242] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;(27) < Aug  4 21:52:25 pservercluster2 xinetd[1483]: logging: fork failed: Cannot allocate memory (errno = 12)
[1596558254] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1596558265] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 353 processes
[1596558269] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 8832 open files (0% of max 2424156)
[1596558843] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596559338] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.47, 3.87, 636.10
[1596559916] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16152.96, 13635.84, 7270.23
[1596559938] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.45, 1.63, 333.89
[1596560036] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16154.65, 14469.27, 8347.46
[1596560156] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16155.18, 15027.26, 9294.21
[1596560537] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(34) < 182.18.184.244 [04/Aug/2020:22:26:53 +0530] TCP 502 0 0 0.001 "10.147.212.49:5501" "0" "0" "0.001"
[1596560657] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596560797] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.58, 5.01, 593.56
[1596561258] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(37) < 182.18.184.244 [04/Aug/2020:22:38:39 +0530] TCP 502 0 0 0.000 "10.147.212.49:5501" "0" "0" "0.000"
[1596561378] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596561397] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16187.18, 12947.75, 6848.43
[1596561517] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16191.68, 14021.34, 7981.38
[1596561637] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16193.30, 14740.08, 8977.11
[1596562337] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 29.19, 40.07, 43.40
[1596562338] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16230.88, 13571.86, 7185.30
[1596562458] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16232.80, 14452.33, 8282.31
[1596562578] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1596562578] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16232.72, 15041.47, 9246.31
[1596562579] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(29) < 182.18.184.244 [04/Aug/2020:23:00:56 +0530] TCP 502 0 0 0.000 "10.147.212.49:5501" "0" "0" "0.000"
[1596562699] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596563087] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 318 processes
[1596563207] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 318 processes
[1596563327] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 316 processes
[1596563756] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.73, 1.25, 222.39
[1596563899] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(24) < 182.18.184.244 [04/Aug/2020:23:25:45 +0530] TCP 502 0 0 0.000 "10.147.212.49:5501" "0" "0" "0.000"
[1596563928] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1596564020] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596564137] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 60.07, 47.78, 44.09
[1596564356] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16186.01, 12723.18, 6444.00
[1596564476] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16192.51, 13871.47, 7626.14
[1596564597] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16194.69, 14640.18, 8665.07
[1596564637] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1596564637] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.98, 2.00, 451.76
[1596564730] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596564855] SERVICE ALERT: pservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.54% system=4.09% iowait=0.00% idle=94.37%
[1596565337] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 34.57, 36.47, 40.57
[1596565579] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.01, 2.46, 524.59