[1592246137] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1592248172] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592248298] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=8.47% system=27.74% iowait=6.52% idle=57.27%
[1592249517] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (830996 kB) free!
[1592249577] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186900 kB) free!
[1592249637] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (176392 kB) free!
[1592249697] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (182860 kB) free!
[1592250597] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 67.4% (16523412 kB) free.
[1592250850] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (628968 kB) free!
[1592251236] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (225016 kB) free!
[1592251356] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240288 kB) free!
[1592251476] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222520 kB) free!
[1592252396] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 204098 MB (19% inode=99%):
[1592252516] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 197298 MB (18% inode=99%):
[1592252636] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 190325 MB (17% inode=99%):
[1592253276] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 33.0% (16255120 kB) free.
[1592254031] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 62.21, 40.24, 34.88
[1592254631] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 45.91, 36.64, 34.73
[1592255584] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 83.2% (41024348 kB) free.
[1592256275] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.9% (42844872 kB) free.
[1592256431] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 20.22, 24.49, 27.72
[1592256635] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1056) < 185.255.8.58 [16/Jun/2020:03:00:13 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1592256755] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(140) < 172.105.59.173 [16/Jun/2020:03:02:22 +0530] TCP 502 0 0 120.002 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1592256875] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592256895] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(126) < 2020/06/16 03:01:22 [error] 12474#0: *94500080 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.62.98, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1592257015] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592257017] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(104) < 182.18.144.246 [16/Jun/2020:03:00:13 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1592257137] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592257631] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 20.18, 35.62, 32.42
[1592257751] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 44.75, 39.41, 34.14
[1592257871] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1592257871] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 36.90, 38.70, 34.59
[1592259375] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.5% (10172236 kB) free.
[1592260576] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (890444 kB) free!
[1592260696] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (296828 kB) free!
[1592260816] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199368 kB) free!
[1592262016] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.5% (10182676 kB) free.
[1592263816] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4678776 kB) free!
[1592263936] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164376 kB) free!
[1592264056] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1592264056] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184296 kB) free!
[1592264978] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.3% (18446412 kB) free.
[1592265390] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 54.3% (13411064 kB) free.
[1592267957] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 74.7% (18309664 kB) free.
[1592268671] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1592269871] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 54.58, 55.62, 50.16
[1592270036] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 226823 MB (21% inode=99%):
[1592270050] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 44.4% (58557868 kB) free.
[1592270636] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 224464 MB (20% inode=99%):
[1592270756] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 255126 MB (23% inode=99%):
[1592273190] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (780628 kB) free!
[1592273310] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (747520 kB) free!
[1592273430] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (720640 kB) free!
[1592274256] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1592277481] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.62, 5.47, 5.29
[1592277601] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.02, 6.27, 5.65
[1592277722] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.76, 5.79, 5.55
[1592278779] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188920 kB) free!
[1592278899] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190372 kB) free!
[1592278904] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1592278908] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592279019] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189088 kB) free!
[1592279022] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1592279025] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1592279036] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=13.32% system=27.32% iowait=0.03% idle=59.33%
[1592279142] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1592279145] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1592279262] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1592279745] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1592280238] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592280368] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1592280498] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1592281620] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 61.8% (15147644 kB) free.
[1592282303] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=14.35% system=27.51% iowait=0.03% idle=58.11%
[1592282697] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4912084 kB) free!
[1592282757] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4914044 kB) free!
[1592282817] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4915196 kB) free.
[1592283557] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4381064 kB) free!
[1592283677] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4164404 kB) free!
[1592283717] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4914068 kB) free!
[1592283778] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4919268 kB) free.
[1592283797] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (3911612 kB) free!
[1592284662] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1592284783] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1592284903] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1592284997] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1723720 kB) free!
[1592285137] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3412892 kB) free!
[1592285145] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1592285220] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191112 kB) free!
[1592285257] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (3018440 kB) free!
[1592285265] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592285341] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188024 kB) free!
[1592285377] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.6% (2615044 kB) free!
[1592285385] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1592285461] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204020 kB) free!
[1592285577] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4913088 kB) free!
[1592285597] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4443140 kB) free!
[1592285637] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4906716 kB) free!
[1592285697] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 20.0% (4907848 kB) free!
[1592285757] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 20.0% (4913520 kB) free!
[1592285977] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (179736 kB) free!
[1592285985] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1592285985] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1592286061] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 67.0% (16411584 kB) free.
[1592286103] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1592286798] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1592286798] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1766684 kB) free!
[1592289357] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4925440 kB) free.
[1592289577] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.4% (15585124 kB) free.
[1592289657] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4908204 kB) free!
[1592289717] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4912624 kB) free!
[1592289777] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 20.0% (4911984 kB) free!
[1592289837] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 20.0% (4909664 kB) free!
[1592289875] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9344432 kB) free!
[1592289995] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9140248 kB) free!
[1592290115] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (8882228 kB) free!
[1592291037] SERVICE FLAPPING ALERT: webserver2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1592291037] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4915780 kB) free.
[1592291337] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4908984 kB) free!
[1592291397] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4911212 kB) free!
[1592291457] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 20.0% (4908240 kB) free!
[1592291517] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 20.0% (4908536 kB) free!
[1592291817] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4916556 kB) free.
[1592291915] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4865092 kB) free!
[1592292117] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4911256 kB) free!
[1592292177] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4913188 kB) free!
[1592292237] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 20.0% (4909764 kB) free!
[1592292297] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 20.0% (4906092 kB) free!
[1592294461] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.4% (4516212 kB) free!
[1592294581] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.0% (4414968 kB) free!
[1592294701] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.6% (4304448 kB) free!
[1592295540] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1592295660] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1592296501] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1592296501] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2271772 kB) free!
[1592296777] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.6% (3954204 kB) free!
[1592296897] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (3735104 kB) free!
[1592297017] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.6% (3512608 kB) free!
[1592297385] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592297598] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1592297997] SERVICE FLAPPING ALERT: webserver2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592298217] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1274332 kB) free!
[1592298420] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.8% (6469252 kB) free!
[1592298540] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.5% (6377428 kB) free!
[1592298660] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.2% (6295912 kB) free!
[1592299384] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9522704 kB) free!
[1592299505] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9425116 kB) free!
[1592299625] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (9301588 kB) free!
[1592299860] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 20.1% (6565300 kB) free.
[1592300017] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1592300017] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.1% (15524548 kB) free.
[1592300075] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(182) < 182.18.144.246 [16/Jun/2020:15:00:12 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1592300195] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592300460] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 18.1% (5922812 kB) free!
[1592300580] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 17.7% (5782976 kB) free!
[1592300701] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 17.4% (5682656 kB) free!
[1592303460] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1592303580] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1592304425] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4558296 kB) free!
[1592305501] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1592305501] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (3084184 kB) free!
[1592306195] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(12) < 172.105.57.57 [16/Jun/2020:16:42:57 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592306315] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592306701] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1592307856] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (2568876 kB) free!
[1592308417] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.6% (3960992 kB) free!
[1592308456] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1757376 kB) free!
[1592308538] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (3675756 kB) free!
[1592308658] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.2% (3411396 kB) free!
[1592309258] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1890072 kB) free!
[1592310787] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592310911] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55264 MB (24% inode=99%):
[1592311058] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.7% (15656496 kB) free.
[1592312701] SERVICE FLAPPING ALERT: mrq4;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1592314507] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47064 MB (20% inode=99%):
[1592314637] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46972 MB (20% inode=99%):
[1592314757] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44901 MB (19% inode=99%):
[1592316198] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.9% (2662004 kB) free!
[1592316557] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 190022 MB (84% inode=99%):
[1592316798] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1875880 kB) free!
[1592317998] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (2998616 kB) free!
[1592318598] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1592318598] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2288128 kB) free!
[1592320658] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.5% (3919248 kB) free!
[1592320778] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.8% (3774484 kB) free!
[1592320898] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.1% (3621744 kB) free!
[1592321050] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (24804052 kB) free!
[1592321170] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.8% (24807988 kB) free!
[1592321290] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.8% (24811272 kB) free!
[1592321499] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.5% (15594724 kB) free.
[1592321522] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.67, 4.82, 4.88
[1592321701] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 66.2% (21685752 kB) free.
[1592322987] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592323107] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1592325584] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1189_n"
[1592325704] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1592325780] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1592325900] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1592326271] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 45.82, 41.45, 42.93
[1592329271] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 61.08, 49.04, 43.81
[1592329871] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 50.00, 38.57, 39.81
[1592329999] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)