[1562265283] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8691016 kB) free!
[1562265345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562265404] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12188580 kB) free.
[1562265465] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562265575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189308 kB) free!
[1562266003] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8645620 kB) free!
[1562266065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562266123] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12191524 kB) free.
[1562266185] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562266723] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8570860 kB) free!
[1562266785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562266843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12182824 kB) free.
[1562266906] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562267443] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8521692 kB) free!
[1562267505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562267563] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12190056 kB) free.
[1562267625] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562267858] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562268163] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8558880 kB) free!
[1562268225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562268283] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12192520 kB) free.
[1562268345] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562268815] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193700 kB) free!
[1562268875] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190872 kB) free!
[1562268883] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8589424 kB) free!
[1562268935] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (196228 kB) free!
[1562268945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562268995] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (195828 kB) free!
[1562269003] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12186500 kB) free.
[1562269065] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562269603] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8587692 kB) free!
[1562269665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562269724] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12190540 kB) free.
[1562269786] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562270323] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8626240 kB) free!
[1562270385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562270443] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12180948 kB) free.
[1562270505] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562271043] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8704168 kB) free!
[1562271105] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562271163] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12187672 kB) free.
[1562271225] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562271438] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.5% (22567908 kB) free.
[1562271454] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.7% (22233944 kB) free.
[1562271465] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.3% (38592536 kB) free.
[1562271501] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.0% (22314804 kB) free.
[1562271521] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 61.9% (15162740 kB) free.
[1562271533] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 76.7% (12463084 kB) free.
[1562271546] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 68.1% (11066296 kB) free.
[1562271606] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.9% (21541840 kB) free.
[1562271623] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.6% (14559916 kB) free.
[1562271763] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8651716 kB) free!
[1562271825] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562271883] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12189116 kB) free.
[1562271945] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562272483] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8621040 kB) free!
[1562272545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562272603] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12186832 kB) free.
[1562272654] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176920 kB) free!
[1562272666] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562272737] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (155804 kB) free!
[1562272757] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166064 kB) free!
[1562272774] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173972 kB) free!
[1562272857] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (175824 kB) free!
[1562272877] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (386984 kB) free!
[1562272894] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178940 kB) free!
[1562272977] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (188784 kB) free!
[1562272997] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (390008 kB) free!
[1562273203] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8651452 kB) free!
[1562273265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562273323] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12189812 kB) free.
[1562273385] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562273924] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8743984 kB) free!
[1562273985] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562274044] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12181976 kB) free.
[1562274105] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562274465] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7404660 kB) free!
[1562274585] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.0% (6401836 kB) free!
[1562274644] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8751420 kB) free!
[1562274705] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562274705] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (5369544 kB) free!
[1562274764] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12187012 kB) free.
[1562274825] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562275305] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (595792 kB) free!
[1562275364] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8728348 kB) free!
[1562275425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562275484] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12188344 kB) free.
[1562275546] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562276084] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8705776 kB) free!
[1562276145] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562276204] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12185804 kB) free.
[1562276265] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562276804] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8657692 kB) free!
[1562276865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562276924] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12188432 kB) free.
[1562276985] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562277333] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.7% (21732816 kB) free.
[1562277347] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1562277347] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.4% (39140084 kB) free.
[1562277368] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 84.3% (13704408 kB) free.
[1562277387] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 69.8% (11347420 kB) free.
[1562277524] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8659376 kB) free!
[1562277585] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562277644] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12180568 kB) free.
[1562277705] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562277937] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191244 kB) free!
[1562278057] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180164 kB) free!
[1562278177] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1562278177] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200728 kB) free!
[1562278245] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8579008 kB) free!
[1562278305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562278365] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12186704 kB) free.
[1562278426] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562278576] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (318984 kB) free!
[1562278696] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (158784 kB) free!
[1562278816] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1562278816] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (194260 kB) free!
[1562278965] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8516268 kB) free!
[1562279025] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562279085] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12191256 kB) free.
[1562279145] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562279685] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8496176 kB) free!
[1562279745] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562279805] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183444 kB) free.
[1562279865] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562280357] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7202164 kB) free!
[1562280405] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8489272 kB) free!
[1562280465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562280477] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.2% (5997920 kB) free!
[1562280525] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183164 kB) free.
[1562280585] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562280597] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (5057084 kB) free!
[1562281058] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 79.58, 63.03, 32.63
[1562281125] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8494804 kB) free!
[1562281178] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 79.92, 68.64, 38.37
[1562281185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562281197] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218524 kB) free!
[1562281231] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15.62, 40.60, 24.20
[1562281245] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12182932 kB) free.
[1562281298] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 13.68, 47.58, 34.37
[1562281306] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562281351] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.60, 27.36, 21.34
[1562281363] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4.27, 30.37, 21.72
[1562281471] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.64, 18.43, 18.80
[1562281483] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.73, 20.37, 19.10
[1562281603] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 13.69, 16.81
[1562281845] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8566460 kB) free!
[1562281898] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.65, 9.03, 19.60
[1562281905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562281965] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12184484 kB) free.
[1562282025] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562282566] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8532408 kB) free!
[1562282625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562282686] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12187476 kB) free.
[1562282745] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562282803] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.47, 4.78
[1562283271] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.41, 0.47, 3.05
[1562283286] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8522676 kB) free!
[1562283345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562283406] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183216 kB) free.
[1562283465] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562284006] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8578452 kB) free!
[1562284065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562284126] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12185548 kB) free.
[1562284186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562284298] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.81, 3.00, 4.29
[1562284306] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562284586] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.9% (22034568 kB) free.
[1562284607] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.5% (39183680 kB) free.
[1562284626] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.8% (14439960 kB) free.
[1562285326] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9197328 kB) free!
[1562285446] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7791248 kB) free!
[1562285506] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562285566] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11910948 kB) free.
[1562285838] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.1% (1639244 kB) free!
[1562285958] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (192432 kB) free!
[1562286078] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180116 kB) free!
[1562286106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562286226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562286346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562286698] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 304.36, 216.14, 101.32
[1562286766] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9013136 kB) free!
[1562286818] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 304.89, 245.55, 126.02
[1562286887] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7306800 kB) free!
[1562286899] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5275 MB (10% inode=97%):
[1562286938] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1562286938] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 280.77, 260.25, 146.12
[1562286946] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1562286946] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562286997] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (1991516 kB) free!
[1562287003] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 302.65, 219.73, 103.15
[1562287007] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11535072 kB) free.
[1562287117] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 52.3% (8501892 kB) free.
[1562287123] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 67.49, 168.51, 99.41
[1562287243] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 9.31, 112.85, 87.39
[1562287499] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6928 MB (13% inode=98%):
[1562287547] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562287667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1562287717] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2977152 kB) free!
[1562287787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562287837] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (148340 kB) free!
[1562287957] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174012 kB) free!
[1562287979] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 574 processes
[1562288039] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 572 processes
[1562288207] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8959572 kB) free!
[1562288215] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.6% (5205672 kB) free!
[1562288327] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7437148 kB) free!
[1562288335] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (3998476 kB) free!
[1562288387] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562288447] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11531204 kB) free.
[1562288455] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2913340 kB) free!
[1562288557] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 61.0% (9908864 kB) free.
[1562288939] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 571 processes
[1562288999] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 568 processes
[1562289043] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1562289043] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.12, 0.45, 12.74
[1562289338] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.08, 3.18, 13.85
[1562289395] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562289587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562289707] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562289827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562289900] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 570 processes
[1562289960] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 574 processes
[1562290243] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.15, 0.23, 3.67
[1562291138] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.77, 2.84, 4.41
[1562291627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562292047] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8495820 kB) free!
[1562292167] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12105700 kB) free.
[1562292227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562292347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562292467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562295078] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1562295655] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1562295767] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8465232 kB) free!
[1562295887] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12100872 kB) free.
[1562296067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562296487] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8417100 kB) free!
[1562296607] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12100660 kB) free.
[1562297207] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8406852 kB) free!
[1562297267] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562297327] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12100940 kB) free.
[1562297387] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562297928] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7860152 kB) free!
[1562298048] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11846480 kB) free.
[1562298157] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.5% (1704420 kB) free!
[1562298278] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155988 kB) free!
[1562298398] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173996 kB) free!
[1562298625] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1562298745] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1562298871] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 138.64, 69.44, 27.78
[1562298991] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 143.57, 94.24, 41.92
[1562299111] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 144.31, 110.86, 54.35
[1562299787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562299843] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 29.76, 10.89, 4.15
[1562299848] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7922492 kB) free!
[1562299907] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562299963] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.52, 7.48, 3.71
[1562299968] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12094800 kB) free.
[1562300507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562300568] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7777664 kB) free!
[1562300627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562300688] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12095680 kB) free.
[1562300911] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.19, 0.85, 10.99
[1562301163] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562301227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562301288] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7409036 kB) free!
[1562301338] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562301347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562301408] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12090284 kB) free.
[1562301714] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1562301834] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1562301948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562301954] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1562302068] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562302111] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.22, 0.99, 3.68
[1562302608] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8603248 kB) free!
[1562302668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562302711] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.82, 114.85, 86.74
[1562302729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12092840 kB) free.
[1562302788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562302826] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (11464564 kB) free!
[1562302831] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.53, 77.30, 76.39
[1562302946] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (21461396 kB) free.
[1562302951] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1562302951] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.35, 52.10, 67.26
[1562303154] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1562303328] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8454264 kB) free!
[1562303388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562303448] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12082416 kB) free.
[1562303508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562303628] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562304028] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 13.9% (2251256 kB) free!
[1562304114] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.33, 5.27, 5.04
[1562304126] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.53, 5.26, 5.02
[1562304146] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (11385616 kB) free!
[1562304148] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 12.7% (2057924 kB) free!
[1562304151] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.66, 1.67, 18.98
[1562304234] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.42, 5.67, 5.21
[1562304246] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.22, 4.96, 4.94
[1562304266] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (21840380 kB) free.
[1562304268] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.5% (1873952 kB) free!
[1562304354] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.04, 5.29, 5.14
[1562304648] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9135084 kB) free!
[1562304768] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7487236 kB) free!
[1562304829] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562304868] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1225528 kB) free!
[1562304888] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11521044 kB) free.
[1562304995] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164128 kB) free!
[1562305115] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177280 kB) free!
[1562305235] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165636 kB) free!
[1562305429] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562305466] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (11078608 kB) free!
[1562305549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562305565] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.24, 5.58, 5.10
[1562305586] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (21579720 kB) free.
[1562305669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562305685] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 4.06, 5.06, 4.97
[1562305951] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.62, 0.63, 3.34
[1562306046] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.21, 6.31, 5.68
[1562306088] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8407592 kB) free!
[1562306138] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.14, 7.43, 7.89
[1562306166] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.30, 6.18, 5.70
[1562306189] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162116 kB) free!
[1562306208] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12072196 kB) free.
[1562306258] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.90, 6.21, 7.38
[1562306286] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.88, 5.66, 5.57
[1562306309] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180760 kB) free!
[1562306378] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.56, 5.68, 7.05
[1562306429] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164748 kB) free!
[1562306786] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (7979328 kB) free!
[1562306846] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173624 kB) free!
[1562306906] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (19070100 kB) free.
[1562306906] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4551588 kB) free!
[1562306966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182220 kB) free!
[1562307027] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178140 kB) free!
[1562307086] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183652 kB) free!
[1562307147] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194428 kB) free!
[1562308009] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8847144 kB) free!
[1562308106] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (9105804 kB) free!
[1562308129] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7137052 kB) free!
[1562308178] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.50, 4.03, 4.76
[1562308187] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (589400 kB) free!
[1562308226] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (19615136 kB) free.
[1562308230] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.0% (14142584 kB) free.
[1562308249] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11242156 kB) free.
[1562308308] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.1% (6775092 kB) free.
[1562308670] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562309031] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 364 processes
[1562309151] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1562309270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562309390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562309426] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (9703588 kB) free!
[1562309449] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9293000 kB) free!
[1562309510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562309546] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (5631704 kB) free!
[1562309569] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7826540 kB) free!
[1562309575] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1562309666] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (17053768 kB) free.
[1562309689] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12066988 kB) free.
[1562309695] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562310266] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (11395628 kB) free!
[1562310386] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (7252564 kB) free!
[1562310506] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (17736072 kB) free.
[1562311706] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (8690084 kB) free!
[1562311827] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (4626372 kB) free!
[1562311947] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (16414316 kB) free.
[1562312089] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9179688 kB) free!
[1562312133] SERVICE ALERT: secondary-server;Nginx Status;OK;HARD;1;NGINX OK -  0.084 sec. response time, Active: 0 (Writing: 0 Reading: 0 Waiting: 0) ReqPerSec: 0.000 ConnPerSec: 0.000 ReqPerConn: 0.000
[1562312209] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7635992 kB) free!
[1562312330] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12042244 kB) free.
[1562312379] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.54, 9.16, 9.01
[1562312499] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.14, 7.49, 8.42
[1562312511] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562312547] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (11721676 kB) free!
[1562312619] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.99, 5.96, 7.75
[1562312667] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (7484028 kB) free!
[1562312787] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (3523516 kB) free!
[1562312929] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7117288 kB) free!
[1562313049] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12031344 kB) free.
[1562313387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (14385024 kB) free.
[1562313819] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1562313819] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.22, 25.58, 15.91
[1562313987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (9719196 kB) free!
[1562314107] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (5728284 kB) free!
[1562314227] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (5970400 kB) free!
[1562314249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8218792 kB) free!
[1562314311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562314369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12033488 kB) free.
[1562314419] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.29, 6.40, 9.96
[1562314431] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562314554] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.01, 4.36, 4.65
[1562314827] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1562314827] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (13965272 kB) free.
[1562314969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7713132 kB) free!
[1562315089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12038860 kB) free.
[1562315427] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (9295964 kB) free!
[1562315507] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (1447100 kB) free!
[1562315547] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (5122004 kB) free!
[1562315551] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.14, 6.05, 6.54
[1562315627] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.6% (1296192 kB) free!
[1562315631] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562315667] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (1214136 kB) free!
[1562315672] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.85, 4.29, 5.83
[1562315689] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7125676 kB) free!
[1562315747] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 15.6% (1149740 kB) free!
[1562315751] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562315792] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.90, 3.16, 5.23
[1562315809] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12029724 kB) free.
[1562316220] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.29, 3.84, 4.43
[1562316267] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (12346020 kB) free!
[1562316347] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (388496 kB) free!
[1562316352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562316392] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.69, 1.75, 3.40
[1562316472] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562316823] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.00, 7.72, 6.17
[1562316943] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.71, 8.04, 6.47
[1562316992] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 119.80, 53.44, 21.79
[1562317009] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7405240 kB) free!
[1562317064] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.22, 8.82, 6.95
[1562317112] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16.77, 36.04, 19.25
[1562317130] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12030504 kB) free.
[1562317232] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.68, 24.31, 17.00
[1562317468] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (3325596 kB) free!
[1562317967] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 23.74, 26.22, 12.17
[1562318068] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13806392 kB) free.
[1562318087] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.13, 18.30, 10.97
[1562318165] SERVICE ALERT: pservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 587 processes
[1562318207] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.22, 12.52, 9.75
[1562318272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562318285] SERVICE ALERT: pservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 587 processes
[1562318392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562318405] SERVICE ALERT: pservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 587 processes
[1562318472] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1562318512] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562318593] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 354 processes
[1562318668] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11626260 kB) free!
[1562318713] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1562318788] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (7624632 kB) free!
[1562318908] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (3620268 kB) free!
[1562319005] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 329 processes
[1562319086] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.8% (10798848 kB) free.
[1562319147] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.2% (5446184 kB) free.
[1562319347] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.1% (6775124 kB) free.
[1562319407] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.56, 0.62, 2.94
[1562319464] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 29.44, 35.64, 19.57
[1562319508] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (14067392 kB) free.
[1562319632] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.03, 0.97, 3.30
[1562319713] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562319747] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4871644 kB) free!
[1562319868] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178872 kB) free!
[1562319913] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 373 processes
[1562319988] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180940 kB) free!
[1562320033] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1562320064] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.16, 8.18, 12.18
[1562320108] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (11027664 kB) free!
[1562320129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7128036 kB) free!
[1562320228] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (6890920 kB) free!
[1562320249] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12012936 kB) free.
[1562320286] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176536 kB) free!
[1562320348] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2800652 kB) free!
[1562320406] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185172 kB) free!
[1562320526] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165924 kB) free!
[1562320607] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.86, 17.89, 10.01
[1562320633] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 372 processes
[1562320664] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 44.19, 23.39, 15.22
[1562320727] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.21, 12.08, 8.84
[1562320753] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 368 processes
[1562320847] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.01, 8.39, 7.88
[1562320873] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 369 processes
[1562320948] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (14058316 kB) free.
[1562321126] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.3% (10200772 kB) free.
[1562321264] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.58, 5.84, 9.53
[1562321447] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1562321447] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 1.60, 4.44
[1562321449] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8292400 kB) free!
[1562321473] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1562321513] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562321548] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (10808260 kB) free!
[1562321569] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12017668 kB) free.
[1562321633] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562321668] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (6661580 kB) free!
[1562321754] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.62, 6.23, 5.53
[1562321788] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (2603720 kB) free!
[1562321864] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 135.49, 34.57, 16.73
[1562321874] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.65, 6.55, 5.74
[1562321994] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.31, 6.62, 5.87
[1562322047] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.42, 11.76, 9.82
[1562322074] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 374 processes
[1562322167] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.80, 8.10, 8.71
[1562322169] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7907656 kB) free!
[1562322194] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 373 processes
[1562322287] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 5.57, 7.71
[1562322290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12019952 kB) free.
[1562322314] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1562322388] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (14292072 kB) free.
[1562322398] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (3234960 kB) free!
[1562322464] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.29, 9.30, 12.56
[1562322632] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.40, 4.24, 5.52
[1562322752] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.71, 3.13, 4.96
[1562322769] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 583 processes
[1562322889] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 580 processes
[1562322988] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (10293460 kB) free!
[1562322998] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3268396 kB) free.
[1562323009] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 322 processes
[1562323109] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (5962944 kB) free!
[1562323194] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.62, 3.97, 4.80
[1562323229] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1802408 kB) free!
[1562323352] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 96.85, 24.31, 10.63
[1562323433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562323473] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.75, 16.48, 9.41
[1562323487] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.64, 1.52, 4.72
[1562323514] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1562323514] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1562323554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562323593] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.72, 11.58, 8.47
[1562323598] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3161276 kB) free!
[1562323674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562323718] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2874128 kB) free!
[1562323794] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.00, 6.45, 5.60
[1562323829] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (14305724 kB) free.
[1562323838] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3124496 kB) free!
[1562323914] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.16, 5.60, 5.39
[1562324034] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1562324034] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.41, 4.83, 5.13
[1562324087] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.36, 7.83, 7.04
[1562324089] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9481164 kB) free!
[1562324193] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.17, 2.40, 4.91
[1562324208] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.49, 5.41, 6.25
[1562324209] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8183592 kB) free!
[1562324328] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.61, 3.82, 5.57
[1562324329] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11988016 kB) free.
[1562324429] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (7780924 kB) free!
[1562324438] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1562324438] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3554208 kB) free.
[1562324549] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (3588472 kB) free!
[1562324634] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.84, 3.43, 4.31
[1562324669] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (20058164 kB) free.
[1562324793] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.35, 22.06, 13.18
[1562324913] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.97, 15.41, 11.83
[1562324928] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.51, 0.96, 3.16
[1562324929] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7875060 kB) free!
[1562325033] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.58, 10.79, 10.57
[1562325049] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11842184 kB) free.
[1562325474] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562325638] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2971856 kB) free!
[1562325758] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3366660 kB) free.
[1562325834] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.51, 5.15, 5.01
[1562325869] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (9935172 kB) free!
[1562325954] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.29, 5.10, 5.00
[1562325989] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (5907252 kB) free!
[1562326109] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (17875988 kB) free.
[1562326233] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.93, 1.21, 3.65
[1562326547] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (725172 kB) free!
[1562326554] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.03, 5.54, 5.31
[1562326667] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (557284 kB) free!
[1562326674] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.98, 5.56, 5.36
[1562326709] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (12597196 kB) free!
[1562326787] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (359124 kB) free!
[1562326795] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.86, 4.96, 5.17
[1562326829] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (8329968 kB) free!
[1562326949] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (19276932 kB) free.
[1562326959] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3218904 kB) free!
[1562327079] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3210496 kB) free!
[1562327126] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (349852 kB) free!
[1562327199] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (3180832 kB) free!
[1562327246] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (301888 kB) free!
[1562327264] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.86, 2.96, 4.33
[1562327274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562327366] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1562327366] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (235556 kB) free!
[1562327395] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562327395] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.09, 4.81, 4.92
[1562327433] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.60, 7.50, 8.31
[1562327553] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.45, 5.16, 7.36
[1562327673] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.58, 3.63, 6.53
[1562327928] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.98, 8.05, 6.29
[1562328048] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.58, 5.57, 5.60
[1562328149] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (11212540 kB) free!
[1562328168] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 3.97, 5.01
[1562328269] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (7019776 kB) free!
[1562328273] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.50, 1.65, 4.00
[1562328390] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (18558560 kB) free.
[1562328399] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3583696 kB) free.
[1562328464] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.55, 11.03, 11.00
[1562328584] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.08, 8.36, 10.02
[1562328594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562328595] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.65, 6.06, 5.45
[1562328649] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9822752 kB) free!
[1562328704] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.80, 6.51, 9.14
[1562328714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562328715] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.70, 4.83, 5.07
[1562328768] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.41, 1.08, 2.96
[1562328769] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8562284 kB) free!
[1562328834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562328835] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.40, 4.40, 4.89
[1562328889] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10413712 kB) free.
[1562329434] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562329473] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.93, 10.34, 6.58
[1562329589] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (10289720 kB) free!
[1562329593] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.43, 7.21, 5.88
[1562329709] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (6052256 kB) free!
[1562329713] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 4.94, 5.22
[1562329829] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (18005432 kB) free.
[1562329968] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 82.95, 37.88, 15.51
[1562330034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562330088] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 87.60, 54.60, 24.35
[1562330114] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 364 processes
[1562330154] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562330208] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 88.41, 65.78, 32.12
[1562330234] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 361 processes
[1562330313] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.49, 1.06, 2.96
[1562330354] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 361 processes
[1562330388] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6829428 kB) free.
[1562330689] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8516796 kB) free!
[1562330755] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562330788] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2501912 kB) free!
[1562330809] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11971080 kB) free.
[1562330875] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562330913] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8.62, 36.86, 25.97
[1562330954] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1562330966] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (2595120 kB) free!
[1562331029] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (7423460 kB) free!
[1562331033] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.09, 25.02, 22.95
[1562331104] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.37, 3.26, 4.52
[1562331149] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (3612100 kB) free!
[1562331153] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.80, 16.89, 20.22
[1562331269] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (15719748 kB) free.
[1562331388] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2029224 kB) free!
[1562331399] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2946400 kB) free!
[1562331408] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.35, 2.78, 13.37
[1562331409] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7109136 kB) free!
[1562331519] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3225404 kB) free!
[1562331529] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11448884 kB) free.
[1562331566] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2318824 kB) free!
[1562331639] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (3095108 kB) free!
[1562331753] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.36, 2.62, 10.80
[1562331835] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.13, 5.81, 5.20
[1562331869] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11225428 kB) free!
[1562331956] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.42, 5.77, 5.27
[1562331989] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (6855184 kB) free!
[1562332075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562332076] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.58, 7.21, 5.85
[1562332109] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (7713420 kB) free!
[1562332129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9743660 kB) free!
[1562332195] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562332239] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3638340 kB) free.
[1562332249] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8344508 kB) free!
[1562332315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562332370] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11959648 kB) free.
[1562332608] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.77, 0.67, 4.13
[1562332709] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (14232756 kB) free.
[1562333276] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.99, 3.59, 4.43
[1562333309] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (9236536 kB) free!
[1562333354] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 367 processes
[1562333429] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (5063884 kB) free!
[1562333474] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 365 processes
[1562333504] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.24, 38.11, 17.17
[1562333550] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (16392388 kB) free.
[1562333553] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.90, 0.91, 3.81
[1562333594] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1562333624] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 9.77, 26.48, 15.45
[1562333744] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.88, 18.62, 13.90
[1562333808] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.69, 5.14, 5.55
[1562333928] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.47, 3.60, 4.94
[1562334039] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2858732 kB) free!
[1562334149] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (11895572 kB) free!
[1562334159] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3399780 kB) free.
[1562334269] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (7584028 kB) free!
[1562334389] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (3773336 kB) free!
[1562334715] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562334944] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 81.49, 37.57, 19.05
[1562334989] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (14655060 kB) free.
[1562335129] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.46, 6.61, 7.63
[1562335249] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.24, 4.53, 6.75
[1562335369] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.51, 3.16, 5.98
[1562335589] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (9714560 kB) free!
[1562335709] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (5617260 kB) free!
[1562335829] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (15490992 kB) free.
[1562335959] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3228340 kB) free!
[1562335969] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.84, 3.36
[1562336079] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3046716 kB) free!
[1562336095] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1562336199] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3685868 kB) free.
[1562336215] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1562336335] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1562336399] SERVICE ALERT: sig2-vm1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 25182 MB (20% inode=96%):
[1562336429] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12533976 kB) free!
[1562336515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562336520] SERVICE ALERT: sig2-vm1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 25176 MB (20% inode=96%):
[1562336549] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (8360408 kB) free!
[1562336569] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.21, 15.44, 13.84
[1562336569] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7645876 kB) free!
[1562336635] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562336640] SERVICE ALERT: sig2-vm1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 25169 MB (20% inode=96%):
[1562336669] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (4012620 kB) free!
[1562336689] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.45, 10.47, 12.22
[1562336689] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11944412 kB) free.
[1562336744] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.89, 5.79, 14.10
[1562336809] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.51, 7.20, 10.81
[1562336935] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562337153] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.86, 13.74, 7.22
[1562337235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562337269] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (14970160 kB) free.
[1562337273] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.04, 9.28, 6.37
[1562337289] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7413672 kB) free!
[1562337356] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562337393] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.11, 6.87, 5.84
[1562337399] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3058840 kB) free!
[1562337409] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11940976 kB) free.
[1562337519] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3519536 kB) free.
[1562337869] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (10325824 kB) free!
[1562337955] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562337989] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (6276324 kB) free!
[1562337993] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 1.44, 3.40
[1562338009] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.68, 3.40
[1562338010] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7408060 kB) free!
[1562338075] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562338109] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (1927332 kB) free!
[1562338130] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11943304 kB) free.
[1562338675] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562338710] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (12903492 kB) free!
[1562338795] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562339209] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.40, 8.88, 9.76
[1562339329] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.48, 6.15, 8.66
[1562339330] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8236836 kB) free!
[1562339395] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562339449] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.94, 4.38, 7.70
[1562339450] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11932776 kB) free.
[1562339515] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562339910] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (4703620 kB) free!
[1562340049] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.88, 4.20
[1562340050] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7911904 kB) free!
[1562340170] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11939164 kB) free.
[1562340510] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (14105880 kB) free.
[1562340716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562340770] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7405880 kB) free!
[1562340836] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562340890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11942140 kB) free.
[1562341110] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (11108076 kB) free!
[1562341230] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (7598348 kB) free!
[1562341350] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (4070324 kB) free!
[1562341436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562341556] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562341849] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.28, 11.65, 8.09
[1562341950] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (13658868 kB) free.
[1562341969] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.64, 7.96, 7.17
[1562342089] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 5.48, 6.36
[1562342090] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8625372 kB) free!
[1562342156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562342210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11934972 kB) free.
[1562342276] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562342550] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (11100644 kB) free!
[1562342594] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562342670] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (7874204 kB) free!
[1562342689] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.95, 3.45
[1562342790] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (4708576 kB) free!
[1562342810] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8574400 kB) free!
[1562342876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562342931] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11941152 kB) free.
[1562342935] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1562342966] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562342996] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562343055] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1562343390] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (13886164 kB) free.
[1562343530] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8314472 kB) free!
[1562343596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562343650] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11928868 kB) free.
[1562343717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562343993] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.10, 16.28, 8.93
[1562344076] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562344113] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.26, 11.09, 7.92
[1562344119] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562344233] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.27, 7.48, 6.99
[1562344250] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7970276 kB) free!
[1562344370] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11933008 kB) free.
[1562344590] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (11943716 kB) free!
[1562344710] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (8637248 kB) free!
[1562344830] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (5332332 kB) free!
[1562344833] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 1.48, 3.93
[1562344916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562344970] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7580492 kB) free!
[1562345036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562345090] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11935408 kB) free.
[1562345431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13827744 kB) free.
[1562345636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562345690] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7282048 kB) free!
[1562345744] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.66, 2.68, 4.46
[1562345756] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562345756] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562345810] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11930412 kB) free.
[1562346031] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (9796524 kB) free!
[1562346151] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (6237804 kB) free!
[1562346271] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13818156 kB) free.
[1562346356] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562346410] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7473880 kB) free!
[1562346477] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562346530] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11935264 kB) free.
[1562346871] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (11426688 kB) free!
[1562346991] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (8075400 kB) free!
[1562347076] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562347111] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (4531616 kB) free!
[1562347130] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7392724 kB) free!
[1562347196] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562347251] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11924528 kB) free.
[1562347489] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 41.49, 27.56, 11.07
[1562347609] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.87, 18.58, 9.78
[1562347711] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (12310464 kB) free!
[1562347729] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.99, 12.52, 8.63
[1562347796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562347916] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562348329] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 1.90, 4.66
[1562348450] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8462924 kB) free!
[1562348516] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562348570] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11932732 kB) free.
[1562348636] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562348911] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13799628 kB) free.
[1562348929] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 39.23, 14.17, 7.40
[1562349049] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.36, 9.52, 6.52
[1562349169] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.86, 6.44, 5.75
[1562349170] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8244048 kB) free!
[1562349236] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562349290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11933476 kB) free.
[1562349356] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562349511] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (10250492 kB) free!
[1562349631] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (6708224 kB) free!
[1562349633] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.13, 10.30, 5.14
[1562349751] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13797988 kB) free.
[1562349753] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.00, 7.02, 4.56
[1562349769] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.20, 1.11, 3.16
[1562349890] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7909908 kB) free!
[1562350010] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11930888 kB) free.
[1562350086] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.17, 4.29, 4.76
[1562350351] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (10083540 kB) free!
[1562350471] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (6449496 kB) free!
[1562350556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562350591] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13799396 kB) free.
[1562350610] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7434944 kB) free!
[1562350676] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562350719] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (187144 kB) free!
[1562350730] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11932296 kB) free.
[1562350839] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (560360 kB) free!
[1562350959] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (559548 kB) free!
[1562350969] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.27, 29.00, 14.50
[1562351089] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.47, 19.50, 12.77
[1562351192] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (10001312 kB) free!
[1562351209] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 13.09, 11.24
[1562351276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562351312] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (6472452 kB) free!
[1562351396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes