[1562092349] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9383928 kB) free!
[1562092403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562092469] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12743220 kB) free.
[1562092481] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.32, 4.37
[1562092523] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562092598] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (2298072 kB) free!
[1562092660] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1562092679] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190120 kB) free!
[1562092721] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164252 kB) free!
[1562092781] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 373 processes
[1562092799] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (212176 kB) free!
[1562092844] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173012 kB) free!
[1562092919] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202432 kB) free!
[1562093069] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9264112 kB) free!
[1562093123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562093189] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12732244 kB) free.
[1562093243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562093789] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9101692 kB) free!
[1562093843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562093909] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12739340 kB) free.
[1562093963] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562094095] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562094509] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8948648 kB) free!
[1562094563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562094629] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12730864 kB) free.
[1562094684] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562095229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8887272 kB) free!
[1562095283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562095349] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12741576 kB) free.
[1562095403] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562095909] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 14.2% (3475952 kB) free!
[1562095949] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8928280 kB) free!
[1562095969] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3476556 kB) free!
[1562096003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562096030] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 14.2% (3480512 kB) free!
[1562096070] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12739332 kB) free.
[1562096090] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 14.2% (3479776 kB) free!
[1562096123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562096669] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8926984 kB) free!
[1562096723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562096789] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12729740 kB) free.
[1562096843] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562097389] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8885160 kB) free!
[1562097443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562097509] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12739408 kB) free.
[1562097564] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562098109] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8791776 kB) free!
[1562098163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562098229] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12729616 kB) free.
[1562098283] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562098829] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8706172 kB) free!
[1562098949] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12739212 kB) free.
[1562099483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562099549] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8657704 kB) free!
[1562099603] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562099669] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12737524 kB) free.
[1562099906] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.5% (2680156 kB) free!
[1562099949] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162240 kB) free!
[1562100026] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.6% (13420976 kB) free.
[1562100069] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (197024 kB) free!
[1562100189] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (280036 kB) free!
[1562100203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562100269] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8728692 kB) free!
[1562100324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562100382] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22055076 kB) free.
[1562100390] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12730848 kB) free.
[1562100429] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 79.0% (19363076 kB) free.
[1562100444] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562100486] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39625760 kB) free.
[1562100989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8795500 kB) free!
[1562101035] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163856 kB) free!
[1562101044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562101109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12739028 kB) free.
[1562101155] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192176 kB) free!
[1562101164] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562101275] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195212 kB) free!
[1562101408] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1562101408] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.1% (21340956 kB) free.
[1562101464] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.3% (22134172 kB) free.
[1562101509] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.6% (14241592 kB) free.
[1562101510] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 362 processes
[1562101589] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173660 kB) free!
[1562101592] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 69.8% (11338472 kB) free.
[1562101636] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.7% (20407412 kB) free.
[1562101709] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8725484 kB) free!
[1562101829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12728368 kB) free.
[1562101895] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 73.97, 63.81, 36.28
[1562102015] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.34, 47.77, 33.92
[1562102116] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (180232 kB) free!
[1562102135] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.51, 32.77, 30.10
[1562102236] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171832 kB) free!
[1562102356] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (150576 kB) free!
[1562102364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562102429] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8657848 kB) free!
[1562102484] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562102549] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12737476 kB) free.
[1562102619] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187348 kB) free!
[1562102735] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.80, 8.59, 17.82
[1562102739] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178256 kB) free!
[1562102847] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187356 kB) free!
[1562102859] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187652 kB) free!
[1562102967] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163804 kB) free!
[1562103003] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562103084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562103087] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175444 kB) free!
[1562103149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8647032 kB) free!
[1562103204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562103266] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 59.35, 15.12, 6.43
[1562103269] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12736168 kB) free.
[1562103281] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562103310] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 360 processes
[1562103386] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 185.14, 77.90, 30.49
[1562103506] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 202.39, 120.00, 51.67
[1562103805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562103869] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8502392 kB) free!
[1562103925] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562103989] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12732152 kB) free.
[1562104525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562104589] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8407896 kB) free!
[1562104645] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562104700] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9118756 kB) free!
[1562104710] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12730896 kB) free.
[1562104820] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (8043768 kB) free!
[1562104941] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.0% (6920772 kB) free!
[1562105135] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.47, 2.65, 4.32
[1562105245] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562105306] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.32, 1.75, 13.18
[1562105309] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8380680 kB) free!
[1562105365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562105403] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.80, 4.56, 5.00
[1562105429] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12722744 kB) free.
[1562105541] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (1652508 kB) free!
[1562105797] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (215144 kB) free!
[1562105917] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (219472 kB) free!
[1562105965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562106003] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.27, 5.92, 5.43
[1562106029] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8211052 kB) free!
[1562106037] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (238028 kB) free!
[1562106086] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562106123] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.97, 6.32, 5.63
[1562106149] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12726364 kB) free.
[1562106243] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.82, 6.42, 5.75
[1562106506] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.23, 1.26, 4.46
[1562106685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562106749] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8030324 kB) free!
[1562106805] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562106869] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12731984 kB) free.
[1562107405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562107525] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562108069] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9452116 kB) free!
[1562108125] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562108189] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12719872 kB) free.
[1562108245] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562108789] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9326764 kB) free!
[1562108845] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562108910] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12728764 kB) free.
[1562108965] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562109509] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9182032 kB) free!
[1562109566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562109629] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12720100 kB) free.
[1562109686] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562110229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9035668 kB) free!
[1562110286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562110349] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12728548 kB) free.
[1562110406] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562110481] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 76.01, 55.60, 26.29
[1562110536] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 185.33, 76.79, 30.15
[1562110601] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.39, 39.92, 24.23
[1562110656] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 204.17, 119.92, 51.61
[1562110721] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.85, 26.72, 21.29
[1562110776] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1562110776] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 206.31, 148.60, 70.40
[1562111006] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562111126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562111246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562111321] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 3.82, 11.26
[1562112149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8795348 kB) free!
[1562112183] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.1% (22724860 kB) free.
[1562112232] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.6% (22939044 kB) free.
[1562112269] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12699620 kB) free.
[1562112281] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39633036 kB) free.
[1562112345] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 85.3% (13866136 kB) free.
[1562112402] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 69.6% (11314088 kB) free.
[1562112521] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.41, 3.35
[1562112911] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 354 processes
[1562112946] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181644 kB) free!
[1562113009] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172240 kB) free!
[1562113066] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (152368 kB) free!
[1562113130] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (181360 kB) free!
[1562113176] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.33, 3.42, 12.24
[1562113186] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (147196 kB) free!
[1562113250] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (182980 kB) free!
[1562113389] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171292 kB) free!
[1562113443] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.14, 4.56, 4.88
[1562113469] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9289628 kB) free!
[1562113510] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186648 kB) free!
[1562113589] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12665272 kB) free.
[1562113630] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171520 kB) free!
[1562113646] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562114111] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 354 processes
[1562114246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562114366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562114486] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562114638] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (207956 kB) free!
[1562114643] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.47, 5.47, 5.01
[1562114667] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3240356 kB) free!
[1562114758] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (162152 kB) free!
[1562114763] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.69, 5.13, 4.95
[1562114787] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (3225028 kB) free!
[1562114789] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9792116 kB) free!
[1562114878] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (158384 kB) free!
[1562114907] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.7% (3202648 kB) free!
[1562114909] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8254556 kB) free!
[1562115029] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12641636 kB) free.
[1562115169] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 551 processes
[1562115229] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 546 processes
[1562115363] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.11, 5.09, 5.04
[1562115483] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.07, 5.26, 5.10
[1562115576] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.24, 3.20, 4.34
[1562115603] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1562115603] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.24, 5.40, 5.15
[1562115650] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1562115650] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 60.7% (9864964 kB) free.
[1562115890] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9252356 kB) free!
[1562116010] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (8186812 kB) free!
[1562116130] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (7132984 kB) free!
[1562116203] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.02, 4.46, 4.83
[1562116229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (8996104 kB) free!
[1562116287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562116349] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12504632 kB) free.
[1562116407] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562116730] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (1715492 kB) free!
[1562116949] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8740952 kB) free!
[1562117007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562117069] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12507732 kB) free.
[1562117127] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562117669] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8567848 kB) free!
[1562117727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562117789] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12500256 kB) free.
[1562117847] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562118003] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.41, 5.22, 5.04
[1562118123] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.60, 5.71, 5.25
[1562118243] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.63, 6.00, 5.41
[1562118389] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8333804 kB) free!
[1562118510] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12281752 kB) free.
[1562118629] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 365 processes
[1562118750] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1562119349] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 359 processes
[1562119469] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1562120043] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.68, 4.67, 4.90
[1562120321] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.66, 100.03, 71.60
[1562120441] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.08, 67.03, 62.95
[1562120561] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.42, 44.92, 55.34
[1562121447] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562121509] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7986320 kB) free!
[1562121567] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562121629] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12492460 kB) free.
[1562121761] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1562121761] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.37, 1.13, 15.42
[1562122107] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1616092 kB) free!
[1562122167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562122229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7795704 kB) free!
[1562122287] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562122349] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12498336 kB) free.
[1562122887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562122949] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7774532 kB) free!
[1562122962] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.20, 4.38
[1562123008] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562123043] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.06, 5.95, 5.49
[1562123069] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12497472 kB) free.
[1562123163] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.33, 5.47, 5.38
[1562123283] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.43, 5.41, 5.37
[1562123607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562123669] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7696836 kB) free!
[1562123727] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562123790] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12494160 kB) free.
[1562123878] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1562123935] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562124055] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1562124327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562124447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562124483] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.80, 4.66, 4.91
[1562124989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9196892 kB) free!
[1562125047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562125109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12502408 kB) free.
[1562125167] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562125176] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562125250] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562125683] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.17, 5.66, 5.18
[1562125709] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9258404 kB) free!
[1562125767] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562125803] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.17, 5.54, 5.20
[1562125829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12347828 kB) free.
[1562125888] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562125923] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.05, 5.10, 5.09
[1562126306] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 75.98, 26.73, 10.51
[1562126426] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 94.08, 50.00, 21.00
[1562126429] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9031376 kB) free!
[1562126487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562126523] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.20, 4.64, 4.78
[1562126546] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 95.95, 65.36, 30.15
[1562126549] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12486920 kB) free.
[1562126562] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.91, 13.70, 4.88
[1562126607] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562126682] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 112.25, 49.30, 18.99
[1562126802] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 119.84, 73.04, 31.37
[1562127149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8755532 kB) free!
[1562127207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562127269] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12491744 kB) free.
[1562127327] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562127746] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.72, 3.64, 12.92
[1562127869] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8401840 kB) free!
[1562127953] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.05, 6.07, 5.38
[1562127990] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12481800 kB) free.
[1562128002] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.30, 2.30, 11.82
[1562128073] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.58, 5.52, 5.26
[1562128193] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.69, 5.37, 5.25
[1562128323] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.84, 5.19, 5.07
[1562128443] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.06, 5.50, 5.20
[1562128527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562128564] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.05, 5.65, 5.29
[1562128589] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8045716 kB) free!
[1562128647] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562128709] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12087820 kB) free.
[1562128946] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.65, 1.71, 4.75
[1562129202] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.51, 3.60
[1562131564] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.41, 4.19, 4.69
[1562131729] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 589 processes
[1562131789] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 590 processes
[1562132248] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562132278] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 40.0% (9802888 kB) free.
[1562132309] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8006928 kB) free!
[1562132368] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562132429] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12464100 kB) free.
[1562132764] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.61, 5.94, 5.16
[1562132884] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.82, 6.89, 5.59
[1562132969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562133004] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.34, 7.28, 5.89
[1562133089] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562133315] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.4% (17860248 kB) free.
[1562133389] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.3% (39555104 kB) free.
[1562133629] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8800164 kB) free!
[1562133671] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.5% (995784 kB) free!
[1562133689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562133750] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12464312 kB) free.
[1562133791] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (658636 kB) free!
[1562133809] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562133869] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.1% (2975776 kB) free!
[1562133912] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (417200 kB) free!
[1562133916] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (158212 kB) free!
[1562133989] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.9% (2906884 kB) free!
[1562134036] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184740 kB) free!
[1562134078] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176440 kB) free!
[1562134109] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2848736 kB) free!
[1562134156] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (144844 kB) free!
[1562134199] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173500 kB) free!
[1562134204] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.41, 3.69, 4.61
[1562134319] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173624 kB) free!
[1562134346] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.30, 8.65, 6.51
[1562134409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562134466] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.29, 9.55, 7.09
[1562134529] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562134587] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1562134587] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.18, 9.97, 7.56
[1562134649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562134709] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 57.7% (14148384 kB) free.
[1562134826] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2871288 kB) free!
[1562134946] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.1% (2779840 kB) free!
[1562135066] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.5% (2686812 kB) free!
[1562135309] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (769124 kB) free!
[1562135429] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (718104 kB) free!
[1562135549] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (531964 kB) free!
[1562135667] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 86.5% (14063232 kB) free.
[1562135713] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.7% (6744552 kB) free.
[1562135779] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 63.4% (15643008 kB) free.
[1562135868] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.7% (22716548 kB) free.
[1562136149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9059416 kB) free!
[1562136269] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12458364 kB) free.
[1562136387] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 91.78, 43.76, 22.07
[1562136393] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171784 kB) free!
[1562136398] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7758408 kB) free!
[1562136449] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562136478] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186432 kB) free!
[1562136513] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (142208 kB) free!
[1562136518] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.7% (7265600 kB) free!
[1562136598] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167952 kB) free!
[1562136633] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1562136633] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (232356 kB) free!
[1562136638] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.8% (6824776 kB) free!
[1562136722] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1562136722] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178872 kB) free!
[1562136869] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8762068 kB) free!
[1562136988] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.51, 11.89, 15.09
[1562136989] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12446884 kB) free.
[1562137204] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.09, 5.79, 5.30
[1562137238] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (500556 kB) free!
[1562137250] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (184344 kB) free!
[1562137324] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.91, 6.54, 5.63
[1562137371] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (168868 kB) free!
[1562137444] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.67, 6.94, 5.89
[1562137491] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180000 kB) free!
[1562137589] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8828140 kB) free!
[1562137649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562137709] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12456124 kB) free.
[1562137769] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562137855] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562137975] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1562138309] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9204856 kB) free!
[1562138369] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562138429] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12447868 kB) free.
[1562138489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562139030] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9070704 kB) free!
[1562139089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562139150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12451000 kB) free.
[1562139210] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562139750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8804004 kB) free!
[1562139809] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562139870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12453644 kB) free.
[1562139929] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1562140002] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562140470] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8559984 kB) free!
[1562140529] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562140588] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.48, 1.60, 3.19
[1562140590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12443896 kB) free.
[1562140649] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562140820] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (12945320 kB) free!
[1562140922] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3328824 kB) free!
[1562140940] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (9027404 kB) free!
[1562141060] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (17644736 kB) free.
[1562141190] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9146556 kB) free!
[1562141249] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562141310] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7771292 kB) free!
[1562141369] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562141430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12438484 kB) free.
[1562141660] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (10520392 kB) free!
[1562141691] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1666956 kB) free!
[1562141780] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.6% (24753752 kB) free.
[1562141969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562142030] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8880576 kB) free!
[1562142090] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562142122] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2152100 kB) free!
[1562142150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12311164 kB) free.
[1562142210] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562142292] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1618888 kB) free!
[1562142750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8562868 kB) free!
[1562142775] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1562142810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562142870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12448600 kB) free.
[1562142895] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562142913] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (606532 kB) free!
[1562142930] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562143033] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (409900 kB) free!
[1562143153] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (212120 kB) free!
[1562143471] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8629736 kB) free!
[1562143530] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562143591] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12440208 kB) free.
[1562143650] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562144191] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8684064 kB) free!
[1562144250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562144311] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12446472 kB) free.
[1562144370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562144644] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.80, 3.28, 4.55
[1562144695] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1562144815] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1562144911] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8647332 kB) free!
[1562144935] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1562144971] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562145031] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12433016 kB) free.
[1562145091] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562145535] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562145631] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8519976 kB) free!
[1562145722] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (3718244 kB) free!
[1562145751] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12440980 kB) free.
[1562146753] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.8% (6826528 kB) free.
[1562146891] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562147011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562147131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562147434] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562147522] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201364 kB) free!
[1562147644] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.74, 6.36, 5.47
[1562147764] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.65, 6.51, 5.64
[1562147780] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (10292952 kB) free!
[1562147884] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.88, 6.92, 5.88
[1562147901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (19866260 kB) free.
[1562148268] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2878908 kB) free!
[1562148388] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.3% (2647644 kB) free!
[1562148500] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13062832 kB) free!
[1562148508] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.8% (2409596 kB) free!
[1562148620] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (9093444 kB) free!
[1562148740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (18852380 kB) free.
[1562148892] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.2% (1651444 kB) free!
[1562148964] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.70, 5.82, 5.33
[1562148988] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1562149084] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.54, 5.81, 5.40
[1562149108] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1242524 kB) free!
[1562149204] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.74, 5.45, 5.31
[1562149340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (12031716 kB) free!
[1562149460] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (8006396 kB) free!
[1562149492] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562149492] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1097740 kB) free!
[1562149580] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (17700264 kB) free.
[1562150131] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562150180] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (11979604 kB) free!
[1562150284] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.90, 4.38, 4.79
[1562150300] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (7905560 kB) free!
[1562150420] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (16745756 kB) free.
[1562150551] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7758972 kB) free!
[1562150671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12428864 kB) free.
[1562150884] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.02, 7.10, 5.89
[1562151004] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.34, 7.50, 6.19
[1562151020] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11819952 kB) free!
[1562151124] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.12, 7.52, 6.37
[1562151140] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (7725136 kB) free!
[1562151260] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (9192192 kB) free!
[1562151604] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 4.02, 3.96, 4.61
[1562151860] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (16012452 kB) free.
[1562151871] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8835128 kB) free!
[1562151931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562151991] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12418936 kB) free.
[1562152052] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562152591] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8397144 kB) free!
[1562152711] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12268456 kB) free.
[1562152753] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (1243060 kB) free!
[1562152873] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.8% (1090532 kB) free!
[1562152993] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.9% (947320 kB) free!
[1562153549] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.1% (3217432 kB) free!
[1562153593] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1562153593] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (200172 kB) free!
[1562154124] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.01, 4.95, 4.94
[1562154149] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174656 kB) free!
[1562154241] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1562154361] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1562154481] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1562154724] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.85, 6.89, 6.01
[1562154844] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.54, 6.85, 6.11
[1562154964] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.48, 6.38, 6.03
[1562155051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562155111] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7835560 kB) free!
[1562155171] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562155193] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.45, 3.95, 4.61
[1562155231] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12402476 kB) free.
[1562155460] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (11926972 kB) free!
[1562155580] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (21429776 kB) free.
[1562155681] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1562155735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1562155771] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562155855] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1562155892] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562156164] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.62, 3.67, 4.42
[1562156431] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8958696 kB) free!
[1562156491] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562156551] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12413720 kB) free.
[1562156611] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1562156711] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1562156780] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (12040848 kB) free!
[1562156900] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (20839992 kB) free.
[1562157151] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8771312 kB) free!
[1562157194] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.3% (6860040 kB) free.
[1562157211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562157271] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12410384 kB) free.
[1562157331] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562157451] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562157508] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.4% (13561320 kB) free.
[1562157722] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562158100] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11597560 kB) free!
[1562158220] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (20186464 kB) free.
[1562158471] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9364424 kB) free!
[1562158592] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7692848 kB) free!
[1562158651] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1562158651] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562158712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12398448 kB) free.
[1562158995] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46381 MB (20% inode=99%):
[1562159115] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46283 MB (20% inode=99%):
[1562159165] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.05, 6.76, 5.76
[1562159235] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46117 MB (20% inode=99%):
[1562159252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562159285] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.77, 6.78, 5.89
[1562159312] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7625136 kB) free!
[1562159372] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562159405] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.06, 6.90, 6.04
[1562159420] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10830044 kB) free!
[1562159432] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11669568 kB) free.
[1562159540] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (19572952 kB) free.
[1562159972] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562160092] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562160212] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562160605] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.44, 2.83, 4.22
[1562160740] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (10664540 kB) free!
[1562160860] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (18994980 kB) free.
[1562160892] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562162060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (11993292 kB) free!
[1562162180] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (8124596 kB) free!
[1562162300] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (18221144 kB) free.
[1562162900] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (13126584 kB) free!
[1562163020] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (9281304 kB) free!
[1562163140] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (18210168 kB) free.
[1562163402] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 78.18, 26.50, 10.19
[1562163522] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 97.37, 50.95, 21.13
[1562163642] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 100.26, 67.37, 30.76
[1562163988] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 80.20, 64.82, 33.24
[1562164108] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 80.63, 70.02, 38.97
[1562164176] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 71.57, 51.41, 25.65
[1562164228] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.30, 47.99, 34.71
[1562164296] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 75.18, 59.37, 31.68
[1562164340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (10225304 kB) free!
[1562164412] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562164416] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 70.42, 63.72, 36.68
[1562164460] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (6398400 kB) free!
[1562164580] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (16729492 kB) free.
[1562164828] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 6.99, 18.50
[1562164833] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7619512 kB) free!
[1562164842] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.22, 3.40, 15.72
[1562164953] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12384036 kB) free.
[1562165180] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12146288 kB) free!
[1562165301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (8109336 kB) free!
[1562165421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (17499604 kB) free.
[1562165594] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1562165616] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.35, 3.51, 12.10
[1562166042] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.43, 4.59
[1562166153] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8761896 kB) free!
[1562166212] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562166273] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12386292 kB) free.
[1562166332] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562166621] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (10175332 kB) free!
[1562166628] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.53, 3.16
[1562166741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (6926324 kB) free!
[1562166861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (16667856 kB) free.
[1562166873] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8438364 kB) free!
[1562166993] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12375468 kB) free.
[1562167416] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.76, 2.88, 4.06
[1562167461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (12907476 kB) free!
[1562167532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562167581] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (9447624 kB) free!
[1562167593] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8042256 kB) free!
[1562167652] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562167701] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (17077992 kB) free.
[1562167713] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12382416 kB) free.
[1562167805] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.94, 6.18, 5.17
[1562167925] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.29, 5.16, 4.92
[1562168252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562168313] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7643168 kB) free!
[1562168372] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562168433] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12373656 kB) free.
[1562168617] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 291.10, 119.54, 46.55
[1562168737] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 321.57, 187.95, 80.47
[1562168857] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1562168857] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 325.98, 233.76, 110.28
[1562168901] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11548360 kB) free!
[1562168972] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1562168980] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 638 processes
[1562169021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (8036112 kB) free!
[1562169033] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7555552 kB) free!
[1562169092] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562169100] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 631 processes
[1562169141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (16909064 kB) free.
[1562169154] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12378168 kB) free.
[1562169220] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 397 processes
[1562169590] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 49.6% (12159520 kB) free.
[1562169692] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562169753] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7560672 kB) free!
[1562169813] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562169873] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12381656 kB) free.
[1562170341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (11530580 kB) free!
[1562170412] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562170461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (7954012 kB) free!
[1562170532] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562170582] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (16754852 kB) free.
[1562171073] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8778824 kB) free!
[1562171132] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562171193] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12373856 kB) free.
[1562171252] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562171525] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562171781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12325548 kB) free!
[1562171793] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8327820 kB) free!
[1562171857] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.44, 2.55, 12.19
[1562171901] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (9538816 kB) free!
[1562171913] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12381248 kB) free.
[1562172021] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (6376560 kB) free!
[1562172452] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1562172513] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7850348 kB) free!
[1562172572] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562172621] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (16487336 kB) free.
[1562172633] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12368632 kB) free.
[1562172892] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (1720196 kB) free!
[1562173172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562173293] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1562173293] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562173657] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.39, 2.49, 3.87
[1562173821] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12255156 kB) free!
[1562173834] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8794964 kB) free!
[1562173892] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562173941] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (9208704 kB) free!
[1562173954] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12378828 kB) free.
[1562174012] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562174061] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (6280860 kB) free!
[1562174554] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8429376 kB) free!
[1562174620] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.1% (22978392 kB) free.
[1562174661] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1562174661] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (15560136 kB) free.
[1562174674] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12382504 kB) free.
[1562174692] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1500308 kB) free!
[1562174746] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.3% (23113096 kB) free.
[1562174842] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39489860 kB) free.
[1562174901] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.7% (14746176 kB) free.
[1562174946] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.5% (23160248 kB) free.
[1562175022] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.2% (14330128 kB) free.
[1562175105] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 66.5% (10816940 kB) free.
[1562175212] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562175261] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11558448 kB) free!
[1562175274] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7956544 kB) free!
[1562175332] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562175381] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (8311184 kB) free!
[1562175394] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12371220 kB) free.
[1562175501] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (16337604 kB) free.
[1562175932] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562175957] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177632 kB) free!
[1562176052] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562176077] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170248 kB) free!
[1562176101] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12185244 kB) free!
[1562176197] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180264 kB) free!
[1562176222] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (9100196 kB) free!
[1562176230] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178488 kB) free!
[1562176342] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (15162856 kB) free.
[1562176350] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (168924 kB) free!
[1562176470] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168508 kB) free!
[1562176594] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8893868 kB) free!
[1562176653] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562176714] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12382700 kB) free.
[1562176773] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562176942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (13137596 kB) free!
[1562177062] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (9811236 kB) free!
[1562177182] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (6395108 kB) free!
[1562177314] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8462280 kB) free!
[1562177434] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12372412 kB) free.
[1562177782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (13453576 kB) free.
[1562177846] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.7% (7236876 kB) free!
[1562177966] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (6082860 kB) free!
[1562177973] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562178035] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7894260 kB) free!
[1562178086] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (4973888 kB) free!
[1562178093] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562178155] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12380196 kB) free.
[1562178382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (7250396 kB) free!
[1562178502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (15460132 kB) free.