[1588012210] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3119240 kB) free!
[1588012330] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.1% (4560952 kB) free.
[1588012753] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8194152 kB) free!
[1588012873] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4424624 kB) free!
[1588012993] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32904892 kB) free.
[1588013593] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7849244 kB) free!
[1588013713] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4067492 kB) free!
[1588013833] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31867592 kB) free.
[1588014376] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199208 kB) free!
[1588014433] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6966576 kB) free!
[1588014436] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196788 kB) free!
[1588014497] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (199632 kB) free!
[1588014553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32785620 kB) free.
[1588014559] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (197248 kB) free!
[1588014730] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (188036 kB) free!
[1588014850] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (715612 kB) free!
[1588014970] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1125224 kB) free!
[1588015153] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9818204 kB) free!
[1588015159] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.2% (17452872 kB) free.
[1588015273] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6337544 kB) free!
[1588015393] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32103276 kB) free.
[1588015993] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8202192 kB) free!
[1588016113] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32795672 kB) free.
[1588016713] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9310536 kB) free!
[1588016833] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32769952 kB) free.
[1588017173] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (21184148 kB) free!
[1588017293] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10286480 kB) free!
[1588017413] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375916 kB) free!
[1588017433] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10196204 kB) free!
[1588017553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32799420 kB) free.
[1588017613] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224556 kB) free!
[1588017733] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225768 kB) free!
[1588017853] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223676 kB) free!
[1588018153] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11318148 kB) free!
[1588018273] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32851452 kB) free.
[1588018570] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (2709396 kB) free!
[1588018873] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11625972 kB) free!
[1588018993] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32863368 kB) free.
[1588019170] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.3% (4766104 kB) free.
[1588019593] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10541420 kB) free!
[1588019653] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.0% (7404764 kB) free!
[1588019713] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30742852 kB) free.
[1588020129] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.1% (9590064 kB) free.
[1588020313] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2135964 kB) free!
[1588020433] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31139304 kB) free.
[1588020729] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190364 kB) free!
[1588020849] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197060 kB) free!
[1588020969] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191728 kB) free!
[1588021033] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2651604 kB) free!
[1588021153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31446328 kB) free.
[1588021207] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 84.4% (41588696 kB) free.
[1588021753] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4210292 kB) free!
[1588021874] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31770116 kB) free.
[1588022222] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.2% (45458720 kB) free.
[1588022473] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4427600 kB) free!
[1588022593] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30606832 kB) free.
[1588023129] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2625) < 203.212.70.170 [28/Apr/2020:03:01:52 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1588023193] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4595840 kB) free!
[1588023249] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588023254] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.3% (44501728 kB) free.
[1588023313] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588023313] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29440924 kB) free.
[1588023391] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1588023435] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(379) < 103.16.101.72 [28/Apr/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1588023442] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(473) < 2020/04/28 03:01:32 [error] 22125#0: *86046535 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588023555] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588023562] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588024513] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15283940 kB) free!
[1588024633] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33132896 kB) free.
[1588025769] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.9% (10277008 kB) free.
[1588026433] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20027424 kB) free!
[1588026553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35303932 kB) free.
[1588026970] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186112 kB) free!
[1588027090] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193100 kB) free!
[1588027210] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1588027210] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203336 kB) free!
[1588028410] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.1% (9343872 kB) free.
[1588028953] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20184464 kB) free!
[1588029074] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35174484 kB) free.
[1588030570] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588031410] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184160 kB) free!
[1588031473] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20571132 kB) free!
[1588031530] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192656 kB) free!
[1588031593] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35287448 kB) free.
[1588031650] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189720 kB) free!
[1588031751] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 60.9% (15018808 kB) free.
[1588031783] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.2% (18912076 kB) free.
[1588031879] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 56.6% (13866388 kB) free.
[1588034593] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17774784 kB) free!
[1588034713] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33840504 kB) free.
[1588035083] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1588035203] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1588035413] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 39.9% (52645980 kB) free.
[1588035951] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.4% (3793736 kB) free!
[1588036071] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.2% (3749248 kB) free!
[1588036191] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3743108 kB) free!
[1588039791] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175704 kB) free!
[1588040650] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1588045426] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.78, 8.20, 6.74
[1588045546] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.31, 6.53, 6.31
[1588045666] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.21, 6.22, 6.20
[1588045679] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (3402100 kB) free!
[1588045781] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 18.4% (4512884 kB) free!
[1588045799] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3353420 kB) free!
[1588045901] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 18.4% (4512000 kB) free!
[1588045919] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (3299532 kB) free!
[1588046021] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 18.4% (4512196 kB) free!
[1588047719] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2298552 kB) free!
[1588048066] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.01, 3.59, 4.47
[1588048591] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1588048667] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.52, 6.90, 5.71
[1588048711] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1588048787] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.94, 6.25, 5.62
[1588048831] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1588048907] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.83, 6.36, 5.73
[1588049083] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1588049184] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183240 kB) free!
[1588049203] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1588049304] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169400 kB) free!
[1588049323] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1588049424] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187532 kB) free!
[1588049431] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1588049923] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1588050031] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1588050151] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1588050271] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1588050313] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17456496 kB) free!
[1588050433] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36450064 kB) free.
[1588050871] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1588050871] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1588051472] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1588051633] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14238712 kB) free!
[1588051753] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34074060 kB) free.
[1588051778] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 64.3% (15749412 kB) free.
[1588051824] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 43.6% (10685972 kB) free.
[1588052672] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1588052792] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1588052912] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1588052953] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12411192 kB) free!
[1588053024] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175304 kB) free!
[1588053073] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33238852 kB) free.
[1588053144] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175496 kB) free!
[1588053163] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.2% (2970580 kB) free!
[1588053264] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193332 kB) free!
[1588053283] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.9% (2665536 kB) free!
[1588053370] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (1430220 kB) free!
[1588053403] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.3% (2317640 kB) free!
[1588053490] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3533784 kB) free.
[1588053673] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9901392 kB) free!
[1588053794] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28185072 kB) free.
[1588053849] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(322) < 103.255.146.164 [28/Apr/2020:11:32:39 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588053969] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588054003] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (819992 kB) free!
[1588054690] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (2931132 kB) free!
[1588054810] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3575572 kB) free.
[1588055410] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2934984 kB) free!
[1588055530] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3752284 kB) free.
[1588055803] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16119744 kB) free.
[1588056193] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17582376 kB) free!
[1588056313] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39517708 kB) free.
[1588056512] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1588056730] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2939332 kB) free!
[1588056851] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3206936 kB) free!
[1588056913] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16524560 kB) free!
[1588056969] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(634) < 103.255.146.164 [28/Apr/2020:12:19:14 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588056971] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3177616 kB) free!
[1588057033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38603416 kB) free.
[1588057089] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588057633] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14310720 kB) free!
[1588057753] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35970124 kB) free.
[1588058353] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11402112 kB) free!
[1588058378] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.7% (4349652 kB) free!
[1588058473] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33366612 kB) free.
[1588058499] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.9% (4133564 kB) free!
[1588058619] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.1% (3943400 kB) free!
[1588058664] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1588058664] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 33.4% (8176756 kB) free.
[1588059073] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2296236 kB) free!
[1588059193] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35227232 kB) free.
[1588059793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4446872 kB) free!
[1588059819] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1966468 kB) free!
[1588059914] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384420 kB) free!
[1588060034] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32238280 kB) free.
[1588060634] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3763480 kB) free!
[1588060754] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34839916 kB) free.
[1588061354] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6644528 kB) free!
[1588061474] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34223216 kB) free.
[1588062074] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9714388 kB) free!
[1588062194] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (21011164 kB) free!
[1588062314] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30139092 kB) free.
[1588062914] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6344264 kB) free!
[1588063034] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31687792 kB) free.
[1588063464] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4617552 kB) free!
[1588063585] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4562168 kB) free!
[1588063634] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6841856 kB) free!
[1588063705] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4491176 kB) free!
[1588063754] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31653624 kB) free.
[1588064354] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7759396 kB) free!
[1588064474] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31348328 kB) free.
[1588064822] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9570928 kB) free!
[1588064912] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1588064942] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9437788 kB) free!
[1588065062] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9304668 kB) free!
[1588065074] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6713860 kB) free!
[1588065195] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31171468 kB) free.
[1588065403] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4476216 kB) free!
[1588065523] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4331464 kB) free!
[1588065643] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.6% (4166456 kB) free!
[1588065707] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.66, 4.35, 4.85
[1588065794] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6841992 kB) free!
[1588065803] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 343 processes
[1588065914] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30590332 kB) free.
[1588065923] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 340 processes
[1588066043] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1588066243] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.5% (16064584 kB) free.
[1588066307] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.92, 9.75, 7.31
[1588066427] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 15.24, 11.93, 8.41
[1588066514] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4910896 kB) free!
[1588066547] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.24, 11.12, 8.54
[1588066634] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17602812 kB) free!
[1588066689] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(153) < 185.255.8.58 [28/Apr/2020:15:00:07 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1588066754] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26684348 kB) free.
[1588066809] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588067354] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2897144 kB) free!
[1588067474] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27480820 kB) free.
[1588068009] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(132) < 103.255.146.164 [28/Apr/2020:15:27:22 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588068074] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2576012 kB) free!
[1588068129] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588068194] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28152008 kB) free.
[1588068505] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2136816 kB) free!
[1588068662] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4688332 kB) free!
[1588068794] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3047292 kB) free!
[1588068914] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.9% (26272408 kB) free!
[1588068947] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.87, 3.57, 4.54
[1588069034] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21052628 kB) free!
[1588069634] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35594168 kB) free.
[1588070171] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (558192 kB) free!
[1588070235] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8986140 kB) free!
[1588070355] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34055920 kB) free.
[1588070747] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.36, 6.14, 5.34
[1588070771] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (1894664 kB) free!
[1588070867] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.36, 5.73, 5.28
[1588070955] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10189660 kB) free!
[1588070987] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1588070987] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.12, 8.57, 6.34
[1588071075] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36086188 kB) free.
[1588071675] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9476920 kB) free!
[1588071795] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20825272 kB) free!
[1588071915] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31590724 kB) free.
[1588072236] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (1722776 kB) free!
[1588072515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6438312 kB) free!
[1588072635] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33853920 kB) free.
[1588073235] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9292740 kB) free!
[1588073355] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4671432 kB) free!
[1588073475] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30138068 kB) free.
[1588073771] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1494012 kB) free!
[1588074075] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4816904 kB) free!
[1588074195] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31878084 kB) free.
[1588074643] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4488096 kB) free!
[1588074763] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4050616 kB) free!
[1588074795] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (4948012 kB) free!
[1588074883] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.0% (3601568 kB) free!
[1588074915] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (503644 kB) free!
[1588074971] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1588074971] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2756076 kB) free!
[1588075036] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32412580 kB) free.
[1588075483] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1588075483] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1358416 kB) free!
[1588075635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5572168 kB) free!
[1588075705] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1588075755] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1246108 kB) free!
[1588075875] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32065056 kB) free.
[1588076312] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1588076408] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9462620 kB) free!
[1588076432] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1588076475] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5547436 kB) free!
[1588076528] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9314552 kB) free!
[1588076553] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1588076595] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1072948 kB) free!
[1588076648] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (9152848 kB) free!
[1588076715] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31906552 kB) free.
[1588077284] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16193368 kB) free.
[1588077315] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6065668 kB) free!
[1588077371] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1119332 kB) free!
[1588077435] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1917000 kB) free!
[1588077555] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30325336 kB) free.
[1588077971] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (2150424 kB) free!
[1588078155] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7115380 kB) free!
[1588078275] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2816104 kB) free!
[1588078395] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29046804 kB) free.
[1588078787] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.46, 3.04, 4.78
[1588078995] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5953616 kB) free!
[1588079115] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16161500 kB) free!
[1588079171] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (941876 kB) free!
[1588079235] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27801100 kB) free.
[1588079771] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3268972 kB) free.
[1588079836] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4495596 kB) free!
[1588079956] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29587164 kB) free.
[1588080371] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3205144 kB) free!
[1588080491] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3202624 kB) free!
[1588080556] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8264508 kB) free!
[1588080612] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1317508 kB) free!
[1588080676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3791896 kB) free!
[1588080796] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29249336 kB) free.
[1588081212] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.8% (2402676 kB) free!
[1588081396] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6083316 kB) free!
[1588081516] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1998608 kB) free!
[1588081636] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28872656 kB) free.
[1588082236] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5583608 kB) free!
[1588082356] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1628480 kB) free!
[1588082476] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27328608 kB) free.
[1588083012] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1356280 kB) free!
[1588083076] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3233440 kB) free!
[1588083196] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28344784 kB) free.
[1588083612] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (2057452 kB) free!
[1588083796] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5035208 kB) free!
[1588083916] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (959720 kB) free!
[1588084037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29340224 kB) free.
[1588084636] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3622944 kB) free!
[1588084756] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374964 kB) free!
[1588084787] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1588084876] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29701992 kB) free.
[1588085476] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5992708 kB) free!
[1588085596] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2116324 kB) free!
[1588085716] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28677244 kB) free.
[1588086316] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6055948 kB) free!
[1588086436] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2183272 kB) free!
[1588086556] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27759048 kB) free.
[1588087156] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4442268 kB) free!
[1588087276] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29826192 kB) free.
[1588087484] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588087876] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8361324 kB) free!
[1588087996] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4342420 kB) free!
[1588088116] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29641260 kB) free.
[1588088323] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1588088443] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1588088621] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 35.4% (8684600 kB) free.
[1588088717] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6827832 kB) free!
[1588088837] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2916776 kB) free!
[1588088957] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29320712 kB) free.
[1588089043] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1588089163] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1588089219] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 65.2% (15976000 kB) free.
[1588089557] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5954100 kB) free!
[1588089612] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (834808 kB) free!
[1588089677] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2155720 kB) free!
[1588089797] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27280596 kB) free.
[1588090212] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.7% (4015492 kB) free.
[1588090397] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5240184 kB) free!
[1588090517] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29239828 kB) free.
[1588091117] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7376296 kB) free!
[1588091237] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3604448 kB) free!
[1588091357] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28280984 kB) free.
[1588091412] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3201224 kB) free!
[1588091532] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2901656 kB) free!
[1588091652] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (2647904 kB) free!
[1588091957] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5302980 kB) free!
[1588092077] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29008636 kB) free.
[1588092252] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3870788 kB) free.
[1588092677] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6602464 kB) free!
[1588092797] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2593740 kB) free!
[1588092852] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3062624 kB) free!
[1588092917] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (25005844 kB) free!
[1588092972] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.4% (2831420 kB) free!
[1588093092] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (2487444 kB) free!
[1588093518] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2243996 kB) free!
[1588093692] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3742940 kB) free.
[1588094118] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14488708 kB) free!
[1588095318] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1588095318] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26694876 kB) free.
[1588096092] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3204172 kB) free!
[1588096212] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2809768 kB) free!
[1588096332] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.5% (2348892 kB) free!
[1588096518] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12791936 kB) free!
[1588096638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35134544 kB) free.
[1588096932] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3434168 kB) free.
[1588096953] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1588097238] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14229264 kB) free!
[1588097358] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36439956 kB) free.
[1588097958] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16196828 kB) free!
[1588098078] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12784956 kB) free!
[1588098198] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34679012 kB) free.