[1583692644] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (21012688 kB) free!
[1583692764] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17126856 kB) free!
[1583692884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45114256 kB) free.
[1583692964] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191488 kB) free!
[1583693034] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (204012 kB) free!
[1583693084] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193944 kB) free!
[1583693103] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186264 kB) free!
[1583693154] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189528 kB) free!
[1583693204] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189952 kB) free!
[1583693223] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176244 kB) free!
[1583693274] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198524 kB) free!
[1583693343] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189452 kB) free!
[1583693478] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(139) < 92.42.104.166 [09/Mar/2020:00:20:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583693484] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20662892 kB) free!
[1583693599] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583693604] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17136596 kB) free!
[1583693724] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44257652 kB) free.
[1583694148] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184108 kB) free!
[1583694210] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178776 kB) free!
[1583694270] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (182152 kB) free!
[1583694324] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20837556 kB) free!
[1583694330] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (181244 kB) free!
[1583694444] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17290088 kB) free!
[1583694565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43051660 kB) free.
[1583694630] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 70.1% (17176688 kB) free.
[1583695164] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16837228 kB) free!
[1583695284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44991020 kB) free.
[1583695884] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21884592 kB) free!
[1583696004] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17497240 kB) free!
[1583696124] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44903384 kB) free.
[1583696724] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (21958820 kB) free!
[1583696844] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18114020 kB) free!
[1583696943] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.8% (3661196 kB) free!
[1583696964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43725412 kB) free.
[1583697434] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240008 kB) free!
[1583697452] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (417736 kB) free!
[1583697543] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193492 kB) free!
[1583697554] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241648 kB) free!
[1583697564] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20879200 kB) free!
[1583697638] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (376844 kB) free!
[1583697674] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239792 kB) free!
[1583697684] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44844528 kB) free.
[1583697758] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (835976 kB) free!
[1583697878] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397472 kB) free!
[1583698246] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 223618 MB (20% inode=99%):
[1583698284] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23849964 kB) free!
[1583698366] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 223618 MB (20% inode=99%):
[1583698404] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19962276 kB) free!
[1583698486] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 223618 MB (20% inode=99%):
[1583698524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44800508 kB) free.
[1583699124] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21317804 kB) free!
[1583699245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44817852 kB) free.
[1583699474] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 21.0% (10326436 kB) free.
[1583699598] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2345) < 103.255.146.164 [09/Mar/2020:02:01:14 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583699719] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583699844] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22644672 kB) free!
[1583699964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (19015664 kB) free!
[1583700084] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42687172 kB) free.
[1583700684] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21628164 kB) free!
[1583700804] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44749480 kB) free.
[1583701182] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234932 kB) free!
[1583701302] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238152 kB) free!
[1583701404] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17518376 kB) free!
[1583701422] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240916 kB) free!
[1583701524] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11900456 kB) free!
[1583701644] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43112488 kB) free.
[1583702022] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.7% (42241864 kB) free.
[1583702244] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16033932 kB) free!
[1583702364] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11165352 kB) free!
[1583702484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41067740 kB) free.
[1583702804] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.1% (12043340 kB) free.
[1583703084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15902228 kB) free!
[1583703086] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(271) < 103.16.101.72 [09/Mar/2020:03:00:31 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583703204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44290076 kB) free.
[1583703207] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583703319] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1255) < 203.212.70.170 [09/Mar/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583703439] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583703501] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(340) < 2020/03/09 03:01:31 [error] 12253#0: *67469019 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583703621] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583703804] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19625296 kB) free!
[1583703924] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15696928 kB) free!
[1583704045] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583704045] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44304108 kB) free.
[1583704644] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18174808 kB) free!
[1583704764] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44297348 kB) free.
[1583705364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19937828 kB) free!
[1583705484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44349832 kB) free.
[1583706084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26339092 kB) free!
[1583706204] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (22737536 kB) free!
[1583706324] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44117796 kB) free.
[1583706924] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23542176 kB) free!
[1583707044] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44629300 kB) free.
[1583707644] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23854496 kB) free!
[1583707764] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45112540 kB) free.
[1583708364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24503744 kB) free!
[1583708484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.3% (46510192 kB) free.
[1583709440] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(168) < 92.42.104.166 [09/Mar/2020:04:45:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583709560] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583710160] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 172.105.59.173 [09/Mar/2020:04:58:58 +0530] TCP 502 0 0 92.064 "10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "31.061, 1.002, 60.001, 0.000"
[1583710280] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(2) < 182.18.144.246 [09/Mar/2020:05:00:02 +0530] TCP 502 0 0 120.001 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1583710400] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583710604] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197988 kB) free!
[1583710688] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.7% (4831784 kB) free!
[1583710724] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182172 kB) free!
[1583710808] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.7% (4827296 kB) free!
[1583710844] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 27.5% (6738596 kB) free.
[1583710928] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.7% (4823464 kB) free!
[1583711275] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.7% (19538952 kB) free.
[1583711343] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 64.4% (15890172 kB) free.
[1583711445] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172788 kB) free!
[1583711565] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (209524 kB) free!
[1583711685] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195192 kB) free!
[1583714543] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 91.9% (121168444 kB) free.
[1583714600] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39210208 kB) free.
[1583714600] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(119) < 92.42.104.166 [09/Mar/2020:06:10:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583714641] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.9% (21298216 kB) free.
[1583714720] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583715286] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 264519 MB (24% inode=99%):
[1583715321] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(144) < 92.42.104.166 [09/Mar/2020:06:25:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583715441] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583716143] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.7% (4127684 kB) free!
[1583716263] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4124380 kB) free!
[1583716383] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (4108432 kB) free!
[1583717241] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(288) < 182.18.144.246 [09/Mar/2020:06:57:06 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583717361] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583717961] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [09/Mar/2020:06:59:42 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1583718082] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583719983] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199040 kB) free!
[1583720481] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(642) < 203.212.70.170 [09/Mar/2020:07:46:43 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1583720601] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583724806] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(150) < 103.16.101.72 [09/Mar/2020:09:02:12 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583724927] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583725222] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(355) < 2020/03/09 09:03:05 [error] 12253#0: *67563635 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583725328] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1869504 kB) free!
[1583725342] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583729222] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.7% (21746336 kB) free.
[1583729261] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.3% (21039844 kB) free.
[1583729484] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21214820 kB) free!
[1583729601] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583729601] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1262) < 103.255.146.164 [09/Mar/2020:10:15:53 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583729604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42109444 kB) free.
[1583729722] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1583729722] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583729823] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198340 kB) free!
[1583729842] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583729876] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185436 kB) free!
[1583729943] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187744 kB) free!
[1583729996] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191532 kB) free!
[1583730063] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194820 kB) free!
[1583730116] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182068 kB) free!
[1583730204] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19118260 kB) free!
[1583730324] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41116492 kB) free.
[1583730442] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1583730460] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583730580] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1583730700] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583730924] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18632192 kB) free!
[1583731044] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40031712 kB) free.
[1583731644] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17411928 kB) free!
[1583731764] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28525332 kB) free.
[1583732364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14618452 kB) free!
[1583732463] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1583732463] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 73.8% (18081080 kB) free.
[1583732484] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12793916 kB) free!
[1583732604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32017480 kB) free.
[1583733101] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1583733400] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 10.0% (2462020 kB) free!
[1583733805] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13149756 kB) free!
[1583733925] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8720608 kB) free!
[1583734045] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33177604 kB) free.
[1583734203] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.1% (3166840 kB) free!
[1583734323] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.1% (2721960 kB) free!
[1583734443] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.2% (2280444 kB) free!
[1583734645] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7899920 kB) free!
[1583734664] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (4349976 kB) free!
[1583734765] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20911512 kB) free!
[1583734784] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (4179776 kB) free!
[1583734885] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29589272 kB) free.
[1583734904] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (4046988 kB) free!
[1583735043] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (471712 kB) free!
[1583735485] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5468672 kB) free!
[1583735516] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 27.3% (6696296 kB) free.
[1583735605] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30464224 kB) free.
[1583735643] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.5% (15387260 kB) free.
[1583736104] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (266488 kB) free!
[1583736205] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7662224 kB) free!
[1583736304] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 89.6% (21972412 kB) free.
[1583736325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3177804 kB) free!
[1583736346] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1583736346] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.2% (20534564 kB) free.
[1583736390] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.1% (56758252 kB) free.
[1583737316] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4739460 kB) free!
[1583737404] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9078260 kB) free!
[1583737436] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4614048 kB) free!
[1583737524] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9063872 kB) free!
[1583737556] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4497000 kB) free!
[1583737645] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (9063668 kB) free!
[1583737807] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1186_n"
[1583737927] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583738793] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13695820 kB) free!
[1583738913] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9396956 kB) free!
[1583739033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43449684 kB) free.
[1583739634] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14508756 kB) free!
[1583739754] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10395408 kB) free!
[1583739874] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44227360 kB) free.
[1583739956] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1583739956] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1914060 kB) free!
[1583740444] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4891204 kB) free!
[1583740474] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16915840 kB) free!
[1583740564] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4716620 kB) free!
[1583740594] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12459020 kB) free!
[1583740684] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4566376 kB) free!
[1583740714] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42102380 kB) free.
[1583741154] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (2312948 kB) free!
[1583741245] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4417176 kB) free!
[1583741274] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (2155288 kB) free!
[1583741314] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15086536 kB) free!
[1583741394] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1998328 kB) free!
[1583741434] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10670304 kB) free!
[1583741554] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41064572 kB) free.
[1583741623] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9152696 kB) free!
[1583741743] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8883976 kB) free!
[1583741863] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8639168 kB) free!
[1583742063] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583742154] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14518308 kB) free!
[1583742274] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10108484 kB) free!
[1583742394] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38972512 kB) free.
[1583742484] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1814748 kB) free!
[1583742663] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4877476 kB) free!
[1583742783] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.9% (4628816 kB) free!
[1583742903] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.3% (4497048 kB) free!
[1583742994] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10829408 kB) free!
[1583743114] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6308488 kB) free!
[1583743235] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2045032 kB) free!
[1583743443] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4448744 kB) free!
[1583743563] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4107580 kB) free!
[1583743663] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4465316 kB) free!
[1583743683] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.6% (3732140 kB) free!
[1583744283] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1583744283] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1783636 kB) free!
[1583744435] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15459852 kB) free!
[1583744703] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1268588 kB) free!
[1583745035] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383448 kB) free!
[1583745674] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8690968 kB) free!
[1583745794] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8331828 kB) free!
[1583745914] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.1% (7953212 kB) free!
[1583746235] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1583746235] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25190176 kB) free!
[1583746515] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.8% (29957740 kB) free.
[1583746522] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(209) < 92.42.104.166 [09/Mar/2020:15:00:22 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583746539] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.1% (21829896 kB) free.
[1583746568] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.7% (20657064 kB) free.
[1583746604] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.6% (15393544 kB) free.
[1583746638] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 56.0% (27606316 kB) free.
[1583746642] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583746669] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 28.4% (13997700 kB) free.
[1583746805] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.5% (3826548 kB) free!
[1583746835] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392412 kB) free!
[1583746925] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.1% (3471268 kB) free!
[1583747045] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.6% (3114768 kB) free!
[1583747174] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (1148384 kB) free!
[1583747284] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 80.8% (19801884 kB) free.
[1583747294] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (1779100 kB) free!
[1583747414] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1560232 kB) free!
[1583747501] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 357 processes
[1583747622] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 352 processes
[1583747645] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 91.9% (22683352 kB) free.
[1583747742] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1583748035] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32560268 kB) free.
[1583748635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4352072 kB) free!
[1583748641] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.6% (20625100 kB) free.
[1583748669] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19512116 kB) free!
[1583748672] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33233628 kB) free.
[1583748678] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8857940 kB) free!
[1583748704] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 67.1% (16457464 kB) free.
[1583748798] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (8651848 kB) free!
[1583748918] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (8446444 kB) free!
[1583749001] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1583749121] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1583749241] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1583749841] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1583749873] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21106536 kB) free!
[1583749979] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1583749993] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42548092 kB) free.
[1583750099] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 330 processes
[1583750183] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1583750219] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1583750303] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1583750423] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1583750593] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14328728 kB) free!
[1583750713] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9709672 kB) free!
[1583750833] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5372704 kB) free!
[1583750909] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3222936 kB) free!
[1583751023] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1583751029] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3392312 kB) free.
[1583751433] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18845820 kB) free!
[1583751623] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1583751743] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1583751864] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1583751918] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4107460 kB) free!
[1583752033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37009824 kB) free.
[1583752222] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 29.5% (14517196 kB) free.
[1583752244] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.2% (2509600 kB) free!
[1583752364] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (2350024 kB) free!
[1583752464] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583752464] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1583752484] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2185916 kB) free!
[1583752633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8960504 kB) free!
[1583752753] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4321008 kB) free!
[1583752873] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396700 kB) free!
[1583752912] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583753473] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14287676 kB) free!
[1583753664] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1583753764] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.3% (20304396 kB) free.
[1583753784] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1583754073] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28619852 kB) free.
[1583754413] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.4% (3453032 kB) free!
[1583754442] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7621908 kB) free!
[1583754533] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3214528 kB) free!
[1583754563] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.5% (7622184 kB) free!
[1583754653] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3033964 kB) free!
[1583754673] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1904908 kB) free!
[1583754683] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.5% (7623244 kB) free!
[1583754793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (416540 kB) free!
[1583754914] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27009492 kB) free.
[1583755230] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2204272 kB) free!
[1583755253] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2162060 kB) free!
[1583755350] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2853396 kB) free!
[1583755470] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3511920 kB) free.
[1583755513] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (386192 kB) free!
[1583755633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (398580 kB) free!
[1583755684] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4718484 kB) free!
[1583755753] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398568 kB) free!
[1583755804] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (4573312 kB) free!
[1583755883] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4112972 kB) free!
[1583755924] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4443820 kB) free!
[1583756953] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31667484 kB) free.
[1583757054] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.0% (15498344 kB) free.
[1583757553] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4439424 kB) free!
[1583757673] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403708 kB) free!
[1583757793] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424404 kB) free!
[1583757870] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3015612 kB) free!
[1583757974] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.9% (4175660 kB) free!
[1583757990] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3559856 kB) free.
[1583758094] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.4% (4048844 kB) free!
[1583758214] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (3903104 kB) free!
[1583758324] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2026204 kB) free!
[1583758590] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (1622432 kB) free!
[1583758710] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.3% (2160652 kB) free!
[1583758814] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1747032 kB) free!
[1583758830] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (2727328 kB) free!
[1583758993] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29126808 kB) free.
[1583759142] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.1% (4441956 kB) free!
[1583759262] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.6% (4313316 kB) free!
[1583759382] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.0% (4170980 kB) free!
[1583759421] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.6% (20136888 kB) free.
[1583759457] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.9% (38878724 kB) free.
[1583759492] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 71.6% (17541264 kB) free.
[1583759593] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2234456 kB) free!
[1583759713] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (390132 kB) free!
[1583759833] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36826680 kB) free.
[1583760030] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.6% (3836292 kB) free.
[1583760434] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9104504 kB) free!
[1583760554] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4214460 kB) free!
[1583760630] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.2% (2307284 kB) free!
[1583760674] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393756 kB) free!
[1583760750] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2909192 kB) free!
[1583760870] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3563868 kB) free.
[1583761182] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1788648 kB) free!
[1583761274] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15732624 kB) free!
[1583761312] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4548840 kB) free!
[1583761432] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4458648 kB) free!
[1583761552] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4355844 kB) free!
[1583761824] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196292 kB) free!
[1583761874] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37273204 kB) free.
[1583761944] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192584 kB) free!
[1583762064] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196660 kB) free!
[1583762474] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8131140 kB) free!
[1583762594] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3401284 kB) free!
[1583762714] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410208 kB) free!
[1583763270] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (2473008 kB) free!
[1583763314] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14936008 kB) free!
[1583763390] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (2957516 kB) free!
[1583763510] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3472720 kB) free.
[1583763914] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32851284 kB) free.
[1583763984] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583764514] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6083180 kB) free!
[1583764552] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2036996 kB) free!
[1583764634] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1075132 kB) free!
[1583764710] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3028984 kB) free!
[1583764754] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378740 kB) free!
[1583764830] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3309464 kB) free.
[1583765354] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13669772 kB) free!
[1583765954] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31961528 kB) free.
[1583766030] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2125892 kB) free!
[1583766063] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9803420 kB) free!
[1583766150] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (2337004 kB) free!
[1583766183] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9676968 kB) free!
[1583766270] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2764252 kB) free!
[1583766303] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (9674392 kB) free!
[1583766554] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4550732 kB) free!
[1583766675] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382496 kB) free!
[1583766795] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31166280 kB) free.
[1583766870] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1583766870] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 35.8% (5825956 kB) free.
[1583767074] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9739060 kB) free!
[1583767194] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9641940 kB) free!
[1583767254] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583767314] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.4% (9551992 kB) free!
[1583767395] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12257468 kB) free!
[1583767515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7372236 kB) free!
[1583767635] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2714048 kB) free!
[1583768235] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (19007972 kB) free!
[1583768835] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38084948 kB) free.
[1583769435] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13510716 kB) free!
[1583769555] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9450816 kB) free!
[1583769675] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (5563928 kB) free!
[1583770275] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17585100 kB) free!
[1583770595] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.2% (21858004 kB) free.
[1583770630] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.0% (20000444 kB) free.
[1583770658] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.7% (20022080 kB) free.
[1583770875] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36477004 kB) free.
[1583771475] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12745872 kB) free!
[1583771595] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8461692 kB) free!
[1583771715] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3922576 kB) free!
[1583772315] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17082452 kB) free!
[1583772916] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31760144 kB) free.
[1583773516] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7528028 kB) free!
[1583773636] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3406804 kB) free!
[1583773756] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37013716 kB) free.
[1583774356] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11930548 kB) free!
[1583774476] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7579640 kB) free!
[1583774596] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3487532 kB) free!
[1583775196] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (16002912 kB) free!
[1583775342] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1583775796] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29188156 kB) free.
[1583776032] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583776396] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6074844 kB) free!
[1583776516] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (2013004 kB) free!
[1583776636] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34957324 kB) free.
[1583777236] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10724924 kB) free!
[1583777356] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6676264 kB) free!
[1583777476] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2727536 kB) free!
[1583777750] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 365 processes
[1583777870] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1583778076] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14486260 kB) free!
[1583778270] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)