[1551983430] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3920348 kB) free!
[1551983813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551983933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551984016] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9245356 kB) free!
[1551984031] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.3% (11474952 kB) free.
[1551984053] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551984136] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16133968 kB) free.
[1551984631] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9013368 kB) free!
[1551984653] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551984751] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.9% (6355216 kB) free!
[1551984773] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551984871] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3652384 kB) free!
[1551984893] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551985471] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11317404 kB) free.
[1551985493] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551985614] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551985734] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551986334] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551986454] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551986537] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8793828 kB) free!
[1551986574] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551986657] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16137900 kB) free.
[1551987174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551987294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551987415] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551988014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551988134] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551988254] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551988854] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551988974] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551989057] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9223796 kB) free!
[1551989071] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3559256 kB) free!
[1551989094] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551989177] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16128564 kB) free.
[1551989191] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11788324 kB) free.
[1551989791] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3485120 kB) free!
[1551989912] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.0% (11813756 kB) free.
[1551990511] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3367656 kB) free!
[1551990631] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.0% (11840116 kB) free.
[1551990895] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551991495] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551991577] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9323924 kB) free!
[1551991615] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551991697] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16133788 kB) free.
[1551991735] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551991831] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.9% (5387276 kB) free!
[1551991951] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.1% (11889584 kB) free.
[1551992335] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551992455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551992551] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4207612 kB) free!
[1551992575] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551992671] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11914256 kB) free.
[1551993175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551993296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551993416] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551993871] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8151348 kB) free!
[1551993991] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.1% (5476052 kB) free!
[1551994016] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551994097] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9726560 kB) free!
[1551994111] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.3% (11957384 kB) free.
[1551994136] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551994217] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16121616 kB) free.
[1551994256] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551994856] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551994976] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551995096] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551995311] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8784348 kB) free!
[1551995431] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.5% (6150888 kB) free!
[1551995551] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3472448 kB) free!
[1551995696] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551995817] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551995937] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551996151] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.3% (11965448 kB) free.
[1551996537] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551996657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551996751] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.4% (5614464 kB) free!
[1551996777] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551996871] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.4% (12000008 kB) free.
[1551997377] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551997497] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551997618] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551998071] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4565376 kB) free!
[1551998191] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.5% (12060140 kB) free.
[1551998217] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551998337] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551998457] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551998791] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4100888 kB) free!
[1551998911] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.6% (12102200 kB) free.
[1551999057] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551999177] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551999297] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551999897] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552000017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552000111] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.0% (5927056 kB) free!
[1552000138] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552000231] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.7% (12173600 kB) free.
[1552000737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552000831] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.1% (5470164 kB) free!
[1552000857] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552000951] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.8% (12215016 kB) free.
[1552000977] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552001551] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4882544 kB) free!
[1552001577] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552001671] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.9% (12251168 kB) free.
[1552001697] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552001817] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552002271] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4290780 kB) free!
[1552002391] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.9% (12282280 kB) free.
[1552002417] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552002538] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552002617] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8888472 kB) free!
[1552002658] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552002737] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16111976 kB) free.
[1552002991] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8446584 kB) free!
[1552003111] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.0% (12311688 kB) free.
[1552003258] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552003378] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552003498] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552004098] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552004218] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552004311] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.8% (5801124 kB) free!
[1552004338] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552004432] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.1% (12375464 kB) free.
[1552005031] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.2% (5033268 kB) free!
[1552005138] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9506008 kB) free!
[1552005151] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.2% (12415132 kB) free.
[1552005258] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16113820 kB) free.
[1552005751] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4605444 kB) free!
[1552005871] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.3% (12448892 kB) free.
[1552006138] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552006738] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552006858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552006978] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552007071] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4428460 kB) free!
[1552007191] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12500988 kB) free.
[1552007336] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 607 processes
[1552007396] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 599 processes
[1552007578] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552007658] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9786136 kB) free!
[1552007698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552007778] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16104684 kB) free.
[1552007818] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552008391] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4851784 kB) free!
[1552008419] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552008511] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.3% (12455640 kB) free.
[1552008539] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552008659] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552008776] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (182856 kB) free!
[1552008896] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (455440 kB) free!
[1552009016] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (297900 kB) free!
[1552009111] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1552009259] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552009379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552009499] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552009712] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.8% (5800036 kB) free!
[1552009832] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12518116 kB) free.
[1552010099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552010178] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9804920 kB) free!
[1552010219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552010299] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16112168 kB) free.
[1552010340] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552010432] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.8% (5333368 kB) free!
[1552010552] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.5% (12549864 kB) free.
[1552010939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552011059] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552011152] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4707804 kB) free!
[1552011179] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552011272] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.5% (12580824 kB) free.
[1552011779] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552011899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552012016] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.8% (15079516 kB) free.
[1552012019] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552012472] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5760620 kB) free!
[1552012592] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.6% (12619264 kB) free.
[1552012619] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552012698] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9543976 kB) free!
[1552012740] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552012818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.7% (16103484 kB) free.
[1552012860] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552013192] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9590072 kB) free!
[1552013312] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.3% (7031892 kB) free!
[1552013432] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4492476 kB) free!
[1552013460] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552013580] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552013700] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552014032] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11310724 kB) free.
[1552014408] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 890 processes
[1552014528] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 917 processes
[1552014564] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 197.85, 302.45, 162.09
[1552014633] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8253384 kB) free!
[1552014648] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 384 processes
[1552014684] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 26.95, 202.48, 142.47
[1552014753] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.6% (5698972 kB) free!
[1552014804] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.84, 135.56, 125.22
[1552014873] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.3% (12482996 kB) free.
[1552014900] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552015473] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9728116 kB) free!
[1552015500] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552015593] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.5% (7164084 kB) free!
[1552015621] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552015713] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11100004 kB) free.
[1552015741] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552016341] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552016419] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8487688 kB) free!
[1552016461] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552016539] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.8% (15182672 kB) free.
[1552016604] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.15, 0.59, 18.31
[1552016913] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7661420 kB) free!
[1552017033] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.1% (4973412 kB) free!
[1552017153] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.3% (11997956 kB) free.
[1552017661] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552017753] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9268956 kB) free!
[1552017781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552017873] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.7% (6742956 kB) free!
[1552017901] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552017993] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.2% (12442952 kB) free.
[1552018404] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.30, 2.90
[1552019072] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 334.49, 94.24, 32.75
[1552019192] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 526.62, 247.21, 96.26
[1552019193] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.4% (7120300 kB) free!
[1552019264] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 554.23, 319.81, 134.11
[1552019312] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 552.72, 349.61, 152.09
[1552019313] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (4445364 kB) free!
[1552019384] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 370.42, 355.27, 171.22
[1552019433] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11733824 kB) free.
[1552019461] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.72, 8.82, 7.64
[1552019504] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 51.90, 238.43, 150.72
[1552019581] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.83, 6.13, 6.79
[1552019701] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 4.25, 6.02
[1552020033] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8680052 kB) free!
[1552020153] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5970104 kB) free!
[1552020273] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3310400 kB) free!
[1552020301] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.78, 1.25, 3.52
[1552020873] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1552020873] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10194280 kB) free.
[1552021413] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1602836 kB) free!
[1552021473] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.2% (6528988 kB) free!
[1552021593] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (3849716 kB) free!
[1552021712] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 0.78, 19.15
[1552021713] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11313792 kB) free.
[1552021904] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.64, 1.94, 13.07
[1552021939] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7816584 kB) free!
[1552022060] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.7% (15122976 kB) free.
[1552022313] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8777356 kB) free!
[1552022433] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.2% (5998592 kB) free!
[1552022554] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3265268 kB) free!
[1552023105] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.83, 1.81, 4.91
[1552023154] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.5% (10583256 kB) free.
[1552023512] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.62, 0.79, 3.43
[1552023754] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7532856 kB) free!
[1552023814] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452144 kB) free!
[1552023874] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4760648 kB) free!
[1552023994] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (1965504 kB) free!
[1552024594] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.8% (8760488 kB) free!
[1552025101] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552025659] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6910784 kB) free!
[1552025701] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552025779] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.8% (5825344 kB) free!
[1552025794] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4226416 kB) free!
[1552025821] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552025899] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12462620 kB) free.
[1552025941] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552026499] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6927532 kB) free!
[1552026541] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552026619] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5850388 kB) free!
[1552026661] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552026739] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4622712 kB) free!
[1552026782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552027340] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6383756 kB) free!
[1552027594] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.6% (7674072 kB) free!
[1552028165] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2457392 kB) free!
[1552028194] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4315516 kB) free!
[1552029994] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.3% (9027588 kB) free!
[1552030340] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.4% (14991416 kB) free.
[1552031165] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (2721152 kB) free!
[1552031194] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4673324 kB) free!
[1552031582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552031614] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2720048 kB) free!
[1552032140] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6677516 kB) free!
[1552032260] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.7% (13659548 kB) free.
[1552032782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552032902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552032909] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21955 MB (10% inode=99%):
[1552032994] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.6% (9161852 kB) free!
[1552033022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552033414] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2415692 kB) free!
[1552033509] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29607 MB (14% inode=99%):
[1552034614] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1552034614] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 25.9% (6346612 kB) free.
[1552034660] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6824676 kB) free!
[1552034765] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.2% (4947648 kB) free.
[1552034780] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.7% (13648036 kB) free.
[1552034794] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2845460 kB) free!
[1552034861] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 332 processes
[1552035309] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22318 MB (10% inode=99%):
[1552035994] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.3% (8996088 kB) free!
[1552036062] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 326 processes
[1552036233] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 340 processes
[1552036509] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1552036509] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29472 MB (14% inode=99%):
[1552036580] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9448240 kB) free!
[1552036622] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552036700] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13240224 kB) free.
[1552037194] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4500544 kB) free!
[1552037709] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22287 MB (10% inode=99%):
[1552037822] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552037942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552038062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552038394] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.7% (9224768 kB) free!
[1552038909] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 30483 MB (14% inode=99%):
[1552039101] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7087480 kB) free!
[1552039221] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.1% (13843852 kB) free.
[1552040013] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 34.4% (5596128 kB) free.
[1552040194] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3647076 kB) free!
[1552040709] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22256 MB (10% inode=99%):
[1552041062] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1552041394] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.4% (9562460 kB) free!
[1552041621] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6166524 kB) free!
[1552041741] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13266832 kB) free.
[1552041909] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 30415 MB (14% inode=99%):
[1552042862] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1552042982] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552043102] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552043194] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233064 kB) free!
[1552044141] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6322436 kB) free!
[1552044261] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.0% (13311676 kB) free.
[1552044394] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.3% (9030112 kB) free!
[1552044814] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3035976 kB) free!
[1552044934] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2970000 kB) free!
[1552045054] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.9% (2910692 kB) free!
[1552045414] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552045809] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.22, 5.33, 5.11
[1552045929] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.06, 5.24, 5.11
[1552045962] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1552046022] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1552046049] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.03, 5.15, 5.10
[1552046082] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1552046142] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1552046194] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4575836 kB) free!
[1552046216] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (2913108 kB) free!
[1552046336] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.5% (2843796 kB) free!
[1552046456] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.1% (2777324 kB) free!
[1552046649] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.69, 4.25, 4.72
[1552046702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552047261] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6275904 kB) free!
[1552047302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552047381] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13255088 kB) free.
[1552047422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552047543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552047981] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5206868 kB) free!
[1552048101] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12454864 kB) free.
[1552049854] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1592452 kB) free!
[1552049901] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5358716 kB) free!
[1552049943] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552050021] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12407112 kB) free.
[1552050056] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1580504 kB) free!
[1552050543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552050663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552050783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552051233] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 340 processes
[1552051565] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4759512 kB) free!
[1552051685] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4719952 kB) free!
[1552051805] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (4681500 kB) free!
[1552052421] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5452060 kB) free!
[1552052541] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12417408 kB) free.
[1552052709] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552053783] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552054341] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5427324 kB) free!
[1552054461] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12346388 kB) free.
[1552054893] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 19.9% (1461040 kB) free!
[1552054983] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552055013] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.8% (1455608 kB) free!
[1552055103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552055134] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 19.7% (1449656 kB) free!
[1552055223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552055742] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 55 users currently logged in
[1552056262] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5910108 kB) free!
[1552056382] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12723644 kB) free.
[1552056394] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1552056909] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22058 MB (10% inode=99%):
[1552057509] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 32285 MB (15% inode=99%):
[1552057623] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552058182] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.7% (5762632 kB) free!
[1552058302] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12923980 kB) free.
[1552058794] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11202096 kB) free.
[1552058823] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552058902] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8045272 kB) free!
[1552058943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552059022] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6817124 kB) free!
[1552059063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1552059142] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12910632 kB) free.
[1552059742] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8416160 kB) free!
[1552059862] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7096860 kB) free!
[1552059983] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12849092 kB) free.
[1552060263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552060582] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8822236 kB) free!
[1552060594] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.2% (6001584 kB) free!
[1552060702] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7930308 kB) free!
[1552060715] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (3459036 kB) free!
[1552060822] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6726320 kB) free!
[1552060835] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12502272 kB) free.
[1552060863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552060983] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552061435] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4807324 kB) free!
[1552061555] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12502032 kB) free.
[1552061583] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552061703] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552061823] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552062022] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12919976 kB) free.
[1552062155] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.2% (6028084 kB) free!
[1552062275] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3505664 kB) free!
[1552062395] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12500412 kB) free.
[1552062423] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552062543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552062622] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6721316 kB) free!
[1552062664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552062742] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12906828 kB) free.
[1552062995] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4503048 kB) free!
[1552063115] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12503928 kB) free.
[1552063342] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8528484 kB) free!
[1552063462] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7319972 kB) free!
[1552063582] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5621780 kB) free!
[1552063715] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.9% (5392184 kB) free!
[1552063835] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (2735800 kB) free!
[1552063864] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552063955] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12505996 kB) free.
[1552064464] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552064555] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3917608 kB) free!
[1552064584] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552064676] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12514844 kB) free.
[1552064783] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1552064783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12907600 kB) free.
[1552065184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552065275] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.3% (5058328 kB) free!
[1552065304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552065395] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11693852 kB) free.
[1552065424] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552065983] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8629784 kB) free!
[1552065995] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.5% (6159320 kB) free!
[1552066025] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552066103] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7468136 kB) free!
[1552066115] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (3426752 kB) free!
[1552066145] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552066223] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6135092 kB) free!
[1552066235] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.5% (12548468 kB) free.
[1552066265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552066835] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4343660 kB) free!
[1552066955] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.5% (12562656 kB) free.
[1552067423] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12909864 kB) free.
[1552067465] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1552067555] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.5% (5192384 kB) free!
[1552067675] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.5% (12584436 kB) free.
[1552068023] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6751168 kB) free!
[1552068065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552068143] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12900648 kB) free.
[1552068185] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552068275] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6186420 kB) free!
[1552068395] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3544788 kB) free!
[1552068516] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.5% (12569732 kB) free.
[1552068743] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8478620 kB) free!
[1552068785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552068863] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7273036 kB) free!
[1552068905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1552068984] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.6% (5701072 kB) free!
[1552069026] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552069115] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4345828 kB) free!
[1552069235] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.5% (12589956 kB) free.
[1552069625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552069745] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes