[1580495929] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189944 kB) free!
[1580496049] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186976 kB) free!
[1580496169] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176160 kB) free!
[1580496339] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188816 kB) free!
[1580496389] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181772 kB) free!
[1580496459] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184084 kB) free!
[1580496509] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194548 kB) free!
[1580496579] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185328 kB) free!
[1580496629] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178556 kB) free!
[1580497369] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 20.5% (5063704 kB) free.
[1580497969] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172044 kB) free!
[1580498089] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170788 kB) free!
[1580498209] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1580498209] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173980 kB) free!
[1580498809] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 49.5% (12215580 kB) free.
[1580499702] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3968720 kB) free!
[1580500009] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176112 kB) free!
[1580500129] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195292 kB) free!
[1580500210] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 49.50.86.152 [01/Feb/2020:01:16:57 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1580500249] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167284 kB) free!
[1580500331] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580500535] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(87) < 2020/02/01 01:17:52 [error] 10216#0: *46782834 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1580500655] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580502629] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.2% (10348168 kB) free.
[1580503580] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4301148 kB) free!
[1580504780] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.2% (40515416 kB) free.
[1580506331] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(158) < 103.16.101.72 [01/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1580506418] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2273) < 203.212.70.170 [01/Feb/2020:03:03:22 +0530] TCP 502 0 0 180.002 "10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.001, 60.000, 0.000"
[1580506453] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580506538] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580506655] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(204) < 2020/02/01 03:00:22 [error] 10211#0: *46811909 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580506776] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580507117] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.8% (27419200 kB) free.
[1580510863] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21790692 kB) free.
[1580510904] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580510908] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/01 04:18:27 [error] 18346#0: *27370700 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1580510943] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39984900 kB) free.
[1580510978] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.2% (22118736 kB) free.
[1580511010] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(16) < 2020/02/01 04:20:07 [error] 18342#0: *27370889 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1580511013] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580511016] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/01 04:20:14 [error] 18342#0: *27370894 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580511019] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/02/01 04:20:19 [error] 18343#0: *27370898 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580511023] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580511627] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(100) < 2020/02/01 04:30:21 [error] 18336#0: *27371691 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580511747] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(20) < 2020/02/01 04:32:21 [error] 18341#0: *27371926 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580511867] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(20) < 2020/02/01 04:34:22 [error] 18345#0: *27372139 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580513429] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (256980 kB) free!
[1580513549] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192328 kB) free!
[1580513670] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188944 kB) free!
[1580513869] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165472 kB) free!
[1580513987] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1719600 kB) free!
[1580513989] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173172 kB) free!
[1580514108] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167352 kB) free!
[1580514109] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182800 kB) free!
[1580514228] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.3% (21894536 kB) free.
[1580514709] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21228224 kB) free.
[1580517318] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26235028 kB) free!
[1580517438] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.5% (25672588 kB) free!
[1580517558] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.1% (25119772 kB) free!
[1580518158] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 36.0% (47502316 kB) free.
[1580518242] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.7% (21737232 kB) free.
[1580521132] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 763 processes
[1580521252] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 761 processes
[1580521372] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 325 processes
[1580521539] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [01/Feb/2020:07:09:28 +0530] TCP 502 0 0 135.043 "10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 15.042, 60.001, 60.000"
[1580521659] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580524388] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580524509] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1580524775] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(85) < 2020/02/01 08:01:22 [error] 10222#0: *46899371 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580524895] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580526109] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1580526695] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(60) < 2020/02/01 08:33:05 [error] 10212#0: *46910530 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580526815] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580528860] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(897) < 103.255.146.164 [01/Feb/2020:09:14:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1580528980] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580529216] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(115) < 2020/02/01 09:15:47 [error] 10222#0: *46925321 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1580529336] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580530125] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1580530245] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 331 processes
[1580530365] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1580535109] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4710028 kB) free!
[1580535229] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4555500 kB) free!
[1580535349] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4242604 kB) free!
[1580535647] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580535731] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1580535768] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580535828] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169540 kB) free!
[1580535851] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1580535888] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1580535948] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175828 kB) free!
[1580536068] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165892 kB) free!
[1580536169] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190520 kB) free!
[1580536289] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195356 kB) free!
[1580536409] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188632 kB) free!
[1580536549] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1706412 kB) free!
[1580538209] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 64.3% (15748416 kB) free.
[1580538580] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(603) < 103.255.146.164 [01/Feb/2020:11:50:42 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1580538700] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(30) < 203.212.70.170 [01/Feb/2020:12:01:15 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1580538821] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(62) < 172.105.59.173 [01/Feb/2020:12:03:15 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1580538936] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2020/02/01 12:01:15 [error] 10207#0: *46986422 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580539056] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580539421] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580539661] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.8% (21428916 kB) free.
[1580539697] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.4% (21912824 kB) free.
[1580540565] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 336 processes
[1580540685] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1580540805] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1580544468] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (4149120 kB) free!
[1580544548] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9244744 kB) free!
[1580544588] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.4% (4043212 kB) free!
[1580544668] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9245156 kB) free!
[1580544708] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (3928656 kB) free!
[1580544788] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (9245180 kB) free!
[1580546508] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1580546508] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (676048 kB) free!
[1580547221] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(153) < 92.42.104.166 [01/Feb/2020:14:21:29 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1580547341] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580550049] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4397492 kB) free!
[1580550169] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (4310016 kB) free!
[1580550188] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4861948 kB) free!
[1580550289] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4226092 kB) free!
[1580550380] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9633248 kB) free!
[1580550500] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9487500 kB) free!
[1580550620] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9314204 kB) free!
[1580550788] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (5140448 kB) free!
[1580551388] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1580551388] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3724292 kB) free!
[1580551606] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1580551726] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1580552689] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2185920 kB) free!
[1580553341] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(14) < 172.105.59.173 [01/Feb/2020:16:02:51 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1580553461] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580553810] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4889588 kB) free!
[1580553930] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.5% (4780524 kB) free!
[1580553967] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 56.9% (14038888 kB) free.
[1580553997] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 56.1% (27652768 kB) free.
[1580554034] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.4% (21912544 kB) free.
[1580554050] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.1% (4673912 kB) free!
[1580554451] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1580554488] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580554508] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580554571] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1580554608] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1580554628] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580554650] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187448 kB) free!
[1580554691] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1580554697] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190508 kB) free!
[1580554728] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1580554817] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167136 kB) free!
[1580554820] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4601500 kB) free!
[1580554938] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198660 kB) free!
[1580555170] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.0% (3938384 kB) free!
[1580555290] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.1% (3234396 kB) free!
[1580555292] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1580555328] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1580555410] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.6% (3102116 kB) free!
[1580555928] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1580556010] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437064 kB) free!
[1580556450] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 83.4% (20431544 kB) free.
[1580556636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1580556636] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9132392 kB) free!
[1580556756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1580556756] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7859948 kB) free!
[1580556876] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 333 processes
[1580556876] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13263528 kB) free.
[1580557476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1580558077] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8352100 kB) free!
[1580558197] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.3% (13922696 kB) free.
[1580558676] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1580559276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1580559396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1580562399] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1580563000] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.7% (7239528 kB) free!
[1580563120] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.7% (7236120 kB) free!
[1580563240] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (7238972 kB) free!
[1580563840] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228928 kB) free!
[1580565687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25907728 kB) free!
[1580565786] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1580565807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45513280 kB) free.
[1580565833] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.1% (18780320 kB) free.
[1580565847] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.4% (35179448 kB) free.
[1580565867] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 26.5% (13038356 kB) free.
[1580565898] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.3% (17469280 kB) free.
[1580567856] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(254) < 2020/02/01 20:03:22 [error] 10217#0: *47174357 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580567861] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [01/Feb/2020:20:02:38 +0530] TCP 502 0 0 61.003 "10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 1.002, 60.000, 0.000"
[1580567976] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580567981] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580568807] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25236460 kB) free!
[1580568927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44251760 kB) free.
[1580569097] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9782352 kB) free!
[1580569217] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9716544 kB) free!
[1580569338] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (9681036 kB) free!
[1580569439] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.3% (2530456 kB) free!
[1580569559] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.1% (2502560 kB) free!
[1580569679] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2462224 kB) free!
[1580571327] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24372944 kB) free!
[1580571437] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4824808 kB) free!
[1580571447] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44114420 kB) free.
[1580571557] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4736144 kB) free!
[1580571678] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4644080 kB) free!
[1580573906] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.9% (21454332 kB) free.
[1580574447] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24062408 kB) free!
[1580574567] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42544964 kB) free.
[1580574678] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2199572 kB) free!
[1580577567] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23078692 kB) free!
[1580577687] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41666400 kB) free.
[1580580437] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (1878020 kB) free!
[1580580557] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237160 kB) free!
[1580580677] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237284 kB) free!
[1580581158] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10243772 kB) free!
[1580581279] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9862308 kB) free!
[1580581399] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9502876 kB) free!