[1558463690] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.8% (13784212 kB) free.
[1558463799] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558464399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558464519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558464639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558465046] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558468839] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1558469164] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.19, 22.04, 10.08
[1558469284] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.44, 18.32, 10.29
[1558469404] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.39, 12.32, 9.06
[1558469439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558469559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558469679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558469812] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 613 processes
[1558469932] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 383 processes
[1558470004] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.32, 127.94, 90.88
[1558470055] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 373 processes
[1558470279] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558470880] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558471000] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558471120] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558471256] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 322 processes
[1558471646] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 273.17, 182.21, 81.43
[1558471766] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 140.86, 178.43, 93.30
[1558471804] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.33, 0.53, 13.24
[1558471886] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.22, 119.47, 82.02
[1558473004] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1558473004] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.27, 3.82
[1558473686] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.51, 0.80, 12.27
[1558474886] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.40, 3.64
[1558475094] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1558475214] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1558476520] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558478920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558479041] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1558481440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558481560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558481680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558483804] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558485880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558486113] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 1 warn out of 609 processes
[1558486173] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 606 processes
[1558486388] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.7% (3364612 kB) free!
[1558486480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558486508] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3328856 kB) free!
[1558486601] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558486628] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (3307544 kB) free!
[1558486721] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558487321] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1558487921] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558488041] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558488161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558488491] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9736156 kB) free!
[1558488611] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8171744 kB) free!
[1558488731] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6081848 kB) free!
[1558488761] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558489204] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.15, 113.63, 83.73
[1558489256] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 357 processes
[1558489324] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.78, 76.07, 73.60
[1558489331] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11558596 kB) free.
[1558489361] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558489444] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.78, 51.07, 64.74
[1558489482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558489602] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558490531] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8810832 kB) free!
[1558490645] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.27, 1.14, 17.93
[1558490651] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7105676 kB) free!
[1558490771] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12169480 kB) free.
[1558491139] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 289.08, 161.48, 67.32
[1558491259] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 294.69, 205.85, 95.00
[1558491379] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 295.30, 235.40, 119.26
[1558491428] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1442192 kB) free!
[1558491971] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9599568 kB) free!
[1558492028] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (3779868 kB) free!
[1558492091] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8044072 kB) free!
[1558492212] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6011272 kB) free!
[1558492256] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 323 processes
[1558492445] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.31, 2.84
[1558492812] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1558492812] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11701712 kB) free.
[1558493779] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.85, 2.00, 14.13
[1558494012] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9064584 kB) free!
[1558494132] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7549808 kB) free!
[1558494252] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12139712 kB) free.
[1558495452] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9776952 kB) free!
[1558495574] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8220856 kB) free!
[1558495579] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.53, 1.73, 3.63
[1558495694] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6263552 kB) free!
[1558496294] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11784668 kB) free.
[1558496457] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 321 processes
[1558497057] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1558497057] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 321 processes
[1558497428] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1558497428] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2018968 kB) free!
[1558497494] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8668384 kB) free!
[1558497614] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7033820 kB) free!
[1558497734] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4964036 kB) free!
[1558498334] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11650296 kB) free.
[1558499534] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8990460 kB) free!
[1558499654] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7263932 kB) free!
[1558499775] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4925148 kB) free!
[1558499802] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558500375] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10946504 kB) free.
[1558500402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558500522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558500642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558500975] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9227876 kB) free!
[1558501095] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7479180 kB) free!
[1558501215] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.0% (5407424 kB) free!
[1558501815] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11582704 kB) free.
[1558502048] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190420 kB) free!
[1558502168] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183208 kB) free!
[1558502272] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 361 processes
[1558502288] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189096 kB) free!
[1558502348] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (189284 kB) free!
[1558502392] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 362 processes
[1558502468] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (160172 kB) free!
[1558502512] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1558502588] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159028 kB) free!
[1558503015] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8511120 kB) free!
[1558503135] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6831532 kB) free!
[1558503255] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4577692 kB) free!
[1558503712] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 361 processes
[1558503832] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1558503855] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11220200 kB) free.
[1558504028] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1558504455] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9851228 kB) free!
[1558504575] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8064908 kB) free!
[1558504696] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6339536 kB) free!
[1558504816] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.7% (4100948 kB) free!
[1558504842] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558504936] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.5% (4042444 kB) free!
[1558505056] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (3977880 kB) free!
[1558505090] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2633032 kB) free!
[1558505210] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.2% (2471984 kB) free!
[1558505330] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.2% (2306696 kB) free!
[1558505442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558505562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558505682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558505791] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1558505911] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1558505930] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1244824 kB) free!
[1558506031] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1558506086] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.02, 6.72, 2.84
[1558506206] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.18, 4.71, 2.57
[1558506341] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.66, 6.00, 5.34
[1558506461] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.30, 6.03, 5.43
[1558506581] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.64, 6.02, 5.51
[1558506882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558507131] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.5% (13408484 kB) free.
[1558507483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558507603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558507723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558508296] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4605700 kB) free!
[1558508457] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558508656] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2084560 kB) free!
[1558508896] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10314288 kB) free.
[1558508923] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558509496] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9457084 kB) free!
[1558509523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558509616] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7384616 kB) free!
[1558509643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1558509736] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5669600 kB) free!
[1558509764] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558510028] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (3201252 kB) free!
[1558510228] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.2% (21366312 kB) free.
[1558510336] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4475980 kB) free!
[1558510381] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 43.7% (21552084 kB) free.
[1558510423] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 96.1% (23718680 kB) free.
[1558510470] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.2% (14826436 kB) free.
[1558510523] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.0% (22068276 kB) free.
[1558510663] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.8% (14926576 kB) free.
[1558510964] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558511564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558511684] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558511804] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558511867] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176524 kB) free!
[1558511987] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174340 kB) free!
[1558512107] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173260 kB) free!
[1558512945] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (734076 kB) free!
[1558513065] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (479628 kB) free!
[1558513185] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (226804 kB) free!
[1558513579] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 258.55, 185.76, 87.54
[1558513605] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558513699] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 258.91, 209.97, 108.32
[1558513819] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 202.31, 213.67, 122.49
[1558513840] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4348248 kB) free!
[1558513960] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (4225052 kB) free!
[1558514080] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (4103468 kB) free!
[1558514205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558514325] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558514431] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1558514445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558514683] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (1442508 kB) free!
[1558514803] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (1001228 kB) free!
[1558514923] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (628144 kB) free!
[1558514931] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 13.7% (2230876 kB) free!
[1558515051] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 12.2% (1989204 kB) free!
[1558515171] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.9% (1774308 kB) free!
[1558515229] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4797116 kB) free!
[1558515245] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 96.21, 169.01, 93.12
[1558515349] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4718004 kB) free!
[1558515365] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.64, 113.31, 81.91
[1558515469] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4642344 kB) free!
[1558515485] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.01, 75.88, 72.00
[1558515619] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.21, 2.93, 19.58
[1558515657] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 323 processes
[1558515771] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1558515771] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (817872 kB) free!
[1558515816] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46701 MB (20% inode=99%):
[1558515880] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1558515880] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2459932 kB) free!
[1558515936] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46616 MB (20% inode=99%):
[1558516056] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46548 MB (20% inode=99%):
[1558516245] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558516845] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558516966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558517086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558517286] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.30, 0.50, 10.66
[1558517419] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.91, 2.09, 4.58
[1558517869] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2389976 kB) free!
[1558517985] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.1% (6845552 kB) free.
[1558518057] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 326 processes
[1558518232] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1558518352] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558518486] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.41, 3.19
[1558518787] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9804384 kB) free!
[1558518886] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558518907] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 33.2% (16351196 kB) free.
[1558519219] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 23.60, 11.78, 6.31
[1558519300] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1558519300] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.7% (23134400 kB) free.
[1558519339] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 25.20, 16.27, 8.62
[1558519358] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.8% (22264588 kB) free.
[1558519459] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1558519459] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 25.58, 19.36, 10.68
[1558519486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558519545] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.2% (14817488 kB) free.
[1558519606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558519608] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 80.3% (13048904 kB) free.
[1558519651] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.9% (14941220 kB) free.
[1558519726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558520298] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167272 kB) free!
[1558520359] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 77.6% (12619648 kB) free.
[1558520966] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175240 kB) free!
[1558521086] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176052 kB) free!
[1558521206] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156748 kB) free!
[1558521259] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.70, 2.72, 4.45
[1558521307] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558521526] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558522126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558522246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1558522366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558522552] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 361 processes
[1558522672] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1558523107] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8892780 kB) free!
[1558523227] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.3% (15416952 kB) free.
[1558523272] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 360 processes
[1558523392] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 360 processes
[1558523512] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1558523985] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (1446376 kB) free!
[1558524105] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 18.1% (1328816 kB) free!
[1558524225] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 16.2% (1193540 kB) free!
[1558524349] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.9% (2740904 kB) free!
[1558524456] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24107 MB (10% inode=99%):
[1558524469] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.5% (2524964 kB) free!
[1558524581] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.36, 3.62, 4.33
[1558524589] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.4% (2343912 kB) free!
[1558524766] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558524825] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1558524825] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (420524 kB) free!
[1558525030] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1558525056] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27308 MB (12% inode=99%):
[1558525189] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1558525189] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1352072 kB) free!
[1558525358] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4622568 kB) free!
[1558525367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558525478] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4587080 kB) free!
[1558525487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558525598] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (4533712 kB) free!
[1558525607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558525616] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.4% (2665508 kB) free!
[1558525736] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2535692 kB) free!
[1558525838] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.9% (14931644 kB) free.
[1558525857] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.9% (2414840 kB) free!
[1558525908] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4789848 kB) free!
[1558525929] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 79.4% (12908220 kB) free.
[1558525997] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.4% (14863532 kB) free.
[1558526028] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4740760 kB) free!
[1558526060] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.8% (6900856 kB) free.
[1558526119] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.3% (23104032 kB) free.
[1558526139] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4892952 kB) free!
[1558526148] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4696088 kB) free!
[1558526171] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.3% (22043768 kB) free.
[1558526227] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9524776 kB) free!
[1558526228] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.0% (20845688 kB) free.
[1558526230] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;SOFT;1;USERS CRITICAL - 13 users currently logged in
[1558526256] SERVICE FLAPPING ALERT: smscapp1;HDD_Home status;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1558526256] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24671 MB (10% inode=99%):
[1558526347] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.8% (15682648 kB) free.
[1558526350] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;SOFT;2;USERS CRITICAL - 13 users currently logged in
[1558526447] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (184760 kB) free!
[1558526470] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1558526474] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558526567] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (183476 kB) free!
[1558526594] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558526665] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (1458800 kB) free!
[1558526687] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1558526687] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181772 kB) free!
[1558526714] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558526779] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.4% (3552248 kB) free!
[1558526785] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 19.7% (1449756 kB) free!
[1558526899] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.2% (3515336 kB) free!
[1558526905] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 19.6% (1441176 kB) free!
[1558527019] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.0% (3466852 kB) free!
[1558527190] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8810700 kB) free!
[1558527287] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 55.0% (8942404 kB) free.
[1558527310] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6826896 kB) free!
[1558527430] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.2% (14407416 kB) free.
[1558527456] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 74322 MB (33% inode=99%):
[1558527670] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1558527712] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 3 warn out of 360 processes
[1558527809] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (152036 kB) free!
[1558527832] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 4 warn out of 362 processes
[1558527887] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (155952 kB) free!
[1558527927] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.9% (2927176 kB) free!
[1558527929] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (151596 kB) free!
[1558527952] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 3 warn out of 365 processes
[1558528007] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (153172 kB) free!
[1558528047] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.8% (2883756 kB) free!
[1558528049] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174264 kB) free!
[1558528127] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165504 kB) free!
[1558528167] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (2844792 kB) free!
[1558528219] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176744 kB) free!
[1558528706] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 96.1% (7071280 kB) free.
[1558528752] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.84, 7.46, 5.33
[1558528768] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192072 kB) free!
[1558528872] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.71, 7.54, 5.62
[1558528993] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.49, 8.16, 6.07
[1558529230] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9111440 kB) free!
[1558529236] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185944 kB) free!
[1558529350] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7256836 kB) free!
[1558529356] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165944 kB) free!
[1558529470] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1558529470] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1558529470] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (5492028 kB) free!
[1558529476] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169168 kB) free!
[1558529947] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9675976 kB) free!
[1558530067] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15311896 kB) free.
[1558530070] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.3% (13964928 kB) free.
[1558530414] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1558530534] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1558531406] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 278.95, 164.56, 69.56
[1558531459] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1558531526] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 281.91, 203.56, 95.37
[1558531646] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.22, 156.33, 92.01
[1558531819] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2885044 kB) free!
[1558532152] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1558532207] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558532593] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.58, 3.47
[1558532761] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.2% (14817596 kB) free.
[1558532807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558532812] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.9% (14933820 kB) free.
[1558532850] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.6% (21479544 kB) free.
[1558532899] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.7% (22961768 kB) free.
[1558532927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558533048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1558533070] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1558533190] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1558533310] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1558533446] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 0.91, 13.75
[1558533619] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 32.3% (7971788 kB) free.
[1558534267] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9197152 kB) free!
[1558534387] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.6% (15056008 kB) free.
[1558534616] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (183704 kB) free!
[1558534646] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.75, 0.93, 4.44
[1558534736] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (181724 kB) free!
[1558534856] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (182248 kB) free!
[1558535069] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1558535448] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1558536048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1558536139] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558536168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558536288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558536486] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 106.30, 173.49, 107.11
[1558536606] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 18.26, 119.41, 95.50
[1558536670] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9417360 kB) free!
[1558536726] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.35, 80.23, 84.04
[1558536790] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 33.1% (16333756 kB) free.
[1558537988] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9266860 kB) free!
[1558538108] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6933564 kB) free!
[1558538228] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.0% (13788516 kB) free.
[1558538256] SERVICE FLAPPING ALERT: smscapp1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558538526] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.21, 0.53, 12.46
[1558538688] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558538769] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1558539288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558539408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558539507] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1558539528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558539726] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.46, 3.74
[1558540390] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8799560 kB) free!
[1558540510] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 30.6% (15084260 kB) free.
[1558541228] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9850108 kB) free!
[1558541348] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7631480 kB) free!
[1558541468] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.2% (13890320 kB) free.
[1558541710] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1558541928] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558542528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558542648] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1558542768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558542910] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9820432 kB) free!
[1558543030] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 30.4% (14976084 kB) free.
[1558543219] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4862896 kB) free!
[1558543339] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4852004 kB) free!
[1558543459] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4837284 kB) free!
[1558543569] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3184432 kB) free!
[1558543689] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (3147272 kB) free!
[1558543809] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (3116736 kB) free!
[1558544456] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1558544468] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9838668 kB) free!
[1558544588] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7792992 kB) free!
[1558544708] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.7% (14120744 kB) free.
[1558545168] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558545768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558545888] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558546008] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558546030] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9632336 kB) free!
[1558546150] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 30.0% (14805212 kB) free.
[1558548308] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9025628 kB) free!
[1558548408] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558548428] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.3% (14422216 kB) free.
[1558549008] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558549128] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558549150] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8691364 kB) free!
[1558549248] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1558549270] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.7% (14165676 kB) free.
[1558549460] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)