[1584297066] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36357724 kB) free.
[1584297515] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196572 kB) free!
[1584297635] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181140 kB) free!
[1584297666] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13716480 kB) free!
[1584297755] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187264 kB) free!
[1584297786] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9921528 kB) free!
[1584297808] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199504 kB) free!
[1584297906] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6268256 kB) free!
[1584298506] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20388076 kB) free!
[1584299108] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36459276 kB) free.
[1584299213] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (205040 kB) free!
[1584299273] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 67.4% (16529720 kB) free.
[1584299708] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15427488 kB) free!
[1584299828] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11746240 kB) free!
[1584299948] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (8142196 kB) free!
[1584300548] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21659124 kB) free!
[1584301148] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36596416 kB) free.
[1584301308] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4828216 kB) free!
[1584301647] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4499296 kB) free!
[1584301748] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15577740 kB) free!
[1584301869] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (11954956 kB) free!
[1584301909] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187820 kB) free!
[1584301989] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8399132 kB) free!
[1584302515] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237680 kB) free!
[1584302589] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20313292 kB) free!
[1584302635] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230136 kB) free!
[1584302755] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239444 kB) free!
[1584303189] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32906368 kB) free.
[1584303587] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.0% (22356916 kB) free!
[1584303707] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.7% (22043472 kB) free!
[1584303789] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13280092 kB) free!
[1584303827] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.0% (22416112 kB) free!
[1584303909] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9802364 kB) free!
[1584304029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36792504 kB) free.
[1584304555] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 53.8% (26521324 kB) free.
[1584304629] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17552624 kB) free!
[1584304650] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(250) < 203.212.70.170 [16/Mar/2020:02:05:09 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1584304749] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (14028856 kB) free!
[1584304770] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584304869] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (10576452 kB) free!
[1584305124] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1584305244] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1584305364] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1584305469] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22413812 kB) free!
[1584306069] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34854396 kB) free.
[1584306217] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243856 kB) free!
[1584306337] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.6% (43153056 kB) free.
[1584306669] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14914052 kB) free!
[1584306789] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11297740 kB) free!
[1584306909] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37007528 kB) free.
[1584307164] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1584307509] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17721552 kB) free!
[1584307630] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14184208 kB) free!
[1584307750] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (10748376 kB) free!
[1584308035] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(252) < 103.16.101.72 [16/Mar/2020:03:00:31 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1584308157] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584308313] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(312) < 2020/03/16 03:01:31 [error] 12257#0: *71297413 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584308350] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20487032 kB) free!
[1584308370] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1725) < 172.105.59.173 [16/Mar/2020:03:01:41 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584308433] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584308491] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584308950] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29292812 kB) free.
[1584309208] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 45.1% (11064080 kB) free.
[1584309353] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 368 processes
[1584309473] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1584309550] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11623988 kB) free!
[1584309670] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37056348 kB) free.
[1584310270] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17423276 kB) free!
[1584310390] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (14034244 kB) free!
[1584310510] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (10745800 kB) free!
[1584311110] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19487296 kB) free!
[1584312310] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2576532 kB) free!
[1584312910] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15773680 kB) free!
[1584313510] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27092416 kB) free.
[1584314711] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17585712 kB) free!
[1584314831] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (14074016 kB) free!
[1584314951] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38245068 kB) free.
[1584315209] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (209104 kB) free!
[1584315329] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191684 kB) free!
[1584315451] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (211388 kB) free!
[1584315551] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19886588 kB) free!
[1584315671] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16501984 kB) free!
[1584315791] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13296492 kB) free!
[1584316309] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 61.3% (15126496 kB) free.
[1584316355] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 78.7% (19285608 kB) free.
[1584316991] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31859824 kB) free.
[1584318191] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21895080 kB) free!
[1584318311] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18503616 kB) free!
[1584318431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38539784 kB) free.
[1584318564] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584318684] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1584318804] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1584319031] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21570476 kB) free!
[1584319151] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18139480 kB) free!
[1584319271] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39446228 kB) free.
[1584319427] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394308 kB) free!
[1584319623] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 92.0% (121341512 kB) free.
[1584319647] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (5368760 kB) free!
[1584319659] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39301988 kB) free.
[1584319694] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.5% (20951648 kB) free.
[1584319871] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20572832 kB) free!
[1584319992] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (16958332 kB) free!
[1584320112] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38125112 kB) free.
[1584320423] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584320544] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584320712] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20086516 kB) free!
[1584320832] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (16624808 kB) free!
[1584320952] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38237560 kB) free.
[1584321109] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.8% (2662336 kB) free!
[1584321229] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.8% (2654784 kB) free!
[1584321349] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2645916 kB) free!
[1584321552] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20542864 kB) free!
[1584321672] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17093844 kB) free!
[1584321792] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38551856 kB) free.
[1584322392] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20812956 kB) free!
[1584322512] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17486772 kB) free!
[1584322632] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38308040 kB) free.
[1584323149] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2291008 kB) free!
[1584323232] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17669068 kB) free!
[1584323352] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14380348 kB) free!
[1584323472] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37365780 kB) free.
[1584324072] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19142756 kB) free!
[1584324192] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38867116 kB) free.
[1584324793] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21901644 kB) free!
[1584324913] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (18599592 kB) free!
[1584325033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38801080 kB) free.
[1584325633] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21785892 kB) free!
[1584325753] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18436696 kB) free!
[1584325873] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38691916 kB) free.
[1584326473] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19693312 kB) free!
[1584326593] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16197024 kB) free!
[1584326713] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38115580 kB) free.
[1584327204] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1584327313] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18201504 kB) free!
[1584327324] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1584327433] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14590904 kB) free!
[1584327444] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1584327553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37401828 kB) free.
[1584328044] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584328153] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18972144 kB) free!
[1584328273] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15366192 kB) free!
[1584328393] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36922692 kB) free.
[1584328993] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18400556 kB) free!
[1584329114] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14835672 kB) free!
[1584329234] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36308244 kB) free.
[1584329245] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1584329365] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1584329485] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1584329834] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15085396 kB) free!
[1584329954] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10972936 kB) free!
[1584330041] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584330041] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584330042] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584330074] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34871100 kB) free.
[1584330162] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 202973 MB (90% inode=99%):
[1584330162] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584330162] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584330674] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12346400 kB) free!
[1584330732] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584330753] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584330772] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584330794] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8204048 kB) free!
[1584330858] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38764 MB (75% inode=98%):
[1584330879] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584330899] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584330914] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34195348 kB) free.
[1584331189] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.6% (20393308 kB) free.
[1584331371] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584331371] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584331483] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584331491] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 201711 MB (89% inode=99%):
[1584331494] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584331514] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10534060 kB) free!
[1584331611] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584331634] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6410060 kB) free!
[1584331754] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584331754] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33780160 kB) free.
[1584332165] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584332291] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.4% (30751480 kB) free.
[1584332354] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11149284 kB) free!
[1584332474] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7090908 kB) free!
[1584332594] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32557056 kB) free.
[1584333194] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10353120 kB) free!
[1584333314] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6150760 kB) free!
[1584333435] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2052880 kB) free!
[1584334635] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15567584 kB) free!
[1584335221] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1584335235] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32906576 kB) free.
[1584335341] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1584335485] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1584335485] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584335835] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5572556 kB) free!
[1584335955] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (994140 kB) free!
[1584336075] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1584336075] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385144 kB) free!
[1584336457] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.8% (53815712 kB) free.
[1584336685] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1584336742] SERVICE ALERT: mrq4;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584336755] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.8% (4372536 kB) free!
[1584336805] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1584336875] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.0% (4176704 kB) free!
[1584336926] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1584336995] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.4% (4015228 kB) free!
[1584337197] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4592784 kB) free!
[1584337317] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (2297972 kB) free!
[1584337437] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2052876 kB) free!
[1584337526] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1584337665] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23217076 kB) free!
[1584337785] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (18574916 kB) free!
[1584337905] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14107220 kB) free!
[1584338126] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1584338195] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2401708 kB) free!
[1584338246] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1584338259] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9601732 kB) free!
[1584338298] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.0% (3198052 kB) free!
[1584338366] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584338379] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9602632 kB) free!
[1584338419] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (1677368 kB) free!
[1584338500] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (8348024 kB) free!
[1584338539] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (240348 kB) free!
[1584338555] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181268 kB) free!
[1584338609] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1584338675] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187512 kB) free!
[1584338729] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1584338795] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190616 kB) free!
[1584338849] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1584338868] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1584338989] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1584339105] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34257580 kB) free.
[1584339449] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1584339705] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8705428 kB) free!
[1584339825] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4371252 kB) free!
[1584339945] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35222916 kB) free.
[1584339995] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 85.9% (21055020 kB) free.
[1584340167] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584340188] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 598 processes
[1584340204] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.8% (19946464 kB) free.
[1584340241] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1584340242] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.7% (20512376 kB) free.
[1584340265] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.0% (18131144 kB) free.
[1584340272] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 669 processes
[1584340283] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 66.1% (32552452 kB) free.
[1584340287] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1584340308] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 614 processes
[1584340344] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3558.20, 3694.89, 1748.68
[1584340361] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1584340392] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 652 processes
[1584340407] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1584340428] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 609 processes
[1584340433] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 324.02, 2031.51, 1267.63
[1584340464] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 482.06, 2473.08, 1536.92
[1584340481] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1584340490] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 111.35, 1596.76, 1149.84
[1584340512] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 587 processes
[1584340545] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10130812 kB) free!
[1584340553] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 46.16, 1360.33, 1114.30
[1584340584] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 67.19, 1655.67, 1350.89
[1584340610] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.96, 1069.53, 1010.82
[1584340612] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.37, 1040.56, 973.79
[1584340665] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5518980 kB) free!
[1584340673] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.43, 911.06, 979.50
[1584340695] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584340717] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.33, 1.17, 4.29
[1584340730] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.22, 716.71, 888.64
[1584340732] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.01, 696.98, 855.98
[1584340742] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 336 processes
[1584340785] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23747220 kB) free!
[1584340822] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.9% (31020308 kB) free.
[1584340833] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(56) < 2020/03/16 12:01:01 [error] 12252#0: *71457592 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584340852] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.95, 467.21, 752.50
[1584340862] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1584340890] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(106) < 172.105.57.57 [16/Mar/2020:12:03:04 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584340953] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584341011] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584341028] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1584341062] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584341081] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1584341184] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.43, 224.42, 709.18
[1584341192] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584341273] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.75, 123.98, 514.32
[1584341322] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38773 MB (75% inode=98%):
[1584341330] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.83, 98.95, 467.14
[1584341385] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375572 kB) free!
[1584341407] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184200 kB) free!
[1584341452] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.30, 64.73, 395.48
[1584341465] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4482728 kB) free!
[1584341527] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165496 kB) free!
[1584341585] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4380980 kB) free!
[1584341610] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(544) < 103.255.146.164 [16/Mar/2020:12:23:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584341647] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (324008 kB) free!
[1584341695] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584341705] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (4266860 kB) free!
[1584341730] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584341784] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.69, 32.47, 372.93
[1584341813] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584341825] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584341922] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584341943] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584341945] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 20.44, 19.41, 19.91
[1584341986] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15180592 kB) free!
[1584342025] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342046] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197064 kB) free!
[1584342049] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38773 MB (75% inode=98%):
[1584342052] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.70, 40.14, 237.24
[1584342062] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1584342071] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584342155] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342166] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182772 kB) free!
[1584342182] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1584342228] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1584342285] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342286] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189004 kB) free!
[1584342300] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342301] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342302] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342348] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1584342428] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1584342428] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584342428] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2390 processes
[1584342468] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1584342555] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342586] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433136 kB) free!
[1584342652] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342677] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 22.52, 21.31, 20.95
[1584342772] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38773 MB (75% inode=98%):
[1584342830] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1584342830] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.3% (20800132 kB) free.
[1584342852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42245348 kB) free.
[1584342873] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1584342873] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.8% (20774708 kB) free.
[1584342890] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.6% (30835196 kB) free.
[1584342918] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342937] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1584342937] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.0% (21824840 kB) free.
[1584342965] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584342967] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584343024] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1584343024] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.5% (30780216 kB) free.
[1584343101] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584343223] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.3% (30686760 kB) free.
[1584343285] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584343409] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 24.43, 24.04, 22.87
[1584343475] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (3754704 kB) free!
[1584343501] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584343509] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584343596] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (875068 kB) free!
[1584343621] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 190540 MB (84% inode=99%):
[1584343630] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2408 processes
[1584343632] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584343716] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187336 kB) free!
[1584343753] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584344036] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179232 kB) free!
[1584344057] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23946196 kB) free!
[1584344137] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7646460 kB) free!
[1584344156] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (741488 kB) free!
[1584344177] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18696656 kB) free!
[1584344257] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 10.4% (5140316 kB) free!
[1584344268] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1584344268] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584344276] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200180 kB) free!
[1584344297] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13792988 kB) free!
[1584344377] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4947780 kB) free!
[1584344977] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 37.0% (18211424 kB) free.
[1584345497] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (747868 kB) free!
[1584346097] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14975724 kB) free!
[1584346407] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 706 processes
[1584346527] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1584346647] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1584346697] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35965764 kB) free.
[1584347146] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4452320 kB) free!
[1584347266] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2355856 kB) free!
[1584347298] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7545824 kB) free!
[1584347386] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194176 kB) free!
[1584347418] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2935580 kB) free!
[1584347538] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390944 kB) free!
[1584347848] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584348086] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9612124 kB) free!
[1584348138] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14160480 kB) free!
[1584348206] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8600240 kB) free!
[1584348326] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (8588556 kB) free!
[1584348738] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34886928 kB) free.
[1584349338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7042188 kB) free!
[1584349458] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2315952 kB) free!
[1584349578] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384180 kB) free!
[1584349596] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.9% (4622288 kB) free!
[1584349716] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.1% (4434312 kB) free!
[1584349836] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.4% (4255144 kB) free!
[1584349905] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (3712308 kB) free!
[1584350025] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3510048 kB) free!
[1584350055] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.8% (21518508 kB) free.
[1584350078] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.4% (49263276 kB) free.
[1584350102] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.4% (21662632 kB) free.
[1584350123] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 61.0% (15050064 kB) free.
[1584350145] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (3315804 kB) free!
[1584350377] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8912952 kB) free!
[1584350497] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8679384 kB) free!
[1584350618] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.2% (8477104 kB) free!
[1584350745] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2246600 kB) free!
[1584351036] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2380852 kB) free!
[1584351182] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351287] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18482040 kB) free!
[1584351312] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351330] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(60) < 92.42.104.166 [16/Mar/2020:15:00:19 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584351345] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.9% (15236792 kB) free.
[1584351407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13633564 kB) free!
[1584351430] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351431] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351431] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351439] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38772 MB (75% inode=98%):
[1584351439] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351450] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584351527] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9049296 kB) free!
[1584351560] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351561] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351561] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351562] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351567] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2398 processes
[1584351631] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351680] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2401 processes
[1584351682] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 169375 MB (75% inode=99%):
[1584351682] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1584351691] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584351761] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351815] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584351880] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1584351887] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.5% (30291728 kB) free.
[1584351926] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3363072 kB) free!
[1584351941] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 25.83, 23.76, 21.30
[1584352000] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1584352048] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 709 processes
[1584352120] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1584352127] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14517448 kB) free!
[1584352169] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 708 processes
[1584352169] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352273] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352289] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 708 processes
[1584352290] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352291] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352291] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352292] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352299] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352403] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352418] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 168510 MB (74% inode=99%):
[1584352421] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352421] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1584352422] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352425] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2415 processes
[1584352491] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352525] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178592 kB) free!
[1584352533] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352545] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352547] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584352548] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2410 processes
[1584352621] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352645] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201440 kB) free!
[1584352651] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(605) < 103.255.146.164 [16/Mar/2020:15:18:49 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584352675] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352720] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1584352727] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32540744 kB) free.
[1584352751] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584352765] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180048 kB) free!
[1584352771] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584352804] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 25.10, 28.05, 25.84
[1584353018] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1584353018] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4460704 kB) free!
[1584353021] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353021] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353029] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353150] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 166121 MB (73% inode=99%):
[1584353151] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353151] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353152] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353158] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2425 processes
[1584353242] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353271] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1584353271] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2417 processes
[1584353273] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584353327] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5016436 kB) free!
[1584353372] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353447] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (422600 kB) free!
[1584353502] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353567] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39231668 kB) free.
[1584353707] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.2% (3965860 kB) free!
[1584353817] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.7% (21144880 kB) free.
[1584353827] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178000 kB) free!
[1584353841] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 24.1% (11886260 kB) free.
[1584353881] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353882] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584353940] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.4% (30256960 kB) free.
[1584353947] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163436 kB) free!
[1584353982] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39287640 kB) free.
[1584354003] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2439 processes
[1584354007] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.4% (21658304 kB) free.
[1584354008] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584354037] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.8% (21771828 kB) free.
[1584354167] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (384720 kB) free!
[1584354246] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32713288 kB) free.
[1584354351] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584354359] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584354468] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584354474] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 162467 MB (72% inode=99%):
[1584354480] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2431 processes
[1584354556] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584354614] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196920 kB) free!
[1584354681] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.8% (29976536 kB) free.
[1584354703] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38772 MB (75% inode=98%):
[1584354735] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176208 kB) free!
[1584354855] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170948 kB) free!
[1584354933] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584355026] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (848444 kB) free!
[1584355146] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (205692 kB) free!
[1584355266] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170600 kB) free!
[1584355457] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (374136 kB) free!
[1584355577] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (383984 kB) free!
[1584355697] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388720 kB) free!
[1584356055] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184160 kB) free!
[1584356176] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171380 kB) free!
[1584356296] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180768 kB) free!
[1584356454] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168628 kB) free!
[1584356574] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 80.8% (19806012 kB) free.
[1584356897] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34742384 kB) free.
[1584357497] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (398260 kB) free!
[1584357617] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394916 kB) free!
[1584357737] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (452076 kB) free!
[1584358337] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18828052 kB) free!
[1584358889] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1584358937] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384212 kB) free!
[1584358982] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.0% (21086904 kB) free.
[1584359005] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.4% (20922448 kB) free.
[1584359028] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.5% (20613644 kB) free.
[1584359055] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16876744 kB) free!
[1584359058] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28788376 kB) free.
[1584359386] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9507688 kB) free!
[1584359506] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9506744 kB) free!
[1584359608] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179660 kB) free!
[1584359626] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (9508008 kB) free!
[1584359728] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192436 kB) free!
[1584359848] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164036 kB) free!
[1584360266] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19392076 kB) free!
[1584360386] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32148568 kB) free.
[1584360945] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.0% (3815072 kB) free!
[1584360987] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1280556 kB) free!
[1584361065] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3569204 kB) free!
[1584361107] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395908 kB) free!
[1584361185] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (3312648 kB) free!
[1584361227] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (443664 kB) free!
[1584361827] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (14029796 kB) free!
[1584361886] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584361913] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362016] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362033] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38772 MB (75% inode=98%):
[1584362037] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1842636 kB) free!
[1584362143] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362146] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362157] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (1666604 kB) free!
[1584362273] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362277] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1493792 kB) free!
[1584362281] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362281] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362289] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362297] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9310664 kB) free!
[1584362385] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.9% (15243668 kB) free.
[1584362395] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584362401] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1584362401] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 136378 MB (60% inode=99%):
[1584362409] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2443 processes
[1584362412] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362417] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9147176 kB) free!
[1584362427] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398680 kB) free!
[1584362537] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (9029940 kB) free!
[1584362542] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362626] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3170936 kB) free!
[1584362672] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584362989] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (2254984 kB) free!
[1584363109] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2896732 kB) free!
[1584363229] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3590216 kB) free.
[1584363354] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.6% (29876908 kB) free.
[1584363445] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584363540] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (247780 kB) free!
[1584363603] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363605] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363611] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363611] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363611] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363619] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363627] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35943316 kB) free.
[1584363660] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (153364 kB) free!
[1584363733] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363735] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363741] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363741] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363741] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363749] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363780] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (376916 kB) free!
[1584363829] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2105584 kB) free!
[1584363843] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363864] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363865] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363871] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363871] SERVICE ALERT: smscapp2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363871] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363879] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584363949] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2606036 kB) free!
[1584363963] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38771 MB (75% inode=98%):
[1584364069] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (3254968 kB) free!
[1584364171] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(122) < 103.255.146.164 [16/Mar/2020:18:31:48 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584364227] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6640476 kB) free!
[1584364291] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584364347] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1794644 kB) free!
[1584364467] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382400 kB) free!
[1584364478] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 133746 MB (59% inode=99%):
[1584364479] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1584364482] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2462 processes
[1584364669] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4093328 kB) free.
[1584364890] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 708 processes
[1584365010] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 708 processes
[1584365067] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16984024 kB) free!
[1584365130] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 707 processes
[1584365491] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(197) < 103.255.146.164 [16/Mar/2020:18:55:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584365611] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584365664] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584365665] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 19.56, 26.02, 23.88
[1584365667] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406576 kB) free!
[1584365671] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2451 processes
[1584365674] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584366174] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.8% (4370456 kB) free!
[1584366267] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13433256 kB) free!
[1584366294] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.8% (4123204 kB) free!
[1584366414] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.2% (3962248 kB) free!
[1584366737] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4615660 kB) free!
[1584366867] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381680 kB) free!
[1584367386] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584367614] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2439288 kB) free!
[1584368067] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35982180 kB) free.
[1584368269] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2844428 kB) free!
[1584368389] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3271404 kB) free.
[1584368667] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8130240 kB) free!
[1584368731] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584368787] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2777668 kB) free!
[1584368907] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (428044 kB) free!
[1584369478] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1584369507] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19450144 kB) free!
[1584369571] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21783216 kB) free.
[1584369614] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.0% (20491192 kB) free.
[1584369645] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 60.9% (30009056 kB) free.
[1584369664] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 32.3% (15941656 kB) free.
[1584369687] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 69.7% (17082692 kB) free.
[1584369717] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.5% (14066652 kB) free.
[1584370107] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382568 kB) free!
[1584370218] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187212 kB) free!
[1584370338] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170648 kB) free!
[1584370458] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190356 kB) free!
[1584370707] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13838468 kB) free!
[1584371131] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584371251] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584371307] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400228 kB) free!
[1584372507] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36347148 kB) free.
[1584373107] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7323972 kB) free!
[1584373227] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3085268 kB) free!
[1584373246] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9825172 kB) free!
[1584373347] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409132 kB) free!
[1584373366] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9370452 kB) free!
[1584373474] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584373486] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (9713224 kB) free!
[1584373601] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584373897] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584373948] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20503908 kB) free!
[1584374154] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584374548] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397332 kB) free!
[1584374675] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584374797] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 12.78, 9.78, 8.89
[1584374851] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584374868] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1584374971] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584375281] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584375284] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584375401] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2455 processes
[1584375404] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584375748] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (36983576 kB) free.
[1584375901] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.1% (20501564 kB) free.
[1584376348] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11426844 kB) free!
[1584376468] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6912812 kB) free!
[1584376588] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2072984 kB) free!
[1584377188] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21065644 kB) free!
[1584377371] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584377491] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1584377611] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1584377788] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385720 kB) free!
[1584378211] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584378388] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13894020 kB) free!
[1584378811] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584378931] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1584379588] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390924 kB) free!
[1584381331] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584381451] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584381886] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4308644 kB) free!
[1584381989] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31408324 kB) free.
[1584382589] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7725488 kB) free!
[1584382709] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3804672 kB) free!
[1584382829] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383388 kB) free!
[1584383251] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584383268] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1584383371] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes