[1567018533] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.68, 6.15, 5.11
[1567018653] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.62, 5.88, 5.14
[1567018774] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.88, 5.47, 5.08
[1567020214] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567020334] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567020701] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 471 processes
[1567020761] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 465 processes
[1567020941] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 293 processes
[1567021415] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (1481104 kB) free!
[1567021533] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.8% (15547940 kB) free!
[1567021535] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (1456080 kB) free!
[1567021541] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 288 processes
[1567021653] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.8% (15546132 kB) free!
[1567021655] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (1433260 kB) free!
[1567021773] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.2% (14760132 kB) free!
[1567022374] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.98, 2.75, 4.45
[1567023934] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567024055] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567025744] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226512 kB) free!
[1567026457] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 101579 MB (45% inode=99%):
[1567026944] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.5% (44126760 kB) free.
[1567028673] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 47.2% (11578640 kB) free.
[1567028782] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 47.7% (11684720 kB) free.
[1567039054] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 409 processes
[1567039055] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 96.5% (33794596 kB) free.
[1567039173] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.5% (32326000 kB) free.
[1567039174] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 408 processes
[1567040374] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1567040494] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1567040561] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 450 processes
[1567040621] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 441 processes
[1567040642] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.9% (11337584 kB) free.
[1567042473] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (2055160 kB) free!
[1567042582] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188308 kB) free!
[1567042594] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183044 kB) free!
[1567042703] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173676 kB) free!
[1567042714] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191700 kB) free!
[1567042824] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183628 kB) free!
[1567043042] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (574500 kB) free!
[1567043162] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (566364 kB) free!
[1567043282] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (538240 kB) free!
[1567056821] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1567056942] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1567057054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 331 processes
[1567057062] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1567057174] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.10, 5.33, 5.04
[1567057174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 331 processes
[1567057294] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.25, 4.77, 4.88
[1567057294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 336 processes
[1567057415] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237780 kB) free!
[1567057535] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227076 kB) free!
[1567057655] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235208 kB) free!
[1567057894] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1567057894] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567058014] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567060062] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1567060663] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 81.81, 30.32, 12.97
[1567060783] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 98.99, 53.98, 23.75
[1567060896] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.12, 8.61, 4.14
[1567060903] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 101.34, 69.77, 33.20
[1567061017] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 77.42, 33.77, 13.90
[1567061139] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 84.03, 51.24, 22.83
[1567062939] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.38, 1.86, 11.35
[1567063062] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1567063182] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1567063302] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1567063303] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.62, 2.90, 13.49
[1567064139] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.95, 1.96, 4.40
[1567065104] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.10, 2.31, 4.07
[1567066423] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 406 processes
[1567066423] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 577 processes
[1567066483] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 552 processes
[1567066543] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567066544] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9426576 kB) free!
[1567066664] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9185576 kB) free!
[1567066723] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.5% (4045480 kB) free!
[1567066783] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.5% (4049292 kB) free!
[1567066785] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8944264 kB) free!
[1567066843] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.5% (4043848 kB) free!
[1567066903] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.5% (4042912 kB) free!
[1567068943] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1567069064] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567069185] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4556380 kB) free!
[1567069902] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1567070012] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 94.0% (23037412 kB) free.
[1567070048] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.2% (19056872 kB) free.
[1567070107] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.2% (20889208 kB) free.
[1567070210] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 48.2% (23735572 kB) free.
[1567070245] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 40.4% (19888648 kB) free.
[1567070300] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21932084 kB) free.
[1567070336] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 60.8% (29962324 kB) free.
[1567070494] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.93, 10.78, 7.52
[1567070614] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.56, 10.66, 7.88
[1567070657] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167952 kB) free!
[1567070707] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188136 kB) free!
[1567070734] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.81, 11.38, 8.47
[1567070777] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184032 kB) free!
[1567070834] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197192 kB) free!
[1567070897] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188000 kB) free!
[1567070904] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176956 kB) free!
[1567070954] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181416 kB) free!
[1567071024] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170648 kB) free!
[1567071145] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181932 kB) free!
[1567071458] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220420 kB) free!
[1567071578] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (223880 kB) free!
[1567071698] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1567071698] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232964 kB) free!
[1567074260] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.5% (17879544 kB) free.
[1567074301] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22433660 kB) free.
[1567074333] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 68.2% (33606996 kB) free.
[1567074370] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.0% (22802700 kB) free.
[1567074703] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1567074823] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1567074823] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1567074864] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190352 kB) free!
[1567074907] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172712 kB) free!
[1567074943] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1567074943] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1567074975] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189676 kB) free!
[1567074984] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176876 kB) free!
[1567075028] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177608 kB) free!
[1567075063] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1567075095] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183248 kB) free!
[1567075104] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1567075104] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173228 kB) free!
[1567075148] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1567075148] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165948 kB) free!
[1567075215] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1567075215] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191676 kB) free!
[1567075423] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 304 processes
[1567075483] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 551 processes
[1567075543] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 557 processes
[1567075544] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1567075614] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9795536 kB) free!
[1567075663] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1567075734] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9600948 kB) free!
[1567075854] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (9389064 kB) free!
[1567076263] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1567076383] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1567076503] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1567076743] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1567076863] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567076984] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567077054] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233164 kB) free!
[1567077703] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 17 users currently logged in
[1567077823] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 17 users currently logged in
[1567077943] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 18 users currently logged in
[1567079257] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46916 MB (20% inode=99%):
[1567079377] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46527 MB (20% inode=99%):
[1567079497] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45901 MB (20% inode=99%):
[1567080334] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.62, 3.10, 4.24
[1567082734] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.38, 5.50, 5.08
[1567082737] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1567082854] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.28, 5.44, 5.11
[1567082974] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.77, 6.89, 5.67
[1567083575] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 316 processes
[1567083695] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1567083943] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1567083943] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 13 users currently logged in
[1567084175] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.03, 4.13, 4.97
[1567085348] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1567085904] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567086015] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567087611] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.8% (18452688 kB) free.
[1567087643] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.2% (23089620 kB) free.
[1567087674] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 70.4% (34718296 kB) free.
[1567087711] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.8% (23246568 kB) free.
[1567089937] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9469876 kB) free!
[1567089943] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1567090015] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.1% (3723624 kB) free!
[1567090057] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9300164 kB) free!
[1567090135] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2366072 kB) free!
[1567090177] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.5% (9102640 kB) free!
[1567090255] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (1035144 kB) free!
[1567091284] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.4% (5608912 kB) free!
[1567091404] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4911048 kB) free!
[1567091524] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4045636 kB) free!
[1567092097] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24567 MB (10% inode=99%):
[1567092583] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1567092697] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25913 MB (11% inode=99%):
[1567092703] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567093697] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.7% (21390068 kB) free.
[1567093733] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 69.1% (34065300 kB) free.
[1567093895] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1567093897] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24640 MB (10% inode=99%):
[1567094305] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165148 kB) free!
[1567094425] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180944 kB) free!
[1567094545] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1567094545] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163612 kB) free!
[1567095577] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4727500 kB) free!
[1567097933] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8065052 kB) free!
[1567098053] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.2% (7503288 kB) free!
[1567098173] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1567098173] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.3% (6546976 kB) free!
[1567098773] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3195012 kB) free!