[1595356971] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(10) < Channel error on connection <0.22414.6> (10.147.212.38:57606 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1595357092] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1595357153] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 1.59, 13.36, 26.17
[1595357691] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(4) < Channel error on connection <0.4424.7> (10.147.212.37:35216 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1595357753] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 49.15, 37.90, 31.76
[1595357812] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(9) < Channel error on connection <0.8296.7> (10.147.212.49:40222 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1595357873] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 19.45, 29.47, 29.37
[1595357933] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(1) < Channel error on connection <0.9368.7> (10.147.212.49:40886 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1595358473] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 36.42, 32.88, 30.07
[1595358593] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 23.06, 28.47, 28.76
[1595359733] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1595359793] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 33.77, 38.01, 34.99
[1595359915] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 19.52, 32.58, 33.46
[1595359969] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (914532 kB) free!
[1595360029] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184488 kB) free!
[1595360035] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 15.27, 28.97, 32.20
[1595360089] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (190724 kB) free!
[1595360149] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (198212 kB) free!
[1595360450] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.8% (17107232 kB) free.
[1595360635] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 56.07, 43.82, 37.20
[1595361235] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 37.17, 32.60, 33.59
[1595361270] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (391128 kB) free!
[1595361390] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391120 kB) free!
[1595361458] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236624 kB) free!
[1595361510] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (462220 kB) free!
[1595361578] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215112 kB) free!
[1595361698] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234024 kB) free!
[1595361835] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 23.16, 25.94, 29.15
[1595361967] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 211565 MB (19% inode=99%):
[1595362087] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 211565 MB (19% inode=99%):
[1595362207] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 211565 MB (19% inode=99%):
[1595363467] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 27.5% (6731084 kB) free.
[1595363498] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.3% (8055400 kB) free!
[1595364098] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 24.6% (12144440 kB) free.
[1595364291] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15918.07, 10428.93, 4605.63
[1595364411] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15990.54, 12272.80, 5987.39
[1595364531] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16000.16, 13506.50, 7201.58
[1595364667] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (298420 kB) free!
[1595364789] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (227856 kB) free!
[1595364909] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (320424 kB) free!
[1595365280] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 54.9% (27033812 kB) free.
[1595365727] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 65.7% (32393000 kB) free.
[1595367067] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(566) < 203.212.70.170 [22/Jul/2020:03:00:24 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1595367188] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(162) < 203.212.70.170 [22/Jul/2020:03:01:35 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1595367308] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595367531] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.82, 14.20, 575.45
[1595367533] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2870924 kB) free!
[1595368131] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.97, 3.56, 302.45
[1595369109] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.2% (10332596 kB) free.
[1595369933] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2084692 kB) free!
[1595370835] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 19.69, 30.02, 30.35
[1595370909] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (2429080 kB) free!
[1595370955] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 37.23, 33.24, 31.50
[1595371029] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (288596 kB) free!
[1595371075] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 28.23, 34.12, 32.23
[1595371149] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1595371149] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (306640 kB) free!
[1595371749] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.2% (10106624 kB) free.
[1595372275] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 19.86, 24.36, 29.31
[1595373475] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 33.80, 32.47, 30.14
[1595373595] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 39.37, 34.80, 31.24
[1595373715] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 36.33, 31.94, 30.44
[1595374688] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 73.3% (17960620 kB) free.
[1595374733] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.4% (12927432 kB) free.
[1595374749] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (204128 kB) free!
[1595374869] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169604 kB) free!
[1595374989] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170100 kB) free!
[1595375515] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 18.61, 23.59, 27.76
[1595376115] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 43.29, 26.90, 26.90
[1595376235] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 36.75, 30.14, 28.06
[1595376835] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 42.80, 31.94, 29.54
[1595376955] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 26.06, 31.54, 29.84
[1595377555] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 34.22, 30.83, 31.08
[1595377675] SERVICE ALERT: mrq4;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 56.74, 38.51, 33.73
[1595377795] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 32.90, 36.28, 33.52
[1595378629] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 58.1% (14251720 kB) free.
[1595379007] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 228287 MB (21% inode=99%):
[1595380710] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 36.0% (47473576 kB) free.
[1595380795] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 72.22, 49.25, 44.66
[1595382175] SERVICE ALERT: pservercluster2;Sensors Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595382295] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595382613] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382733] SERVICE ALERT: sservercluster1;HDD_Root status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595382734] SERVICE ALERT: sservercluster1;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595382739] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382789] SERVICE ALERT: sservercluster1;Memory;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595382826] SERVICE ALERT: sservercluster1;CPU-STATISTICS;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595382846] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382863] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382864] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382869] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382956] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382957] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382976] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382994] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595382999] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383000] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383048] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383086] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383087] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383106] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383109] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383130] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383133] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (225228 kB) free!
[1595383178] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383217] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383239] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383253] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (204712 kB) free!
[1595383260] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383308] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383369] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595383373] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (306308 kB) free!
[1595383399] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595384286] SERVICE ALERT: sservercluster1;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595384896] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595389795] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1595390127] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1595390247] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1595394268] SERVICE ALERT: sservercluster1;HDD_Root status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595394665] SERVICE ALERT: sservercluster1;CPU Load;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595394708] SERVICE ALERT: sservercluster1;CPU_Procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595394809] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595394878] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595395275] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595395318] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595395374] SERVICE ALERT: sservercluster1;Sensors Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595395984] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595395999] SERVICE ALERT: sservercluster1;Memory;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595396186] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.0% (3364728 kB) free!
[1595396306] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3054828 kB) free!
[1595396426] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.2% (2744156 kB) free!
[1595396584] SERVICE ALERT: sservercluster1;Sensors Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595396609] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595397026] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1445040 kB) free!
[1595397194] SERVICE FLAPPING ALERT: sservercluster1;Sensors Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1595397194] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595397229] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4710792 kB) free!
[1595397349] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4461772 kB) free!
[1595397469] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (4212832 kB) free!
[1595397622] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1595397680] SERVICE ALERT: sservercluster1;CPU Load;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595397723] SERVICE ALERT: sservercluster1;CPU_Procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595398101] SERVICE ALERT: sservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=5.13% system=11.40% iowait=0.03% idle=83.44%
[1595398106] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1595398222] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595398280] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1595398280] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 282.18, 2795.06, 2003.96
[1595398323] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1595398323] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1595398395] SERVICE ALERT: sservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1595398478] SERVICE ALERT: sservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39190 MB (76% inode=99%):
[1595398594] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 109330 MB (49% inode=99%):
[1595398599] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 362 processes
[1595399247] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1595399269] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1863608 kB) free!
[1595399367] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1595399426] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 51.1% (11464168 kB) free.
[1595399480] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.78, 52.95, 552.13
[1595399888] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4633700 kB) free!
[1595400008] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4500824 kB) free!
[1595400080] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.53, 9.27, 290.49
[1595400128] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4383224 kB) free!
[1595401194] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1595401314] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1595401727] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9784980 kB) free!
[1595401847] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9626824 kB) free!
[1595401967] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (9449520 kB) free!
[1595402528] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1890728 kB) free!
[1595405567] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4815264 kB) free!
[1595406217] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15956.54, 13707.79, 7462.06
[1595406311] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.2% (6278552 kB) free!
[1595406337] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15957.53, 14452.31, 8492.19
[1595406431] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.0% (6213896 kB) free!
[1595406465] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16229.09, 15045.06, 9487.85
[1595406468] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595406551] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.8% (6156964 kB) free!
[1595406593] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=7.88% system=16.42% iowait=0.00% idle=75.70%
[1595407151] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 20.5% (6701108 kB) free.
[1595407226] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4226184 kB) free!
[1595407346] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.6% (3959252 kB) free!
[1595407466] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.5% (3703308 kB) free!
[1595407752] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 18.9% (6199488 kB) free!
[1595407872] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 18.7% (6124528 kB) free!
[1595407880] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9528612 kB) free!
[1595407992] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.5% (6060928 kB) free!
[1595407995] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595408001] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9451460 kB) free!
[1595408088] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595408121] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9353588 kB) free!
[1595408125] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595408194] SERVICE ALERT: sservercluster1;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595408199] SERVICE ALERT: sservercluster1;Total Processes;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595408218] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595408245] SERVICE ALERT: sservercluster1;Sensors Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595408311] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595408316] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595408321] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595408329] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595408348] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595408436] SERVICE ALERT: sservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=5.54% system=10.76% iowait=0.00% idle=83.70%
[1595408436] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1595408441] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 99377 MB (45% inode=99%):
[1595408449] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 362 processes
[1595408666] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1595408666] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1502612 kB) free!
[1595408708] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(49) < 203.212.70.170 [22/Jul/2020:14:31:23 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 0.000, 0.001"
[1595408828] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595408845] SERVICE ALERT: sservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1595408865] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.90, 9.51, 719.42
[1595408948] SERVICE ALERT: sservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39190 MB (76% inode=99%):
[1595409465] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.05, 4.28, 378.68
[1595409723] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595410467] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 51.5% (11573064 kB) free.
[1595410628] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(520) < 172.105.62.98 [22/Jul/2020:15:01:31 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1595410748] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595411480] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595413767] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1595413836] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595413887] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1595413956] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1595414250] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414353] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414370] SERVICE ALERT: sservercluster1;Sensors Status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595414427] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414431] SERVICE ALERT: sservercluster1;CPU-STATISTICS;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595414441] SERVICE ALERT: sservercluster1;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595414454] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414478] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414480] SERVICE ALERT: sservercluster1;CPU Load;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595414495] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414528] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414551] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414552] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414561] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414566] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414574] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1595414600] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414603] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414648] SERVICE ALERT: sservercluster1;CPU_Procs;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595414673] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414677] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414686] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414691] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414699] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414721] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4577360 kB) free!
[1595414725] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414773] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595414806] SERVICE ALERT: sservercluster1;Current Users;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595415087] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1595415208] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1595415411] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595415793] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (3026460 kB) free!
[1595415873] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595415877] SERVICE ALERT: sservercluster1;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595416477] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595416478] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595417007] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 326 processes
[1595417127] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1595417667] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4456716 kB) free!
[1595417787] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4217372 kB) free!
[1595417907] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.9% (4018732 kB) free!
[1595418277] SERVICE FLAPPING ALERT: sservercluster1;CPU-STATISTICS;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1595418277] SERVICE ALERT: sservercluster1;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595418278] SERVICE FLAPPING ALERT: sservercluster1;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1595418278] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595418878] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595418882] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595419107] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2223520 kB) free!
[1595419759] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595419798] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595419889] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595419928] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420019] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420029] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420058] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420096] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420105] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420156] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420159] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420226] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420235] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420248] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420275] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420286] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420289] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420298] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420356] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420365] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420378] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420405] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420416] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420428] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420436] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420508] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420535] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420558] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420566] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595420696] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595421100] SERVICE FLAPPING ALERT: sservercluster1;Sensors Status;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1595421140] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595421507] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 51.8% (11622840 kB) free.
[1595421750] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595421973] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1595422011] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1595422089] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 395 processes
[1595422156] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 395 processes
[1595422165] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595422216] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 36550 MB (71% inode=98%):
[1595422308] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1595422350] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1595422350] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.48, 3.12, 3.23
[1595422358] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 160420 MB (71% inode=99%):
[1595422419] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595422463] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=7.62% system=14.63% iowait=0.00% idle=77.74%
[1595422496] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 13536 open files (0% of max 2424156)
[1595422578] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595422616] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595422694] SERVICE ALERT: pservercluster2;CPU_Procs;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595422703] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595422741] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595422756] SERVICE ALERT: pservercluster2;Total Processes;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595422770] SERVICE ALERT: pservercluster2;Sensors Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595422824] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595422826] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595422828] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595422866] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1595422866] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595422886] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595422900] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595422918] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595422954] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595422956] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595422960] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595422968] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423016] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423029] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423030] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423048] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423073] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423080] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595423086] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423098] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423106] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423159] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423178] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423198] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1595423210] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595423219] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 159900 MB (71% inode=99%):
[1595423226] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1595423279] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595423323] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=7.43% system=14.18% iowait=0.00% idle=78.39%
[1595423346] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 13152 open files (0% of max 2424156)
[1595423554] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1595423554] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1595423616] SERVICE FLAPPING ALERT: pservercluster2;Total Processes;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1595423616] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 395 processes
[1595423630] SERVICE FLAPPING ALERT: pservercluster2;Sensors Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1595423630] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595423686] SERVICE FLAPPING ALERT: pservercluster2;HDD_Root status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1595423686] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 36550 MB (71% inode=98%):
[1595423778] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1595423778] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1595423810] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.98, 3.76, 3.48
[1595424028] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1595424028] SERVICE ALERT: sservercluster1;CPU_Procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595424633] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595425403] SERVICE ALERT: sservercluster1;HDD_Root status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595425833] SERVICE ALERT: sservercluster1;CPU_Procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595426008] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595426439] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595426608] SERVICE ALERT: sservercluster1;HDD_Root status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595426678] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595427213] SERVICE FLAPPING ALERT: sservercluster1;HDD_Root status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1595427213] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595427283] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595427813] SERVICE ALERT: sservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39189 MB (76% inode=99%):
[1595427887] SERVICE ALERT: sservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=4.73% system=10.29% iowait=0.01% idle=84.97%
[1595427891] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 88644 MB (40% inode=99%):
[1595427899] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 360 processes
[1595427925] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.03, 2.28, 2.30
[1595428239] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1595428266] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1595428295] SERVICE ALERT: sservercluster1;Sensors Status;UNKNOWN;HARD;3;NRPE: Unable to read output
[1595428483] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595428491] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595428611] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 88396 MB (40% inode=99%):
[1595428895] SERVICE ALERT: sservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1595430292] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595430325] SERVICE ALERT: sservercluster1;CPU Load;UNKNOWN;SOFT;1;Error opening /usr/bin/uptime
[1595430417] SERVICE ALERT: sservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=4.54% system=9.32% iowait=0.00% idle=86.14%
[1595430445] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.04, 2.28, 2.20
[1595431107] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1595431698] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 373 processes
[1595431818] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1595431993] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 47.7% (15609048 kB) free.
[1595432869] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.3% (20657668 kB) free.
[1595433595] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 30.32, 39.19, 44.49
[1595434017] SERVICE FLAPPING ALERT: sservercluster1;CPU-STATISTICS;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1595434045] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 13635.30, 10951.11, 5551.30
[1595434165] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 13638.52, 11840.48, 6531.93
[1595434285] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3997.83, 9461.72, 6356.23
[1595434409] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2743612 kB) free!
[1595434499] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 654 processes
[1595434619] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 372 processes
[1595434795] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 54.27, 53.05, 48.18
[1595434866] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1595434928] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 55.7% (13653508 kB) free.
[1595434954] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595435016] SERVICE FLAPPING ALERT: pservercluster2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595435030] SERVICE FLAPPING ALERT: pservercluster2;Sensors Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595435086] SERVICE FLAPPING ALERT: pservercluster2;HDD_Root status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595435178] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595435210] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595435318] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1462 processes
[1595435528] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193856 kB) free!
[1595435609] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2413792 kB) free!
[1595435648] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182092 kB) free!
[1595435768] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185620 kB) free!
[1595436595] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 29.17, 41.16, 44.43
[1595436685] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.69, 5.72, 481.44
[1595436919] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595436995] SERVICE ALERT: sservercluster2;Sensors Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595437049] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437097] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437115] SERVICE ALERT: sservercluster2;Sensors Status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595437141] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437179] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437182] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437227] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437228] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437240] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437245] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437271] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437278] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437302] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595437319] SERVICE ALERT: sservercluster2;Current Users;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595437357] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437358] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437370] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437401] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437406] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 31393 MB (61% inode=98%):
[1595437432] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437439] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1595437478] SERVICE ALERT: sservercluster2;CPU_Procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595437500] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437569] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595437795] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1595437795] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.48, 50.07, 46.06
[1595438008] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595438088] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595438138] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595438168] SERVICE ALERT: sservercluster2;Memory;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595438268] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595438395] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 22.53, 36.96, 41.78
[1595438439] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1595438778] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1595438778] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595439213] SERVICE FLAPPING ALERT: sservercluster1;HDD_Root status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595439284] SERVICE FLAPPING ALERT: sservercluster1;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1595440195] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.41, 48.03, 44.12
[1595440401] SERVICE ALERT: sservercluster2;CPU Load;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595441011] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595441273] SERVICE ALERT: sservercluster2;HDD_Root status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595441395] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 36.80, 38.66, 41.36
[1595441609] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.3% (21161560 kB) free.
[1595441883] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.