[1557340314] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557340435] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557340555] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557342224] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 250.92, 106.08, 40.08
[1557342344] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 272.14, 162.11, 68.61
[1557342367] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1557342464] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 274.81, 199.54, 93.66
[1557342914] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1557344264] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 0.90, 14.63
[1557344756] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1557345464] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.32, 4.23
[1557345785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8777456 kB) free!
[1557345905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.7% (13663608 kB) free.
[1557345956] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557345967] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 40.82, 17.00, 7.19
[1557346076] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557346087] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.19, 12.06, 6.57
[1557346196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557346207] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.29, 8.56, 5.96
[1557346807] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.37, 2.77, 3.97
[1557347396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557347996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557348117] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557348716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557348836] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557349436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557349514] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 95.20, 22.77, 7.86
[1557349556] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557349634] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 253.80, 107.37, 40.67
[1557349754] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 275.10, 163.96, 69.50
[1557350105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9742912 kB) free!
[1557350156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557350226] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.7% (14622612 kB) free.
[1557350276] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557350825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9647096 kB) free!
[1557350945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.7% (14620784 kB) free.
[1557351477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557351545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9218432 kB) free!
[1557351554] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.16, 1.09, 16.32
[1557351597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557351665] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.6% (14610140 kB) free.
[1557352197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557352265] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9059952 kB) free!
[1557352317] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557352385] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.7% (14613996 kB) free.
[1557352754] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 0.43, 4.73
[1557352917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557353037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557353637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557353757] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1557354358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1557354478] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557355078] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557355198] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557355798] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557355918] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557356518] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557356638] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557357239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557357359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1557357959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557358079] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557358679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557358799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557358919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557361664] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 252.76, 135.30, 54.93
[1557361784] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 258.49, 176.34, 79.72
[1557361904] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 171.10, 183.08, 94.64
[1557361985] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9778272 kB) free!
[1557362105] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.4% (14504340 kB) free.
[1557363704] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 0.80, 13.85
[1557364904] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.28, 4.02
[1557364919] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1557367123] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1557367243] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1557367807] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 316.32, 245.48, 121.48
[1557367927] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 316.47, 269.00, 145.14
[1557368047] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 61.08, 195.19, 133.78
[1557368705] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9719864 kB) free!
[1557368825] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.1% (14343796 kB) free.
[1557369719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557370320] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557370440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557370447] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.91, 1.82, 11.72
[1557370560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557370625] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8896404 kB) free!
[1557370745] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.0% (13792864 kB) free.
[1557371160] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557371647] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.47, 1.70, 4.46
[1557371760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557371880] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557372000] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1557372000] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557372545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8489028 kB) free!
[1557372600] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557372665] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13393788 kB) free.
[1557372863] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 11.1% (2718764 kB) free!
[1557372923] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3495784 kB) free!
[1557372983] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 14.3% (3493304 kB) free!
[1557373043] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 14.2% (3491228 kB) free!
[1557373201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557373321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557373441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557374041] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557374465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8300544 kB) free!
[1557374585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.3% (13460284 kB) free.
[1557374641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557374647] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.86, 5.35, 5.92
[1557374761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557374767] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.74, 4.17, 5.41
[1557374881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557374887] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.89, 3.45, 5.00
[1557375481] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557376082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557376202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1557376322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557376922] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557377522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557377642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557377721] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196252 kB) free!
[1557377762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557377764] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176592 kB) free!
[1557377841] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163616 kB) free!
[1557377884] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (205692 kB) free!
[1557377962] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197276 kB) free!
[1557378004] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189780 kB) free!
[1557378362] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1557378785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7693632 kB) free!
[1557378905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13225484 kB) free.
[1557378962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557379083] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557379203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557379643] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 53.5% (13113048 kB) free.
[1557379803] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1557380146] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.0% (2594516 kB) free!
[1557380164] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4716136 kB) free!
[1557380266] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.8% (2400568 kB) free!
[1557380284] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4594896 kB) free!
[1557380386] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.5% (2194364 kB) free!
[1557380403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557380404] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4472780 kB) free!
[1557380523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557380600] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.81, 6.22, 5.34
[1557380643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557380705] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7793196 kB) free!
[1557380720] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.83, 6.92, 5.71
[1557380825] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13245312 kB) free.
[1557380840] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.34, 6.85, 5.82
[1557380887] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.26, 11.80, 9.14
[1557380987] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1062828 kB) free!
[1557381007] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.84, 8.50, 8.25
[1557381128] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.75, 6.23, 7.45
[1557381243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557381623] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.7% (14424604 kB) free.
[1557381728] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1557381728] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.21, 2.85, 5.00
[1557381801] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.9% (22689664 kB) free.
[1557381843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557381860] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.6% (22439992 kB) free.
[1557381964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557382084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557382144] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1557382625] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8688932 kB) free!
[1557382745] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.1% (13847244 kB) free.
[1557383404] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2426948 kB) free!
[1557384358] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.6% (22453240 kB) free.
[1557384544] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1557384545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8665016 kB) free!
[1557384665] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.7% (13670644 kB) free.
[1557386465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9487992 kB) free!
[1557386489] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 349 processes
[1557386585] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7504364 kB) free!
[1557386609] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 349 processes
[1557386706] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13103868 kB) free.
[1557386729] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1557386763] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186840 kB) free!
[1557386832] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (183712 kB) free!
[1557386883] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 35.6% (8719524 kB) free.
[1557386952] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (235412 kB) free!
[1557387072] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (240456 kB) free!
[1557387329] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557387905] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9493900 kB) free!
[1557387956] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 84.4% (13726412 kB) free.
[1557388025] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7656020 kB) free!
[1557388145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13345132 kB) free.
[1557388231] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.1% (2935028 kB) free!
[1557388351] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (2715780 kB) free!
[1557388471] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.4% (2509400 kB) free!
[1557388640] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.43, 4.64, 4.96
[1557389071] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1557389071] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1462680 kB) free!
[1557389883] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4658244 kB) free!
[1557389945] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8583196 kB) free!
[1557390003] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4549072 kB) free!
[1557390065] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6282540 kB) free!
[1557390123] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4436240 kB) free!
[1557390185] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12060632 kB) free.
[1557390440] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.93, 6.02, 5.13
[1557390554] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 127.30, 51.98, 19.65
[1557390560] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.82, 6.21, 5.30
[1557390674] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 140.29, 81.90, 34.54
[1557390680] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.55, 5.82, 5.26
[1557390794] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 142.58, 102.06, 47.66
[1557391280] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1557391280] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.24, 4.50, 4.88
[1557391385] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7917248 kB) free!
[1557391505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11613900 kB) free.
[1557391904] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 284.51, 194.09, 88.26
[1557391994] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 2.48, 14.28
[1557392024] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 134.74, 184.44, 98.68
[1557392144] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 18.62, 123.57, 86.77
[1557392221] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1557392221] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.3% (22142748 kB) free.
[1557392271] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.3% (14508932 kB) free.
[1557392598] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46294 MB (20% inode=99%):
[1557392705] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8390520 kB) free!
[1557392718] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46138 MB (20% inode=99%):
[1557392825] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6632672 kB) free!
[1557392838] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45749 MB (20% inode=99%):
[1557392945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11997876 kB) free.
[1557393128] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557393194] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.73, 0.65, 4.31
[1557393484] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557393944] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 0.74, 12.80
[1557394145] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8605016 kB) free!
[1557394265] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6867328 kB) free!
[1557394328] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 137.07, 37.48, 13.82
[1557394385] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11909928 kB) free.
[1557394448] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 241.34, 110.26, 43.36
[1557394569] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 255.53, 159.07, 69.36
[1557395144] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.49, 0.33, 3.73
[1557395585] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9390216 kB) free!
[1557395705] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7780604 kB) free!
[1557395825] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5924780 kB) free!
[1557395885] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1557396426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11803028 kB) free.
[1557396485] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557396605] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557396725] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557396783] HOST ALERT: db1;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1557396791] SERVICE ALERT: db1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'db1' was out of bounds)
[1557396884] SERVICE ALERT: db1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'db1' was out of bounds)
[1557396894] SERVICE ALERT: db1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'db1' was out of bounds)
[1557396969] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.38, 1.75, 13.99
[1557397026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9871208 kB) free!
[1557397056] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'db1' was out of bounds)
[1557397145] SERVICE ALERT: db1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'db1' was out of bounds)
[1557397146] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8190144 kB) free!
[1557397266] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.1% (6436780 kB) free!
[1557397295] SERVICE ALERT: db1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'db1' was out of bounds)
[1557397320] SERVICE ALERT: db1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'db1' was out of bounds)
[1557397320] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'db1' was out of bounds)
[1557397320] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'db1' was out of bounds)
[1557397350] HOST ALERT: db1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.47)
[1557397394] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 283.83, 167.28, 70.56
[1557397514] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 286.86, 207.01, 96.85
[1557397634] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1557397634] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 147.61, 197.66, 107.86
[1557397834] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1557397834] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 47.5% (23425672 kB) free.
[1557398664] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4793520 kB) free!
[1557398769] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.77, 1.69, 3.43
[1557398784] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4741156 kB) free!
[1557398905] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4693816 kB) free!
[1557399434] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 1.04, 16.99
[1557399439] SERVICE ALERT: webserver1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1557399459] HOST ALERT: webserver1;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1557399518] SERVICE ALERT: webserver1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'webserver1' was out of bounds)
[1557399553] SERVICE ALERT: webserver1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'webserver1' was out of bounds)
[1557399562] SERVICE ALERT: webserver1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'webserver1' was out of bounds)
[1557399603] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3119612 kB) free!
[1557399624] SERVICE ALERT: webserver1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'webserver1' was out of bounds)
[1557399674] SERVICE ALERT: webserver1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'webserver1' was out of bounds)
[1557399723] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3046768 kB) free!
[1557399750] SERVICE ALERT: webserver1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'webserver1' was out of bounds)
[1557399812] SERVICE ALERT: webserver1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'webserver1' was out of bounds)
[1557399844] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.3% (2971004 kB) free!
[1557399864] SERVICE ALERT: webserver1;Nginx_status;CRITICAL;HARD;1;(Return code of 255 for service 'Nginx_status' on host 'webserver1' was out of bounds)
[1557399876] SERVICE ALERT: webserver1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'webserver1' was out of bounds)
[1557399935] SERVICE ALERT: webserver1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'webserver1' was out of bounds)
[1557400026] HOST ALERT: webserver1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1557400159] SERVICE ALERT: webserver1;Total Processes;OK;HARD;1;PROCS OK: 372 processes
[1557400238] SERVICE ALERT: webserver1;CPU Load;OK;HARD;1;OK - load average: 0.05, 0.18, 0.10
[1557400273] SERVICE ALERT: webserver1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1557400325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557400344] SERVICE ALERT: webserver1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 40986 MB (80% inode=99%):
[1557400394] SERVICE ALERT: webserver1;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1557400470] SERVICE ALERT: webserver1;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1557400532] SERVICE ALERT: webserver1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 224832 MB (99% inode=99%):
[1557400543] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.6% (13585528 kB) free.
[1557400585] SERVICE ALERT: webserver1;Nginx_status;OK;HARD;1;NGINX OK -  0.275 sec. response time, Active: 1 (Writing: 1 Reading: 0 Waiting: 0) ReqPerSec: -47.099 ConnPerSec: -20.929 ReqPerConn: 1.250
[1557400595] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.5% (22181980 kB) free.
[1557400596] SERVICE ALERT: webserver1;Memory;OK;HARD;1;OK - 94.1% (23024348 kB) free.
[1557400634] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.42, 4.98
[1557400655] SERVICE ALERT: webserver1;Sensors Status;OK;HARD;1;SENSORS OK
[1557400671] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1557400925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557401045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557401165] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1557401165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557401601] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (3709532 kB) free!
[1557401721] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.9% (3689100 kB) free!
[1557401805] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.9% (3405360 kB) free!
[1557401821] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194556 kB) free!
[1557401841] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.8% (3658544 kB) free!
[1557401925] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.6% (3095032 kB) free!
[1557401941] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (212676 kB) free!
[1557402045] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (3054496 kB) free!
[1557402061] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164796 kB) free!
[1557402680] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557402944] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 354.17, 184.52, 74.17
[1557402969] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.43, 77.59, 38.10
[1557403064] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 47.89, 123.53, 65.20
[1557403089] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.60, 52.59, 33.72
[1557403184] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.65, 82.75, 57.33
[1557403209] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1557403209] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.51, 35.73, 29.83
[1557403809] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.82, 6.24, 16.39
[1557404384] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1557404384] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 1.78, 15.94
[1557404445] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1557404445] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (926572 kB) free!
[1557405438] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.6% (23344864 kB) free.
[1557405475] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.1% (22817388 kB) free.
[1557405523] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.3% (22367136 kB) free.
[1557405544] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1557405584] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.25, 0.27, 4.57
[1557405609] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.19, 1.95, 3.93
[1557406144] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1557406671] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3070136 kB) free!
[1557406791] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (2974996 kB) free!
[1557406911] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (2857820 kB) free!
[1557408043] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557408111] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1516832 kB) free!
[1557408485] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (183952 kB) free!
[1557408605] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194804 kB) free!
[1557408725] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200840 kB) free!
[1557408929] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 345 processes
[1557408962] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (183980 kB) free!
[1557409049] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 344 processes
[1557409082] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (170324 kB) free!
[1557409169] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 344 processes
[1557409202] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169724 kB) free!
[1557409769] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1557410005] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.9% (22526764 kB) free.
[1557410046] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.7% (14586156 kB) free.
[1557410086] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 84.4% (13716556 kB) free.
[1557410766] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1557411366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557411434] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1557411486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557411606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557412184] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 24.29, 16.27, 6.61
[1557412209] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 73.98, 27.23, 10.86
[1557412304] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.59, 11.02, 5.86
[1557412329] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.37, 18.82, 9.76
[1557412424] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.73, 7.48, 5.19
[1557412449] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.97, 13.51, 8.92
[1557413024] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 1.32, 2.90
[1557413133] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 11.0% (809032 kB) free!
[1557413253] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (668936 kB) free!
[1557413373] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (523576 kB) free!
[1557413449] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 95.5% (7020664 kB) free.
[1557413624] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 121.71, 82.69, 38.62
[1557413649] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.92, 2.07, 3.76
[1557413744] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 41.68, 72.02, 40.55
[1557413864] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.07, 48.35, 35.69
[1557415064] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.27, 1.96, 11.95
[1557415134] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1557415206] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1557415413] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1557415806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557415926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557416047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557416264] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.31, 3.51
[1557416647] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1557417247] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1557417367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557417487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557418254] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (1414840 kB) free!
[1557418374] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.3% (1273232 kB) free!
[1557418494] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (1130756 kB) free!
[1557418687] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1557419094] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1557419094] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (436540 kB) free!
[1557419287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557419408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557419528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557419617] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 95.3% (7009132 kB) free.
[1557420343] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 280 processes
[1557420463] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 4 warn out of 280 processes
[1557420583] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 280 processes
[1557423128] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557423729] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1557423849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1557423969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557424422] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (1434472 kB) free!
[1557424449] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1557424542] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.6% (1296708 kB) free!
[1557424662] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 15.7% (1153868 kB) free!
[1557425262] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (459468 kB) free!
[1557425443] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1557425563] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes