[1588617008] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13302764 kB) free!
[1588618967] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4685056 kB) free!
[1588619027] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (1619432 kB) free!
[1588619087] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (170524 kB) free!
[1588619147] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (190300 kB) free!
[1588619791] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (1493528 kB) free!
[1588619911] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (1498860 kB) free!
[1588620008] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32060776 kB) free.
[1588620031] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1495728 kB) free!
[1588620048] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 78.5% (19233988 kB) free.
[1588620608] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7024000 kB) free!
[1588620728] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34632224 kB) free.
[1588621328] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11322408 kB) free!
[1588621448] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7457152 kB) free!
[1588621568] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34291316 kB) free.
[1588622168] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13037412 kB) free!
[1588622288] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9272032 kB) free!
[1588622408] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34218212 kB) free.
[1588622431] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3322532 kB) free.
[1588623008] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12435272 kB) free!
[1588623031] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3181296 kB) free!
[1588623128] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8688540 kB) free!
[1588623151] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3321028 kB) free.
[1588623249] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34128376 kB) free.
[1588623848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11653856 kB) free!
[1588623968] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8108120 kB) free!
[1588624088] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31453692 kB) free.
[1588624351] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3107480 kB) free!
[1588624471] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3299740 kB) free.
[1588624688] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9644784 kB) free!
[1588624808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34048244 kB) free.
[1588625071] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2602284 kB) free!
[1588625191] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3282744 kB) free.
[1588625408] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12628712 kB) free!
[1588625528] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8970488 kB) free!
[1588625648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31769460 kB) free.
[1588625791] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3255152 kB) free!
[1588625911] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3185360 kB) free!
[1588626031] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3144368 kB) free!
[1588626248] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10662908 kB) free!
[1588626368] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34000728 kB) free.
[1588626491] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.2% (42996732 kB) free.
[1588626631] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1270512 kB) free!
[1588626906] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (44866668 kB) free.
[1588626968] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12757316 kB) free!
[1588627088] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34061968 kB) free.
[1588627232] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3685836 kB) free.
[1588627688] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14000128 kB) free!
[1588627808] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13378068 kB) free!
[1588627881] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(305) < 103.16.101.72 [05/May/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1588627929] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31330380 kB) free.
[1588628001] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 103.16.101.72 [05/May/2020:03:01:32 +0530] TCP 502 0 0 60.000 "10.147.212.37:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1588628016] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(388) < 2020/05/05 03:01:32 [error] 22135#0: *86847867 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588628121] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588628136] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588628196] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1588628303] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2845) < 203.212.70.170 [05/May/2020:03:02:32 +0530] TCP 502 0 0 120.002 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1588628424] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588628528] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11627980 kB) free!
[1588628648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (33011876 kB) free.
[1588629248] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12529984 kB) free!
[1588629368] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33959136 kB) free.
[1588629968] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13332356 kB) free!
[1588630088] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33959836 kB) free.
[1588630688] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13879808 kB) free!
[1588630694] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.5% (10170892 kB) free.
[1588630808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33899812 kB) free.
[1588631408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14312820 kB) free!
[1588631528] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1588631528] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33884424 kB) free.
[1588631894] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (225992 kB) free!
[1588632014] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (239900 kB) free!
[1588632128] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14423340 kB) free!
[1588632134] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (217372 kB) free!
[1588632248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33879380 kB) free.
[1588632848] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14538716 kB) free!
[1588632968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33892032 kB) free.
[1588633334] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.3% (9621168 kB) free.
[1588633569] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15063000 kB) free!
[1588633689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33820656 kB) free.
[1588634289] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10805404 kB) free!
[1588634409] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27697300 kB) free.
[1588635009] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11441832 kB) free!
[1588635129] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (19100676 kB) free!
[1588635135] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184512 kB) free!
[1588635249] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30609932 kB) free.
[1588635255] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (222476 kB) free!
[1588635375] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1588635375] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190864 kB) free!
[1588635849] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5225540 kB) free!
[1588635969] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31062956 kB) free.
[1588636433] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 54.7% (13497760 kB) free.
[1588636569] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26240368 kB) free!
[1588636587] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.7% (18552332 kB) free.
[1588636689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30745320 kB) free.
[1588636803] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 55.4% (13582948 kB) free.
[1588637889] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12521524 kB) free!
[1588638009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34468748 kB) free.
[1588638609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14404552 kB) free!
[1588638632] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588638729] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33935356 kB) free.
[1588639330] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12117264 kB) free!
[1588639450] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33516132 kB) free.
[1588640225] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26290096 kB) free!
[1588640345] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.9% (22300992 kB) free!
[1588640465] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.4% (17651172 kB) free!
[1588640650] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12597628 kB) free!
[1588640770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40929116 kB) free.
[1588641065] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (558812 kB) free!
[1588641234] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178724 kB) free!
[1588641354] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178984 kB) free!
[1588641474] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188424 kB) free!
[1588641665] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 65.6% (86518644 kB) free.
[1588642570] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23766412 kB) free!
[1588642690] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42783836 kB) free.
[1588643890] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24401368 kB) free!
[1588644010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43898512 kB) free.
[1588644396] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1588644516] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1588644610] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22319444 kB) free!
[1588644636] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1588644730] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41870288 kB) free.
[1588645575] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588645931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22512688 kB) free!
[1588646051] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41718796 kB) free.
[1588647251] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22400972 kB) free!
[1588647371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41851460 kB) free.
[1588647971] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20886216 kB) free!
[1588648091] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40835668 kB) free.
[1588648691] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19063324 kB) free!
[1588648811] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39389408 kB) free.
[1588649038] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 16.66, 7.88, 5.20
[1588649158] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 14.61, 10.14, 6.33
[1588649278] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.76, 8.09, 6.03
[1588650011] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16800680 kB) free!
[1588650131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12265132 kB) free!
[1588650251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36643812 kB) free.
[1588650604] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (226076 kB) free!
[1588650724] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (215556 kB) free!
[1588650844] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (214248 kB) free!
[1588650851] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12798544 kB) free!
[1588650971] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30829504 kB) free.
[1588651571] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2100368 kB) free!
[1588651692] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36058096 kB) free.
[1588652291] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10413124 kB) free!
[1588652411] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6200064 kB) free!
[1588652436] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1588652531] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32529168 kB) free.
[1588653131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7441336 kB) free!
[1588653251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33117500 kB) free.
[1588653632] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3218268 kB) free!
[1588653752] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3685916 kB) free.
[1588653851] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7030920 kB) free!
[1588653971] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18083684 kB) free!
[1588653987] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189532 kB) free!
[1588654004] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1588654091] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28597412 kB) free.
[1588654107] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187956 kB) free!
[1588654125] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1588654227] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (235040 kB) free!
[1588654245] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1588654691] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4554376 kB) free!
[1588654811] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29510544 kB) free.
[1588654952] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3155928 kB) free!
[1588655072] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3470616 kB) free.
[1588655411] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4129768 kB) free!
[1588655531] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29323556 kB) free.
[1588656045] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1595580 kB) free!
[1588656131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4900944 kB) free!
[1588656251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (626824 kB) free!
[1588656272] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2856836 kB) free!
[1588656304] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4418444 kB) free!
[1588656372] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24589236 kB) free!
[1588656392] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3316088 kB) free.
[1588656424] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4089340 kB) free!
[1588656544] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.4% (3683248 kB) free!
[1588656645] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (1670892 kB) free!
[1588656972] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392336 kB) free!
[1588656992] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3099724 kB) free!
[1588657112] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3454220 kB) free.
[1588657144] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1180488 kB) free!
[1588657300] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 48.9% (11996364 kB) free.
[1588657572] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16685932 kB) free!
[1588657713] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3215108 kB) free!
[1588657833] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2590796 kB) free!
[1588657845] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (450784 kB) free!
[1588657953] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (3099808 kB) free!
[1588658772] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1588658772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27762376 kB) free.
[1588659023] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(698) < 103.255.146.164 [05/May/2020:11:30:44 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588659143] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588659972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11711960 kB) free!
[1588660092] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (7022320 kB) free!
[1588660212] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34310232 kB) free.
[1588660353] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1416552 kB) free!
[1588660745] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.0% (16166236 kB) free.
[1588660812] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8591664 kB) free!
[1588660900] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.1% (4691440 kB) free!
[1588660932] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4036496 kB) free!
[1588660953] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.2% (2311176 kB) free!
[1588661021] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.1% (4430160 kB) free!
[1588661052] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396132 kB) free!
[1588661141] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.0% (4174472 kB) free!
[1588662252] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19891920 kB) free!
[1588662306] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9340472 kB) free!
[1588662341] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (930220 kB) free!
[1588662426] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8927144 kB) free!
[1588662546] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8600212 kB) free!
[1588662852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33358204 kB) free.
[1588662932] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1588663052] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1588663353] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1588663353] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (934124 kB) free!
[1588663453] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9336472 kB) free!
[1588663573] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4921280 kB) free!
[1588663693] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32075888 kB) free.
[1588664293] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8317472 kB) free!
[1588664413] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3595816 kB) free!
[1588664427] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (4829928 kB) free!
[1588664533] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (438288 kB) free!
[1588664553] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (2974600 kB) free!
[1588664946] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3708864 kB) free!
[1588665733] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23168816 kB) free!
[1588666333] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399128 kB) free!
[1588666746] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4406564 kB) free!
[1588666827] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1883600 kB) free!
[1588666866] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.8% (3773240 kB) free!
[1588666986] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3139228 kB) free!
[1588667533] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23320676 kB) free!
[1588667586] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1588667586] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (179856 kB) free!
[1588668133] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404448 kB) free!
[1588668491] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9267972 kB) free!
[1588668611] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (9138224 kB) free!
[1588668731] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8923992 kB) free!
[1588668753] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1050588 kB) free!
[1588669333] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15948040 kB) free!
[1588669343] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(452) < 103.255.146.164 [05/May/2020:14:31:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588669353] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (1860916 kB) free!
[1588669463] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588669933] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (453432 kB) free!
[1588669953] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3427504 kB) free.
[1588670371] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (6870292 kB) free!
[1588670491] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.7% (6549080 kB) free!
[1588670534] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25538244 kB) free!
[1588670553] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (2006012 kB) free!
[1588670611] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 18.0% (6297396 kB) free!
[1588670664] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(378) < 103.255.146.164 [05/May/2020:14:54:10 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588670673] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.3% (2321316 kB) free!
[1588670784] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588670793] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (2505636 kB) free!
[1588671134] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2867160 kB) free!
[1588671187] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16263868 kB) free.
[1588671211] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.7% (31757756 kB) free.
[1588671283] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4569196 kB) free!
[1588671384] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(108) < 185.255.8.58 [05/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1588671404] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (4087788 kB) free!
[1588671504] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588671524] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3671388 kB) free!
[1588671646] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 13.0% (2105332 kB) free!
[1588671731] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4299868 kB) free!
[1588671734] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15198524 kB) free!
[1588672104] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(345) < 103.255.146.164 [05/May/2020:15:12:04 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588672124] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.4% (22801820 kB) free.
[1588672224] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588672334] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44486476 kB) free.
[1588672934] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17299196 kB) free!
[1588673054] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12601152 kB) free!
[1588673174] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8429192 kB) free!
[1588673774] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (22019124 kB) free!
[1588673793] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (889464 kB) free!
[1588674374] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37087324 kB) free.
[1588674393] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3317228 kB) free.
[1588674478] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.76, 3.09, 4.70
[1588674624] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(256) < 103.255.146.164 [05/May/2020:15:51:52 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588674744] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588674974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11969032 kB) free!
[1588674994] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3169704 kB) free!
[1588675094] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7758648 kB) free!
[1588675114] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3158472 kB) free!
[1588675214] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40690548 kB) free.
[1588675234] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3116300 kB) free!
[1588675814] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13132620 kB) free!
[1588675934] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8673524 kB) free!
[1588675944] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(674) < 103.255.146.164 [05/May/2020:16:17:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588676054] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4456624 kB) free!
[1588676065] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588676252] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1588676278] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.34, 6.76, 5.55
[1588676372] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1588676398] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.35, 6.83, 5.74
[1588676446] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175476 kB) free!
[1588676518] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.12, 8.49, 6.46
[1588676654] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16258460 kB) free!
[1588676973] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 335 processes
[1588677034] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1250352 kB) free!
[1588677093] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1588677255] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32331320 kB) free.
[1588677855] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5960116 kB) free!
[1588677975] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1752604 kB) free!
[1588678095] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381160 kB) free!
[1588678387] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.3% (3445056 kB) free!
[1588678508] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.1% (2945320 kB) free!
[1588678536] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/05 17:01:07 [error] 22134#0: *86911869 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1588678628] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.5% (2351644 kB) free!
[1588678656] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588678834] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (2487396 kB) free!
[1588679228] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (239820 kB) free!
[1588679257] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/05 17:12:29 [error] 22132#0: *86913446 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1588679295] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32393768 kB) free.
[1588679377] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588679895] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6647296 kB) free!
[1588680015] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2309940 kB) free!
[1588680135] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393052 kB) free!
[1588681234] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.8% (4844352 kB) free.
[1588681335] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30508852 kB) free.
[1588681834] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1097372 kB) free!
[1588681935] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4873120 kB) free!
[1588681954] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (2743836 kB) free!
[1588682055] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (417664 kB) free!
[1588682074] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.8% (4849736 kB) free.
[1588682175] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35951872 kB) free.
[1588682228] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.8% (16341128 kB) free.
[1588682675] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (1607736 kB) free!
[1588682775] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9818032 kB) free!
[1588682795] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3507696 kB) free.
[1588682896] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5214552 kB) free!
[1588683016] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (971008 kB) free!
[1588683395] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (560648 kB) free!
[1588683515] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.2% (1819568 kB) free!
[1588683616] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18778664 kB) free!
[1588683635] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.7% (4987160 kB) free.
[1588684216] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35381280 kB) free.
[1588684816] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6935784 kB) free!
[1588684835] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3153732 kB) free!
[1588684936] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2193724 kB) free!
[1588684955] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.8% (4357768 kB) free.
[1588685056] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414032 kB) free!
[1588685493] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1588685518] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.25, 3.85, 4.47
[1588685555] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3217188 kB) free!
[1588685613] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1588685675] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2423708 kB) free!
[1588685795] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3732244 kB) free.
[1588686256] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16556064 kB) free!
[1588686856] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385956 kB) free!
[1588687318] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.49, 8.72, 6.52
[1588687438] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.02, 7.85, 6.49
[1588687456] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21689832 kB) free!
[1588687558] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.95, 6.82, 6.27
[1588687595] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (791284 kB) free!
[1588687715] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.6% (1877484 kB) free!
[1588687835] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3338644 kB) free.
[1588688056] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382912 kB) free!
[1588688656] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14825168 kB) free!
[1588689256] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37799328 kB) free.
[1588689359] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.49, 4.72, 4.93
[1588689856] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10973064 kB) free!
[1588689959] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.07, 6.00, 5.45
[1588689976] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6442324 kB) free!
[1588690079] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.45, 5.91, 5.50
[1588690097] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2134028 kB) free!
[1588690199] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1588690199] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.23, 5.96, 5.56
[1588690235] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (662708 kB) free!
[1588690355] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.6% (1728716 kB) free!
[1588690475] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2944348 kB) free!
[1588690697] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17436320 kB) free!
[1588691013] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1588691075] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3312384 kB) free.
[1588691133] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1588691228] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1588691297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38064024 kB) free.
[1588691675] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3246608 kB) free!
[1588691795] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3246240 kB) free!
[1588691897] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10931600 kB) free!
[1588691915] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (563232 kB) free!
[1588692017] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6463752 kB) free!
[1588692137] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2179348 kB) free!
[1588692737] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18216116 kB) free!
[1588693337] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38585708 kB) free.
[1588693715] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3164728 kB) free!
[1588693937] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11866760 kB) free!
[1588694057] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7519680 kB) free!
[1588694177] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3478360 kB) free!
[1588694399] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.85, 2.40, 3.79
[1588694777] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17440352 kB) free!
[1588695377] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38911324 kB) free.
[1588695515] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.2% (5076960 kB) free.
[1588695933] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1588695977] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12874640 kB) free!
[1588696053] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1588696097] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8803208 kB) free!
[1588696217] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4803908 kB) free!
[1588696818] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17782836 kB) free!
[1588697418] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37781468 kB) free.
[1588698018] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11323152 kB) free!
[1588698138] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6802048 kB) free!
[1588698258] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2498932 kB) free!
[1588698858] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17407088 kB) free!
[1588699115] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3096980 kB) free!
[1588699235] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3266144 kB) free.
[1588699458] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34418520 kB) free.
[1588700058] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10489724 kB) free!
[1588700178] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6567100 kB) free!
[1588700298] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2712700 kB) free!
[1588700898] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14831976 kB) free!
[1588701498] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28733608 kB) free.
[1588702098] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6027740 kB) free!
[1588702199] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1588702218] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2431348 kB) free!
[1588702338] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35504612 kB) free.
[1588702835] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2884336 kB) free!
[1588702939] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12016836 kB) free!
[1588702955] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.0% (4548236 kB) free.
[1588703059] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8400748 kB) free!
[1588703179] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5383592 kB) free!