[1575225131] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 67.3% (33168820 kB) free.
[1575225447] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164380 kB) free!
[1575225475] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2019/12/02 00:05:24 [error] 24130#0: *11751557 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575225477] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39610636 kB) free.
[1575225535] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (246164 kB) free!
[1575225567] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (212708 kB) free!
[1575225596] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(26) < 2019/12/02 00:08:57 [error] 24130#0: *11752448 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575225655] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (228980 kB) free!
[1575225687] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203748 kB) free!
[1575225716] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575225775] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (216608 kB) free!
[1575226316] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(50) < 2019/12/02 00:17:46 [error] 24130#0: *11754849 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575226436] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575227036] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2019/12/02 00:31:11 [error] 24130#0: *11758472 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575227156] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(30) < 2019/12/02 00:34:39 [error] 24130#0: *11759376 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575227276] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575227846] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.8% (5621648 kB) free.
[1575227877] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(53) < 2019/12/02 00:43:17 [error] 24130#0: *11761692 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575227997] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575228447] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184108 kB) free!
[1575228567] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183540 kB) free!
[1575228597] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(56) < 2019/12/02 00:59:30 [error] 24130#0: *11766066 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575228687] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186684 kB) free!
[1575228717] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575229317] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(49) < 2019/12/02 01:08:13 [error] 24130#0: *11768428 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575229437] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575230037] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(31) < 2019/12/02 01:21:16 [error] 24132#0: *11771967 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575230158] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(29) < 2019/12/02 01:25:09 [error] 24125#0: *11772984 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575230278] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575230878] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(51) < 2019/12/02 01:33:53 [error] 24132#0: *11775372 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575230998] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575231598] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2019/12/02 01:46:38 [error] 24130#0: *11778833 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575231718] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(30) < 2019/12/02 01:50:41 [error] 24130#0: *11779892 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575231839] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575232438] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(46) < 2019/12/02 01:58:48 [error] 24130#0: *11782114 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575232558] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575233158] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(57) < 2019/12/02 02:15:45 [error] 24130#0: *11786654 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575233278] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575233617] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 46.1% (22714264 kB) free.
[1575233878] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(50) < 2019/12/02 02:23:45 [error] 24130#0: *11788783 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575233998] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575234087] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.2% (5466664 kB) free.
[1575234167] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 74.0% (36462336 kB) free.
[1575234599] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2019/12/02 02:36:53 [error] 24134#0: *11792363 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575234719] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(29) < 2019/12/02 02:41:40 [error] 24130#0: *11793625 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575234839] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575235439] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(46) < 2019/12/02 02:49:15 [error] 24130#0: *11795800 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575235559] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575236159] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2019/12/02 03:00:54 [error] 24132#0: *11798870 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575236279] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(29) < 2019/12/02 03:07:18 [error] 24130#0: *11800586 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575236400] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575236999] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(52) < 2019/12/02 03:14:17 [error] 24130#0: *11802470 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575237119] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575237719] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/12/02 03:25:11 [error] 24130#0: *11805484 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575237839] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(25) < 2019/12/02 03:32:28 [error] 24130#0: *11807419 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575237959] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575238560] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(51) < 2019/12/02 03:45:38 [error] 24130#0: *11811114 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575238680] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575239280] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(61) < 2019/12/02 03:57:33 [error] 24130#0: *11814225 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575239400] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575240000] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(59) < 2019/12/02 04:06:22 [error] 24127#0: *11816680 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575240120] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575240720] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2019/12/02 04:14:53 [error] 24132#0: *11819142 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575240840] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(26) < 2019/12/02 04:23:14 [error] 24130#0: *11821371 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575240961] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575241560] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(197) < 2019/12/02 04:35:59 [error] 24129#0: *11825039 upstream timed out (110: Connection timed out) while connecting to upstream, client: 202.65.140.29, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575241680] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(13) < 2019/12/02 04:36:09 [error] 24135#0: *11825261 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575241800] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575242400] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(51) < 2019/12/02 04:47:42 [error] 24130#0: *11828717 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575242520] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575243087] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4875592 kB) free!
[1575243120] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2019/12/02 04:55:29 [error] 24130#0: *11830949 connect() failed (111: Connection refused) while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575243208] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4867304 kB) free!
[1575243241] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575243328] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (4768312 kB) free!
[1575243840] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(52) < 2019/12/02 05:12:58 [error] 24130#0: *11835626 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575243960] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575244560] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2019/12/02 05:20:18 [error] 24130#0: *11837759 connect() failed (111: Connection refused) while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575244680] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575245280] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2019/12/02 05:30:40 [error] 24130#0: *11840469 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575245400] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(28) < 2019/12/02 05:38:29 [error] 24130#0: *11842597 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575245520] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575246121] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(52) < 2019/12/02 05:46:42 [error] 24130#0: *11844820 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575246241] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575246841] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2019/12/02 05:54:07 [error] 24130#0: *11846887 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575246928] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171632 kB) free!
[1575246961] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(25) < 2019/12/02 06:04:43 [error] 24130#0: *11849738 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575247081] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575247287] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.5% (21190188 kB) free.
[1575247375] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 71.5% (17520708 kB) free.
[1575247558] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.9% (30251512 kB) free.
[1575247681] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(72) < 2019/12/02 06:17:45 [error] 24131#0: *11853353 connect() failed (111: Connection refused) while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575247801] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575248158] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11845484 kB) free!
[1575248278] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7769856 kB) free!
[1575248398] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 31.7% (41733304 kB) free.
[1575248402] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(31) < 2019/12/02 06:29:52 [error] 24130#0: *11856546 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575248522] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575249122] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(54) < 2019/12/02 06:36:07 [error] 24130#0: *11858260 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575249242] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(20) < 2019/12/02 06:43:32 [error] 24130#0: *11860254 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575249362] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575249962] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2019/12/02 06:54:34 [error] 24130#0: *11863253 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575250082] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575250464] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(177) < 2019/12/02 07:02:15 [error] 26848#0: *13952328 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:156/17, bytes from/to upstream:17/156
[1575250585] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575250635] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.4% (20816856 kB) free.
[1575250672] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.7% (14262180 kB) free.
[1575250683] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(50) < 2019/12/02 07:01:05 [error] 24130#0: *11865104 connect() failed (111: Connection refused) while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575250803] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(29) < 2019/12/02 07:08:56 [error] 24132#0: *11867124 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575250923] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575251277] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9797652 kB) free!
[1575251397] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9796944 kB) free!
[1575251517] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9796792 kB) free!
[1575251523] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(34) < 2019/12/02 07:19:56 [error] 24130#0: *11870104 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575251643] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575251784] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/02 07:23:34 [error] 26851#0: *13960996 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:196/17, bytes from/to upstream:17/196
[1575251904] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/02 07:27:19 [error] 26846#0: *13962585 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:115/17, bytes from/to upstream:17/115
[1575252024] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575252243] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(59) < 2019/12/02 07:26:49 [error] 24130#0: *11871928 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575252363] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(21) < 2019/12/02 07:35:21 [error] 24130#0: *11874252 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575252483] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575252624] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/02 07:36:06 [error] 26848#0: *13966201 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:160/17, bytes from/to upstream:17/160
[1575252744] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575253084] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2019/12/02 07:44:17 [error] 24130#0: *11876662 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575253204] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575253344] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/02 07:44:52 [error] 26851#0: *13969765 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:156/17, bytes from/to upstream:17/156
[1575253465] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575253804] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(76) < 2019/12/02 07:59:51 [error] 24130#0: *11880942 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575253924] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575254064] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/02 07:57:24 [error] 26845#0: *13974806 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:147/17, bytes from/to upstream:17/147
[1575254184] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575254524] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(22) < 2019/12/02 08:09:38 [error] 24130#0: *11883610 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575254644] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575255244] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(57) < 2019/12/02 08:16:59 [error] 24130#0: *11885610 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575255365] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(24) < 2019/12/02 08:25:56 [error] 24130#0: *11888037 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575255485] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575256085] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(35) < 2019/12/02 08:33:45 [error] 24127#0: *11890164 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575256205] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(27) < 2019/12/02 08:39:26 [error] 24130#0: *11891691 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575256325] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(23) < 2019/12/02 08:40:59 [error] 24130#0: *11892115 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575264446] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.7% (3136260 kB) free!
[1575264496] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25486172 kB) free!
[1575264567] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.4% (3068464 kB) free!
[1575264616] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.9% (53974028 kB) free.
[1575264687] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2825864 kB) free!
[1575265287] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172852 kB) free!
[1575265317] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4699948 kB) free!
[1575265417] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (1164752 kB) free!
[1575265537] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228728 kB) free!
[1575265657] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223288 kB) free!
[1575265917] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (5128408 kB) free!
[1575265967] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (1257656 kB) free!
[1575266087] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (224676 kB) free!
[1575266207] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242576 kB) free!
[1575266257] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 52.9% (26066096 kB) free.
[1575266517] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3261568 kB) free!
[1575266631] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.3% (20563356 kB) free.
[1575266674] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575266674] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 73.9% (36408556 kB) free.
[1575266713] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.3% (43046180 kB) free.
[1575266857] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9298700 kB) free!
[1575266977] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9149100 kB) free!
[1575267098] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (9016424 kB) free!
[1575267234] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (329272 kB) free!
[1575267354] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (248684 kB) free!
[1575267474] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1575267474] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192576 kB) free!
[1575269772] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 388 processes
[1575269892] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 384 processes
[1575270012] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 384 processes
[1575270698] SERVICE FLAPPING ALERT: qservercluster2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1575270698] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 45.5% (22421780 kB) free.
[1575271887] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (4318528 kB) free!
[1575271918] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1575272007] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.6% (4070480 kB) free!
[1575272038] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1575272127] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (3881624 kB) free!
[1575272575] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4669292 kB) free!
[1575272695] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (4519356 kB) free!
[1575272816] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (4362532 kB) free!
[1575273910] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22845600 kB) free.
[1575273928] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1949404 kB) free!
[1575275084] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1575275204] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1575275216] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2118680 kB) free!
[1575276313] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.1% (3975788 kB) free!
[1575276433] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (4093656 kB) free!
[1575276553] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180648 kB) free!
[1575276970] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22823764 kB) free.
[1575277011] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.2% (22850360 kB) free.
[1575277059] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 80.6% (19761828 kB) free.
[1575277477] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575277816] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21272516 kB) free!
[1575277936] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44796900 kB) free.
[1575278436] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4551800 kB) free!
[1575278556] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4483092 kB) free!
[1575278676] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4388744 kB) free!
[1575279098] SERVICE FLAPPING ALERT: qservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1575280336] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22204036 kB) free!
[1575280456] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42900288 kB) free.
[1575280477] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9506152 kB) free!
[1575280598] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8665440 kB) free!
[1575280718] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (8332512 kB) free!
[1575281076] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2366312 kB) free!
[1575281177] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4687656 kB) free!
[1575281297] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4596660 kB) free!
[1575281417] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (4517776 kB) free!
[1575282856] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18501756 kB) free!
[1575282976] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39601504 kB) free.
[1575283030] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.3% (22874664 kB) free.
[1575283817] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (1329448 kB) free!
[1575283898] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9394892 kB) free!
[1575284018] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9296028 kB) free!
[1575284138] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (9189824 kB) free!
[1575284318] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4762520 kB) free!
[1575284776] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16360780 kB) free!
[1575284896] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (36980592 kB) free.
[1575286672] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.7% (22884852 kB) free.
[1575286715] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39478672 kB) free.
[1575287063] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1575287183] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1575288482] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4284316 kB) free!
[1575288496] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24319408 kB) free!
[1575288602] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (4201548 kB) free!
[1575288616] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46282992 kB) free.
[1575288722] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (4103464 kB) free!
[1575288939] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4628044 kB) free!
[1575289515] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9720348 kB) free!
[1575289635] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9511768 kB) free!
[1575289755] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9313144 kB) free!
[1575289816] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22033088 kB) free!
[1575289936] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44285176 kB) free.
[1575290136] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 73.7% (36317340 kB) free.
[1575290412] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 385 processes
[1575290522] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (647976 kB) free!
[1575290532] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 382 processes
[1575290814] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4793708 kB) free!
[1575290934] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4633336 kB) free!
[1575291054] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4463336 kB) free!
[1575291136] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19185000 kB) free!
[1575291257] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43078268 kB) free.
[1575291467] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4620888 kB) free!
[1575291587] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (4461256 kB) free!
[1575291707] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (4320236 kB) free!
[1575291856] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17106704 kB) free!
[1575291976] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40048744 kB) free.
[1575292155] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4649320 kB) free!
[1575292431] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46703 MB (20% inode=99%):
[1575292551] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46487 MB (20% inode=99%):
[1575292671] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44431 MB (19% inode=99%):
[1575292854] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1911332 kB) free!
[1575293176] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17021992 kB) free!
[1575293296] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40372388 kB) free.
[1575293715] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.3% (22772484 kB) free.
[1575293756] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.1% (22563016 kB) free.
[1575293804] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1575293806] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.2% (45946292 kB) free.
[1575293896] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14784700 kB) free!
[1575293925] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1575294016] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37783524 kB) free.
[1575294108] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2071024 kB) free!
[1575295216] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15051712 kB) free!
[1575295336] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37892000 kB) free.
[1575295936] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14468856 kB) free!
[1575296056] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37381836 kB) free.
[1575296166] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170628 kB) free!
[1575296271] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22323 MB (9% inode=99%):
[1575296286] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197792 kB) free!
[1575296406] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1575296406] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197700 kB) free!
[1575296656] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12448648 kB) free!
[1575296776] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35074428 kB) free.
[1575297132] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 385 processes
[1575297253] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 385 processes
[1575297361] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 79.5% (19492088 kB) free.
[1575297373] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 383 processes
[1575297383] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1575297399] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.0% (22300312 kB) free.
[1575297471] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25453 MB (11% inode=99%):
[1575297503] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1575297976] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11030436 kB) free!
[1575298071] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1575298071] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 96473 MB (42% inode=99%):
[1575298096] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43899552 kB) free.
[1575298491] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1489 processes
[1575298584] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(36) < 2019/12/02 20:26:00 [error] 26843#0: *14259498 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:152/17, bytes from/to upstream:17/152
[1575298607] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191612 kB) free!
[1575298704] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(176) < 2019/12/02 20:28:22 [error] 26842#0: *14260824 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575298727] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176520 kB) free!
[1575298824] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575298847] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188484 kB) free!
[1575298928] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575299296] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23369568 kB) free!
[1575299416] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45585348 kB) free.
[1575299528] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2019/12/02 20:42:07 [error] 24130#0: *12085520 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575299648] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(60) < 2019/12/02 20:43:49 [error] 24130#0: *12085809 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575299768] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575300925] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1575301043] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.3% (23104980 kB) free.
[1575301568] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(111) < 2019/12/02 21:09:17 [error] 24130#0: *12092733 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575301689] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575302888] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(112) < 2019/12/02 21:35:03 [error] 24130#0: *12099718 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575302925] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9625556 kB) free!
[1575303008] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575303046] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9624432 kB) free!
[1575303166] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (9623648 kB) free!
[1575304208] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(105) < 2019/12/02 22:00:08 [error] 24132#0: *12106672 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575304328] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(7) < 2019/12/02 22:00:30 [error] 24132#0: *12106607 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575304449] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575305648] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2019/12/02 22:23:30 [error] 24130#0: *12113015 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575305768] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(39) < 2019/12/02 22:26:07 [error] 24130#0: *12113687 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575305888] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(4) < 2019/12/02 22:26:15 [error] 24130#0: *12113669 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575306016] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25875056 kB) free!
[1575306136] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45168668 kB) free.
[1575306488] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575307088] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2019/12/02 22:47:34 [error] 24130#0: *12119485 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575307208] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(26) < 2019/12/02 22:49:41 [error] 24134#0: *12120171 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575307328] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1575307328] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(27) < 2019/12/02 22:51:27 [error] 24130#0: *12120560 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575307336] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25893268 kB) free!
[1575307456] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45588448 kB) free.
[1575307928] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575308656] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26329600 kB) free!
[1575308776] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (45014860 kB) free.
[1575309128] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(83) < 2019/12/02 23:16:15 [error] 24130#0: *12127326 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575309248] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575309471] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575309976] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26297916 kB) free!
[1575310047] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575310096] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44259316 kB) free.
[1575310448] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(82) < 2019/12/02 23:41:13 [error] 24130#0: *12134094 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575310568] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575311296] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25825516 kB) free!