[1546887199] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 626 processes
[1546887319] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 627 processes
[1546887330] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.40, 28.81, 18.56
[1546887439] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 358 processes
[1546887450] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.68, 19.37, 16.34
[1546887570] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1546887570] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 13.17, 14.43
[1546888770] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.48, 0.78, 4.37
[1546897170] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.43, 13.33, 14.38
[1546897290] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.67, 9.16, 12.73
[1546897410] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 6.25, 11.23
[1546898610] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.45, 0.64, 3.47
[1546901610] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 132.55, 71.93, 27.91
[1546901730] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 18.40, 48.36, 24.61
[1546901850] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.05, 32.56, 21.70
[1546902450] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.85, 4.94, 11.67
[1546903650] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.71, 0.77, 3.66
[1546906650] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 20.21, 42.79, 21.05
[1546906770] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.34, 28.83, 18.57
[1546906890] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.09, 19.50, 16.39
[1546908090] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.48, 0.89, 4.91
[1546911690] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 5.74, 32.73, 19.26
[1546911810] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.19, 22.03, 16.97
[1546911930] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 14.90, 14.97
[1546913130] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.76, 4.51
[1546916730] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.87, 21.43, 18.93
[1546916850] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.67, 14.51, 16.70
[1546916970] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 9.90, 14.75
[1546918170] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.68, 0.83, 4.54
[1546918600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546918720] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546922156] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546922276] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546922396] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546925920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1546926040] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1546926641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1546926761] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1546928970] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1546929456] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1546929576] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1546930170] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.64, 13.74, 16.26
[1546930290] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.56, 9.40, 14.36
[1546930410] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 6.58, 12.73
[1546931561] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546931610] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.71, 0.84, 4.02
[1546931681] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546931725] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.4% (2496784 kB) free!
[1546931801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546932401] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1546933601] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546933721] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1546937921] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1546938041] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1546938161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1546938762] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1546938762] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1546943610] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 46.05, 59.37, 26.61
[1546943730] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6.64, 39.88, 23.44
[1546943850] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.50, 26.99, 20.72
[1546944450] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 4.19, 11.18
[1546944762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546944882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1546945650] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.28, 0.53, 3.48
[1546945796] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 21 users currently logged in
[1546946397] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 19 users currently logged in
[1546947225] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 314 processes
[1546947282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546947345] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1546947402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546947523] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1546948197] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 21 users currently logged in
[1546949322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546949397] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1546949397] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 19 users currently logged in
[1546949442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546949562] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1546950162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1546950282] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1546950883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1546951003] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546951003] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1546951536] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1546951656] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1546951776] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1546952125] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1631424 kB) free!
[1546952647] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 43 users currently logged in
[1546952707] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 43 users currently logged in
[1546952767] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 43 users currently logged in
[1546952827] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 43 users currently logged in
[1546952997] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 21 users currently logged in
[1546953403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546953427] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 37 users currently logged in
[1546953523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546953597] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546953643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546954176] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1546954243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546957843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546957963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546958083] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546960197] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 17 users currently logged in
[1546960317] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 17 users currently logged in
[1546960437] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 17 users currently logged in
[1546961037] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1546965883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546966003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546966123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546967924] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1546970506] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.09, 10.93, 10.43
[1546970626] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.14, 10.59, 10.37
[1546970746] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.65, 10.23, 10.27