[1576780835] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9668940 kB) free!
[1576780955] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9332424 kB) free!
[1576781075] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (9572144 kB) free!
[1576782327] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2506036 kB) free!
[1576782928] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177772 kB) free!
[1576783528] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.1% (10147976 kB) free.
[1576784128] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181292 kB) free!
[1576784156] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1576784191] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1576784248] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183476 kB) free!
[1576784368] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191964 kB) free!
[1576784675] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4467628 kB) free!
[1576789274] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219740 kB) free!
[1576789394] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243068 kB) free!
[1576789514] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.1% (43895560 kB) free.
[1576789851] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576789971] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1576790541] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1576795768] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576799571] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576799692] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576800852] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1327.77, 807.54, 344.61
[1576800972] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1339.36, 984.12, 465.45
[1576801092] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 325.58, 770.46, 454.43
[1576801168] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2618336 kB) free!
[1576801556] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (2700460 kB) free!
[1576801692] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.55, 104.12, 238.58
[1576802156] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175292 kB) free!
[1576802368] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (952304 kB) free!
[1576802826] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.9% (22036128 kB) free.
[1576803356] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.8% (20304812 kB) free.
[1576803391] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.3% (20426376 kB) free.
[1576807491] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576807611] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1576809676] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.29716.1628> (10.147.212.158:43526 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1576809796] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1576811811] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576811931] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576814401] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 306 processes
[1576819818] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.8% (22400436 kB) free.
[1576819886] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.6% (38756968 kB) free.
[1576822228] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (930128 kB) free!
[1576822348] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (813784 kB) free!
[1576822468] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (722148 kB) free!
[1576824875] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(757) < 103.255.146.164 [20/Dec/2019:12:17:32 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576824995] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576829196] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1318) < 103.255.146.164 [20/Dec/2019:13:29:55 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576829316] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576830957] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4423448 kB) free!
[1576831077] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (4305832 kB) free!
[1576831197] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (4199560 kB) free!
[1576832401] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1576833264] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 335 processes
[1576833384] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1576834197] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2057460 kB) free!
[1576835317] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.8% (20177032 kB) free.
[1576835368] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22428996 kB) free.
[1576835791] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4745844 kB) free!
[1576835911] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4658532 kB) free!
[1576836031] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4563148 kB) free!
[1576837731] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1576837851] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576838496] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9591552 kB) free!
[1576838616] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9579028 kB) free!
[1576838736] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9874020 kB) free!
[1576838927] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.7% (4112212 kB) free!
[1576839031] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2228440 kB) free!
[1576839047] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (3994440 kB) free!
[1576839167] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (3896984 kB) free!
[1576839768] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (1034132 kB) free!
[1576840251] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576840371] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576840536] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239532 kB) free!
[1576840585] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 344 processes
[1576840705] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1576841114] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9214236 kB) free!
[1576841234] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9105884 kB) free!
[1576841354] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9013952 kB) free!
[1576841826] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4633748 kB) free!
[1576841921] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.9% (4869248 kB) free!
[1576841946] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4535780 kB) free!
[1576842041] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.8% (4851128 kB) free!
[1576842066] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (4419376 kB) free!
[1576842162] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.7% (4839664 kB) free!
[1576843910] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235640 kB) free!
[1576844030] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238412 kB) free!
[1576844150] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248180 kB) free!
[1576844607] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6241668 kB) free!
[1576844728] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6218440 kB) free!
[1576844848] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8407756 kB) free!
[1576844902] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 220446 MB (20% inode=99%):
[1576844954] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4353412 kB) free!
[1576845022] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 220446 MB (20% inode=99%):
[1576845142] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 220446 MB (20% inode=99%):
[1576845350] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.6% (28871200 kB) free.
[1576845666] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2161252 kB) free!
[1576847404] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (27680092 kB) free.
[1576847443] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (7152280 kB) free!
[1576847445] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1576847445] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 49.1% (24194396 kB) free.
[1576847465] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 24.5% (12072136 kB) free.
[1576847490] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 55.5% (13599900 kB) free.
[1576847513] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 54.5% (13364656 kB) free.
[1576847535] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.7% (12019544 kB) free.
[1576848142] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 293955 MB (27% inode=99%):
[1576848505] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 414 processes
[1576848625] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 409 processes
[1576848745] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 408 processes
[1576849935] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.7% (3379452 kB) free!
[1576850055] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.3% (3290160 kB) free!
[1576850175] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2094912 kB) free!
[1576852352] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1576852472] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1576852592] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 466 processes
[1576852971] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576853091] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576854270] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 569 processes
[1576854390] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 575 processes
[1576854510] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 578 processes
[1576855110] SERVICE ALERT: webservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1576858456] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.6% (20876168 kB) free.
[1576858845] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576864491] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576864611] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 402 processes
[1576864731] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1576865378] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4866056 kB) free!
[1576865498] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4860420 kB) free!
[1576865618] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4849692 kB) free!