[1546109307] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 360 processes
[1546109430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546109550] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546109907] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 364 processes
[1546112875] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 14 users currently logged in
[1546116507] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 360 processes
[1546121907] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 360 processes
[1546128507] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1546130307] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 361 processes
[1546136114] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.10, 13.58, 7.73
[1546136234] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.79, 9.24, 6.85
[1546136307] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 362 processes
[1546136354] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.33, 6.27, 6.05
[1546136954] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.53, 1.24, 3.39
[1546144707] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 359 processes
[1546148129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9332908 kB) free!
[1546148249] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7768028 kB) free!
[1546148369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10651512 kB) free.
[1546148955] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.47, 63.83, 45.09
[1546149075] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.77, 42.81, 39.66
[1546149150] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546149195] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.24, 28.70, 34.87
[1546149270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546149390] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546149581] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (3134348 kB) free!
[1546149795] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 4.14, 18.42
[1546150675] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546150795] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546150915] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546151595] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.34, 2.47, 4.95
[1546152342] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.34, 8.84, 9.51
[1546154742] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.08, 11.12, 10.31
[1546154862] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.67, 11.58, 10.57
[1546154982] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.26, 12.19, 10.91
[1546156169] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9481340 kB) free!
[1546156289] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7905392 kB) free!
[1546156409] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10626400 kB) free.
[1546158390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546158510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546158630] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1546160609] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9111700 kB) free!
[1546160729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10633852 kB) free.
[1546161031] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546161151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546161271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546161871] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546164446] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1546164566] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1546165182] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.81, 8.19, 9.34
[1546168182] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.32, 10.89, 10.10
[1546168302] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.20, 11.27, 10.33
[1546168422] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.17, 11.54, 10.53
[1546168884] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (177596 kB) free!
[1546169004] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (177676 kB) free!
[1546169124] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179340 kB) free!
[1546173115] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546173222] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.59, 9.46, 9.96
[1546174181] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174584 kB) free!
[1546176222] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.19, 10.50, 10.03
[1546176343] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.70, 10.26, 10.01
[1546176463] SERVICE FLAPPING ALERT: qservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1546176463] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.85, 10.83, 10.24
[1546182271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546182391] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1546185463] SERVICE FLAPPING ALERT: qservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1546185716] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546185836] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546185956] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546187907] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 363 processes
[1546188463] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.17, 7.69, 9.88
[1546189529] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8313172 kB) free!
[1546189649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10609796 kB) free.
[1546191356] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 14 users currently logged in