[1544469641] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 310 processes
[1544469762] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 311 processes
[1544469882] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 311 processes
[1544475586] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544475706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544475826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544476426] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1544501082] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 311 processes
[1544502826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544502947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544503067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544504867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544504987] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544505107] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544507222] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 45 users currently logged in
[1544507282] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 45 users currently logged in
[1544507342] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 44 users currently logged in
[1544507402] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 44 users currently logged in
[1544507881] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1544508481] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1544508601] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 337 processes
[1544508707] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1544508722] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1544508827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1544508947] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1544509628] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 41.2% (6703344 kB) free.
[1544510082] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1544510522] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1544510682] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 322 processes
[1544510802] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1544511122] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 339 processes
[1544511242] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 339 processes
[1544511362] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1544511362] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 339 processes
[1544511403] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1544511523] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544511643] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544517947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1544518067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1544520468] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1544520588] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1544520708] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1544522762] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1544524438] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 16.4% (2657748 kB) free!
[1544526962] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1544527308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1544527428] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1544527548] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1544532549] SERVICE ALERT: sig3;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'sig3' was out of bounds)
[1544532567] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1544532616] SERVICE ALERT: sig3;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig3' was out of bounds)
[1544532631] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1544532669] SERVICE ALERT: sig3;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 220138 MB (97% inode=99%):
[1544532736] SERVICE ALERT: sig3;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1544532948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1544532962] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 343 processes
[1544533069] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1544533082] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 343 processes
[1544533202] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 346 processes
[1544533243] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1544533900] SERVICE ALERT: sig4;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544533930] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1544533931] SERVICE ALERT: sig4;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig4' was out of bounds)
[1544533940] SERVICE ALERT: sig4;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig4' was out of bounds)
[1544533948] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4' was out of bounds)
[1544533948] SERVICE ALERT: sig4;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4' was out of bounds)
[1544534021] SERVICE ALERT: sig4;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 222674 MB (99% inode=99%):
[1544534025] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1544534051] SERVICE ALERT: sig4;Sensors Status;OK;HARD;1;SENSORS OK
[1544534060] SERVICE ALERT: sig4;Total Processes;OK;HARD;1;PROCS OK: 305 processes
[1544534068] SERVICE ALERT: sig4;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1544534068] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 308 processes
[1544534668] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1544535043] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 324 processes
[1544535164] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 324 processes
[1544535284] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 321 processes
[1544536484] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1544536771] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544536772] HOST ALERT: sig4;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.56)
[1544536780] SERVICE ALERT: sig4;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig4' was out of bounds)
[1544536836] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1544536891] SERVICE ALERT: sig4;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37641 MB (73% inode=98%):
[1544536901] SERVICE ALERT: sig4;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.17, 0.06
[1544537501] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544537510] SERVICE ALERT: sig4;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544537531] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1544537621] SERVICE ALERT: sig4;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37603 MB (73% inode=98%):
[1544537624] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1544537630] SERVICE ALERT: sig4;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.13, 0.05
[1544539485] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1544539485] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 311 processes
[1544539803] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 37 users currently logged in
[1544540085] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1544542669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544542789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544542909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544543509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1544548350] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 207.42, 299.23, 136.46
[1544548470] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 28.90, 200.60, 120.06
[1544548590] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.80, 134.54, 105.61
[1544550390] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.98, 1.19, 15.89
[1544550709] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544550830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1544551485] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1544551590] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.65, 0.80, 4.98