[1552588244] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (149708 kB) free!
[1552588364] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169500 kB) free!
[1552588604] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552588724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552588844] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1552588965] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 54.5% (8863892 kB) free.
[1552589038] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1552589038] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.0% (13326920 kB) free.
[1552589444] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552589564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552589638] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5582572 kB) free!
[1552589685] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552589758] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.0% (13327880 kB) free.
[1552590284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552590359] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7804880 kB) free!
[1552590404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552590479] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6485060 kB) free!
[1552590524] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552590599] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13334696 kB) free.
[1552590765] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (1356940 kB) free!
[1552590885] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.5% (9188504 kB) free.
[1552591199] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8355084 kB) free!
[1552591319] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7086992 kB) free!
[1552591439] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.6% (5237944 kB) free!
[1552591724] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552592324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552592444] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552592639] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.0% (13329316 kB) free.
[1552593045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552593165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552593285] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552593839] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5512188 kB) free!
[1552593885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552593959] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13146504 kB) free.
[1552594005] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552594125] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552594593] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11754876 kB) free.
[1552595159] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.7% (5784344 kB) free!
[1552595193] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9445972 kB) free!
[1552595279] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13252608 kB) free.
[1552595313] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8926656 kB) free!
[1552595325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552595433] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1552595433] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.9% (8335404 kB) free!
[1552595879] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8215388 kB) free!
[1552595925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552595999] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6853660 kB) free!
[1552596034] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (274348 kB) free!
[1552596046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552596119] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5016040 kB) free!
[1552596166] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552596766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552596886] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552597486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552597606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552597726] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552597919] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10379112 kB) free.
[1552598326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552598447] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552598567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552599167] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552599719] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8195344 kB) free!
[1552599767] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552599840] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7142340 kB) free!
[1552599887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552599960] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5605924 kB) free!
[1552600007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552601160] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13212928 kB) free.
[1552601207] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552601760] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6949808 kB) free!
[1552601808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552601880] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5189680 kB) free!
[1552601928] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552602000] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12701004 kB) free.
[1552602528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552602600] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5158804 kB) free!
[1552602648] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552602720] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12923892 kB) free.
[1552603248] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552603368] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552603920] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8657060 kB) free!
[1552603968] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552604040] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7360764 kB) free!
[1552604088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552604161] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5829944 kB) free!
[1552604209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552605361] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13168992 kB) free.
[1552605409] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552605914] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 33358 MB (15% inode=99%):
[1552605961] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5885100 kB) free!
[1552606009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552606081] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13176020 kB) free.
[1552606129] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552606681] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7911464 kB) free!
[1552606729] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552606801] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6616592 kB) free!
[1552606849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552606921] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13173968 kB) free.
[1552606969] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552607114] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21064 MB (10% inode=99%):
[1552607434] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1552607521] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8511372 kB) free!
[1552607569] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552607641] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7210880 kB) free!
[1552607690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552607714] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 31831 MB (15% inode=99%):
[1552607761] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5595292 kB) free!
[1552607810] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552608314] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1552608314] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21582 MB (10% inode=99%):
[1552608410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552608530] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552608634] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.6% (5234776 kB) free!
[1552608650] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552608914] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29840 MB (14% inode=99%):
[1552608961] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13168500 kB) free.
[1552609234] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226484 kB) free!
[1552609561] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5424916 kB) free!
[1552609681] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13177772 kB) free.
[1552609850] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552610281] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7715212 kB) free!
[1552610401] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6412864 kB) free!
[1552610450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552610521] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13168940 kB) free.
[1552610570] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552611121] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8259040 kB) free!
[1552611171] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552611241] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6965328 kB) free!
[1552611291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552611361] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5176884 kB) free!
[1552611411] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552612011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552612131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552612251] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552612514] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21730 MB (10% inode=99%):
[1552612561] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13164384 kB) free.
[1552613114] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25880 MB (12% inode=99%):
[1552613162] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5076212 kB) free!
[1552613282] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13173372 kB) free.
[1552613451] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552613882] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7473476 kB) free!
[1552614002] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6055040 kB) free!
[1552614052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552614122] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13164944 kB) free.
[1552614172] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552614722] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8072508 kB) free!
[1552614772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552614842] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6772028 kB) free!
[1552614892] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552614962] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4843036 kB) free!
[1552615012] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552615562] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.5% (8625672 kB) free!
[1552615612] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552615732] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552615852] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552616162] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13162708 kB) free.
[1552616714] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22875 MB (10% inode=99%):
[1552617052] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1552617314] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26055 MB (12% inode=99%):
[1552617363] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8552040 kB) free!
[1552617483] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7268964 kB) free!
[1552617603] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.6% (5709768 kB) free!
[1552617652] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1552617772] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552618372] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552618492] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552618612] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552618803] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9957576 kB) free.
[1552619812] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552619893] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4755384 kB) free!
[1552619932] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552620013] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4752484 kB) free!
[1552620133] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4749560 kB) free!
[1552620532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552620603] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7040684 kB) free!
[1552620652] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552620718] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 558 processes
[1552620723] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5240656 kB) free!
[1552620772] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552620778] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 557 processes
[1552620843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12741792 kB) free.
[1552621372] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552621492] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552621613] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552622043] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.7% (5259568 kB) free!
[1552622164] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12834520 kB) free.
[1552622763] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5002272 kB) free!
[1552622813] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552622883] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12488436 kB) free.
[1552623413] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552623533] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552624083] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5666292 kB) free!
[1552624133] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552624203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13097764 kB) free.
[1552624253] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552624373] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552624803] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4452064 kB) free!
[1552624923] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12512024 kB) free.
[1552626123] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4555244 kB) free!
[1552626173] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552626243] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12677116 kB) free.
[1552626314] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21581 MB (10% inode=99%):
[1552626773] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552626893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552627013] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552627443] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4295388 kB) free!
[1552627563] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12529060 kB) free.
[1552628114] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 30792 MB (14% inode=99%):
[1552628533] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198908 kB) free!
[1552628635] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192176 kB) free!
[1552628755] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196440 kB) free!
[1552628763] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4227016 kB) free!
[1552628813] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552628875] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166272 kB) free!
[1552628883] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12349528 kB) free.
[1552629414] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552629534] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552629654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552630083] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4147388 kB) free!
[1552630203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12245856 kB) free.
[1552630803] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552631403] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4206728 kB) free!
[1552631454] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552631523] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12190252 kB) free.
[1552632054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552632174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552632295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552632723] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4182484 kB) free!
[1552632843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12642248 kB) free.
[1552633443] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4545292 kB) free!
[1552633514] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21802 MB (10% inode=99%):
[1552633563] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12581608 kB) free.
[1552634163] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4135268 kB) free!
[1552634283] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11500108 kB) free.
[1552635483] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4115888 kB) free!
[1552635603] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12140464 kB) free.
[1552636804] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4031144 kB) free!
[1552636924] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12191232 kB) free.
[1552638124] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4224824 kB) free!
[1552638244] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12363760 kB) free.
[1552638352] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552638472] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1552638592] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1552638844] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (3281960 kB) free!
[1552638895] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552638964] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11292436 kB) free.
[1552639495] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552639615] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552639735] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552640114] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1552640392] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1552640764] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4809544 kB) free!
[1552640791] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4721812 kB) free!
[1552640884] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12251504 kB) free.
[1552640911] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4651804 kB) free!
[1552641003] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 26.0% (4218640 kB) free.
[1552641031] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4455864 kB) free!
[1552641182] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 28.08, 31.77, 14.36
[1552641302] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.30, 21.44, 12.68
[1552641314] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25022 MB (11% inode=99%):
[1552641413] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51355 MB (22% inode=99%):
[1552641422] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.43, 14.64, 11.25
[1552641485] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3582116 kB) free!
[1552641536] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552641586] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.36, 10.10, 10.96
[1552641605] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11821572 kB) free.
[1552641706] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.46, 7.32, 9.84
[1552641826] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.55, 5.42, 8.84
[1552642136] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552642256] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552642376] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552642622] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.48, 0.75, 3.47
[1552642805] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3906568 kB) free!
[1552642899] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 18.7% (1374028 kB) free!
[1552642925] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12116048 kB) free.
[1552643019] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 18.3% (1345424 kB) free!
[1552643026] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.56, 1.57, 3.48
[1552643114] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22450 MB (10% inode=99%):
[1552643139] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 17.9% (1317756 kB) free!
[1552643453] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 20.0% (1471232 kB) free!
[1552643525] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2914592 kB) free!
[1552643573] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.7% (1451880 kB) free!
[1552643576] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552643645] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11302660 kB) free.
[1552643693] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 19.5% (1432776 kB) free!
[1552644176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552644296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552644417] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552644845] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3274764 kB) free!
[1552644965] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11400500 kB) free.
[1552645333] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.00, 5.75, 5.08
[1552645333] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.9% (22271500 kB) free.
[1552645453] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.28, 5.68, 5.14
[1552645573] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.99, 5.83, 5.26
[1552645803] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (3009172 kB) free!
[1552645923] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.5% (2848980 kB) free!
[1552646043] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.6% (2694668 kB) free!
[1552646139] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (643560 kB) free!
[1552646166] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4151356 kB) free!
[1552646258] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552646286] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12001720 kB) free.
[1552646378] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1552646498] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1552646693] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (382608 kB) free!
[1552646886] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3957344 kB) free!
[1552647006] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11958796 kB) free.
[1552647044] SERVICE ALERT: sig2-vm1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552647164] SERVICE ALERT: sig2-vm1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1552647243] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (782716 kB) free!
[1552647285] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2830292 kB) free!
[1552647387] HOST ALERT: sig5-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.192)
[1552647405] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.6% (2695112 kB) free!
[1552647451] HOST ALERT: sig5-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1552647525] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.8% (2569288 kB) free!
[1552647606] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3366544 kB) free!
[1552647686] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.58, 16.04, 7.16
[1552647726] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11735336 kB) free.
[1552647806] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.99, 10.81, 6.32
[1552647893] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 45.2% (3327408 kB) free.
[1552647926] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 7.28, 5.57
[1552647973] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.44, 4.29, 4.65
[1552648022] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 24.74, 27.05, 11.93
[1552648142] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.47, 18.14, 10.50
[1552648236] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 332 processes
[1552648262] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.51, 12.17, 9.24
[1552648326] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3476016 kB) free!
[1552648426] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.49, 15.84, 13.90
[1552648446] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11682964 kB) free.
[1552648526] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.20, 1.22, 3.07
[1552648546] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.41, 11.08, 12.39
[1552648666] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.72, 7.98, 11.10
[1552648836] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 334 processes
[1552649462] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1552649462] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.41, 0.57, 2.81
[1552649646] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4130976 kB) free!
[1552649766] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698916 kB) free.
[1552649866] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1552649866] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.38, 1.52, 4.07
[1552650367] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3750688 kB) free!
[1552650417] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552650487] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11500504 kB) free.
[1552650939] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 95.9% (7050904 kB) free.
[1552651017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552651087] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3768060 kB) free!
[1552651137] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552651207] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11499744 kB) free.
[1552651298] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1552651444] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1552651444] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.6% (14882448 kB) free.
[1552651737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552651857] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1552652325] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1398548 kB) free!
[1552652407] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230164 kB) free!
[1552652457] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1552652498] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1552652527] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239424 kB) free!
[1552652577] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552652647] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13148004 kB) free.
[1552652698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552652925] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.7% (14259736 kB) free.
[1552653248] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4703748 kB) free!
[1552653298] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1552653368] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12839148 kB) free.
[1552653373] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.06, 5.29, 5.05
[1552653493] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.91, 5.18, 5.04
[1552653613] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.60, 4.87, 4.94
[1552653898] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1552653968] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4220248 kB) free!
[1552654018] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1552654088] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.6% (12598020 kB) free.
[1552654618] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552654738] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1552655288] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6216432 kB) free!
[1552655338] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552655408] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13228628 kB) free.
[1552655458] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552655579] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552656008] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5846428 kB) free!
[1552656128] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13214776 kB) free.
[1552656728] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5466976 kB) free!
[1552656779] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552656848] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13042172 kB) free.
[1552657379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552657449] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4913004 kB) free!
[1552657499] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552657569] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13199272 kB) free.
[1552658099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552658169] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5375180 kB) free!
[1552658219] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552658289] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13220904 kB) free.
[1552658819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552658889] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5980916 kB) free!
[1552658939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552659009] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13198056 kB) free.
[1552659060] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552659244] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1552659609] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8070492 kB) free!
[1552659659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552659729] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6888292 kB) free!
[1552659779] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552659849] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5096496 kB) free!
[1552659899] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552660263] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552660499] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552660619] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552660666] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552660739] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552661049] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13201988 kB) free.
[1552661649] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5862692 kB) free!
[1552661717] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 310 processes
[1552661770] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13216780 kB) free.
[1552661837] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 309 processes
[1552661939] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552661957] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552662369] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8179924 kB) free!
[1552662489] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (6981156 kB) free!
[1552662540] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552662609] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5387640 kB) free!
[1552662660] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552663260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552663380] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552663809] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1552663809] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10629564 kB) free.
[1552663980] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552664049] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.53, 5.60, 5.08
[1552664100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552664169] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.05, 5.74, 5.19
[1552664220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552664289] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.28, 5.53, 5.18
[1552665009] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7393244 kB) free!
[1552665129] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6107624 kB) free!
[1552665249] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13213676 kB) free.
[1552665421] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552665849] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8132260 kB) free!
[1552665914] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29744 MB (14% inode=99%):
[1552665970] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6928588 kB) free!
[1552666021] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552666090] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.6% (5247628 kB) free!
[1552666141] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552666741] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552666861] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552666981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552667290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13207820 kB) free.
[1552667714] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22478 MB (10% inode=99%):
[1552667890] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5942232 kB) free!
[1552668010] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13183080 kB) free.
[1552668610] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5121128 kB) free!
[1552668730] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12740024 kB) free.
[1552669381] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552669689] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.77, 4.41, 4.75
[1552669930] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8520072 kB) free!
[1552669981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552670050] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7316344 kB) free!
[1552670101] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552670171] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5973072 kB) free!
[1552670221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552671371] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13182496 kB) free.
[1552671421] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552671971] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6555796 kB) free!
[1552672021] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552672091] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13200364 kB) free.
[1552672141] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552672691] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8561092 kB) free!
[1552672742] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552672811] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7329612 kB) free!
[1552672862] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552672931] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5979476 kB) free!
[1552672982] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552674131] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13195872 kB) free.
[1552674182] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes