[1567536283] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3370304 kB) free!
[1567536403] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3369944 kB) free!
[1567536523] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.7% (3369292 kB) free!
[1567538402] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.1% (10146464 kB) free.
[1567539348] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 293 processes
[1567539458] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (3002060 kB) free!
[1567539523] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1981988 kB) free!
[1567539574] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1213444 kB) free!
[1567539578] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (251256 kB) free!
[1567539602] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.4% (3056212 kB) free!
[1567539694] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386832 kB) free!
[1567539698] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (243364 kB) free!
[1567539722] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (2272972 kB) free!
[1567539814] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1068428 kB) free!
[1567539842] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1770136 kB) free!
[1567540548] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 289 processes
[1567543096] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 90.33, 35.02, 13.35
[1567543216] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 102.71, 58.02, 24.40
[1567543336] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 104.24, 73.36, 34.10
[1567543542] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 54.61, 63.44, 32.59
[1567543662] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.79, 42.64, 28.71
[1567543723] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.9% (2662732 kB) free!
[1567543782] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.34, 28.64, 25.27
[1567543825] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2672784 kB) free!
[1567544204] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236424 kB) free!
[1567544324] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225188 kB) free!
[1567544382] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.43, 4.38, 13.55
[1567544444] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227128 kB) free!
[1567544536] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.84, 3.08, 13.99
[1567545044] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.0% (43865024 kB) free.
[1567545582] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.71, 0.78, 4.19
[1567545736] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.18, 1.11, 4.57
[1567546225] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1255740 kB) free!
[1567546723] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1567546723] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 39.4% (9645956 kB) free.
[1567546825] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.8% (10004792 kB) free.
[1567555723] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1567557098] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.9% (33573396 kB) free.
[1567559614] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 52.8% (69664488 kB) free.
[1567563345] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 353 processes
[1567563465] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 349 processes
[1567563523] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.2% (3224480 kB) free!
[1567563585] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 349 processes
[1567563626] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.0% (2701584 kB) free!
[1567563643] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.1% (3210988 kB) free!
[1567563746] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.0% (2687696 kB) free!
[1567563763] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (3196276 kB) free!
[1567563866] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2672696 kB) free!
[1567564185] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1567565666] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2433196 kB) free!
[1567566824] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1567566945] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567568563] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2346308 kB) free!
[1567572944] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567573064] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567575097] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1567575468] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.51, 5.67, 5.05
[1567575588] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.93, 5.49, 5.06
[1567575708] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.86, 5.00, 4.94
[1567576297] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1567576417] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1567577987] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 313 processes
[1567578107] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1567578143] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 93.8% (22982476 kB) free.
[1567578183] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.4% (22407148 kB) free.
[1567578225] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.9% (22914956 kB) free.
[1567578302] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.3% (39590540 kB) free.
[1567578349] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.2% (23337448 kB) free.
[1567578697] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1567578707] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1567578826] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1567578827] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1567578836] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196300 kB) free!
[1567578907] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (226132 kB) free!
[1567578946] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1567578947] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1567578956] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197184 kB) free!
[1567579027] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225232 kB) free!
[1567579066] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1567579077] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189448 kB) free!
[1567579148] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217488 kB) free!
[1567579893] SERVICE ALERT: webservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: /home 77608 MB (34% inode=99%):
[1567580508] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.57, 5.47, 5.24
[1567580628] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.37, 5.75, 5.37
[1567580748] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.06, 6.07, 5.52
[1567581948] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.85, 3.49, 4.35
[1567583017] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 317 processes
[1567583137] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1567586738] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1567586858] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1567587044] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9832976 kB) free!
[1567587165] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9564664 kB) free!
[1567587285] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9296840 kB) free!
[1567588415] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 28711 MB (20% inode=99%):
[1567588535] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 28710 MB (20% inode=99%):
[1567588563] SERVICE ALERT: pservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 566 processes
[1567588655] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 28710 MB (20% inode=99%):
[1567588683] SERVICE ALERT: pservercluster2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 561 processes
[1567588803] SERVICE ALERT: pservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 560 processes
[1567589404] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 355 processes
[1567589855] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45064 MB (32% inode=99%):
[1567590285] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4838812 kB) free!
[1567590458] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 320 processes
[1567590578] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1567591018] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4645612 kB) free!
[1567591065] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 570 processes
[1567591078] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 12.4% (3050284 kB) free!
[1567591125] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 579 processes
[1567591138] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 5.9% (1444092 kB) free!
[1567591185] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 569 processes
[1567591198] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 1.5% (358156 kB) free!
[1567591662] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.8% (22152252 kB) free.
[1567591667] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1567591676] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.8% (40795464 kB) free.
[1567591692] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.8% (47202684 kB) free.
[1567592747] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 309 processes
[1567594185] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 570 processes
[1567594245] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 573 processes
[1567594305] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 570 processes
[1567594365] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 0 warn out of 566 processes
[1567594665] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 565 processes
[1567594666] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187604 kB) free!
[1567594786] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (285972 kB) free!
[1567594906] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181212 kB) free!
[1567595283] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8929300 kB) free!
[1567595403] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8333748 kB) free!
[1567595523] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (7687156 kB) free!
[1567596123] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4259788 kB) free!
[1567597983] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4896012 kB) free!
[1567598103] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4801124 kB) free!
[1567598223] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (4722716 kB) free!
[1567598257] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 357 processes
[1567598378] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 325 processes
[1567598498] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1567598754] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4461920 kB) free!
[1567598857] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 353 processes
[1567598874] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (4349856 kB) free!
[1567598994] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4203920 kB) free!
[1567600297] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1567601223] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2191436 kB) free!
[1567601394] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2431560 kB) free!
[1567603897] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1567604913] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 21.9% (5374948 kB) free.
[1567605575] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1567605695] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1567610264] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567610385] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567612692] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9670608 kB) free!
[1567612812] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9587568 kB) free!
[1567612932] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.3% (9517472 kB) free!
[1567615184] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567615304] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes