[1581532533] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.2% (29197404 kB) free.
[1581532596] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581532716] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 366 processes
[1581532741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14111964 kB) free!
[1581532861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36105524 kB) free.
[1581533028] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195556 kB) free!
[1581533053] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198424 kB) free!
[1581533068] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181196 kB) free!
[1581533148] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181252 kB) free!
[1581533174] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197784 kB) free!
[1581533188] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168248 kB) free!
[1581533268] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174740 kB) free!
[1581533294] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204120 kB) free!
[1581533308] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184732 kB) free!
[1581533461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15213180 kB) free!
[1581533581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36807088 kB) free.
[1581534181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15805592 kB) free!
[1581534301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36804404 kB) free.
[1581534508] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2518268 kB) free!
[1581534901] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16089168 kB) free!
[1581535021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36772916 kB) free.
[1581535108] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1581535108] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196236 kB) free!
[1581535621] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16339040 kB) free!
[1581535709] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.9% (10594188 kB) free.
[1581535741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36763756 kB) free.
[1581536309] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.4% (2814016 kB) free!
[1581536341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15615184 kB) free!
[1581536429] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (1258196 kB) free!
[1581536461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35260228 kB) free.
[1581536549] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183460 kB) free!
[1581537188] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.2% (5534332 kB) free!
[1581537308] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4464908 kB) free!
[1581537428] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3557588 kB) free!
[1581537661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17432544 kB) free!
[1581537781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36660840 kB) free.
[1581538381] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16242484 kB) free!
[1581538501] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36422900 kB) free.
[1581539701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16591920 kB) free!
[1581539821] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35310628 kB) free.
[1581539868] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 52.8% (12951904 kB) free.
[1581540769] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (332760 kB) free!
[1581540889] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214496 kB) free!
[1581540912] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3990340 kB) free!
[1581541009] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223872 kB) free!
[1581541021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17239764 kB) free!
[1581541032] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (3993228 kB) free!
[1581541141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35286568 kB) free.
[1581541152] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 79.5% (39155608 kB) free.
[1581542209] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.8% (41769856 kB) free.
[1581543033] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(202) < 103.16.101.72 [13/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581543154] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581543348] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(249) < 2020/02/13 03:01:32 [error] 19834#0: *52944609 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581543360] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1905) < 172.105.59.173 [13/Feb/2020:03:02:22 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581543468] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581543481] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581543541] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16821460 kB) free!
[1581543662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33455708 kB) free.
[1581544460] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.2% (13508340 kB) free!
[1581547080] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 172.105.59.173 [13/Feb/2020:04:01:56 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581547200] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581547949] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581550669] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192840 kB) free!
[1581550789] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171968 kB) free!
[1581550909] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184340 kB) free!
[1581551294] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.9% (20571688 kB) free.
[1581551549] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.6% (11260412 kB) free.
[1581554060] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12615752 kB) free!
[1581555261] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 31.6% (41614032 kB) free.
[1581556864] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.5% (39664572 kB) free.
[1581556888] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.4% (21910708 kB) free.
[1581558149] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (217160 kB) free!
[1581558269] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201576 kB) free!
[1581558389] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192600 kB) free!
[1581560956] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1581561076] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581561468] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2020/02/13 08:00:56 [error] 19836#0: *53006493 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581561588] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581561600] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 49.50.86.152 [13/Feb/2020:08:01:57 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1581561720] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581562316] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21939944 kB) free.
[1581568159] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(10) < webmachine error: path="/api/whoami"
[1581568279] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1581570953] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.3% (4485484 kB) free!
[1581571073] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.6% (4303912 kB) free!
[1581571193] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.7% (4095936 kB) free!
[1581572012] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 708 processes
[1581572133] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1581572253] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 705 processes
[1581572393] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2397344 kB) free!
[1581572526] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4390076 kB) free!
[1581572646] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4810060 kB) free!
[1581572766] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4594172 kB) free!
[1581573516] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3647588 kB) free!
[1581573636] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3338764 kB) free!
[1581573756] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.6% (3046188 kB) free!
[1581573966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181396 kB) free!
[1581574053] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1581574356] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1373540 kB) free!
[1581575133] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9600544 kB) free!
[1581575253] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.3% (9024672 kB) free!
[1581575373] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.1% (8423320 kB) free!
[1581575467] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.7% (6248140 kB) free!
[1581575556] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.3% (15570344 kB) free.
[1581575587] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (2608672 kB) free!
[1581575707] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246584 kB) free!
[1581575974] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.0% (30549396 kB) free.
[1581575988] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(55) < 2020/02/13 12:01:03 [error] 19840#0: *53098281 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581576108] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581576120] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(66) < 172.105.59.173 [13/Feb/2020:12:03:12 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581576185] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 13.4% (3307272 kB) free!
[1581576240] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581576305] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 90.9% (22427572 kB) free.
[1581576386] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1581576506] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1581576626] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1581577227] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1581577332] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39267728 kB) free.
[1581577373] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.3% (21293296 kB) free.
[1581577395] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21640460 kB) free.
[1581579176] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184832 kB) free!
[1581579296] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172444 kB) free!
[1581579409] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9157344 kB) free!
[1581579416] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1581579416] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172360 kB) free!
[1581579494] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (4270696 kB) free!
[1581579529] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8848104 kB) free!
[1581579614] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4177816 kB) free!
[1581579650] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (8545624 kB) free!
[1581579734] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (4067884 kB) free!
[1581580861] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18431376 kB) free!
[1581580897] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4441304 kB) free!
[1581580981] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39330504 kB) free.
[1581581017] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4220732 kB) free!
[1581581137] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.3% (3993132 kB) free!
[1581581449] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4176188 kB) free!
[1581581581] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15986832 kB) free!
[1581581701] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38452320 kB) free.
[1581582134] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2228456 kB) free!
[1581582301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15325508 kB) free!
[1581582337] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2293504 kB) free!
[1581582421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34159180 kB) free.
[1581583021] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10433676 kB) free!
[1581583141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45675080 kB) free.
[1581583956] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3540196 kB) free!
[1581584076] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.5% (3261932 kB) free!
[1581584092] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (3134648 kB) free!
[1581584196] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.3% (2992596 kB) free!
[1581584212] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3137944 kB) free!
[1581584332] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (3136280 kB) free!
[1581584796] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1581584796] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1714892 kB) free!
[1581586174] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8619536 kB) free!
[1581586294] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.7% (29901956 kB) free.
[1581586440] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(178) < 92.42.104.166 [13/Feb/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581586560] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581586597] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.6% (15626060 kB) free.
[1581586653] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 708 processes
[1581586741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18307184 kB) free!
[1581586773] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 709 processes
[1581586861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42327248 kB) free.
[1581586893] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 712 processes
[1581586937] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9432404 kB) free!
[1581587057] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9419624 kB) free!
[1581587177] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 20.5% (10102704 kB) free.
[1581587461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19654888 kB) free!
[1581587581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43109200 kB) free.
[1581587777] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (1429880 kB) free!
[1581587898] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238232 kB) free!
[1581587952] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9595380 kB) free!
[1581588018] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246576 kB) free!
[1581588072] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9432596 kB) free!
[1581588181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20683356 kB) free!
[1581588192] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (9282268 kB) free!
[1581588301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (16756880 kB) free!
[1581588416] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 28.3% (6979432 kB) free.
[1581588421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40811260 kB) free.
[1581588618] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 22.6% (11132296 kB) free.
[1581589016] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175192 kB) free!
[1581589021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17103916 kB) free!
[1581589136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193780 kB) free!
[1581589141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40478024 kB) free.
[1581589218] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (218804 kB) free!
[1581589256] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169988 kB) free!
[1581589338] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226924 kB) free!
[1581589458] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228484 kB) free!
[1581589741] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11116324 kB) free!
[1581589861] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6758428 kB) free!
[1581589908] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(22) < 2020/02/13 16:01:02 [error] 19827#0: *53224623 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581589981] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2795416 kB) free!
[1581590028] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581590160] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [13/Feb/2020:16:01:55 +0530] TCP 502 0 0 63.009 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 3.008, 60.000"
[1581590209] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1581590281] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581590330] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581591181] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17810884 kB) free!
[1581591782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36202040 kB) free.
[1581591792] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4834332 kB) free!
[1581592382] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6005312 kB) free!
[1581592502] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2191620 kB) free!
[1581592622] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1581592622] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14759616 kB) free!
[1581592827] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1581592893] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1581592947] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1581592995] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.0% (4646724 kB) free!
[1581593068] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1581593115] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.6% (4563872 kB) free!
[1581593222] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (8150352 kB) free!
[1581593235] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.3% (4490020 kB) free!
[1581593493] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1581593613] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581593822] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23748620 kB) free!
[1581593929] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1581594049] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581594422] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44343504 kB) free.
[1581594997] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4177696 kB) free!
[1581595022] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19134460 kB) free!
[1581595117] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.3% (3893416 kB) free!
[1581595142] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15210320 kB) free!
[1581595238] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.9% (3560472 kB) free!
[1581595262] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (44055644 kB) free.
[1581595635] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1478020 kB) free!
[1581595838] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1663808 kB) free!
[1581595862] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17863904 kB) free!
[1581595898] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (2464424 kB) free!
[1581595982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (14028120 kB) free!
[1581596018] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (2744072 kB) free!
[1581596102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42161240 kB) free.
[1581596138] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (3103976 kB) free!
[1581596494] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9852652 kB) free!
[1581596614] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9507032 kB) free!
[1581596702] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18236884 kB) free!
[1581596734] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.6% (9179172 kB) free!
[1581596738] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3977176 kB) free.
[1581596823] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14183520 kB) free!
[1581596943] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (10309248 kB) free!
[1581597038] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.8% (15683276 kB) free.
[1581597338] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1855156 kB) free!
[1581597340] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.6% (29361424 kB) free.
[1581597458] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.6% (2054680 kB) free!
[1581597543] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17634084 kB) free!
[1581597578] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (2202760 kB) free!
[1581598143] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26621244 kB) free.
[1581598178] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1581598178] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3550324 kB) free.
[1581598694] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.1% (21241220 kB) free.
[1581598726] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.0% (42864944 kB) free.
[1581598759] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39300944 kB) free.
[1581598793] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.2% (43460892 kB) free.
[1581598841] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.7% (21995804 kB) free.
[1581598930] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.3% (22130820 kB) free.
[1581598967] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21641244 kB) free.
[1581599297] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188224 kB) free!
[1581599343] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11734544 kB) free!
[1581599365] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.0% (6390956 kB) free!
[1581599378] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (2685080 kB) free!
[1581599417] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184016 kB) free!
[1581599463] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8047760 kB) free!
[1581599485] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (1465360 kB) free!
[1581599498] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2905564 kB) free!
[1581599537] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183700 kB) free!
[1581599583] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36076052 kB) free.
[1581599605] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235020 kB) free!
[1581599618] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3139124 kB) free!
[1581600183] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11261344 kB) free!
[1581600218] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.7% (4335504 kB) free.
[1581600303] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7426532 kB) free!
[1581600423] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35963168 kB) free.
[1581600828] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(62) < 2020/02/13 18:55:44 [error] 19833#0: *53323338 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581600948] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581601023] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11240584 kB) free!
[1581601143] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7644144 kB) free!
[1581601263] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35594640 kB) free.
[1581601418] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.2% (2140476 kB) free!
[1581601538] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.0% (2268744 kB) free!
[1581601658] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2636928 kB) free!
[1581601863] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11570176 kB) free!
[1581601984] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7655056 kB) free!
[1581602104] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35918628 kB) free.
[1581602258] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3288252 kB) free.
[1581602704] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11798060 kB) free!
[1581602824] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8153524 kB) free!
[1581602944] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35397752 kB) free.
[1581603544] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12879692 kB) free!
[1581603664] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8895424 kB) free!
[1581603784] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35090372 kB) free.
[1581604384] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10082728 kB) free!
[1581604504] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5918216 kB) free!
[1581604549] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(106) < 2020/02/13 20:01:21 [error] 19832#0: *53348772 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581604624] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (13096600 kB) free!
[1581604669] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581604681] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 203.212.70.170 [13/Feb/2020:20:03:21 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581604801] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581605824] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20895104 kB) free!
[1581606424] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32893888 kB) free.
[1581607025] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10544528 kB) free!
[1581607145] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6836000 kB) free!
[1581607265] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32687292 kB) free.
[1581607844] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581607865] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9932632 kB) free!
[1581607985] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6191200 kB) free!
[1581608105] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2673292 kB) free!
[1581609305] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22631500 kB) free!
[1581609905] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31995088 kB) free.
[1581610337] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581610505] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9838716 kB) free!
[1581610625] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6247464 kB) free!
[1581610745] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31899336 kB) free.
[1581611345] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11745384 kB) free!
[1581611465] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7947536 kB) free!
[1581611585] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4076436 kB) free!
[1581612786] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19562564 kB) free!
[1581613386] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29599704 kB) free.
[1581613658] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581613986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6536252 kB) free!
[1581614106] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2794828 kB) free!
[1581614226] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29392036 kB) free.
[1581614826] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6741716 kB) free!
[1581614946] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2954840 kB) free!
[1581615066] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28308064 kB) free.
[1581615666] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5892036 kB) free!
[1581615786] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6611836 kB) free!
[1581615906] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27468756 kB) free.
[1581616506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5716728 kB) free!
[1581616626] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30565884 kB) free.
[1581617140] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6852784 kB) free!
[1581617226] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8770580 kB) free!
[1581617261] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (222376 kB) free!
[1581617346] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (5073344 kB) free!
[1581617381] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216088 kB) free!
[1581617467] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30545728 kB) free.
[1581618066] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7328708 kB) free!
[1581618186] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3644188 kB) free!
[1581618261] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1377808 kB) free!
[1581618306] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27350016 kB) free.
[1581618381] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2905048 kB) free!
[1581618501] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3014700 kB) free!