[1566671734] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 95.9% (47273716 kB) free.
[1566672813] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175912 kB) free!
[1566672933] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188740 kB) free!
[1566673053] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194928 kB) free!
[1566675033] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 368 processes
[1566675093] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1566675332] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 409 processes
[1566675453] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1566675453] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 49.4% (12202600 kB) free.
[1566677853] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 94.6% (124725100 kB) free.
[1566679654] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.4% (3300364 kB) free!
[1566679774] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164232 kB) free!
[1566679894] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165512 kB) free!
[1566680013] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4319996 kB) free!
[1566680133] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (1230572 kB) free!
[1566680253] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 94.7% (46653476 kB) free.
[1566682535] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199824 kB) free!
[1566682536] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166500 kB) free!
[1566682655] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 43.6% (10689440 kB) free.
[1566682656] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.7% (15856712 kB) free.
[1566683494] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1566683494] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22855716 kB) free.
[1566691894] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1566693452] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 409 processes
[1566693573] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 409 processes
[1566693694] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181684 kB) free!
[1566693814] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177084 kB) free!
[1566693934] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189856 kB) free!
[1566695134] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.9% (20945808 kB) free.
[1566698793] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 375 processes
[1566698853] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1566706655] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4909328 kB) free!
[1566706776] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4871312 kB) free!
[1566706896] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4852016 kB) free!
[1566710372] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566710492] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1566711934] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4843452 kB) free!
[1566712055] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4819032 kB) free!
[1566712175] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4806616 kB) free!
[1566715892] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566716012] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1566716496] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2417268 kB) free!
[1566717096] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (3101672 kB) free!
[1566717212] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566717332] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566720096] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2442028 kB) free!
[1566723332] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566723452] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1566723623] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1566723623] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.0% (22799576 kB) free.
[1566724175] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22950076 kB) free.
[1566726452] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1566726572] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566730831] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1566739536] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 299 processes
[1566739656] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1566743972] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566744093] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1566748656] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4892100 kB) free!
[1566748777] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4883608 kB) free!
[1566748897] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4874616 kB) free!