[1584470020] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19760988 kB) free!
[1584470225] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (403100 kB) free!
[1584470269] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202396 kB) free!
[1584470345] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192136 kB) free!
[1584470389] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191016 kB) free!
[1584470465] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196164 kB) free!
[1584470509] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183540 kB) free!
[1584470621] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (11134720 kB) free!
[1584471221] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17285364 kB) free!
[1584471656] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.9% (3399568 kB) free!
[1584471716] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (1784556 kB) free!
[1584471776] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (176196 kB) free!
[1584471821] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34400616 kB) free.
[1584471836] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (180992 kB) free!
[1584472421] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13502440 kB) free!
[1584472541] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9772760 kB) free!
[1584472661] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6250572 kB) free!
[1584472736] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 72.1% (17659840 kB) free.
[1584473261] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20246756 kB) free!
[1584473786] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.6% (10264144 kB) free.
[1584473861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35242168 kB) free.
[1584474386] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (2077980 kB) free!
[1584474461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14708356 kB) free!
[1584474506] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (508308 kB) free!
[1584474581] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11068936 kB) free!
[1584474626] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199816 kB) free!
[1584474701] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7523916 kB) free!
[1584475138] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (210968 kB) free!
[1584475258] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236156 kB) free!
[1584475301] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20031528 kB) free!
[1584475378] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228416 kB) free!
[1584475901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33807340 kB) free.
[1584476341] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.2% (18725192 kB) free!
[1584476461] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.8% (18230204 kB) free!
[1584476501] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14142972 kB) free!
[1584476581] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.2% (18728628 kB) free!
[1584476621] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10649640 kB) free!
[1584476702] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1584476741] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (7313200 kB) free!
[1584476822] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1584477178] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.0% (27605004 kB) free.
[1584477342] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18677864 kB) free!
[1584477579] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(201) < 103.13.97.241 [18/Mar/2020:02:02:59 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1584477700] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584477942] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30958020 kB) free.
[1584478542] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12519100 kB) free!
[1584478662] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9220536 kB) free!
[1584478782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36499988 kB) free.
[1584478868] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9755052 kB) free!
[1584478988] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 12.8% (6303820 kB) free!
[1584479108] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2029584 kB) free!
[1584479382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16154628 kB) free!
[1584479502] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12729464 kB) free!
[1584479622] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9458832 kB) free!
[1584480065] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 48.8% (11967028 kB) free.
[1584480222] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19891128 kB) free!
[1584480308] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.3% (44499652 kB) free.
[1584480822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30311984 kB) free.
[1584481422] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12241668 kB) free!
[1584481542] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37409260 kB) free.
[1584481899] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(191) < 92.42.104.166 [18/Mar/2020:03:15:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584482019] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584482142] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18569540 kB) free!
[1584482262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15120136 kB) free!
[1584482382] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11836988 kB) free!
[1584482982] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20097328 kB) free!
[1584483583] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29709580 kB) free.
[1584484183] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12255632 kB) free!
[1584484303] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37305248 kB) free.
[1584484903] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17714608 kB) free!
[1584485023] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14421096 kB) free!
[1584485143] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37359348 kB) free.
[1584485743] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20409360 kB) free!
[1584485863] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17135296 kB) free!
[1584485983] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13962180 kB) free!
[1584487183] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26605080 kB) free.
[1584487865] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197136 kB) free!
[1584487985] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168428 kB) free!
[1584488106] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171456 kB) free!
[1584488383] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14301208 kB) free!
[1584488503] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10275780 kB) free!
[1584488623] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37073320 kB) free.
[1584489224] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16294676 kB) free!
[1584489344] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12108072 kB) free!
[1584489464] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (8017124 kB) free!
[1584489626] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.8% (11292112 kB) free.
[1584489709] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.9% (18609276 kB) free.
[1584490064] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17745528 kB) free!
[1584490664] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27827104 kB) free.
[1584491864] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17842384 kB) free!
[1584491984] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (14116324 kB) free!
[1584492104] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38823240 kB) free.
[1584492704] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19633496 kB) free!
[1584492781] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (837616 kB) free!
[1584492824] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16052288 kB) free!
[1584492944] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38522636 kB) free.
[1584493544] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18696504 kB) free!
[1584493664] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15211840 kB) free!
[1584493784] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38072608 kB) free.
[1584493981] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.9% (40770040 kB) free.
[1584494384] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18404324 kB) free!
[1584494505] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14641968 kB) free!
[1584494625] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37376808 kB) free.
[1584495225] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17538932 kB) free!
[1584495345] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37724276 kB) free.
[1584495536] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.4% (4272652 kB) free!
[1584495596] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 17.4% (4273732 kB) free!
[1584495656] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 17.4% (4271068 kB) free!
[1584495717] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.4% (4270400 kB) free!
[1584495945] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20798760 kB) free!
[1584496065] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17308876 kB) free!
[1584496185] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37559596 kB) free.
[1584496785] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19781548 kB) free!
[1584496905] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16146416 kB) free!
[1584497025] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37357936 kB) free.
[1584497625] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18566124 kB) free!
[1584497745] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14576592 kB) free!
[1584497865] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36937016 kB) free.
[1584498465] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18253612 kB) free!
[1584498585] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14753956 kB) free!
[1584498705] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36518284 kB) free.
[1584499306] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15499392 kB) free!
[1584499426] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (11515660 kB) free!
[1584499546] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35797560 kB) free.
[1584500146] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16332660 kB) free!
[1584500266] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12536332 kB) free!
[1584500386] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35258732 kB) free.
[1584500986] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15187316 kB) free!
[1584501106] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (11414336 kB) free!
[1584501226] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33097896 kB) free.
[1584501826] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14306648 kB) free!
[1584501946] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10788732 kB) free!
[1584502066] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27314492 kB) free.
[1584502666] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6620820 kB) free!
[1584502786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29723296 kB) free.
[1584502826] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (3868132 kB) free!
[1584502946] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.4% (3313156 kB) free!
[1584503066] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (3542544 kB) free!
[1584503386] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10350188 kB) free!
[1584503506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6649516 kB) free!
[1584503561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.9% (55206204 kB) free.
[1584503643] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39485284 kB) free.
[1584503706] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 80.4% (19698568 kB) free.
[1584504266] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194156 kB) free!
[1584504470] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.6% (20376056 kB) free.
[1584504767] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584505967] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24287572 kB) free!
[1584505994] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1132.70, 973.07, 520.11
[1584506087] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19870512 kB) free!
[1584506114] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1132.10, 1045.85, 602.77
[1584506207] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15622308 kB) free!
[1584506235] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1189.47, 1095.10, 674.55
[1584506807] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12342364 kB) free!
[1584506875] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1723068 kB) free!
[1584506918] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.63, 6.85, 5.41
[1584506923] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3237968 kB) free!
[1584506995] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (1609176 kB) free!
[1584507038] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.66, 6.59, 5.50
[1584507043] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (3239836 kB) free!
[1584507116] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1584507116] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1508860 kB) free!
[1584507158] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.04, 7.14, 5.83
[1584507163] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (3243856 kB) free!
[1584507273] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/18 10:14:34 [error] 12257#0: *72612260 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1584507394] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584507407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33143768 kB) free.
[1584507606] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.9% (4386276 kB) free!
[1584507713] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4099.32, 1569.25, 576.39
[1584507726] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.3% (4237384 kB) free!
[1584507772] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4528.41, 2157.86, 841.59
[1584507825] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584507833] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1388.10, 1628.27, 737.05
[1584507847] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.2% (3979280 kB) free!
[1584507892] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 612.70, 1444.59, 739.81
[1584507945] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584507953] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 189.60, 1090.46, 648.03
[1584508008] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9731648 kB) free!
[1584508011] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1179.42, 3487.24, 1973.87
[1584508012] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 84.98, 967.47, 650.42
[1584508069] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1584508128] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5439864 kB) free!
[1584508131] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 160.82, 2334.11, 1734.80
[1584508189] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1584508248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36579828 kB) free.
[1584508251] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.58, 1562.45, 1524.68
[1584508309] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1584508358] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.04, 3.17, 4.30
[1584508381] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25629184 kB) free!
[1584508419] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 172.105.59.173 [18/Mar/2020:10:35:03 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 0.001, 0.000"
[1584508447] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1964992 kB) free!
[1584508501] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (28552696 kB) free.
[1584508540] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584508545] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(2) < {writer,send_failed,{error,timeout}}
[1584508553] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.69, 148.35, 340.67
[1584508612] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.27, 131.82, 341.94
[1584508665] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584508848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11023956 kB) free!
[1584508968] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6955560 kB) free!
[1584509088] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1584509088] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3086516 kB) free!
[1584509451] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.73, 30.69, 420.47
[1584509688] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15011336 kB) free!
[1584509701] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.6% (23227668 kB) free!
[1584509737] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21621788 kB) free.
[1584509799] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.1% (20750596 kB) free.
[1584509821] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26508768 kB) free.
[1584510108] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4494036 kB) free!
[1584510228] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4268236 kB) free!
[1584510288] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29963244 kB) free.
[1584510348] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.7% (3976884 kB) free!
[1584510421] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (21218108 kB) free!
[1584510541] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.9% (26167824 kB) free!
[1584510661] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (28814092 kB) free.
[1584510888] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6537444 kB) free!
[1584511008] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2178708 kB) free!
[1584511128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33590636 kB) free.
[1584511448] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.7% (6765068 kB) free!
[1584511548] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1289020 kB) free!
[1584511569] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.7% (6764156 kB) free!
[1584511635] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 26.56, 23.81, 37.83
[1584511689] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (6762720 kB) free!
[1584511728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9240812 kB) free!
[1584511848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (5042568 kB) free!
[1584511861] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.5% (21770968 kB) free!
[1584511968] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1010052 kB) free!
[1584511982] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.7% (22057812 kB) free!
[1584512102] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.1% (27772700 kB) free.
[1584512227] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 45.9% (60524936 kB) free.
[1584512277] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.2% (15094268 kB) free.
[1584512429] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 13.2% (3269252 kB) free!
[1584512509] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (4034400 kB) free!
[1584512549] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 11.4% (2808388 kB) free!
[1584512629] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.7% (3838820 kB) free!
[1584512669] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2352852 kB) free!
[1584512702] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.6% (24564140 kB) free!
[1584512749] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (3605260 kB) free!
[1584512822] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.8% (26125864 kB) free!
[1584512942] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.4% (21676924 kB) free!
[1584513349] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2307420 kB) free!
[1584513403] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.8% (2922764 kB) free!
[1584513435] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 19.22, 20.77, 25.52
[1584513523] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (1580340 kB) free!
[1584513644] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191884 kB) free!
[1584513763] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1059548 kB) free!
[1584514039] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7079392 kB) free!
[1584514089] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2499104 kB) free!
[1584514159] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3028288 kB) free!
[1584514175] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.6% (14079160 kB) free.
[1584514237] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.8% (20551248 kB) free.
[1584514279] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17610044 kB) free!
[1584514297] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1584514297] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39469504 kB) free.
[1584514349] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.7% (20412528 kB) free.
[1584514469] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.0% (22718204 kB) free.
[1584514879] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385240 kB) free!
[1584515235] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 33.57, 34.82, 31.14
[1584515356] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 33.58, 33.97, 31.25
[1584515476] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1584515476] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 23.39, 30.47, 30.34
[1584515554] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180672 kB) free!
[1584515674] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177840 kB) free!
[1584515713] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178844 kB) free!
[1584515794] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185248 kB) free!
[1584515815] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584515818] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584515819] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584515822] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584515833] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197932 kB) free!
[1584515937] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584515938] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 173691 MB (77% inode=99%):
[1584515939] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584515942] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2477 processes
[1584515953] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186944 kB) free!
[1584515979] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3242164 kB) free!
[1584516064] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584516081] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 28.05, 29.02, 29.41
[1584516089] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584516099] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (3248292 kB) free!
[1584516107] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8650532 kB) free!
[1584516188] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584516192] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584516209] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584516219] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1584516219] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (3243640 kB) free!
[1584516227] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8651036 kB) free!
[1584516318] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584516329] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1584516346] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168924 kB) free!
[1584516347] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (8649124 kB) free!
[1584516422] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1584516448] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584516466] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187740 kB) free!
[1584516541] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584516542] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1584516545] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584516548] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584516549] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584516552] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584516587] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179008 kB) free!
[1584516646] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195072 kB) free!
[1584516661] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2486 processes
[1584516666] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584516668] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 177487 MB (78% inode=99%):
[1584516670] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584516673] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2484 processes
[1584516679] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14823056 kB) free!
[1584516683] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 63.94, 53.24, 39.10
[1584516766] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195356 kB) free!
[1584516803] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 29.57, 43.80, 37.36
[1584516887] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1584516887] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190856 kB) free!
[1584516908] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8830416 kB) free!
[1584516923] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 21.89, 36.84, 35.63
[1584516929] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1584517028] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8541776 kB) free!
[1584517142] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1584517142] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (10591584 kB) free!
[1584517148] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8284908 kB) free!
[1584517262] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1584517279] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391328 kB) free!
[1584517419] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1325508 kB) free!
[1584517452] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 84.4% (111298736 kB) free.
[1584517523] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 24.83, 24.34, 29.07
[1584517570] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.2% (30169436 kB) free.
[1584517612] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.2% (20292556 kB) free.
[1584517647] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1584517647] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21652004 kB) free.
[1584517691] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.2% (20877100 kB) free.
[1584517745] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.9% (55173788 kB) free.
[1584517798] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.7% (20505988 kB) free.
[1584517838] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.5% (14056604 kB) free.
[1584518139] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(88) < 103.255.146.164 [18/Mar/2020:13:25:39 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584518259] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(35) < 103.255.146.164 [18/Mar/2020:13:25:42 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584518379] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584518747] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (1542840 kB) free!
[1584518948] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4614128 kB) free!
[1584519257] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (45301984 kB) free.
[1584519314] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.4% (38644580 kB) free.
[1584519418] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (369944 kB) free!
[1584519538] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190292 kB) free!
[1584519658] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169948 kB) free!
[1584520112] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.1% (20514056 kB) free.
[1584520148] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15483364 kB) free!
[1584520268] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11083200 kB) free!
[1584520388] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (6881012 kB) free!
[1584521316] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170924 kB) free!
[1584521436] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177316 kB) free!
[1584521556] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189016 kB) free!
[1584521886] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (3916296 kB) free!
[1584521979] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(144) < 103.255.146.164 [18/Mar/2020:14:26:56 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584522006] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (3674996 kB) free!
[1584522046] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (3147140 kB) free!
[1584522100] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584522126] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (3449944 kB) free!
[1584522166] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (3195800 kB) free!
[1584522188] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14203452 kB) free!
[1584522286] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (3159240 kB) free!
[1584522318] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.9% (7838452 kB) free!
[1584522438] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.6% (7678348 kB) free!
[1584522558] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.4% (7569640 kB) free!
[1584522788] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385796 kB) free!
[1584522907] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9285496 kB) free!
[1584523027] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9155096 kB) free!
[1584523063] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.9% (20460372 kB) free.
[1584523096] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 42.4% (55841676 kB) free.
[1584523147] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (8996832 kB) free!
[1584523309] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 334 processes
[1584523327] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1340024 kB) free!
[1584523696] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24024756 kB) free!
[1584523816] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19599512 kB) free!
[1584523927] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.7% (15195840 kB) free.
[1584523936] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15294408 kB) free!
[1584524669] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3640404 kB) free!
[1584524789] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.1% (22730728 kB) free.
[1584525137] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420380 kB) free!
[1584525558] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3062396 kB) free!
[1584525737] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22667172 kB) free!
[1584525933] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584525989] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584526063] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584526066] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.1% (2733868 kB) free!
[1584526119] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584526148] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4908932 kB) free!
[1584526186] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.3% (2532824 kB) free!
[1584526193] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584526207] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175796 kB) free!
[1584526249] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584526271] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584526275] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584526306] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1487832 kB) free!
[1584526309] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 344 processes
[1584526327] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195904 kB) free!
[1584526337] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393112 kB) free!
[1584526392] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2484 processes
[1584526403] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584526447] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177936 kB) free!
[1584526794] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 31.60, 28.68, 25.20
[1584526850] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.6% (29879516 kB) free.
[1584526937] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20928068 kB) free!
[1584527254] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584527499] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(203) < 103.255.146.164 [18/Mar/2020:16:01:27 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584527537] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383076 kB) free!
[1584527619] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584527687] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (220640 kB) free!
[1584528109] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1584528137] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17676380 kB) free!
[1584528220] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 183.82.105.122 [18/Mar/2020:16:09:24 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584528340] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(2) < 183.82.105.122 [18/Mar/2020:16:14:24 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584528460] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(247) < 103.255.146.164 [18/Mar/2020:16:16:41 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584528498] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4366696 kB) free!
[1584528618] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4193148 kB) free!
[1584528737] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392080 kB) free!
[1584528738] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (4061672 kB) free!
[1584528832] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2578104 kB) free!
[1584528952] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3220672 kB) free!
[1584529054] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.7% (4587380 kB) free!
[1584529072] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3970312 kB) free.
[1584529174] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.1% (4447500 kB) free!
[1584529294] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.5% (4288640 kB) free!
[1584529672] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2983660 kB) free!
[1584529792] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3629800 kB) free.
[1584530120] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 70.2% (17316820 kB) free.
[1584530151] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 22.4% (11014588 kB) free.
[1584530182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33075140 kB) free.
[1584530215] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 55.4% (27325080 kB) free.
[1584530242] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 73.5% (18024760 kB) free.
[1584530261] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584530276] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.6% (14076024 kB) free.
[1584530392] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2532556 kB) free!
[1584530494] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2353872 kB) free!
[1584530509] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1584530509] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 346 processes
[1584530512] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (1344708 kB) free!
[1584530538] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2167732 kB) free!
[1584530632] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 34.9% (5672040 kB) free.
[1584530726] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176088 kB) free!
[1584530846] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188492 kB) free!
[1584530848] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186664 kB) free!
[1584530966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189608 kB) free!
[1584530968] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199144 kB) free!
[1584531088] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181240 kB) free!
[1584531709] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1584531988] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4572224 kB) free!
[1584532108] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (390244 kB) free!
[1584532228] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392248 kB) free!
[1584532309] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 342 processes
[1584532619] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9319512 kB) free!
[1584532739] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (8295712 kB) free!
[1584532828] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21331064 kB) free!
[1584532859] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (7707144 kB) free!
[1584533428] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402464 kB) free!
[1584533527] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4469664 kB) free!
[1584533647] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4328688 kB) free!
[1584533767] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.6% (4177968 kB) free!
[1584534029] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23061472 kB) free!
[1584534109] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1584534629] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395856 kB) free!
[1584534660] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584534679] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584534782] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584534790] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584534794] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584534802] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584534805] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584534912] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584534920] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584534924] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584534927] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584534968] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.6% (15173492 kB) free.
[1584535042] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584535054] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584535229] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (25493736 kB) free!
[1584535259] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4146276 kB) free!
[1584535278] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584535282] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584535400] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 140761 MB (62% inode=99%):
[1584535401] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584535403] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2495 processes
[1584535531] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584535535] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584535661] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584535661] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(398) < 103.255.146.164 [18/Mar/2020:18:13:33 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584535664] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584535782] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584535829] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (475140 kB) free!
[1584536008] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584536009] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584536012] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584536138] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584536139] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584536142] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584536265] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584536267] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 140013 MB (62% inode=99%):
[1584536267] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584536268] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2501 processes
[1584536268] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2501 processes
[1584536385] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584536429] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26712800 kB) free.
[1584536515] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.8% (21771324 kB) free.
[1584536566] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.7% (20407620 kB) free.
[1584536633] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.8% (36868092 kB) free.
[1584536681] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.7% (20771736 kB) free.
[1584536721] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584536721] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.9% (29511728 kB) free.
[1584536726] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.6% (22454244 kB) free.
[1584536809] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1584536929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1584536995] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584537029] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (470392 kB) free!
[1584537049] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1584537120] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584537149] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (387912 kB) free!
[1584537166] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172452 kB) free!
[1584537269] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432244 kB) free!
[1584537286] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181540 kB) free!
[1584537331] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584537406] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192948 kB) free!
[1584537461] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584537591] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584537649] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1584537869] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18127452 kB) free!
[1584538192] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.9% (29526364 kB) free.
[1584538194] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584538432] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.1% (1797444 kB) free!
[1584538469] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383532 kB) free!
[1584538552] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.3% (2491020 kB) free!
[1584538614] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9072536 kB) free!
[1584538642] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38771 MB (75% inode=98%):
[1584538654] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584538672] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3400844 kB) free.
[1584538734] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.1% (13861960 kB) free.
[1584538794] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 17.09, 26.64, 30.06
[1584538849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1584538914] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 15.31, 22.87, 28.27
[1584538969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1584539041] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8516420 kB) free!
[1584539069] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18173616 kB) free!
[1584539089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1584539161] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (8192732 kB) free!
[1584539272] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (1724428 kB) free!
[1584539276] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584539281] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (8752720 kB) free!
[1584539392] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2577140 kB) free!
[1584539512] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3366180 kB) free.
[1584539669] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (430904 kB) free!
[1584539689] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1584539689] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1584540269] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20159704 kB) free!
[1584540289] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1584540481] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4927276 kB) free!
[1584540712] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2866132 kB) free!
[1584540832] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3276180 kB) free.
[1584540866] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9578436 kB) free!
[1584540869] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419588 kB) free!
[1584540986] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9397744 kB) free!
[1584541106] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9265016 kB) free!
[1584541469] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24166460 kB) free!
[1584542069] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408924 kB) free!
[1584542670] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28466048 kB) free.
[1584542726] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584542953] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.3% (20074048 kB) free.
[1584543003] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.2% (35111068 kB) free.
[1584543270] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (492744 kB) free!
[1584543390] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403124 kB) free!
[1584543476] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16829532 kB) free!
[1584544076] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41938448 kB) free.
[1584544164] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186308 kB) free!
[1584544209] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (7143140 kB) free!
[1584544285] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187100 kB) free!
[1584544309] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584544329] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6776460 kB) free!
[1584544405] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191088 kB) free!
[1584544433] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3036036 kB) free!
[1584544450] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.4% (6620356 kB) free!
[1584544553] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3304520 kB) free.
[1584544676] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16216544 kB) free!
[1584544796] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12303364 kB) free!
[1584544916] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (8249152 kB) free!
[1584545050] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4447356 kB) free!
[1584545126] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4666004 kB) free!
[1584545246] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4622680 kB) free!
[1584545366] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (4566052 kB) free!
[1584545524] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584545525] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584545654] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584545655] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584545784] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584545785] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584545852] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584545864] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584545982] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584545990] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584546103] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38770 MB (75% inode=98%):
[1584546109] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584546468] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584546567] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180828 kB) free!
[1584546588] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 97278 MB (43% inode=99%):
[1584546602] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584546716] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17501088 kB) free!
[1584546732] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584546859] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.5% (29316212 kB) free.
[1584547024] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.0% (19980244 kB) free.
[1584547084] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39428412 kB) free.
[1584547106] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4688380 kB) free!
[1584547191] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1584547191] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.6% (20503056 kB) free.
[1584547228] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 40.8% (20121036 kB) free.
[1584547316] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394424 kB) free!
[1584547628] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193400 kB) free!
[1584547748] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176696 kB) free!
[1584547797] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192456 kB) free!
[1584547868] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178668 kB) free!
[1584547917] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20140628 kB) free!
[1584547917] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165252 kB) free!
[1584548037] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176268 kB) free!
[1584548062] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584548192] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584548292] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (8009476 kB) free!
[1584548323] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584548394] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584548412] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (7770800 kB) free!
[1584548509] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1584548514] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584548517] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386852 kB) free!
[1584548533] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (7672184 kB) free!
[1584548562] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.8% (19943316 kB) free.
[1584548592] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17541168 kB) free!
[1584548598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42087140 kB) free.
[1584548629] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584548679] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.4% (29286064 kB) free.
[1584548712] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.7% (20767508 kB) free.
[1584549316] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195236 kB) free!
[1584549384] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 7.01, 20.59, 20.63
[1584549386] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584549437] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168764 kB) free!
[1584549557] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170676 kB) free!
[1584549798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17363240 kB) free!
[1584549918] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13429312 kB) free!
[1584550038] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9548440 kB) free!
[1584550933] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4174400 kB) free!
[1584551029] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1584551149] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1584551269] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584551689] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584553038] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14664808 kB) free!
[1584553070] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1584553190] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1584553310] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584553638] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385348 kB) free!
[1584554238] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16382672 kB) free!
[1584554510] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1584554838] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (435820 kB) free!
[1584555439] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16534792 kB) free!
[1584555710] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1584555830] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1584556039] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385744 kB) free!