[1549997175] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10366020 kB) free.
[1549997226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549997346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549997467] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549997775] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9107760 kB) free!
[1549997895] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.7% (5290468 kB) free!
[1549998015] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.7% (9203916 kB) free!
[1549999215] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10370244 kB) free.
[1549999266] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549999386] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549999506] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549999815] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7724912 kB) free!
[1549999935] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10372260 kB) free.
[1550000106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550000226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1550000347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1550000535] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7869476 kB) free!
[1550000656] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10379704 kB) free.
[1550001255] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9652160 kB) free!
[1550001375] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.4% (6088068 kB) free!
[1550001495] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.9% (9805960 kB) free!
[1550002147] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1550002747] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1550002867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550002987] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1550003295] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10400548 kB) free.
[1550003587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1550003708] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1550003895] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5776080 kB) free!
[1550004015] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9754788 kB) free!
[1550004135] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.2% (6998608 kB) free!
[1550004307] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550004427] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1550005027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1550005147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550005267] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1550005867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1550005987] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1550006108] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550006536] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10397504 kB) free.
[1550006708] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1550007308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1550007428] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1550007548] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1550008148] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1550008748] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1550008868] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1550008936] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.6% (6680688 kB) free!
[1550009056] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10401728 kB) free.
[1550009468] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1550009589] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550009656] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8737936 kB) free!
[1550009709] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1550009776] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10407588 kB) free.
[1550009956] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (1118976 kB) free!
[1550010076] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172940 kB) free!
[1550010196] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168284 kB) free!
[1550011509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1550011749] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550011791] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550011805] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1550011879] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1550011887] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550011915] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550011921] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1550011950] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012009] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012015] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012017] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012041] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012045] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1550012045] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012073] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1550012109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550012141] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1550012145] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012171] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012175] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012203] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012229] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1550012275] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012301] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1550012578] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.32, 5.37, 5.08
[1550012596] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550012698] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.23, 5.38, 5.12
[1550012819] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.57, 5.43, 5.17
[1550012829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1550012901] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550012949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1550013005] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7620 MB (14% inode=98%):
[1550013070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550013196] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 96.1% (23554056 kB) free.
[1550013209] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 370 processes
[1550013245] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 14 users currently logged in
[1550013375] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1550013376] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.5% (6659040 kB) free!
[1550013403] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1550013475] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 68968 MB (30% inode=99%):
[1550013496] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10481492 kB) free.
[1550014019] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.04, 4.79, 4.94
[1550014270] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1550014541] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.60, 2.33, 18.46
[1550014696] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9595136 kB) free!
[1550014816] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.2% (6519632 kB) free!
[1550014870] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1550014936] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10570332 kB) free.
[1550014990] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1550015110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550016136] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9012916 kB) free!
[1550016256] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.5% (10600244 kB) free.
[1550016310] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1550016341] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.74, 1.75, 4.17
[1550016856] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1550016911] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1550017031] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1550017151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1550018056] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.6% (7679132 kB) free!
[1550018176] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10628936 kB) free.
[1550019376] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8141212 kB) free!
[1550019496] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10634428 kB) free.
[1550020096] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8395628 kB) free!
[1550020216] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10632176 kB) free.
[1550021351] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1550021951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1550022016] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7741364 kB) free!
[1550022071] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1550022136] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10286644 kB) free.
[1550022671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1550022736] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9615928 kB) free!
[1550022791] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1550022856] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9227904 kB) free!
[1550022911] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550022976] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10519632 kB) free.
[1550023576] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8057896 kB) free!
[1550023697] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10516116 kB) free.
[1550024112] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1550024296] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.2% (6995992 kB) free!
[1550024416] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9533764 kB) free!
[1550024536] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.6% (6707332 kB) free!
[1550024712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1550024832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550024952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550025736] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10060648 kB) free.
[1550026152] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1550026692] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4671804 kB) free!
[1550026752] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1550026787] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9350804 kB) free!
[1550026812] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4564212 kB) free!
[1550026872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1550026907] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12499764 kB) free.
[1550026932] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4453340 kB) free!
[1550026936] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8718492 kB) free!
[1550026993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1550027056] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10433128 kB) free.
[1550027596] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550027716] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550027836] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550028193] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1550028256] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8332984 kB) free!
[1550028376] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10401692 kB) free.
[1550028793] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1550028913] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1550028976] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.9% (7812900 kB) free!
[1550029097] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10279720 kB) free.
[1550029332] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2388204 kB) free!
[1550029513] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550029633] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1550029696] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.0% (6910336 kB) free!
[1550029753] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1550029816] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9549880 kB) free!
[1550029936] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.9% (6372684 kB) free!
[1550031107] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8486636 kB) free!
[1550031136] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1550031136] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.5% (10100628 kB) free.
[1550031227] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11942364 kB) free.
[1550031736] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.4% (7583508 kB) free!
[1550031827] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8944152 kB) free!
[1550031856] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10041400 kB) free.
[1550031947] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12055336 kB) free.
[1550032456] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.4% (6088824 kB) free!
[1550032547] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9138448 kB) free!
[1550032577] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 68.9% (33958616 kB) free.
[1550032642] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 49.8% (24538384 kB) free.
[1550032691] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.4% (22806992 kB) free.
[1550033236] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550033353] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1550033953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1550034073] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1550034674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550034794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1550034914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550035037] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550035157] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550035229] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3165872 kB) free!
[1550035277] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550035349] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3143668 kB) free!
[1550035469] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (3122348 kB) free!
[1550038056] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.3% (15162100 kB) free.
[1550038277] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1550038277] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550040314] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1550040914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1550041034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1550041154] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1550041754] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1550041874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1550041877] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550041995] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1550041997] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550042117] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550042176] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9645912 kB) free!
[1550042296] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1550042296] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.5% (11567320 kB) free.
[1550042596] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4817800 kB) free!
[1550042716] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4727528 kB) free!
[1550042836] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (4631952 kB) free!
[1550042896] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8260068 kB) free!
[1550043016] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11077208 kB) free.
[1550043616] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8910708 kB) free!
[1550043737] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.2% (10931088 kB) free.
[1550044336] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9573084 kB) free!
[1550044346] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3178760 kB) free!
[1550044456] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.4% (6584236 kB) free!
[1550044466] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3144508 kB) free!
[1550044564] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 535 processes
[1550044576] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.7% (9713968 kB) free!
[1550044586] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.1% (3107376 kB) free!
[1550044624] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 542 processes
[1550045385] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 96.0% (23523588 kB) free.
[1550045444] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 68.9% (33959756 kB) free.
[1550045504] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.2% (14659984 kB) free.
[1550046195] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1550046796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550046916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1550046917] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550047036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550047045] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9427872 kB) free!
[1550047165] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13058284 kB) free.
[1550047664] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.9% (2739352 kB) free!
[1550047785] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (2714832 kB) free!
[1550047905] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.3% (2492192 kB) free!
[1550048741] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.71, 9.21, 7.98
[1550048836] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1550048861] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.65, 6.68, 7.20
[1550048981] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.55, 5.26, 6.61
[1550049436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1550049556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550049581] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.16, 2.42, 4.42
[1550049677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1550050905] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (928984 kB) free!
[1550051023] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.3% (15171696 kB) free.
[1550051981] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.19, 7.82, 8.18
[1550052101] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.63, 5.79, 7.40
[1550052221] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.81, 4.47, 6.72
[1550052677] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1550052821] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1550052821] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.77, 2.20, 4.42
[1550052917] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550053037] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550053157] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550053277] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1550053397] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1550053518] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1550054117] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1550054237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1550054357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1550054365] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9050524 kB) free!
[1550054485] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12723536 kB) free.
[1550055221] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 186.57, 58.04, 21.38
[1550055341] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 26.56, 39.36, 18.98
[1550055461] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.44, 26.99, 16.91
[1550056061] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.12, 5.74, 10.04
[1550056157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1550056520] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 18.8% (3054116 kB) free!
[1550056640] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 18.4% (2994512 kB) free!
[1550056757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1550056760] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 18.0% (2922744 kB) free!
[1550056878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1550056885] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8639276 kB) free!
[1550056998] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1550057005] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11997492 kB) free.
[1550057261] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.55, 1.83, 4.19
[1550057605] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9511512 kB) free!
[1550057725] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12192904 kB) free.
[1550058798] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1550058921] SERVICE ALERT: sig1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550058921] SERVICE ALERT: sig1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550058923] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1550058971] SERVICE ALERT: sig1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig1' was out of bounds)
[1550058986] SERVICE ALERT: sig1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig1' was out of bounds)
[1550058987] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig1' was out of bounds)
[1550059041] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.14, 0.05
[1550059041] SERVICE ALERT: sig1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1550059045] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1550059061] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.85, 7.20, 8.21
[1550059091] SERVICE ALERT: sig1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37134 MB (72% inode=98%):
[1550059106] SERVICE ALERT: sig1;Sensors Status;OK;HARD;1;SENSORS OK
[1550059107] SERVICE ALERT: sig1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 160993 MB (71% inode=99%):
[1550059157] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1550059160] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.5% (15358956 kB) free.
[1550059181] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.45, 5.33, 7.41
[1550059301] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.44, 4.05, 6.69
[1550059348] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550059468] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550059588] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550059901] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.81, 2.23, 4.43
[1550060111] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46806 MB (20% inode=99%):
[1550060188] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550060231] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46804 MB (20% inode=99%):
[1550060275] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47117 MB (20% inode=99%):
[1550060351] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46804 MB (20% inode=99%):
[1550060395] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47099 MB (20% inode=99%):
[1550060515] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46936 MB (20% inode=99%):
[1550061551] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47364 MB (21% inode=99%):
[1550061715] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50592 MB (22% inode=99%):
[1550062151] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46710 MB (20% inode=99%):
[1550062271] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48321 MB (21% inode=99%):
[1550062301] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 28.97, 10.34, 4.97
[1550062339] SERVICE ALERT: db2;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'db2' was out of bounds)
[1550062342] HOST ALERT: db2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.59)
[1550062363] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 344 processes
[1550062365] SERVICE ALERT: db2;Mariadb status;CRITICAL;HARD;1;(Return code of 255 for service 'Mariadb status' on host 'db2' was out of bounds)
[1550062368] SERVICE ALERT: db2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'db2' was out of bounds)
[1550062398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;1;CPU CRITICAL: 2 crit, 0 warn out of 314 processes
[1550062400] HOST ALERT: db2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1550062421] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.03, 7.34, 4.52
[1550062459] SERVICE ALERT: db2;CPU Load;OK;HARD;1;OK - load average: 1.10, 0.39, 0.14
[1550062485] SERVICE ALERT: db2;Mariadb status;OK;HARD;1;PROCS OK: 1 process with command name 'mysqld'
[1550062488] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 59652 MB (26% inode=99%):
[1550062963] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 329 processes
[1550063725] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7960844 kB) free!
[1550063845] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 37.3% (18377956 kB) free.
[1550064113] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2950336 kB) free!
[1550064233] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2913936 kB) free!
[1550064353] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.9% (2739828 kB) free!
[1550064419] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.17, 5.11, 5.03
[1550064445] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6011312 kB) free!
[1550064539] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.93, 5.04, 5.02
[1550064565] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 33.5% (16492064 kB) free.
[1550064659] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.71, 5.00, 5.01
[1550064826] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (2963444 kB) free!
[1550064946] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.0% (2756324 kB) free!
[1550065066] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.6% (2529760 kB) free!
[1550065165] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.7% (5769588 kB) free!
[1550065285] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.2% (14408296 kB) free.
[1550065757] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1550065856] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8386440 kB) free!
[1550065859] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.34, 4.73, 4.95
[1550065976] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.6% (12636144 kB) free.
[1550066267] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1580084 kB) free!
[1550066486] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6902832 kB) free!
[1550066606] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 34.3% (16881672 kB) free.
[1550067176] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8554552 kB) free!
[1550067297] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.1% (12852160 kB) free.
[1550067353] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1623824 kB) free!
[1550068553] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.9% (1777808 kB) free!
[1550068798] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (3011564 kB) free!
[1550068918] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.0% (2933032 kB) free!
[1550069038] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (2859776 kB) free!
[1550069153] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1550069153] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1614384 kB) free!
[1550069638] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2455452 kB) free!
[1550070296] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8684208 kB) free!
[1550070318] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1550070416] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.9% (12771436 kB) free.
[1550070610] SERVICE ALERT: sigqscluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1550070731] SERVICE ALERT: sigqscluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 334 processes
[1550070851] SERVICE ALERT: sigqscluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 334 processes
[1550071421] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550071518] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 303 processes
[1550071916] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.6% (23430704 kB) free.
[1550072013] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.5% (15202028 kB) free.
[1550072066] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.3% (14833288 kB) free.
[1550072118] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 300 processes
[1550072216] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8074596 kB) free!
[1550072238] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 300 processes
[1550072336] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.3% (11975664 kB) free.
[1550072358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 299 processes
[1550072478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550072598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 298 processes
[1550072718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550072838] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 297 processes
[1550072959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 297 processes
[1550073079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 303 processes
[1550073199] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 300 processes
[1550073319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 306 processes
[1550073439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 303 processes
[1550073559] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 303 processes
[1550073679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 306 processes
[1550073799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550073919] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1550074040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 299 processes
[1550074160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 300 processes
[1550074280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 305 processes
[1550074400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550074520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 305 processes
[1550074640] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550074760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550074880] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 305 processes
[1550075001] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550075121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 306 processes
[1550075241] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1550075336] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8758188 kB) free!
[1550075361] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 303 processes
[1550075456] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.2% (5525724 kB) free!
[1550075474] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (1275208 kB) free!
[1550075481] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 304 processes
[1550075524] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 564 processes
[1550075534] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181592 kB) free!
[1550075576] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10433220 kB) free.
[1550075592] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 559 processes
[1550075594] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (193404 kB) free!
[1550075601] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 300 processes
[1550075654] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (199916 kB) free!
[1550075721] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 304 processes
[1550075841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 300 processes
[1550075961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 300 processes
[1550076082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550076202] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 299 processes
[1550076322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550076442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 299 processes
[1550076562] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 298 processes
[1550076682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 303 processes
[1550076776] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.2% (7993844 kB) free!
[1550076802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 300 processes
[1550076880] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2981408 kB) free!
[1550076896] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11684008 kB) free.
[1550076922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550077000] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2976764 kB) free!
[1550077042] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 299 processes
[1550077120] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.3% (2970468 kB) free!
[1550077163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 300 processes
[1550077283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550077403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550077523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 305 processes
[1550077643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550077763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550077782] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 61.0% (14947908 kB) free.
[1550077837] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.3% (14832652 kB) free.
[1550077883] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 304 processes
[1550078003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550078123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 306 processes
[1550078244] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 300 processes
[1550078364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550078484] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 304 processes
[1550078604] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 300 processes
[1550078724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 306 processes
[1550078844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 303 processes
[1550078964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 303 processes
[1550079084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 305 processes
[1550079204] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550079325] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 305 processes
[1550079445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550079565] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 301 processes
[1550079646] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2845732 kB) free!
[1550079685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 304 processes
[1550079766] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2840068 kB) free!
[1550079805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550079886] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.5% (2840140 kB) free!
[1550079925] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 303 processes
[1550080045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550080165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550080286] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 303 processes
[1550080406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 300 processes
[1550080526] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 304 processes
[1550080646] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550080766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550080886] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 307 processes
[1550081006] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550081126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 305 processes
[1550081247] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 301 processes
[1550081367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550081487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 304 processes
[1550081607] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 300 processes
[1550081727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 304 processes
[1550081847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550081967] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550082087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 304 processes
[1550082207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550082328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 307 processes
[1550082448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 1 warn out of 305 processes
[1550082568] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 305 processes