[1562783455] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1562783575] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562783689] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562783752] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1562783757] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11278656 kB) free!
[1562783806] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (1102032 kB) free!
[1562783876] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1562783877] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (20441276 kB) free.
[1562783895] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7904252 kB) free!
[1562783926] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196048 kB) free!
[1562783941] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1562783994] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.42, 5.84, 5.40
[1562784015] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12133880 kB) free.
[1562784020] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182724 kB) free!
[1562784046] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182112 kB) free!
[1562784117] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.24, 6.07, 5.55
[1562784237] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.39, 5.76, 5.49
[1562784501] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1562784541] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562784615] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7631664 kB) free!
[1562784661] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562784735] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12142676 kB) free.
[1562785261] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562785335] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7202076 kB) free!
[1562785382] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562785455] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12143688 kB) free.
[1562785981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562786101] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562786655] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8240012 kB) free!
[1562786701] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562786775] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12136252 kB) free.
[1562786821] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562787194] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 11.9% (2917804 kB) free!
[1562787254] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 11.9% (2917284 kB) free!
[1562787314] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 11.9% (2915396 kB) free!
[1562787375] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 11.9% (2915412 kB) free!
[1562787375] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7956088 kB) free!
[1562787421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562787495] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12142216 kB) free.
[1562787541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562787775] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1562787895] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562788077] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12502456 kB) free!
[1562788095] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8044736 kB) free!
[1562788141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562788197] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (20950068 kB) free.
[1562788215] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12131392 kB) free.
[1562788261] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562788495] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562788615] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562788815] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8132208 kB) free!
[1562788862] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562788935] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12140468 kB) free.
[1562788982] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562789535] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8274284 kB) free!
[1562789582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562789655] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12142132 kB) free.
[1562789702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562789997] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12485220 kB) free!
[1562790117] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (20600232 kB) free.
[1562790255] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8382192 kB) free!
[1562790302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562790376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12134124 kB) free.
[1562790422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562790717] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12795512 kB) free!
[1562790837] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562790837] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (21044088 kB) free.
[1562790975] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8454156 kB) free!
[1562790993] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173992 kB) free!
[1562791022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562791095] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12141672 kB) free.
[1562791113] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 47.6% (7729240 kB) free.
[1562791143] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562791179] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (192124 kB) free!
[1562791299] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (190000 kB) free!
[1562791419] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (235692 kB) free!
[1562791695] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8563596 kB) free!
[1562791742] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562791815] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12132620 kB) free.
[1562791862] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562792415] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8573168 kB) free!
[1562792462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562792535] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12140148 kB) free.
[1562792582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562793135] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8580424 kB) free!
[1562793182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562793255] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12140508 kB) free.
[1562793302] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562793404] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 69.3% (17111740 kB) free.
[1562793479] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 81.8% (20059420 kB) free.
[1562793552] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.7% (13926000 kB) free.
[1562793612] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22012848 kB) free.
[1562793658] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.7% (14581304 kB) free.
[1562793720] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 69.2% (11249552 kB) free.
[1562793855] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8564204 kB) free!
[1562793902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562793975] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12131724 kB) free.
[1562794004] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (680244 kB) free!
[1562794022] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562794125] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195128 kB) free!
[1562794245] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187432 kB) free!
[1562794265] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3246416 kB) free!
[1562794385] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (1514828 kB) free!
[1562794437] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12303864 kB) free!
[1562794505] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (405268 kB) free!
[1562794557] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (20288796 kB) free.
[1562794575] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8641208 kB) free!
[1562794623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562794696] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12139668 kB) free.
[1562794743] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562795037] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562795157] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12513624 kB) free!
[1562795277] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (20566284 kB) free.
[1562795295] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8743496 kB) free!
[1562795343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562795415] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12136772 kB) free.
[1562795463] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562795877] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (12775140 kB) free!
[1562795997] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (20843860 kB) free.
[1562796015] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8819396 kB) free!
[1562796063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562796135] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12037876 kB) free.
[1562796183] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562796597] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (13143524 kB) free!
[1562796717] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (21058176 kB) free.
[1562796735] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8863484 kB) free!
[1562796783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562796855] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7206492 kB) free!
[1562796903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562796975] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12136036 kB) free.
[1562797453] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 542 processes
[1562797489] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.77, 3.96, 4.67
[1562797504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562797513] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 540 processes
[1562797575] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7278540 kB) free!
[1562797624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562797695] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12137708 kB) free.
[1562798224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562798295] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7325916 kB) free!
[1562798344] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562798415] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12130908 kB) free.
[1562798637] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.61, 4.81, 5.00
[1562798944] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562799016] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7313608 kB) free!
[1562799064] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562799136] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12139076 kB) free.
[1562799664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562799736] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7215168 kB) free!
[1562799784] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562799856] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12142548 kB) free.
[1562800385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562800456] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7171804 kB) free!
[1562800505] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562800576] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12132824 kB) free.
[1562801105] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562801225] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562801637] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.37, 5.17, 5.04
[1562801757] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.42, 5.19, 5.06
[1562801776] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8907144 kB) free!
[1562801825] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562801877] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.36, 5.00, 5.01
[1562801896] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7213292 kB) free!
[1562801945] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562802016] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12130520 kB) free.
[1562802545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562802665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562802786] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562803817] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8841516 kB) free!
[1562803937] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7167936 kB) free!
[1562803986] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562804057] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11530332 kB) free.
[1562804586] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562804706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562804826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562805257] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8636052 kB) free!
[1562805377] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7683376 kB) free!
[1562805426] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1562805426] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562805497] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11240792 kB) free.
[1562806026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562806147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562806267] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562806697] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8638196 kB) free!
[1562806817] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12059056 kB) free.
[1562806867] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562807417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8676404 kB) free!
[1562807467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562807537] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12056324 kB) free.
[1562807587] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562807689] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.06, 22.13, 18.64
[1562807809] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.41, 15.89, 16.78
[1562807929] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.51, 11.83, 15.18
[1562808137] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8706900 kB) free!
[1562808187] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562808258] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12056256 kB) free.
[1562808307] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562808857] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8725572 kB) free!
[1562808907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562808917] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1562808977] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11950488 kB) free.
[1562809027] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562809090] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 325 processes
[1562809577] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8821444 kB) free!
[1562809628] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562809697] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7233956 kB) free!
[1562809729] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.25, 0.38, 3.08
[1562809748] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562809817] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12046652 kB) free.
[1562810348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562810417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7096204 kB) free!
[1562810468] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562810537] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12052800 kB) free.
[1562812317] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11596488 kB) free!
[1562812437] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20237736 kB) free.
[1562812868] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562812937] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7477952 kB) free!
[1562812989] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562813058] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12052272 kB) free.
[1562813588] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562813657] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7569384 kB) free!
[1562813708] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562813777] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12043280 kB) free.
[1562814308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562814377] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7504720 kB) free!
[1562814428] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562814497] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12048368 kB) free.
[1562815028] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562815097] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7626120 kB) free!
[1562815148] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1562815217] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12041124 kB) free.
[1562815437] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (11516128 kB) free!
[1562815557] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (20190124 kB) free.
[1562815748] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562815817] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7571908 kB) free!
[1562815869] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562815937] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12048244 kB) free.
[1562816290] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 325 processes
[1562816338] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220304 kB) free!
[1562816458] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219152 kB) free!
[1562816468] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562816578] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230708 kB) free!
[1562816588] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562817188] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562817308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562817428] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562817722] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.92, 7.21, 5.19
[1562817737] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9559528 kB) free!
[1562817842] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.92, 6.28, 5.10
[1562817858] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8204704 kB) free!
[1562817962] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.50, 5.81, 5.07
[1562817978] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12028168 kB) free.
[1562818015] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562818135] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562818517] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 353 processes
[1562818557] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (9827992 kB) free!
[1562818637] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1562818678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (19335068 kB) free.
[1562818735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1562818757] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1562818845] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4138060 kB) free!
[1562818855] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562819034] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (1455964 kB) free!
[1562819154] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (1231400 kB) free!
[1562819274] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (1039688 kB) free!
[1562819445] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175796 kB) free!
[1562819874] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.7% (6745928 kB) free.
[1562819877] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (7331780 kB) free!
[1562819992] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.0% (18249252 kB) free.
[1562819997] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (17162108 kB) free.
[1562820055] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39731524 kB) free.
[1562820102] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.1% (14484664 kB) free.
[1562820558] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.7% (2558548 kB) free!
[1562820597] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10909180 kB) free!
[1562820599] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (880916 kB) free!
[1562820678] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.0% (2270532 kB) free!
[1562820702] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (165068 kB) free!
[1562820717] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (7021612 kB) free!
[1562820719] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168564 kB) free!
[1562820798] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.2% (1982632 kB) free!
[1562820822] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (151584 kB) free!
[1562820837] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (17177548 kB) free.
[1562820839] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1562820839] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179496 kB) free!
[1562820942] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (150048 kB) free!
[1562820978] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9009568 kB) free!
[1562821098] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7340800 kB) free!
[1562821218] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11904568 kB) free.
[1562821398] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (388320 kB) free!
[1562821437] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (11097564 kB) free!
[1562821557] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (7064312 kB) free!
[1562821628] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562821677] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (17081428 kB) free.
[1562822228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1562822277] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (11112628 kB) free!
[1562822348] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562822397] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (6988412 kB) free!
[1562822517] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (16295712 kB) free.
[1562823068] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4815168 kB) free!
[1562823117] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10906844 kB) free!
[1562823188] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4766844 kB) free!
[1562823237] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (6667260 kB) free!
[1562823308] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.2% (4715848 kB) free!
[1562823358] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (2644792 kB) free!
[1562823958] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (11374884 kB) free!
[1562824148] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562824269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562824389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562824558] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (5723328 kB) free!
[1562824693] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183028 kB) free!
[1562824813] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188240 kB) free!
[1562824818] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9388160 kB) free!
[1562824933] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187668 kB) free!
[1562824938] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8060216 kB) free!
[1562825058] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11993292 kB) free.
[1562825274] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (385472 kB) free!
[1562825394] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (121700 kB) free!
[1562825408] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.91, 5.91, 5.48
[1562825418] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163620 kB) free!
[1562825515] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (128016 kB) free!
[1562825528] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.90, 5.94, 5.54
[1562825538] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174044 kB) free!
[1562825589] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21904820 kB) free.
[1562825642] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.4% (20200080 kB) free.
[1562825648] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.80, 5.91, 5.58
[1562825658] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7579956 kB) free!
[1562825658] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178296 kB) free!
[1562825698] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1562825698] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.3% (6789404 kB) free.
[1562825758] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1562825758] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (9893344 kB) free!
[1562825759] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.4% (13873880 kB) free.
[1562825778] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11980464 kB) free.
[1562825823] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14648660 kB) free.
[1562825886] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.6% (21221852 kB) free.
[1562826194] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172536 kB) free!
[1562826253] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191516 kB) free!
[1562826314] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183360 kB) free!
[1562826358] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2814912 kB) free!
[1562826373] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191072 kB) free!
[1562826378] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7691200 kB) free!
[1562826429] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (150144 kB) free!
[1562826434] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181084 kB) free!
[1562826493] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176592 kB) free!
[1562826498] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11984148 kB) free.
[1562826499] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183748 kB) free!
[1562826531] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.68, 6.89, 5.78
[1562826549] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (158064 kB) free!
[1562826619] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179028 kB) free!
[1562826651] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.36, 6.78, 5.89
[1562826669] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1562826669] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (184192 kB) free!
[1562826739] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173108 kB) free!
[1562826772] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.17, 5.90, 5.68
[1562826848] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.92, 4.21, 4.71
[1562826958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (18654980 kB) free.
[1562826972] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.32, 7.51, 5.05
[1562827092] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.81, 7.83, 5.46
[1562827212] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.49, 5.99, 5.10
[1562827372] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.39, 4.67, 4.71
[1562827389] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562827698] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8758844 kB) free!
[1562827812] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.85, 5.16, 4.90
[1562827818] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7075820 kB) free!
[1562827939] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11239024 kB) free.
[1562827989] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562828110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562828158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (8129004 kB) free!
[1562828165] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (158792 kB) free!
[1562828230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562828278] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (19951724 kB) free.
[1562828285] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166012 kB) free!
[1562828405] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1562828405] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167152 kB) free!
[1562828530] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (160116 kB) free!
[1562828648] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.70, 5.96, 5.07
[1562828651] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173392 kB) free!
[1562828768] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 4.63, 5.44, 4.99
[1562828771] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166324 kB) free!
[1562828830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562829055] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4301644 kB) free!
[1562829138] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8909952 kB) free!
[1562829175] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3411112 kB) free!
[1562829258] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7378220 kB) free!
[1562829296] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2663620 kB) free!
[1562829378] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11474964 kB) free.
[1562829430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562829478] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (9620976 kB) free!
[1562829550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562829598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (17300292 kB) free.
[1562829612] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.47, 5.96, 5.53
[1562829670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562829732] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.49, 5.88, 5.56
[1562829852] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.58, 6.24, 5.73
[1562830093] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1562830093] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (2828004 kB) free!
[1562830270] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562830578] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9190496 kB) free!
[1562830693] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1263908 kB) free!
[1562830698] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7915564 kB) free!
[1562830798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10884268 kB) free!
[1562830805] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.0% (4875312 kB) free.
[1562830818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11876140 kB) free.
[1562830870] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562830919] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (6843672 kB) free!
[1562830991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562831039] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (17966536 kB) free.
[1562831111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562831639] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (12956540 kB) free!
[1562831711] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562831759] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (8769484 kB) free!
[1562831879] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (4817696 kB) free!
[1562832005] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.6% (2374164 kB) free!
[1562832018] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9210680 kB) free!
[1562832125] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 13.1% (2123384 kB) free!
[1562832138] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7935392 kB) free!
[1562832245] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.7% (1902616 kB) free!
[1562832253] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1562832253] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.24, 3.03, 4.32
[1562832258] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11960444 kB) free.
[1562832311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562832431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562832479] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (15146452 kB) free.
[1562832551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562832772] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.47, 5.87, 5.40
[1562832845] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (449332 kB) free!
[1562832892] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.80, 6.13, 5.55
[1562832908] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2275732 kB) free!
[1562833012] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.22, 6.09, 5.60
[1562833035] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.6% (21481328 kB) free.
[1562833079] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (9894784 kB) free!
[1562833082] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 68.2% (16826976 kB) free.
[1562833124] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 80.9% (19834524 kB) free.
[1562833199] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (5679556 kB) free!
[1562833226] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.6% (13914792 kB) free.
[1562833311] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39820876 kB) free.
[1562833319] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (16244100 kB) free.
[1562833360] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.8% (14590260 kB) free.
[1562833415] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 64.9% (10555508 kB) free.
[1562833458] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9321144 kB) free!
[1562833462] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 77.6% (19027016 kB) free.
[1562833508] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (695388 kB) free!
[1562833578] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7994940 kB) free!
[1562833628] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (453532 kB) free!
[1562833686] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (688736 kB) free!
[1562833698] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11942216 kB) free.
[1562833748] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (191288 kB) free!
[1562833806] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173580 kB) free!
[1562833919] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12837600 kB) free!
[1562833926] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179212 kB) free!
[1562834039] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (8588124 kB) free!
[1562834159] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (4474140 kB) free!
[1562834571] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171340 kB) free!
[1562834691] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165572 kB) free!
[1562834759] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (15197060 kB) free.
[1562834811] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (154700 kB) free!
[1562834898] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9108876 kB) free!
[1562835018] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7771628 kB) free!
[1562835055] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1562835138] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11793412 kB) free.
[1562835175] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1562835295] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1562835359] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10883984 kB) free!
[1562835479] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (6505564 kB) free!
[1562835599] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (2526616 kB) free!
[1562836200] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (13443992 kB) free.
[1562836338] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9469720 kB) free!
[1562836458] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8264484 kB) free!
[1562836578] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11921680 kB) free.
[1562836800] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (9037396 kB) free!
[1562836920] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (4834188 kB) free!
[1562837040] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (861092 kB) free!
[1562837640] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (12065656 kB) free!
[1562837756] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1562837778] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9840908 kB) free!
[1562837898] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.6% (8649784 kB) free!
[1562837929] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (4129760 kB) free!
[1562837948] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6834068 kB) free.
[1562838018] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (6989592 kB) free!
[1562838049] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169840 kB) free!
[1562838050] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14343.06, 9707.27, 4361.21
[1562838169] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (260896 kB) free!
[1562838170] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14388.34, 11258.46, 5577.85
[1562838270] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186528 kB) free!
[1562838290] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14395.27, 12296.69, 6647.08
[1562838390] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167824 kB) free!
[1562838510] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (2955740 kB) free!
[1562838551] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562838618] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11248980 kB) free.
[1562838716] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7896172 kB) free!
[1562838836] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.0% (7376308 kB) free!
[1562838840] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2812084 kB) free!
[1562838957] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.8% (6823388 kB) free!
[1562839012] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.55, 3.98, 4.70
[1562839110] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190772 kB) free!
[1562839151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562839236] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3175724 kB) free!
[1562839271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562839356] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (3005872 kB) free!
[1562839392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562839440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (13950772 kB) free.
[1562839476] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.4% (2825404 kB) free!
[1562839557] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4441680 kB) free!
[1562839818] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8551240 kB) free!
[1562839938] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6891900 kB) free!
[1562839992] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562840040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (11038340 kB) free!
[1562840058] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11139852 kB) free.
[1562840160] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (6874760 kB) free!
[1562840278] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.58, 3.58, 4.58
[1562840280] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2884296 kB) free!
[1562840592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562840676] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (943688 kB) free!
[1562840712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562840812] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.81, 10.68, 7.34
[1562840832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562840880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (13646664 kB) free.
[1562840932] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.44, 10.19, 7.58
[1562841052] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.72, 8.99, 7.46
[1562841253] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.04, 7.22, 6.09
[1562841258] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9112004 kB) free!
[1562841373] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.62, 4.87, 5.36
[1562841378] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7796772 kB) free!
[1562841480] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (10610812 kB) free!
[1562841494] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 3.30, 4.73
[1562841498] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11800132 kB) free.
[1562841600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (6246300 kB) free!
[1562841720] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (2206004 kB) free!
[1562842093] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562842126] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 69.5% (17155504 kB) free.
[1562842192] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 73.5% (18017732 kB) free.
[1562842227] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 61.6% (40538260 kB) free.
[1562842270] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1562842270] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39826424 kB) free.
[1562842323] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.0% (13975644 kB) free.
[1562842357] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.9% (14620176 kB) free.
[1562842405] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1562842405] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22010664 kB) free.
[1562842699] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9338244 kB) free!
[1562842735] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197580 kB) free!
[1562842757] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1562842793] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.0% (2691808 kB) free!
[1562842819] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8135616 kB) free!
[1562842855] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195760 kB) free!
[1562842913] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (716788 kB) free!
[1562842939] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11893280 kB) free.
[1562842975] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178372 kB) free!
[1562843033] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173112 kB) free!
[1562843357] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1562843568] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (163516 kB) free!
[1562843689] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164748 kB) free!
[1562843809] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (153900 kB) free!
[1562843957] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1562844075] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 8.3% (2026692 kB) free!
[1562844139] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9564248 kB) free!
[1562844259] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8341272 kB) free!
[1562844290] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.91, 0.52, 15.95
[1562844379] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11871828 kB) free.
[1562844548] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (1340380 kB) free!
[1562844669] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (1149736 kB) free!
[1562844789] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (937024 kB) free!
[1562845158] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1562845158] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1562845389] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (132364 kB) free!
[1562845490] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 0.45, 4.72
[1562846179] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9165384 kB) free!
[1562846300] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7936496 kB) free!
[1562846420] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11861752 kB) free.
[1562847620] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9574912 kB) free!
[1562847740] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8378716 kB) free!
[1562847860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11865820 kB) free.
[1562847962] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.67, 3.86, 4.75
[1562848158] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 354 processes
[1562848390] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.5% (6803408 kB) free.
[1562849060] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9842504 kB) free!
[1562849180] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8596212 kB) free!
[1562849300] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10233288 kB) free.
[1562849358] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1562849900] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8430224 kB) free!
[1562850021] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11849644 kB) free.
[1562850123] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (3133256 kB) free!
[1562850243] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (2969948 kB) free!
[1562850363] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.4% (2836096 kB) free!
[1562850558] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1562850620] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8551732 kB) free!
[1562850675] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1562850740] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6909204 kB) free!
[1562850805] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.8% (2895708 kB) free!
[1562850860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11853628 kB) free.
[1562850925] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190304 kB) free!
[1562850956] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1562851032] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562851045] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174916 kB) free!
[1562851093] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.30, 7.46, 5.42
[1562851214] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.12, 5.78, 5.08
[1562851308] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1562851334] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 3.90, 4.47
[1562852060] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8359548 kB) free!
[1562852163] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1466136 kB) free!
[1562852180] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11845640 kB) free.
[1562852438] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562852475] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9842748 kB) free!
[1562852595] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8850264 kB) free!
[1562852715] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (7799288 kB) free!
[1562852780] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8339216 kB) free!
[1562852833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562852900] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11855388 kB) free.
[1562852953] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562853175] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562853233] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562853410] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562853500] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7142740 kB) free!
[1562853620] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11839076 kB) free.
[1562853652] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14281.31, 8772.39, 3757.97
[1562853869] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1562853915] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2383372 kB) free!
[1562853989] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1562854153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562854273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562854334] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.53, 6.11, 5.24
[1562854454] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.33, 4.46, 4.74
[1562854821] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8379704 kB) free!
[1562854873] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562854941] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11850236 kB) free.
[1562854993] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562855541] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8225644 kB) free!
[1562855594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562855661] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11851804 kB) free.
[1562855714] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562856261] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8113248 kB) free!
[1562856314] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562856381] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11838500 kB) free.
[1562856434] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562856638] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (11407784 kB) free!
[1562856759] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (23315760 kB) free.
[1562856981] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7841868 kB) free!
[1562857034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562857045] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562857101] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11845728 kB) free.
[1562857154] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562857390] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1562857701] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7699248 kB) free!
[1562857754] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562857821] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11834444 kB) free.
[1562857874] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562857874] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562857945] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.1% (19036836 kB) free.
[1562857980] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.7% (20262804 kB) free.
[1562858027] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 73.8% (36391488 kB) free.
[1562858054] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.83, 6.97, 5.10
[1562858070] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.6% (13915720 kB) free.
[1562858108] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.2% (21378052 kB) free.
[1562858174] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.85, 5.22, 4.70
[1562858325] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.2% (14668668 kB) free.
[1562858421] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7584812 kB) free!
[1562858541] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11845416 kB) free.
[1562858545] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (1563772 kB) free!
[1562858558] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11671708 kB) free!
[1562858665] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185964 kB) free!
[1562858678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (23294348 kB) free.
[1562858785] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196604 kB) free!
[1562859075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562859142] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7258700 kB) free!
[1562859189] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4448376 kB) free!
[1562859195] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562859262] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11845444 kB) free.
[1562859309] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4392440 kB) free!
[1562859429] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4314028 kB) free!
[1562859558] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1562859795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562859862] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6926500 kB) free!
[1562859915] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562859956] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1562859982] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11836620 kB) free.
[1562859985] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 64.1% (15808796 kB) free.
[1562860125] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (2554836 kB) free!
[1562860245] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.8% (2079976 kB) free!
[1562860252] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.40, 1.41, 15.53
[1562860365] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1604816 kB) free!
[1562860478] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12243656 kB) free!
[1562860515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562860582] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7206488 kB) free!
[1562860598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (23086924 kB) free.
[1562860635] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562860702] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11842208 kB) free.
[1562861235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562861302] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8424152 kB) free!
[1562861355] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562861422] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11833552 kB) free.
[1562861452] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.42, 4.55
[1562861956] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562862076] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562862109] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1562862622] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8372932 kB) free!
[1562862676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562862742] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11840284 kB) free.
[1562862796] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562862998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (11788844 kB) free!
[1562863118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.4% (22597260 kB) free.
[1562863342] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8339584 kB) free!
[1562863396] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562863462] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11840480 kB) free.
[1562863476] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562863516] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562864063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8340588 kB) free!
[1562864116] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562864183] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11833824 kB) free.
[1562864236] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562864783] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8369944 kB) free!
[1562864836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562864903] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11841020 kB) free.
[1562864957] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562865429] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2445620 kB) free!
[1562865503] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8374272 kB) free!
[1562865556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562865558] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562865623] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11824720 kB) free.
[1562865676] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562866223] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8382012 kB) free!
[1562866276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562866343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11829132 kB) free.
[1562866396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562866943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8439048 kB) free!
[1562866996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562867063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11831948 kB) free.
[1562867116] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562867663] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8514628 kB) free!
[1562867716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562867784] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6964528 kB) free!
[1562867837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562867904] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11827504 kB) free.
[1562868436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562868504] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7192056 kB) free!
[1562868556] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562868624] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11822156 kB) free.
[1562869156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562869224] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7255348 kB) free!
[1562869253] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 7368.42, 4433.82, 1880.61
[1562869276] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562869344] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11827356 kB) free.
[1562869373] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1082.03, 3051.17, 1688.02
[1562869493] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 146.22, 2041.72, 1483.39