[1554402714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554402834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1554402954] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554403026] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10087172 kB) free.
[1554403521] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3131308 kB) free!
[1554403554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554403627] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4093440 kB) free!
[1554403674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554403747] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (2782576 kB) free!
[1554403794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554403867] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (994388 kB) free!
[1554404121] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.1% (9429616 kB) free!
[1554404467] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5971028 kB) free!
[1554404995] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1554405067] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10078624 kB) free.
[1554405321] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4120244 kB) free!
[1554405595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554405667] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3531996 kB) free!
[1554405715] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554405787] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (2244608 kB) free!
[1554405835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554405907] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4105920 kB) free!
[1554406507] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.6% (5217576 kB) free!
[1554406521] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.3% (8015688 kB) free!
[1554407035] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1554407107] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10075100 kB) free.
[1554407635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554407707] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3077200 kB) free!
[1554407721] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3155296 kB) free!
[1554407755] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554407827] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (1523864 kB) free!
[1554407876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554407948] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10080592 kB) free.
[1554408547] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4826256 kB) free!
[1554408667] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3510304 kB) free!
[1554408787] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2238448 kB) free!
[1554409076] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1554409387] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7216956 kB) free!
[1554409676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554409796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554409916] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1554409987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10081764 kB) free.
[1554410516] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554410587] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4250568 kB) free!
[1554410636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554410707] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2959580 kB) free!
[1554410757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554410827] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1332768 kB) free!
[1554411427] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6350012 kB) free!
[1554411957] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1554412027] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10079636 kB) free.
[1554412521] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.1% (5478596 kB) free!
[1554412557] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554412628] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3724664 kB) free!
[1554412677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554412748] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (2435628 kB) free!
[1554412797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554412868] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (386736 kB) free!
[1554413468] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5268048 kB) free!
[1554413997] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1554414068] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10081208 kB) free.
[1554414598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554414668] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3076016 kB) free!
[1554414718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554414788] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (1488204 kB) free!
[1554414838] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554414908] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10071672 kB) free.
[1554415508] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4716272 kB) free!
[1554415628] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3386992 kB) free!
[1554415748] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2005904 kB) free!
[1554415794] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46411 MB (20% inode=99%):
[1554415914] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 58453 MB (26% inode=99%):
[1554416038] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1554416121] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3113616 kB) free!
[1554416349] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.1% (6931360 kB) free!
[1554416638] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554416721] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5111508 kB) free!
[1554416758] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554416878] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1554416949] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10078428 kB) free.
[1554417478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554417549] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4188040 kB) free!
[1554417599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554417669] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2881492 kB) free!
[1554417719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554417789] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1197980 kB) free!
[1554418389] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6187372 kB) free!
[1554418919] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1554418989] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10075552 kB) free.
[1554419519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554419589] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3567896 kB) free!
[1554419639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554419709] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (2300644 kB) free!
[1554419759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554419829] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248684 kB) free!
[1554420321] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1784884 kB) free!
[1554420429] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5253200 kB) free!
[1554420960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1554421029] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10072836 kB) free.
[1554421521] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.5% (5659816 kB) free!
[1554421560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554421630] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3072492 kB) free!
[1554421680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1554421750] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (1501848 kB) free!
[1554421800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1554421870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10096072 kB) free.
[1554422470] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4921548 kB) free!
[1554422590] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3676048 kB) free!
[1554422710] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2383228 kB) free!
[1554423000] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1554423310] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7464324 kB) free!
[1554423600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554423720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554423841] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1554423910] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10090676 kB) free.
[1554424440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1554424510] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4505744 kB) free!
[1554424560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554424630] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (3230576 kB) free!
[1554424680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554424750] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (1698892 kB) free!
[1554425351] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6597276 kB) free!
[1554425880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1554425951] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10096144 kB) free.
[1554426480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1554426551] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (1649404 kB) free!
[1554426600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554426671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10292748 kB) free.
[1554426721] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1554427271] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2411992 kB) free!
[1554427320] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554427391] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (437012 kB) free!
[1554427440] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1554427511] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.6% (9650876 kB) free!
[1554427521] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4845588 kB) free!
[1554428040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554428111] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (826768 kB) free!
[1554428160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554428281] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554428476] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3257720 kB) free!
[1554428596] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3225308 kB) free!
[1554428716] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.6% (3191436 kB) free!
[1554428721] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.0% (8877068 kB) free!
[1554429848] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 20.0% (1469412 kB) free!
[1554429911] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.5% (7620676 kB) free!
[1554429921] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4681496 kB) free!
[1554429968] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 19.9% (1466560 kB) free!
[1554430088] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 19.9% (1463468 kB) free!
[1554430512] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10107112 kB) free.
[1554431112] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (1941160 kB) free!
[1554431121] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.7% (8732768 kB) free!
[1554431232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10056984 kB) free.
[1554431832] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2415028 kB) free!
[1554431952] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (455352 kB) free!
[1554432072] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10062908 kB) free.
[1554432322] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3999648 kB) free!
[1554432469] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1554432589] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1554432672] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4119556 kB) free!
[1554432709] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1554432792] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2836800 kB) free!
[1554432912] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (1110684 kB) free!
[1554433082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1554433522] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.3% (8041780 kB) free!
[1554433682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1554433802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554433922] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1554434113] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7913968 kB) free!
[1554434509] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554434522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1554434642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1554434713] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10073328 kB) free.
[1554434722] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2679804 kB) free!
[1554434762] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1554435313] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2945568 kB) free!
[1554435322] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.6% (8678412 kB) free!
[1554435362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1554435433] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (1275240 kB) free!
[1554435483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1554435553] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10087472 kB) free.
[1554435603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1554435709] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554435829] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554435949] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1554436153] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (1982072 kB) free!
[1554436273] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10094676 kB) free.
[1554436522] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4556024 kB) free!
[1554436549] SERVICE FLAPPING ALERT: secondary-server;Current Users;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1554436549] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554436873] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2509984 kB) free!
[1554436993] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (618856 kB) free!
[1554437113] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.7% (9694716 kB) free!
[1554437714] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1376504 kB) free!
[1554437722] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.5% (8612320 kB) free!
[1554438603] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1554438914] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5144192 kB) free!
[1554438922] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4104024 kB) free!
[1554439203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1554439315] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46084 MB (20% inode=99%):
[1554439323] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1554439435] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57345 MB (25% inode=99%):
[1554439549] SERVICE ALERT: secondary-server;Current Users;CRITICAL;SOFT;1;USERS CRITICAL - 11 users currently logged in
[1554439669] SERVICE ALERT: secondary-server;Current Users;CRITICAL;SOFT;2;USERS CRITICAL - 11 users currently logged in
[1554439789] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1554439924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1554440044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1554440114] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10155532 kB) free.
[1554440123] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.9% (8821084 kB) free!
[1554440164] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1554440714] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2027388 kB) free!
[1554440835] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227272 kB) free!
[1554440955] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.1% (9409692 kB) free!
[1554441323] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4073248 kB) free!
[1554441555] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (605396 kB) free!
[1554442516] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1564684 kB) free!
[1554442523] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.4% (9089444 kB) free!
[1554442755] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.0% (5413808 kB) free!
[1554443351] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (701864 kB) free!
[1554443395] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554443515] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554443635] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554443955] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (694204 kB) free!
[1554444323] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (1545868 kB) free!
[1554444923] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.0% (8862364 kB) free!
[1554445155] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6200484 kB) free!
[1554445435] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554446123] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4646840 kB) free!
[1554446355] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (434332 kB) free!
[1554446578] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554446626] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (305752 kB) free!
[1554446635] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1554446676] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197724 kB) free!
[1554446698] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554446755] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1554446818] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554446875] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1554447324] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.8% (8300892 kB) free!
[1554447573] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23201 MB (10% inode=99%):
[1554448155] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7938108 kB) free!
[1554448756] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10109436 kB) free.
[1554448789] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1554449032] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554449124] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3149680 kB) free!
[1554449152] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554449272] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554449356] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2051716 kB) free!
[1554449476] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232196 kB) free!
[1554449596] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.0% (9356068 kB) free!
[1554450196] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (768820 kB) free!
[1554450235] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47021 MB (20% inode=99%):
[1554450324] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.9% (8812264 kB) free!
[1554450355] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56978 MB (25% inode=99%):
[1554450488] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (736100 kB) free!
[1554450955] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46643 MB (20% inode=99%):
[1554450964] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1554451076] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57977 MB (25% inode=99%):
[1554451286] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (4597784 kB) free!
[1554451426] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 26.8% (6559900 kB) free.
[1554451476] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 26.7% (6551228 kB) free.
[1554451524] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3435476 kB) free!
[1554451996] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7378140 kB) free!
[1554452596] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10035984 kB) free.
[1554452724] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.1% (9412424 kB) free!
[1554452818] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554453196] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (1659600 kB) free!
[1554453316] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10089604 kB) free.
[1554453716] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4697244 kB) free!
[1554453776] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3998508 kB) free!
[1554453826] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199960 kB) free!
[1554453836] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.3% (4001404 kB) free!
[1554453896] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.3% (4002188 kB) free!
[1554453916] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2247556 kB) free!
[1554453946] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198512 kB) free!
[1554454018] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554454038] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (248564 kB) free!
[1554454066] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194984 kB) free!
[1554454138] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554454158] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.2% (9441720 kB) free!
[1554454196] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 7.4% (1807352 kB) free!
[1554454258] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554454524] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2657148 kB) free!
[1554454683] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 423 processes
[1554454743] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 427 processes
[1554454758] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (881772 kB) free!
[1554455724] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.7% (8724664 kB) free!
[1554456365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1554456559] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7394020 kB) free!
[1554456924] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4092772 kB) free!
[1554456965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554457085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554457159] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10015364 kB) free.
[1554457205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554457759] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1793168 kB) free!
[1554457879] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10111144 kB) free.
[1554458124] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9509828 kB) free!
[1554458158] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 327 processes
[1554458479] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2262172 kB) free!
[1554458599] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (264164 kB) free!
[1554458719] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.4% (9552996 kB) free!
[1554459243] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 421 processes
[1554459303] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 419 processes
[1554459320] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (992172 kB) free!
[1554459366] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 420 processes
[1554459427] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 0 warn out of 420 processes
[1554459597] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 41.9% (10262936 kB) free.
[1554459686] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188780 kB) free!
[1554459727] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 423 processes
[1554459924] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3274196 kB) free!
[1554460027] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 421 processes
[1554460087] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 424 processes
[1554460205] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1554460387] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 421 processes
[1554460447] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 423 processes
[1554460497] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (2176564 kB) free!
[1554460507] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 2 crit, 0 warn out of 426 processes
[1554460520] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.9% (5843456 kB) free!
[1554460557] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (1477576 kB) free!
[1554460567] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 2 crit, 0 warn out of 426 processes
[1554460617] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 3.8% (932620 kB) free!
[1554460677] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 24.3% (5947832 kB) free.
[1554460805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554460858] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554460867] SERVICE FLAPPING ALERT: webserver2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1554460867] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 423 processes
[1554460925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1554461045] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1554461045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554461076] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192748 kB) free!
[1554461124] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.2% (9442424 kB) free!
[1554461196] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187112 kB) free!
[1554461316] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192580 kB) free!
[1554461322] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19420 MB (8% inode=99%):
[1554461720] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (686916 kB) free!
[1554462325] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4705088 kB) free!
[1554462475] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45786 MB (20% inode=99%):
[1554462596] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 58063 MB (25% inode=99%):
[1554463520] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.8% (8277048 kB) free!
[1554464121] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10070868 kB) free.
[1554464125] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.4% (9058456 kB) free!
[1554464721] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2406268 kB) free!
[1554464758] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 325 processes
[1554464841] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (600096 kB) free!
[1554464961] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.9% (9827160 kB) free!
[1554465325] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4242828 kB) free!
[1554465561] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (1648772 kB) free!
[1554465595] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45759 MB (20% inode=99%):
[1554465715] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 58036 MB (25% inode=99%):
[1554466267] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 3 warn out of 469 processes
[1554466327] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 2 warn out of 481 processes
[1554466377] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (206808 kB) free!
[1554466387] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 465 processes
[1554466437] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (213176 kB) free!
[1554466446] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1554466497] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.9% (210572 kB) free!
[1554466516] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.7% (14250272 kB) free.
[1554466557] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (206572 kB) free!
[1554466687] SERVICE FLAPPING ALERT: webserver2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554466688] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 96.2% (7072708 kB) free.
[1554466752] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 96.2% (7075804 kB) free.
[1554466761] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5389764 kB) free!
[1554466908] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 96.2% (7073720 kB) free.
[1554467046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554467166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554467286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554467287] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 460 processes
[1554467347] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 7 crit, 0 warn out of 468 processes
[1554467407] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 7 crit, 0 warn out of 462 processes
[1554467467] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 7 crit, 0 warn out of 466 processes
[1554467725] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.6% (8156436 kB) free!
[1554467961] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10017532 kB) free.
[1554468367] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 460 processes
[1554468562] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (1988688 kB) free!
[1554468667] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 453 processes
[1554468682] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (222980 kB) free!
[1554468715] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45669 MB (20% inode=99%):
[1554468728] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 462 processes
[1554468802] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.8% (9760388 kB) free!
[1554468835] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56935 MB (25% inode=99%):
[1554468925] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3509356 kB) free!
[1554469075] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554469402] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (1407056 kB) free!
[1554470272] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554470287] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1554470457] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 30.4% (7458360 kB) free.
[1554470603] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5570492 kB) free!
[1554470725] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9492552 kB) free!
[1554470872] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554470887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554470992] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554471007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554471112] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554471127] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1554471712] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554471727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554471803] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (389448 kB) free!
[1554471835] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45433 MB (20% inode=99%):
[1554471848] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554471925] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4749248 kB) free!
[1554471955] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57600 MB (25% inode=99%):
[1554471968] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554473003] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5111044 kB) free!
[1554473725] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9907756 kB) free.
[1554474087] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (3774444 kB) free!
[1554474203] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213108 kB) free!
[1554474325] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7616124 kB) free!
[1554474445] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (4889048 kB) free!
[1554474565] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2035048 kB) free!
[1554474753] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (2921768 kB) free!
[1554475165] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10140304 kB) free.
[1554475765] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9056368 kB) free!
[1554475885] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.0% (6414044 kB) free!
[1554475887] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186400 kB) free!
[1554476003] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7369612 kB) free!
[1554476006] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3661784 kB) free!
[1554476512] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554476632] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554476752] SERVICE FLAPPING ALERT: db1;Current Users;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1554476752] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554476875] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554476973] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 72500 MB (32% inode=99%):
[1554476995] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554477087] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 83.7% (20513904 kB) free.
[1554477115] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554477203] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (1721444 kB) free!
[1554477206] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10139360 kB) free.
[1554477368] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554477806] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8465020 kB) free!
[1554477926] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.9% (5857376 kB) free!
[1554477968] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554478046] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3164576 kB) free!
[1554478088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554478208] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554478403] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5898992 kB) free!
[1554479246] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10031284 kB) free.
[1554479515] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554479603] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233248 kB) free!
[1554479752] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554479846] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8223944 kB) free!
[1554479966] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.2% (5520452 kB) free!
[1554480086] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2859092 kB) free!
[1554481208] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554481287] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.3% (9988956 kB) free.
[1554481315] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554481403] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (7013384 kB) free!
[1554481435] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554481555] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554481808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554481887] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8078172 kB) free!
[1554481928] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554482007] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.2% (5507152 kB) free!
[1554482049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554482127] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2839984 kB) free!
[1554482603] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (1480276 kB) free!
[1554482755] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1554482755] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554483327] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9488780 kB) free!
[1554483803] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5171248 kB) free!
[1554485049] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1554485127] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2882180 kB) free!
[1554485603] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2388108 kB) free!
[1554485649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554485756] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46519 MB (20% inode=99%):
[1554485769] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1554485876] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56535 MB (25% inode=99%):
[1554486327] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9489972 kB) free!
[1554486370] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554486490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554486610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554486803] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6602040 kB) free!
[1554488003] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3236964 kB) free!
[1554488128] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3189220 kB) free!
[1554488158] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 318 processes
[1554488410] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1554488603] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.1% (9413700 kB) free!
[1554488752] SERVICE FLAPPING ALERT: db1;Current Users;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)