[1554661898] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4815972 kB) free!
[1554661999] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (861544 kB) free!
[1554662072] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192160 kB) free!
[1554662119] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (192880 kB) free!
[1554662239] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155860 kB) free!
[1554662316] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1554662391] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8767252 kB) free!
[1554662916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1554662991] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2224988 kB) free!
[1554663036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554663156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1554663245] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45255 MB (20% inode=99%):
[1554663273] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3762520 kB) free!
[1554663365] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55495 MB (24% inode=99%):
[1554663591] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.2% (8500928 kB) free!
[1554663873] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187040 kB) free!
[1554664039] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 86.1% (13990772 kB) free.
[1554664191] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (430884 kB) free!
[1554664356] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1554664792] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.7% (5785536 kB) free!
[1554664956] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554665076] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554665197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554665240] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178664 kB) free!
[1554665360] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (170588 kB) free!
[1554665480] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (187916 kB) free!
[1554665992] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3159140 kB) free!
[1554666080] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1554666080] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 46.8% (7607700 kB) free.
[1554666273] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (3947808 kB) free!
[1554666397] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1554666592] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8762368 kB) free!
[1554666873] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (253628 kB) free!
[1554666965] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554666997] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554667117] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554667192] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (1451344 kB) free!
[1554667237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554667473] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.0% (14553404 kB) free.
[1554667559] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 320 processes
[1554667792] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7518744 kB) free!
[1554668073] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.9% (3916408 kB) free!
[1554668194] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181024 kB) free!
[1554668314] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188076 kB) free!
[1554668392] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218964 kB) free!
[1554668914] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4265752 kB) free!
[1554669514] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 23.2% (5733808 kB) free.
[1554669592] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8780736 kB) free!
[1554669637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554669698] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.9% (8808236 kB) free!
[1554670193] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2570068 kB) free!
[1554670237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554670357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554670477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554670793] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8749244 kB) free!
[1554671393] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (893480 kB) free!
[1554671498] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (1509780 kB) free!
[1554671677] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1554671914] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193852 kB) free!
[1554671993] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6313904 kB) free!
[1554672034] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (2414200 kB) free!
[1554672154] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.1% (3714148 kB) free!
[1554672278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554672365] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46144 MB (20% inode=99%):
[1554672398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554672485] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46144 MB (20% inode=99%):
[1554672518] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554672605] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55554 MB (24% inode=99%):
[1554672754] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2262012 kB) free!
[1554673193] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3227416 kB) free!
[1554673355] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 21.2% (5226072 kB) free.
[1554673718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1554673793] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8750080 kB) free!
[1554673955] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175148 kB) free!
[1554674075] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (549048 kB) free!
[1554674195] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (567768 kB) free!
[1554674318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554674393] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (1994980 kB) free!
[1554674438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1554674559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554674993] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.6% (8174176 kB) free!
[1554675594] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235192 kB) free!
[1554675605] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45313 MB (20% inode=99%):
[1554675725] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55552 MB (24% inode=99%):
[1554675759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1554676194] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5065032 kB) free!
[1554676359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554676479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554676599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554677394] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229244 kB) free!
[1554677480] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554678698] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10238800 kB) free.
[1554679194] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.2% (7973240 kB) free!
[1554679299] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9222820 kB) free!
[1554679419] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9062560 kB) free!
[1554679539] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.1% (8924728 kB) free!
[1554679794] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222192 kB) free!
[1554681594] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7742864 kB) free!
[1554682194] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233616 kB) free!
[1554682925] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47103 MB (20% inode=99%):
[1554683045] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47103 MB (20% inode=99%):
[1554683139] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4652108 kB) free!
[1554683165] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47103 MB (20% inode=99%):
[1554683394] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.4% (8591380 kB) free!
[1554683765] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51201 MB (22% inode=99%):
[1554683994] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2690656 kB) free!
[1554684399] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1554684595] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.4% (8590444 kB) free!
[1554684965] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46111 MB (20% inode=99%):
[1554684999] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554685085] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46111 MB (20% inode=99%):
[1554685120] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554685195] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (816344 kB) free!
[1554685205] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46111 MB (20% inode=99%):
[1554685240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554685539] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.0% (6888452 kB) free!
[1554685595] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554685795] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6371332 kB) free!
[1554685805] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1554685805] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51201 MB (22% inode=99%):
[1554686440] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1554686995] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3352004 kB) free!
[1554687005] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47102 MB (20% inode=99%):
[1554687040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554687125] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47102 MB (20% inode=99%):
[1554687160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554687245] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47102 MB (20% inode=99%):
[1554687280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554687595] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.5% (8614752 kB) free!
[1554687845] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51200 MB (22% inode=99%):
[1554688195] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (1521320 kB) free!
[1554688481] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1554688796] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7472248 kB) free!
[1554689045] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47102 MB (20% inode=99%):
[1554689081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554689139] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4649084 kB) free!
[1554689165] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47102 MB (20% inode=99%):
[1554689201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554689285] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47102 MB (20% inode=99%):
[1554689321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554689396] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219080 kB) free!
[1554689885] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53248 MB (23% inode=99%):
[1554690521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554690596] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.5% (8636100 kB) free!
[1554691121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554691196] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2256236 kB) free!
[1554691241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554691362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554691796] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.4% (8062944 kB) free!
[1554692157] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554692277] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554692396] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215884 kB) free!
[1554692397] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554692562] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554692997] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554693162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554693282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554693402] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554693596] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7924736 kB) free!
[1554694002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554694085] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47100 MB (20% inode=99%):
[1554694122] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554694196] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233248 kB) free!
[1554694205] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45051 MB (20% inode=99%):
[1554694243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554694325] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55288 MB (24% inode=99%):
[1554694925] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47048 MB (20% inode=99%):
[1554695045] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46568 MB (20% inode=99%):
[1554695139] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.3% (11505204 kB) free.
[1554695166] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55240 MB (24% inode=99%):
[1554695739] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9270792 kB) free!
[1554695765] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47046 MB (20% inode=99%):
[1554695859] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.7% (8720444 kB) free!
[1554695885] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45832 MB (20% inode=99%):
[1554695979] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.4% (8095948 kB) free!
[1554695996] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.7% (8246508 kB) free!
[1554696005] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54727 MB (24% inode=99%):
[1554696579] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2932928 kB) free!
[1554696596] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239356 kB) free!
[1554696605] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47018 MB (20% inode=99%):
[1554696725] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55209 MB (24% inode=99%):
[1554697325] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47043 MB (20% inode=99%):
[1554697446] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47039 MB (20% inode=99%):
[1554697566] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55222 MB (24% inode=99%):
[1554697797] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.7% (5781572 kB) free!
[1554698166] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47013 MB (20% inode=99%):
[1554698286] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46502 MB (20% inode=99%):
[1554698379] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9977132 kB) free.
[1554698397] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554698406] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55204 MB (24% inode=99%):
[1554698517] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554698637] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554698979] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7207008 kB) free!
[1554698997] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (325460 kB) free!
[1554699006] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47010 MB (20% inode=99%):
[1554699099] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.2% (6520184 kB) free!
[1554699126] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45988 MB (20% inode=99%):
[1554699219] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.6% (5707000 kB) free!
[1554699246] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54660 MB (24% inode=99%):
[1554699569] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (2171540 kB) free!
[1554699689] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (2154904 kB) free!
[1554699809] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2143720 kB) free!
[1554699819] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (442804 kB) free!
[1554699847] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43678 MB (19% inode=99%):
[1554699967] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55171 MB (24% inode=99%):
[1554700197] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5376220 kB) free!
[1554700419] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.7% (6740352 kB) free!
[1554700432] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554700552] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554700567] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46974 MB (20% inode=99%):
[1554700673] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554700687] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45952 MB (20% inode=99%):
[1554700807] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55161 MB (24% inode=99%):
[1554701019] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (315416 kB) free!
[1554701397] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (440792 kB) free!
[1554701407] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45012 MB (20% inode=99%):
[1554701527] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45017 MB (20% inode=99%):
[1554701620] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.7% (6772584 kB) free!
[1554701647] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54906 MB (24% inode=99%):
[1554702220] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1318036 kB) free!
[1554702248] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46984 MB (20% inode=99%):
[1554702368] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43399 MB (19% inode=99%):
[1554702488] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54344 MB (24% inode=99%):
[1554702598] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6026264 kB) free!
[1554703088] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46391 MB (20% inode=99%):
[1554703208] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55059 MB (24% inode=99%):
[1554703798] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (296636 kB) free!
[1554703808] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46737 MB (20% inode=99%):
[1554703928] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45651 MB (20% inode=99%):
[1554704020] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.7% (8242944 kB) free!
[1554704048] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554704048] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54611 MB (24% inode=99%):
[1554704119] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1554704239] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1554704273] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554704553] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 24.6% (3998092 kB) free.
[1554704649] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46675 MB (20% inode=99%):
[1554704769] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44626 MB (19% inode=99%):
[1554704889] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53840 MB (23% inode=99%):
[1554704998] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5683472 kB) free!
[1554705043] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1554705220] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3223260 kB) free!
[1554705489] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45648 MB (20% inode=99%):
[1554705609] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54861 MB (24% inode=99%):
[1554705753] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3230208 kB) free!
[1554705814] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24709 MB (10% inode=99%):
[1554705873] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3138148 kB) free!
[1554705993] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (3059064 kB) free!
[1554706199] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224772 kB) free!
[1554706209] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44712 MB (19% inode=99%):
[1554706329] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44589 MB (19% inode=99%):
[1554706414] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 28546 MB (12% inode=99%):
[1554706449] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54346 MB (24% inode=99%):
[1554707050] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46131 MB (20% inode=99%):
[1554707170] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43058 MB (19% inode=99%):
[1554707290] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53295 MB (23% inode=99%):
[1554707399] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5684472 kB) free!
[1554707890] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44079 MB (19% inode=99%):
[1554708011] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54316 MB (24% inode=99%):
[1554708393] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1344832 kB) free!
[1554708600] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239860 kB) free!
[1554708610] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46122 MB (20% inode=99%):
[1554708730] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43052 MB (19% inode=99%):
[1554708850] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53799 MB (23% inode=99%):
[1554709450] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46119 MB (20% inode=99%):
[1554709570] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54307 MB (24% inode=99%):
[1554709801] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5318284 kB) free!
[1554710159] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 326 processes
[1554710171] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46307 MB (20% inode=99%):
[1554710291] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42201 MB (18% inode=99%):
[1554710411] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53789 MB (23% inode=99%):
[1554710444] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1554710759] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 324 processes
[1554711001] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241380 kB) free!
[1554711011] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43512 MB (19% inode=99%):
[1554711044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554711131] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43512 MB (19% inode=99%):
[1554711164] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554711251] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53774 MB (23% inode=99%):
[1554711284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554711852] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43842 MB (19% inode=99%):
[1554711972] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43842 MB (19% inode=99%):
[1554712092] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53572 MB (23% inode=99%):
[1554712202] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6856540 kB) free!
[1554712692] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45352 MB (20% inode=99%):
[1554712802] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226856 kB) free!
[1554712812] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43304 MB (19% inode=99%):
[1554712932] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53281 MB (23% inode=99%):
[1554713532] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45345 MB (20% inode=99%):
[1554713652] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53790 MB (23% inode=99%):
[1554713685] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554714252] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45341 MB (20% inode=99%):
[1554714285] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554714372] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43294 MB (19% inode=99%):
[1554714405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554714492] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53530 MB (23% inode=99%):
[1554714525] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1554714525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554714602] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6805264 kB) free!
[1554715093] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45336 MB (20% inode=99%):
[1554715202] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227740 kB) free!
[1554715213] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42265 MB (18% inode=99%):
[1554715333] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53020 MB (23% inode=99%):
[1554715673] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554715793] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554715913] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554715933] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45340 MB (20% inode=99%):
[1554716020] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1554716053] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53526 MB (23% inode=99%):
[1554716654] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42316 MB (18% inode=99%):
[1554716774] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42316 MB (18% inode=99%):
[1554716894] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53521 MB (23% inode=99%):
[1554717003] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6596228 kB) free!
[1554717220] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.6% (6712848 kB) free!
[1554717494] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45491 MB (20% inode=99%):
[1554717603] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228064 kB) free!
[1554717614] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43443 MB (19% inode=99%):
[1554717734] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54095 MB (24% inode=99%):
[1554717820] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3099592 kB) free!
[1554718335] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44723 MB (19% inode=99%):
[1554718455] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44659 MB (19% inode=99%):
[1554718575] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54160 MB (24% inode=99%):
[1554719014] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24692 MB (10% inode=99%):
[1554719176] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46096 MB (20% inode=99%):
[1554719195] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.4% (3307864 kB) free!
[1554719296] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54218 MB (24% inode=99%):
[1554719319] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 23.6% (5794376 kB) free.
[1554719403] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6216592 kB) free!
[1554719614] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24750 MB (11% inode=99%):
[1554719896] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45960 MB (20% inode=99%):
[1554720003] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233076 kB) free!
[1554720017] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44682 MB (19% inode=99%):
[1554720137] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 59491 MB (26% inode=99%):
[1554721725] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554721803] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5591260 kB) free!
[1554722014] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23888 MB (10% inode=99%):
[1554722326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554722403] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235632 kB) free!
[1554722446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554722513] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554722566] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554723080] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3251528 kB) free!
[1554723166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554723201] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3147968 kB) free!
[1554723286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554723321] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.9% (3074076 kB) free!
[1554723407] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554723701] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 391 processes
[1554723821] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 392 processes
[1554724203] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.3% (8545356 kB) free!
[1554724258] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554724378] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554724498] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554724595] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185136 kB) free!
[1554724803] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (1112216 kB) free!
[1554725320] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (243268 kB) free!
[1554725440] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206588 kB) free!
[1554725560] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (209936 kB) free!
[1554725620] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5141240 kB) free!
[1554726003] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5455696 kB) free!
[1554726113] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554726220] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1554726220] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3645784 kB) free!
[1554726233] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554726353] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554726604] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225732 kB) free!
[1554727360] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 72.8% (17841352 kB) free.
[1554727420] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.1% (8908884 kB) free!
[1554727494] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 117741 MB (53% inode=99%):
[1554728404] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7128656 kB) free!
[1554728620] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4628916 kB) free!
[1554729004] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228904 kB) free!
[1554729237] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554729298] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554729353] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554730521] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1549824 kB) free!
[1554730805] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.8% (7789324 kB) free!
[1554731405] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218896 kB) free!
[1554733205] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.0% (7881364 kB) free!
[1554733806] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228672 kB) free!
[1554734208] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1554734406] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6083544 kB) free!
[1554734808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554734929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554735006] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232148 kB) free!
[1554735049] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1554735649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1554735770] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554735890] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1554736490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554736610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554736730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1554736806] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8769920 kB) free!
[1554737407] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (687180 kB) free!
[1554738530] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1554738607] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.7% (8241628 kB) free!
[1554739130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554739207] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226780 kB) free!
[1554739250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1554739371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554739420] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1554740407] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7341060 kB) free!
[1554740793] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 33.8% (5497020 kB) free.
[1554741007] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228900 kB) free!
[1554741557] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1554741678] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1554741847] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 553 processes
[1554741907] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 555 processes
[1554742668] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4653548 kB) free!
[1554742808] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7260804 kB) free!
[1554743408] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235044 kB) free!
[1554745208] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8781936 kB) free!
[1554745559] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 321 processes
[1554745808] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (925188 kB) free!
[1554746159] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 320 processes
[1554747008] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7116524 kB) free!
[1554747609] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241612 kB) free!