[1591556637] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.61, 50.21, 44.83
[1591557837] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 35.09, 42.24, 44.83
[1591558493] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195276 kB) free!
[1591558553] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186404 kB) free!
[1591558614] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (185560 kB) free!
[1591558674] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (191616 kB) free!
[1591558974] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 72.2% (17689652 kB) free.
[1591560086] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4522180 kB) free!
[1591560129] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (663196 kB) free!
[1591560206] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (224976 kB) free!
[1591560249] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396856 kB) free!
[1591560326] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228124 kB) free!
[1591560369] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393424 kB) free!
[1591561795] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 213129 MB (19% inode=99%):
[1591561915] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 213129 MB (19% inode=99%):
[1591562035] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 213129 MB (19% inode=99%):
[1591562126] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 31.1% (15350212 kB) free.
[1591563524] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 78.1% (38491112 kB) free.
[1591564929] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.1% (42438680 kB) free.
[1591565764] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3391) < 203.212.70.170 [08/Jun/2020:03:03:32 +0530] TCP 502 0 0 180.001 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 60.000, 0.000"
[1591565803] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(488) < 2020/06/08 03:01:32 [error] 12485#0: *93720533 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1591565884] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591565907] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(363) < 103.16.101.72 [08/Jun/2020:03:01:32 +0530] TCP 502 0 0 60.001 "10.147.212.61:4000, backend3" "0, 0" "0, 0" "60.001, 0.000"
[1591565923] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591566028] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591566484] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(77) < 182.18.144.246 [08/Jun/2020:03:14:23 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591566604] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591569122] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (3413776 kB) free!
[1591569722] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (242536 kB) free!
[1591571037] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 13.65, 16.96, 23.34
[1591571522] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.6% (9453836 kB) free.
[1591572722] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (183940 kB) free!
[1591572842] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186076 kB) free!
[1591572965] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1591572965] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173468 kB) free!
[1591573389] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 72.3% (17716492 kB) free.
[1591573457] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.1% (9395784 kB) free.
[1591574057] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197840 kB) free!
[1591574177] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188396 kB) free!
[1591574297] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203092 kB) free!
[1591574638] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 46.14, 34.73, 28.95
[1591574758] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 24.63, 29.69, 27.76
[1591575358] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 40.05, 33.03, 29.97
[1591575478] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 49.56, 36.36, 31.42
[1591575598] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 36.94, 34.79, 31.34
[1591576198] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 21.93, 26.63, 29.09
[1591576737] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 72.1% (17684152 kB) free.
[1591576798] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 43.07, 40.98, 36.14
[1591576918] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 25.41, 33.96, 34.09
[1591577038] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1591577038] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 38.40, 36.89, 35.16
[1591578835] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 274052 MB (25% inode=99%):
[1591578969] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 37.0% (48719660 kB) free.
[1591581838] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.68, 46.98, 45.45
[1591582438] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 27.38, 33.03, 39.67
[1591583638] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 43.46, 50.00, 47.11
[1591583765] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1591587639] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.70, 6.56, 5.79
[1591587759] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.46, 5.94, 5.66
[1591587879] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.74, 5.92, 5.66
[1591593537] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4360780 kB) free!
[1591593657] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4161448 kB) free!
[1591593777] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (3913784 kB) free!
[1591594072] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 46.50, 179.96, 109.13
[1591594192] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16.86, 124.45, 97.37
[1591594312] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.40, 87.80, 87.21
[1591594377] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 22.0% (5387268 kB) free.
[1591594438] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1591594753] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (2033344 kB) free!
[1591594873] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (1418960 kB) free!
[1591594965] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.67, 7.60, 5.36
[1591594977] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3591092 kB) free!
[1591594993] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (893972 kB) free!
[1591595085] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.52, 5.76, 4.95
[1591595097] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.3% (3271848 kB) free!
[1591595217] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (2945992 kB) free!
[1591595512] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 11.89, 14.30, 34.47
[1591595589] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4676884 kB) free!
[1591595709] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.4% (5007000 kB) free.
[1591595817] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1591595817] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1481064 kB) free!
[1591596005] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(14) < 203.212.70.170 [08/Jun/2020:11:25:28 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.000, 0.000"
[1591596112] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 15.19, 13.72, 24.45
[1591596125] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591596309] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4452660 kB) free!
[1591596429] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4367764 kB) free!
[1591596549] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4278700 kB) free!
[1591597749] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178820 kB) free!
[1591597922] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1591597993] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.0% (15932300 kB) free.
[1591598043] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 13.01, 12.94, 14.11
[1591598525] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(13) < 203.212.70.170 [08/Jun/2020:12:08:35 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1591598529] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9468572 kB) free!
[1591598532] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1591598646] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591598649] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9108004 kB) free!
[1591598652] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1591598769] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.8% (8765392 kB) free!
[1591598772] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1591599252] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1591599377] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 14.41, 14.74, 14.43
[1591600569] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3781264 kB) free!
[1591600572] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 488 processes
[1591600723] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/06/08 12:44:38 [error] 12489#0: *93738245 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /favicon.ico HTTP/1.1", host: "qservercluster1.com:8091"
[1591600838] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1591600843] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591600958] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1591603475] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 20.0% (6533548 kB) free!
[1591603595] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.6% (6412060 kB) free!
[1591603715] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.2% (6298388 kB) free!
[1591604172] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1591604292] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1591604379] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1591604412] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1591604500] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1591604620] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1591604924] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9090564 kB) free!
[1591605012] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1591605012] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1591605044] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (8859556 kB) free!
[1591605164] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (8603500 kB) free!
[1591605193] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2890824 kB) free!
[1591605220] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1591605313] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.2% (2293236 kB) free!
[1591605434] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1764312 kB) free!
[1591605635] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 57.6% (14121088 kB) free.
[1591607045] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(39) < 203.212.70.170 [08/Jun/2020:14:29:48 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591607165] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591607217] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591607915] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2844936 kB) free!
[1591608164] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4187316 kB) free!
[1591608612] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1591608732] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1591608852] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1591608965] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 185.255.8.58 [08/Jun/2020:15:00:12 +0530] TCP 502 0 0 0.001 "10.147.212.142:13001, 10.147.212.154:13001, 10.147.212.143:13001, 10.147.212.155:13001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 0.000, 0.000"
[1591609034] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16128056 kB) free.
[1591609086] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591609115] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 10.5% (3430132 kB) free!
[1591609452] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1591609715] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1591609715] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2861848 kB) free!
[1591610652] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 386 processes
[1591610885] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(26) < 203.212.70.170 [08/Jun/2020:15:34:52 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591611005] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591611252] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1591611372] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1591611492] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1591611635] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.2% (3978416 kB) free!
[1591611756] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.2% (3735528 kB) free!
[1591611758] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1591611876] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.8% (3376044 kB) free!
[1591611878] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1591612476] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2013780 kB) free!
[1591612805] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(12) < 172.105.57.57 [08/Jun/2020:16:06:26 +0530] TCP 502 0 0 180.001 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 60.000, 0.000"
[1591612925] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591613685] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.33, 8.07, 6.06
[1591613805] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.56, 8.69, 6.54
[1591613892] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1591613925] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.40, 7.98, 6.55
[1591614089] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(5) < Channel error on connection <0.23858.4001> (10.147.212.158:55560 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1591614209] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1591615125] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.76, 3.45, 4.69
[1591615634] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.1% (3846264 kB) free!
[1591615692] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1591615754] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3603628 kB) free!
[1591615875] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3336216 kB) free!
[1591616475] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1591616475] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1789524 kB) free!
[1591619292] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1591619475] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16190624 kB) free.
[1591621115] SERVICE FLAPPING ALERT: mrq4;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591627692] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1591628475] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1591629369] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (25120060 kB) free!
[1591629489] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.1% (25125636 kB) free!
[1591629609] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.1% (25123828 kB) free!
[1591629675] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2988948 kB) free!
[1591629795] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.6% (2615052 kB) free!
[1591629915] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2191292 kB) free!
[1591630092] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1591630212] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1591630332] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1591630515] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16116228 kB) free.
[1591630715] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 62.6% (20493200 kB) free.
[1591633332] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1591635279] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.07, 4.68, 4.99
[1591636039] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 25.0% (4062332 kB) free.
[1591636601] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46441 MB (20% inode=99%):
[1591636721] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46415 MB (20% inode=99%):
[1591636841] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45387 MB (20% inode=99%):
[1591637638] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 44.29, 40.00, 43.51
[1591638238] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 62.93, 52.02, 48.21
[1591640038] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 24.25, 39.36, 44.63
[1591640441] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47794 MB (21% inode=99%):
[1591640638] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1591640638] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.62, 47.00, 45.85