[1555785322] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9615116 kB) free!
[1555785375] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1555785442] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13012904 kB) free.
[1555785540] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555785594] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1555785975] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555786050] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 327 processes
[1555786095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555786215] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1555786776] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212288 kB) free!
[1555786815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555786936] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1555787056] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1555787656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1555787776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555787896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1555789776] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.1% (5465888 kB) free!
[1555790376] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226592 kB) free!
[1555792176] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.0% (6880428 kB) free!
[1555792776] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11774912 kB) free.
[1555793377] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229420 kB) free!
[1555793497] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (473832 kB) free!
[1555793617] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9838284 kB) free!
[1555794217] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10034020 kB) free.
[1555794817] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (2585516 kB) free!
[1555794937] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (1513512 kB) free!
[1555795057] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.4% (11510608 kB) free.
[1555796896] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1555797496] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555797616] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1555798658] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215676 kB) free!
[1555798778] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234728 kB) free!
[1555798898] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224320 kB) free!
[1555799242] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9366500 kB) free!
[1555799362] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.6% (13624684 kB) free.
[1555799498] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11237056 kB) free.
[1555800017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555800098] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2659468 kB) free!
[1555800137] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555800218] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (248840 kB) free!
[1555800257] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1555800338] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235116 kB) free!
[1555801062] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2449952 kB) free!
[1555803257] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1555803350] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 511.06, 329.18, 144.34
[1555803470] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 513.76, 390.43, 189.20
[1555803590] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 147.39, 317.34, 188.75
[1555803718] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 199.67, 49.79, 17.75
[1555803838] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 462.84, 200.04, 76.70
[1555803958] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 498.43, 300.57, 128.49
[1555804136] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 39.83, 8.56, 3.04
[1555804240] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5121 MB (10% inode=97%):
[1555804256] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 432.16, 168.96, 62.49
[1555804376] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 491.03, 279.87, 116.12
[1555804419] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 658 processes
[1555804458] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555804538] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.3% (11963812 kB) free.
[1555804540] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 659 processes
[1555804578] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555804660] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 325 processes
[1555804698] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1555804840] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7011 MB (13% inode=98%):
[1555804965] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.0% (2967884 kB) free!
[1555805138] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (1666684 kB) free!
[1555805258] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11791920 kB) free.
[1555805898] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1555805990] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.73, 0.73, 14.81
[1555806018] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555806138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1555806562] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9162892 kB) free!
[1555806683] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12773932 kB) free.
[1555806776] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 0.57, 16.07
[1555806958] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.36, 2.23, 15.31
[1555807058] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9855692 kB) free!
[1555807178] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.0% (6884404 kB) free!
[1555807190] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.81, 0.76, 4.61
[1555807298] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3942240 kB) free!
[1555807976] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.36, 4.66
[1555808758] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.08, 2.25, 4.21
[1555809082] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8960208 kB) free!
[1555809138] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1555809202] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12680472 kB) free.
[1555809738] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1555809858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555809978] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555810298] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11779352 kB) free.
[1555810402] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9059000 kB) free!
[1555810522] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.6% (12608884 kB) free.
[1555811722] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9460628 kB) free!
[1555811779] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1555811779] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1555811843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12868816 kB) free.
[1555812379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555812499] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555812619] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1555813219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555813339] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555813459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555813898] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9449472 kB) free!
[1555814018] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.2% (6987188 kB) free!
[1555814139] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4482644 kB) free!
[1555814242] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9536328 kB) free!
[1555814362] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13028200 kB) free.
[1555815562] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9227724 kB) free!
[1555815682] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12666060 kB) free.
[1555816459] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1555816966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203264 kB) free!
[1555817059] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555817179] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555817299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1555817899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555817950] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 526 processes
[1555818010] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 528 processes
[1555818019] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555818083] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9485868 kB) free!
[1555818139] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1555818203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13009992 kB) free.
[1555818740] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555818860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1555818980] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555819403] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9121008 kB) free!
[1555819523] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12711624 kB) free.
[1555820739] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.9% (5364764 kB) free!
[1555820780] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1555821340] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11239560 kB) free.
[1555821380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555821500] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555821568] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 426.46, 245.85, 102.56
[1555821621] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555821688] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 432.35, 307.89, 142.67
[1555821768] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2824324 kB) free!
[1555821808] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 159.15, 270.56, 150.80
[1555821923] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9291920 kB) free!
[1555822043] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12956728 kB) free.
[1555822047] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4755108 kB) free!
[1555822167] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4730896 kB) free!
[1555822287] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (4696860 kB) free!
[1555823243] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9537320 kB) free!
[1555823364] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13119248 kB) free.
[1555823568] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2217292 kB) free!
[1555824182] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 96.3% (23752912 kB) free.
[1555824208] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.11, 0.33, 11.67
[1555824241] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.0% (22552288 kB) free.
[1555824320] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (14918928 kB) free.
[1555824376] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.4% (22887700 kB) free.
[1555824563] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9830852 kB) free!
[1555824683] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13037876 kB) free.
[1555825408] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.32, 3.51
[1555825883] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9547100 kB) free!
[1555826003] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12981216 kB) free.
[1555826141] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239132 kB) free!
[1555826261] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10136076 kB) free.
[1555826422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1555827022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555827087] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2323292 kB) free!
[1555827142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555827164] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.2% (22353168 kB) free.
[1555827203] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9116408 kB) free!
[1555827262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555827323] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12892188 kB) free.
[1555828061] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (241800 kB) free!
[1555828181] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9240152 kB) free!
[1555828301] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.4% (6125404 kB) free!
[1555828524] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8983484 kB) free!
[1555828644] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12825888 kB) free.
[1555828902] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.1% (10888220 kB) free.
[1555829062] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1555829390] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 139.91, 32.49, 11.24
[1555829502] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (1177644 kB) free!
[1555829510] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 444.19, 184.40, 69.48
[1555829622] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226756 kB) free!
[1555829630] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 485.08, 286.05, 120.67
[1555829662] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555829742] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238344 kB) free!
[1555829783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555829844] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8538776 kB) free!
[1555829903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1555829964] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12443332 kB) free.
[1555830342] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.1% (7945372 kB) free!
[1555830503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555830623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555830743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555830942] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236408 kB) free!
[1555831164] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8530892 kB) free!
[1555831284] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.6% (12608584 kB) free.
[1555831542] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9846028 kB) free!
[1555831558] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 251.23, 68.05, 24.36
[1555831678] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 443.18, 202.01, 78.75
[1555831798] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 468.98, 291.64, 126.53
[1555832142] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2195424 kB) free!
[1555832484] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8103000 kB) free!
[1555832543] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1555832604] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183828 kB) free.
[1555832630] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.39, 0.75, 11.38
[1555833143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555833247] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 36.7% (18064884 kB) free.
[1555833264] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555833384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555833805] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7878572 kB) free!
[1555833830] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.94, 0.87, 3.75
[1555833852] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9555336 kB) free!
[1555833925] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11823080 kB) free.
[1555833972] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.0% (6424412 kB) free!
[1555834092] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3491908 kB) free!
[1555834525] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9744820 kB) free!
[1555834645] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8228712 kB) free!
[1555834765] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12062660 kB) free.
[1555834798] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.91, 1.91, 14.36
[1555834860] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47134 MB (20% inode=99%):
[1555834980] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47071 MB (20% inode=99%):
[1555835100] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47071 MB (20% inode=99%):
[1555835184] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1555835700] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47554 MB (21% inode=99%):
[1555835725] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (3210864 kB) free!
[1555835784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555835845] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (3112668 kB) free!
[1555835904] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555835965] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (3036596 kB) free!
[1555835965] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8641500 kB) free!
[1555836024] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555836085] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12283292 kB) free.
[1555836300] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46977 MB (20% inode=99%):
[1555836420] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46881 MB (20% inode=99%):
[1555836540] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46721 MB (20% inode=99%):
[1555836599] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.41, 1.64, 3.62
[1555837285] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8269340 kB) free!
[1555837405] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12078900 kB) free.
[1555838005] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9636532 kB) free!
[1555838071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 36.1% (17813260 kB) free.
[1555838267] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.6% (15087500 kB) free.
[1555838365] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1565060 kB) free!
[1555838994] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1555839386] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193076 kB) free!
[1555839477] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.2% (5029988 kB) free!
[1555839506] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186784 kB) free!
[1555839540] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555839597] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (1859732 kB) free!
[1555839626] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169192 kB) free!
[1555839660] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555839718] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (221388 kB) free!
[1555839723] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555839843] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555839876] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9115976 kB) free!
[1555839964] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555839996] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13107536 kB) free.
[1555840824] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1555841196] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8754004 kB) free!
[1555841316] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12782688 kB) free.
[1555841340] SERVICE FLAPPING ALERT: smscapp1;HDD_Home status;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1555841340] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 151725 MB (67% inode=99%):
[1555841365] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 68.6% (11154340 kB) free.
[1555841367] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.3% (15437644 kB) free.
[1555841416] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22001768 kB) free.
[1555841424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1555841545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555841665] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1555842265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555842385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555842505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555842572] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.7% (6240576 kB) free!
[1555842692] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3474376 kB) free!
[1555842812] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.9% (6376228 kB) free!
[1555843412] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220272 kB) free!
[1555844164] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555844394] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1555845812] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.5% (10093600 kB) free.
[1555846117] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9238548 kB) free!
[1555846237] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13051840 kB) free.
[1555846412] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4220784 kB) free!
[1555846532] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (2984036 kB) free!
[1555846652] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (1413352 kB) free!
[1555847317] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.9% (15220732 kB) free.
[1555847437] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9661376 kB) free!
[1555847557] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13276764 kB) free.
[1555848505] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1555848521] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9210084 kB) free!
[1555848540] SERVICE FLAPPING ALERT: smscapp1;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1555848641] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.4% (6126400 kB) free!
[1555848761] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3195824 kB) free!
[1555849105] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555849185] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3079592 kB) free!
[1555849199] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 51.64, 23.99, 10.12
[1555849226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555849305] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2996580 kB) free!
[1555849319] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.37, 16.56, 9.07
[1555849346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555849357] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8190564 kB) free!
[1555849361] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.9% (5883272 kB) free!
[1555849425] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.0% (2926452 kB) free!
[1555849439] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.56, 11.58, 8.16
[1555849477] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12338256 kB) free.
[1555849961] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225520 kB) free!
[1555850008] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 59.68, 237.60, 154.01
[1555850025] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.3% (13697012 kB) free.
[1555850047] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 294 processes
[1555850128] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8.30, 159.08, 135.37
[1555850248] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.37, 106.55, 119.00
[1555850639] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.86, 1.90, 3.49
[1555850677] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9073888 kB) free!
[1555850798] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13000612 kB) free.
[1555851205] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4880816 kB) free!
[1555851325] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.9% (4871228 kB) free!
[1555851445] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.8% (4861904 kB) free!
[1555851565] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (3046832 kB) free!
[1555851685] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (2974704 kB) free!
[1555851805] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.9% (2913944 kB) free!
[1555851997] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9551068 kB) free!
[1555852048] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.10, 0.46, 17.34
[1555852117] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13081836 kB) free.
[1555853248] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.23, 4.94
[1555853917] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8139784 kB) free!
[1555854037] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12154260 kB) free.
[1555854805] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1568036 kB) free!
[1555854830] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 40.12, 9.06, 3.38
[1555854951] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 426.60, 167.17, 62.01
[1555855071] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 478.29, 272.80, 113.47
[1555855237] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9411232 kB) free!
[1555855346] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1555855357] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12985356 kB) free.
[1555855361] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.0% (7868540 kB) free!
[1555855946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555856066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555856186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555856558] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9259552 kB) free!
[1555856561] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2891232 kB) free!
[1555856678] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12927748 kB) free.
[1555857776] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 330.29, 316.75, 156.79
[1555857878] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9683672 kB) free!
[1555857896] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 44.73, 212.02, 137.81
[1555857998] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8063876 kB) free!
[1555858016] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.17, 141.93, 121.12
[1555858071] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.23, 1.46, 11.44
[1555858118] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12156044 kB) free.
[1555859271] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.50, 0.71, 3.76
[1555859318] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8900460 kB) free!
[1555859438] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12771364 kB) free.
[1555859816] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 0.70, 17.76
[1555860386] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1555860638] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9697980 kB) free!
[1555860759] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8254900 kB) free!
[1555860848] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 293 processes
[1555860879] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12310540 kB) free.
[1555860986] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555861106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555861226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555861439] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 465.08, 324.53, 149.18
[1555861559] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 466.12, 371.45, 187.63
[1555861616] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.37, 2.87
[1555861679] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 428.86, 394.98, 218.87
[1555861794] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1555862079] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8380984 kB) free!
[1555862199] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12421076 kB) free.
[1555862994] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1555863399] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9609864 kB) free!
[1555863519] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8210412 kB) free!
[1555863626] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1555863639] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12345724 kB) free.
[1555864079] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.07, 1.92, 18.15
[1555864226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555864346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555864361] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.2% (6501492 kB) free!
[1555864466] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555864839] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9048488 kB) free!
[1555864959] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12542096 kB) free.
[1555864961] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (371672 kB) free!
[1555865880] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.70, 1.98, 4.29
[1555866159] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9767008 kB) free!
[1555866279] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8452148 kB) free!
[1555866399] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12494588 kB) free.
[1555867599] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8654948 kB) free!
[1555867616] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 28.68, 6.29, 2.31
[1555867719] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12485668 kB) free.
[1555867736] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 17.90, 16.53, 7.14
[1555867856] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.65, 11.16, 6.31
[1555868456] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1555868456] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 1.84, 3.48
[1555868666] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1555868919] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9393756 kB) free!
[1555869039] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7765816 kB) free!
[1555869159] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11736960 kB) free.
[1555869266] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555869386] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555869506] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555869761] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.0% (9379144 kB) free!
[1555870360] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8264848 kB) free!
[1555870480] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12129396 kB) free.
[1555870961] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (330676 kB) free!
[1555871306] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes