[1573583513] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 95.1% (46867184 kB) free.
[1573583846] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.1% (2954412 kB) free!
[1573583872] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194728 kB) free!
[1573583885] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189132 kB) free!
[1573583968] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190136 kB) free!
[1573583993] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190844 kB) free!
[1573584005] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185796 kB) free!
[1573584088] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182408 kB) free!
[1573584113] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204792 kB) free!
[1573584125] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184120 kB) free!
[1573585176] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(188) < 2019/11/13 00:28:40 [error] 23646#0: *2551715 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1573585296] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573585516] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186832 kB) free!
[1573585577] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 60.7% (14874196 kB) free.
[1573585861] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.2% (20774856 kB) free.
[1573585888] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.4% (22149316 kB) free.
[1573585913] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.7% (22963108 kB) free.
[1573585926] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 58.7% (14378472 kB) free.
[1573585985] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 86.5% (114038112 kB) free.
[1573586461] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4302216 kB) free!
[1573586559] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/13 00:45:39 [error] 24125#0: *5949492 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573586581] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (989932 kB) free!
[1573586679] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573586701] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1573586701] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186872 kB) free!
[1573587301] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 35.7% (8813780 kB) free.
[1573587901] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191652 kB) free!
[1573588022] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195220 kB) free!
[1573588142] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185536 kB) free!
[1573588724] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (6177972 kB) free!
[1573588844] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.3% (5086788 kB) free!
[1573588964] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2784928 kB) free!
[1573589079] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/13 01:30:40 [error] 24133#0: *5957917 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573589199] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573591796] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2144608 kB) free!
[1573591916] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (220728 kB) free!
[1573592036] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228848 kB) free!
[1573592288] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.7% (5296240 kB) free!
[1573592409] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 88.6% (43679500 kB) free.
[1573592942] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 79.2% (19540028 kB) free.
[1573593164] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.6% (40222376 kB) free.
[1573593237] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.0% (47292240 kB) free.
[1573594476] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1573594596] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1573594896] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1271) < 2019/11/13 03:11:22 [error] 23645#0: *2565679 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573595016] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573595187] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1203) < 103.255.146.164 [13/Nov/2019:03:10:12 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1573595200] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1595) < 2019/11/13 03:12:12 [error] 24124#0: *5978454 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573595225] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1371) < 49.50.64.122 [13/Nov/2019:03:11:22 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1573595308] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573595320] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573595342] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (328256 kB) free!
[1573595345] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573595462] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (324924 kB) free!
[1573595582] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (328736 kB) free!
[1573602396] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573602516] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1573603716] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1573603836] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1573603982] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.5% (23077500 kB) free.
[1573605182] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188488 kB) free!
[1573605302] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183140 kB) free!
[1573605423] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178996 kB) free!
[1573605688] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164280 kB) free!
[1573605713] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165268 kB) free!
[1573605727] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168800 kB) free!
[1573605808] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170528 kB) free!
[1573605833] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199024 kB) free!
[1573605847] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179392 kB) free!
[1573605928] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164344 kB) free!
[1573605953] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.9% (22766920 kB) free.
[1573605967] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200360 kB) free!
[1573606528] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.2% (21366540 kB) free.
[1573606567] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 54.5% (13368508 kB) free.
[1573608996] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.9% (22931008 kB) free.
[1573611711] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 299 processes
[1573611831] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1573615001] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4022576 kB) free!
[1573615121] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.0% (3824788 kB) free!
[1573615241] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.4% (3686856 kB) free!
[1573615842] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 40.2% (9016732 kB) free.
[1573616801] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1573620636] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573620756] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573621505] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25743376 kB) free!
[1573621601] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4630100 kB) free!
[1573621625] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45159580 kB) free.
[1573621721] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4518792 kB) free!
[1573621841] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4426812 kB) free!
[1573621842] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3612356 kB) free!
[1573621962] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3354212 kB) free!
[1573622082] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3033228 kB) free!
[1573622225] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25035000 kB) free!
[1573622345] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44334984 kB) free.
[1573622682] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1573622682] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1661732 kB) free!
[1573622945] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23860252 kB) free!
[1573623065] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43434972 kB) free.
[1573623367] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4833760 kB) free!
[1573623487] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4405392 kB) free!
[1573623520] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2019/11/13 11:08:28 [error] 24130#0: *6068212 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573623607] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (3872268 kB) free!
[1573623640] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(5) < 2019/11/13 11:10:10 [error] 24134#0: *6068327 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573623641] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1245140 kB) free!
[1573623665] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21554324 kB) free!
[1573623760] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(8) < 2019/11/13 11:12:40 [error] 24130#0: *6068403 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573623786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42168504 kB) free.
[1573624026] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.6% (19160940 kB) free.
[1573624124] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 39.7% (8914224 kB) free.
[1573624385] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20209972 kB) free!
[1573624505] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39125136 kB) free.
[1573624632] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167444 kB) free!
[1573624752] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174388 kB) free!
[1573624872] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1573624872] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172220 kB) free!
[1573624960] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573625705] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18811244 kB) free!
[1573625825] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39095584 kB) free.
[1573625932] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4441280 kB) free!
[1573626008] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2138140 kB) free!
[1573626052] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4121776 kB) free!
[1573626172] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.0% (3817272 kB) free!
[1573626216] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(224) < 2019/11/13 11:49:54 [error] 23644#0: *2607376 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1573626336] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573626425] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17667104 kB) free!
[1573626545] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38308464 kB) free.
[1573626772] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 39.0% (8758152 kB) free.
[1573626937] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/13 12:03:38 [error] 23647#0: *2609633 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:3029/17, bytes from/to upstream:17/3029
[1573627057] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/11/13 12:06:09 [error] 23647#0: *2609915 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:684/17, bytes from/to upstream:17/684
[1573627145] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15494856 kB) free!
[1573627177] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2019/11/13 12:07:39 [error] 23642#0: *2610092 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:3390/17, bytes from/to upstream:17/3390
[1573627265] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37146252 kB) free.
[1573627303] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.9% (18484136 kB) free.
[1573627357] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 61.7% (15116044 kB) free.
[1573627364] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9800140 kB) free!
[1573627484] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9096216 kB) free!
[1573627528] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4782220 kB) free!
[1573627604] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (9085212 kB) free!
[1573627648] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4625752 kB) free!
[1573627768] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4501680 kB) free!
[1573627865] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16376856 kB) free!
[1573627912] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176988 kB) free!
[1573627960] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/13 12:13:54 [error] 24132#0: *6079903 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573627985] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36160468 kB) free.
[1573628033] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186760 kB) free!
[1573628080] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573628153] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171992 kB) free!
[1573628377] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573628585] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16252176 kB) free!
[1573628705] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12113868 kB) free!
[1573628825] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33605912 kB) free.
[1573628977] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1125) < 2019/11/13 12:39:32 [error] 23647#0: *2616077 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573629097] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(27) < 2019/11/13 12:40:10 [error] 23646#0: *2616403 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573629218] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573629426] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11853840 kB) free!
[1573629508] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1028) < 103.255.146.164 [13/Nov/2019:12:39:01 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1573629546] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33554464 kB) free.
[1573629568] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1724084 kB) free!
[1573629628] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573630146] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12666352 kB) free!
[1573630266] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33552764 kB) free.
[1573630866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11309348 kB) free!
[1573630986] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32572372 kB) free.
[1573631037] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8499540 kB) free!
[1573631153] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4402484 kB) free!
[1573631157] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.8% (8257876 kB) free!
[1573631205] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4893192 kB) free!
[1573631273] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (4315960 kB) free!
[1573631277] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.3% (8045056 kB) free!
[1573631393] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4233944 kB) free!
[1573631522] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.5% (20599464 kB) free.
[1573631567] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.4% (22654972 kB) free.
[1573631586] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13228080 kB) free!
[1573631600] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.3% (40070024 kB) free.
[1573631646] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 46.7% (61523804 kB) free.
[1573632172] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (3908120 kB) free!
[1573632292] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (3677388 kB) free!
[1573632412] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.3% (3435936 kB) free!
[1573632726] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182764 kB) free!
[1573632846] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173448 kB) free!
[1573632966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186748 kB) free!
[1573633077] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4640672 kB) free!
[1573633612] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1111088 kB) free!
[1573633793] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1960812 kB) free!
[1573634563] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (4027236 kB) free!
[1573634683] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.3% (3750232 kB) free!
[1573634803] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (3526588 kB) free!
[1573634859] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.7% (22633112 kB) free.
[1573634891] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 39.6% (8891680 kB) free.
[1573634929] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 39.1% (19272368 kB) free.
[1573634957] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.6% (22442584 kB) free.
[1573636003] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1573636003] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1612576 kB) free!
[1573636695] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4077856 kB) free!
[1573636815] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.8% (3782164 kB) free!
[1573636935] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.6% (3507024 kB) free!
[1573637046] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23104172 kB) free!
[1573637166] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46818996 kB) free.
[1573637266] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177840 kB) free!
[1573637386] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186768 kB) free!
[1573637506] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190304 kB) free!
[1573637535] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 39.4% (8849928 kB) free.
[1573637766] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23938416 kB) free!
[1573637886] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.0% (47420072 kB) free.
[1573638280] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2019/11/13 15:09:25 [error] 24131#0: *6111800 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573638400] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573638486] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24328128 kB) free!
[1573638525] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.8% (20437676 kB) free.
[1573638558] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 60.9% (14918580 kB) free.
[1573638606] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20257436 kB) free!
[1573638727] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44600704 kB) free.
[1573639134] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172860 kB) free!
[1573639254] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185792 kB) free!
[1573639326] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21913076 kB) free!
[1573639374] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196800 kB) free!
[1573639404] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (8031228 kB) free!
[1573639446] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36925772 kB) free.
[1573639524] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (8019088 kB) free!
[1573639644] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (7405200 kB) free!
[1573639860] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.8% (4852900 kB) free!
[1573639972] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4208384 kB) free!
[1573639981] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.7% (4835936 kB) free!
[1573640046] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14334784 kB) free!
[1573640092] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (4051756 kB) free!
[1573640101] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.7% (4816464 kB) free!
[1573640166] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36930308 kB) free.
[1573640212] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (3910468 kB) free!
[1573640766] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15453012 kB) free!
[1573640886] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11227972 kB) free!
[1573641006] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35642616 kB) free.
[1573641412] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1558096 kB) free!
[1573641606] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11361024 kB) free!
[1573641608] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9278468 kB) free!
[1573641726] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16849052 kB) free!
[1573641728] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (9055924 kB) free!
[1573641846] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32403664 kB) free.
[1573641848] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (8837428 kB) free!
[1573642446] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9369024 kB) free!
[1573642566] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33925904 kB) free.
[1573642644] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3668856 kB) free!
[1573642935] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.2% (3189276 kB) free!
[1573643055] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2902956 kB) free!
[1573643166] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11916436 kB) free!
[1573643175] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.8% (2654396 kB) free!
[1573643287] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7790632 kB) free!
[1573643407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32374780 kB) free.
[1573643775] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1333692 kB) free!
[1573644007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9723520 kB) free!
[1573644127] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.8% (23512692 kB) free!
[1573644247] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30088408 kB) free.
[1573644563] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4790336 kB) free!
[1573644684] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (4391228 kB) free!
[1573644804] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4153276 kB) free!
[1573644847] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7400724 kB) free!
[1573644967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31680288 kB) free.
[1573645448] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4935236 kB) free!
[1573645567] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8438688 kB) free!
[1573645687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4150652 kB) free!
[1573645807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30112892 kB) free.
[1573646004] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1932552 kB) free!
[1573646407] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6780664 kB) free!
[1573646527] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2565140 kB) free!
[1573646550] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 68.9% (17010740 kB) free.
[1573646575] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 66.7% (16348580 kB) free.
[1573646606] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 37.3% (18397180 kB) free.
[1573646638] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.5% (36705028 kB) free.
[1573646647] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28099680 kB) free.
[1573646693] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 59.3% (14528368 kB) free.
[1573646727] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 39.4% (8847288 kB) free.
[1573647153] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184040 kB) free!
[1573647247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4468796 kB) free!
[1573647273] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179564 kB) free!
[1573647367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28969368 kB) free.
[1573647393] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171992 kB) free!
[1573647794] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3165456 kB) free!
[1573647914] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3650860 kB) free.
[1573647967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8234768 kB) free!
[1573648087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (4006864 kB) free!
[1573648207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28349604 kB) free.
[1573648515] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3150844 kB) free!
[1573648635] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3593580 kB) free.
[1573648807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6472416 kB) free!
[1573648927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2192668 kB) free!
[1573649047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27346588 kB) free.
[1573649647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3643512 kB) free!
[1573649767] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389476 kB) free!
[1573649813] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 45.4% (59896292 kB) free.
[1573649835] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3239092 kB) free!
[1573649861] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.0% (18993460 kB) free.
[1573649955] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3699944 kB) free.
[1573650131] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1573650251] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1573650371] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1573650461] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169960 kB) free!
[1573650532] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (246792 kB) free!
[1573650582] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185608 kB) free!
[1573650653] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235612 kB) free!
[1573650702] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199392 kB) free!
[1573650773] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (284544 kB) free!
[1573650820] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (264588 kB) free!
[1573650940] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (270468 kB) free!
[1573651060] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172264 kB) free!
[1573651155] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2995372 kB) free!
[1573651275] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3363952 kB) free.
[1573651373] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 36.9% (18195136 kB) free.
[1573651767] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4841464 kB) free!
[1573651888] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4633240 kB) free!
[1573652008] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (4526808 kB) free!
[1573652475] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (2419532 kB) free!
[1573652595] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.4% (2830672 kB) free!
[1573652642] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9661236 kB) free!
[1573652704] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4274440 kB) free!
[1573652715] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (3217728 kB) free!
[1573652762] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 20.5% (10084484 kB) free.
[1573652824] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (4131736 kB) free!
[1573652945] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.3% (3988012 kB) free!
[1573653362] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9560172 kB) free!
[1573653482] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9041924 kB) free!
[1573653602] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (9025184 kB) free!
[1573653901] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2440160 kB) free!
[1573653915] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4197404 kB) free.
[1573654060] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 36.6% (5951232 kB) free.
[1573654535] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573654628] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24005576 kB) free!
[1573654745] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2324864 kB) free!
[1573654748] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.8% (48558016 kB) free.
[1573655135] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3977924 kB) free!
[1573655255] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.1% (3842092 kB) free!
[1573655375] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.5% (3707120 kB) free!
[1573655524] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22509728 kB) free.
[1573655561] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 58.5% (14341212 kB) free.
[1573655586] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 92.4% (22653880 kB) free.
[1573655608] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2251560 kB) free!
[1573655771] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 13 users currently logged in
[1573657175] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2109276 kB) free!
[1573659134] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163676 kB) free!
[1573659254] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175164 kB) free!
[1573659374] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188616 kB) free!
[1573659576] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 37.7% (8459008 kB) free.
[1573660202] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4925672 kB) free!
[1573660224] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.7% (22482784 kB) free.
[1573660257] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.7% (22863368 kB) free.
[1573660286] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.4% (40101956 kB) free.
[1573661462] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184516 kB) free!
[1573661582] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166004 kB) free!
[1573661702] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185304 kB) free!
[1573661792] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22828912 kB) free.
[1573661948] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26263500 kB) free!
[1573662068] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.0% (47401632 kB) free.
[1573662773] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9753668 kB) free!
[1573662893] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9713324 kB) free!
[1573663013] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (9676812 kB) free!
[1573663600] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/13 22:16:01 [error] 24130#0: *6191924 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:53/17, bytes from/to upstream:17/53
[1573663720] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573664920] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/13 22:30:14 [error] 24130#0: *6194732 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573665041] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573665764] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1573666913] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234640 kB) free!
[1573666964] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4793356 kB) free!
[1573666985] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3484224 kB) free!
[1573667033] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243608 kB) free!
[1573667084] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (4763176 kB) free!
[1573667105] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4980656 kB) free!
[1573667153] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240548 kB) free!
[1573667205] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (4724104 kB) free!
[1573667225] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401876 kB) free!
[1573667440] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/13 23:15:15 [error] 24130#0: *6201879 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573667560] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573667649] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1573667753] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.4% (7596344 kB) free!
[1573667769] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1573667889] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1573668160] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(71) < 2019/11/13 23:29:43 [error] 24130#0: *6202432 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573668280] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573668405] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 21.3% (5211980 kB) free.
[1573668489] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1573669268] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21068368 kB) free!
[1573669388] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18248616 kB) free!
[1573669508] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37422520 kB) free.