[1581100210] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1686316 kB) free!
[1581100450] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1581100762] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581100767] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.2% (29165504 kB) free.
[1581100838] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581101038] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178560 kB) free!
[1581101158] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180916 kB) free!
[1581101278] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195232 kB) free!
[1581102330] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192004 kB) free!
[1581102390] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.6% (17049720 kB) free.
[1581103696] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1581104050] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.5% (9753064 kB) free.
[1581104650] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1733928 kB) free!
[1581104770] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167384 kB) free!
[1581104891] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204844 kB) free!
[1581105069] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7542572 kB) free!
[1581105189] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.4% (6608620 kB) free!
[1581105309] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (5693716 kB) free!
[1581105909] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (1468820 kB) free!
[1581108562] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1581108682] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1581108802] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1581108978] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (255988 kB) free!
[1581109078] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.1% (10320948 kB) free.
[1581109099] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240604 kB) free!
[1581109219] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227656 kB) free!
[1581109819] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.3% (42510628 kB) free.
[1581110002] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1581111144] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(247) < 103.16.101.72 [08/Feb/2020:03:00:31 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581111202] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581111266] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581111322] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1581111462] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(301) < 2020/02/08 03:01:32 [error] 10207#0: *50538604 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581111472] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2480) < 172.105.59.173 [08/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581111582] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581111592] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581112238] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25926152 kB) free!
[1581112358] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38790152 kB) free.
[1581112959] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25975252 kB) free!
[1581113079] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.6% (49543360 kB) free.
[1581115182] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/02/08 04:00:55 [error] 10222#0: *50551462 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581115303] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581118678] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188132 kB) free!
[1581118798] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197008 kB) free!
[1581118919] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172976 kB) free!
[1581119138] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.8% (21021580 kB) free.
[1581119291] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.8% (11300168 kB) free.
[1581120922] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581123010] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 31.4% (41365360 kB) free.
[1581123323] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581123332] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39790696 kB) free.
[1581123381] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.7% (21980512 kB) free.
[1581123443] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581124091] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186820 kB) free!
[1581124211] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187296 kB) free!
[1581124332] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186088 kB) free!
[1581125739] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22052608 kB) free.
[1581129103] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(41) < 2020/02/08 08:00:50 [error] 10221#0: *50626476 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581129223] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581139479] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22437408 kB) free!
[1581139600] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17640328 kB) free!
[1581139720] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13298792 kB) free!
[1581140670] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4077684 kB) free!
[1581140790] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (3741548 kB) free!
[1581140910] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.1% (3398364 kB) free!
[1581141350] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4859284 kB) free!
[1581141470] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.0% (3457760 kB) free!
[1581141511] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1647376 kB) free!
[1581141520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37949188 kB) free.
[1581141590] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (3387468 kB) free!
[1581142167] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9802888 kB) free!
[1581142190] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1931204 kB) free!
[1581142287] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.5% (9115980 kB) free!
[1581142408] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.3% (8548128 kB) free!
[1581142744] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (4291944 kB) free!
[1581142864] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3929188 kB) free!
[1581142984] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3607272 kB) free!
[1581143320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23451896 kB) free!
[1581143440] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.7% (25942364 kB) free!
[1581143560] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43466824 kB) free.
[1581143584] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2020120 kB) free!
[1581143608] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.8% (30456500 kB) free.
[1581143788] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (4052260 kB) free!
[1581143896] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1581143908] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.9% (3899940 kB) free!
[1581143911] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.1% (15733308 kB) free.
[1581143977] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194104 kB) free!
[1581144017] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1581144028] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3723596 kB) free!
[1581144097] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179544 kB) free!
[1581144137] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1581144160] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 86.9% (21310080 kB) free.
[1581144160] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19517696 kB) free!
[1581144199] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.6% (21377632 kB) free.
[1581144280] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43666344 kB) free.
[1581144294] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581144338] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.4% (7573148 kB) free!
[1581144339] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190072 kB) free!
[1581144414] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1581144458] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8491012 kB) free!
[1581144459] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185208 kB) free!
[1581144578] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (8491224 kB) free!
[1581144579] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172024 kB) free!
[1581144737] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1581144767] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.2% (4711364 kB) free!
[1581144784] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22848060 kB) free.
[1581144880] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19215364 kB) free!
[1581144888] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.3% (4496096 kB) free!
[1581145000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43049324 kB) free.
[1581145008] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.4% (4265112 kB) free!
[1581145228] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1567512 kB) free!
[1581145601] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19111908 kB) free!
[1581145721] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15237748 kB) free!
[1581145841] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40958076 kB) free.
[1581146208] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2321252 kB) free!
[1581146420] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9310644 kB) free!
[1581146441] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16031436 kB) free!
[1581146537] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1581146540] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8968904 kB) free!
[1581146561] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12218020 kB) free!
[1581146657] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1581146660] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (8655760 kB) free!
[1581146681] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37294852 kB) free.
[1581146777] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1581146788] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 638 processes
[1581146808] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 73.4% (17984312 kB) free.
[1581146908] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 635 processes
[1581146927] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1581146992] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(55) < 172.105.59.173 [08/Feb/2020:12:57:02 +0530] TCP 502 0 0 181.001 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "1.001, 60.000, 60.000, 60.000"
[1581147028] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 367 processes
[1581147047] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1581147112] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581147167] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.9% (21547568 kB) free.
[1581147205] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.3% (21164216 kB) free.
[1581147223] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(73) < 2020/02/08 12:55:16 [error] 10219#0: *50742689 recv() failed (104: Connection reset by peer) while proxying connection, client: 124.153.112.226, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1581147241] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 376 processes
[1581147281] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13164600 kB) free!
[1581147343] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581147362] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1581147401] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37413224 kB) free.
[1581147578] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4946444 kB) free!
[1581147865] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(210) < 103.16.101.72 [08/Feb/2020:13:10:14 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581147943] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(311) < 2020/02/08 13:11:22 [error] 10214#0: *50750036 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581147985] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581148063] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581148410] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188556 kB) free!
[1581148460] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4324640 kB) free!
[1581148530] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186436 kB) free!
[1581148650] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192272 kB) free!
[1581149010] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (3583580 kB) free!
[1581149131] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.0% (3448532 kB) free!
[1581149251] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.5% (3333716 kB) free!
[1581149777] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1581149777] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1581149850] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1581149850] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 61.2% (15002668 kB) free.
[1581150377] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1581150451] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1581150451] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2472068 kB) free!
[1581150497] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581150617] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1581151050] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189288 kB) free!
[1581151170] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189024 kB) free!
[1581151290] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174496 kB) free!
[1581152311] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4237096 kB) free!
[1581152378] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3099560 kB) free!
[1581152431] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4064564 kB) free!
[1581152498] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 34.5% (5601636 kB) free.
[1581152552] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.3% (3893336 kB) free!
[1581152819] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17454204 kB) free!
[1581152939] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42762844 kB) free.
[1581153539] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18849016 kB) free!
[1581153617] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1581153659] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15189256 kB) free!
[1581153752] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2005404 kB) free!
[1581153779] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40977316 kB) free.
[1581154218] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1581154290] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 33.9% (8309620 kB) free.
[1581154312] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(181) < 92.42.104.166 [08/Feb/2020:15:00:22 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581154338] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1581154352] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1581154352] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.1% (15738824 kB) free.
[1581154379] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17436004 kB) free!
[1581154432] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581154458] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1581154499] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41674780 kB) free.
[1581154891] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.9% (2681628 kB) free!
[1581155011] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.7% (2625504 kB) free!
[1581155058] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1581155099] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19856548 kB) free!
[1581155131] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (2579728 kB) free!
[1581155219] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15948056 kB) free!
[1581155339] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41059920 kB) free.
[1581155437] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9540720 kB) free!
[1581155557] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9467636 kB) free!
[1581155568] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4427048 kB) free!
[1581155677] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (9408540 kB) free!
[1581155688] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (4301036 kB) free!
[1581155731] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2289680 kB) free!
[1581155809] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4135168 kB) free!
[1581155940] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18466272 kB) free!
[1581156060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (14450984 kB) free!
[1581156180] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40432520 kB) free.
[1581156780] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15834372 kB) free!
[1581156900] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11832884 kB) free!
[1581157009] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2448388 kB) free!
[1581157020] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38075736 kB) free.
[1581157620] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13196620 kB) free!
[1581157740] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9359676 kB) free!
[1581157860] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36305356 kB) free.
[1581158175] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 49.0% (12092628 kB) free.
[1581158216] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 34.8% (17166752 kB) free.
[1581158244] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 21.6% (10651820 kB) free.
[1581158263] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(99) < 2020/02/08 16:01:26 [error] 10221#0: *50822355 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581158274] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1581158274] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 62.6% (15340416 kB) free.
[1581158300] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.3% (20669592 kB) free.
[1581158384] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581158460] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12248948 kB) free!
[1581158580] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8271888 kB) free!
[1581158700] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35340516 kB) free.
[1581158778] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187636 kB) free!
[1581158808] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4905476 kB) free!
[1581158818] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (223032 kB) free!
[1581158898] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168732 kB) free!
[1581158928] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.6% (4811872 kB) free!
[1581158938] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238728 kB) free!
[1581159018] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193376 kB) free!
[1581159048] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.2% (4710568 kB) free!
[1581159058] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239004 kB) free!
[1581159300] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10150680 kB) free!
[1581159420] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36429176 kB) free.
[1581160020] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5586408 kB) free!
[1581160140] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (519492 kB) free!
[1581160261] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24143288 kB) free!
[1581160861] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (10747592 kB) free!
[1581161461] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21450376 kB) free!
[1581162018] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 21.2% (5243780 kB) free.
[1581162058] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (7348676 kB) free!
[1581162061] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1581162061] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37262008 kB) free.
[1581162098] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2852340 kB) free!
[1581162218] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2791360 kB) free!
[1581162338] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (2593620 kB) free!
[1581162618] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189104 kB) free!
[1581162658] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1581162658] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228368 kB) free!
[1581162661] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8308984 kB) free!
[1581162738] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195272 kB) free!
[1581162781] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3332704 kB) free!
[1581162859] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191220 kB) free!
[1581162901] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402012 kB) free!
[1581162938] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.6% (4316852 kB) free.
[1581163198] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.4% (21564988 kB) free.
[1581163233] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36785184 kB) free.
[1581163267] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 46.3% (22830432 kB) free.
[1581163477] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4733576 kB) free!
[1581163808] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181188 kB) free!
[1581163869] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4835988 kB) free!
[1581163928] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190700 kB) free!
[1581163989] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (1173684 kB) free!
[1581164048] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167568 kB) free!
[1581164109] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220008 kB) free!
[1581164448] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2270856 kB) free!
[1581165509] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1581165752] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581166458] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581167278] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1581168037] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23050752 kB) free!
[1581168157] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.5% (57330404 kB) free.
[1581170591] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581170656] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 45.9% (22641316 kB) free.
[1581170687] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 32.6% (16042412 kB) free.
[1581170711] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581170717] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 71.8% (17719324 kB) free.
[1581170743] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21785832 kB) free.
[1581172433] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(411) < 172.105.59.173 [08/Feb/2020:20:02:08 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581172553] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581173117] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (3577564 kB) free!
[1581173237] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.4% (3551416 kB) free!
[1581173357] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.3% (3530136 kB) free!
[1581173557] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581173958] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2402052 kB) free!
[1581176059] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1581178599] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.0% (21476396 kB) free.
[1581180157] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15392044 kB) free!
[1581180277] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (43018916 kB) free.
[1581180878] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18465728 kB) free!
[1581180998] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44776224 kB) free.
[1581181598] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19689972 kB) free!
[1581181718] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44680588 kB) free.
[1581182318] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20985072 kB) free!
[1581182438] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44593660 kB) free.
[1581183038] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22658252 kB) free!
[1581183158] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44614620 kB) free.
[1581183758] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22900476 kB) free!
[1581183878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44585728 kB) free.
[1581184009] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1581184478] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23142196 kB) free!
[1581184598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44602252 kB) free.
[1581185198] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22942564 kB) free!
[1581185318] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44583188 kB) free.
[1581185608] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239164 kB) free!
[1581185728] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (211736 kB) free!
[1581185849] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229764 kB) free!
[1581185918] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21136596 kB) free!
[1581186010] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9157540 kB) free!
[1581186039] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44458760 kB) free.
[1581186130] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3358020 kB) free!
[1581186250] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3000056 kB) free!