[1548632133] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 310 processes
[1548632253] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1548644298] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 334 processes
[1548653621] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548653741] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548653861] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548654471] SERVICE ALERT: db1;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5574 MB (10% inode=98%):
[1548655661] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548655781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548655901] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548658301] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548660130] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 305 processes
[1548660250] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1548661396] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548661516] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548661636] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548661653] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 308 processes
[1548661773] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 308 processes
[1548661893] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 310 processes
[1548661901] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548662021] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548662514] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548662621] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548662635] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548662741] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548662862] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548667662] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548667782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548667902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548668034] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548668154] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548668274] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548668874] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548669102] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548669702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548669822] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548669942] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548672694] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 317 processes
[1548672942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548673062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548673182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548673294] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 315 processes
[1548673674] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1548673794] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1548673914] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1548674494] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 316 processes
[1548674514] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1548674514] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548675094] SERVICE FLAPPING ALERT: webservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1548675094] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 315 processes
[1548675114] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548675234] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548675355] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1548675954] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1548676074] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548676194] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548676782] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1548676782] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548676794] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548677394] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548677514] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 329 processes
[1548677546] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1548677634] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548677666] SERVICE ALERT: sig5;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 338 processes
[1548677786] SERVICE ALERT: sig5;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1548677815] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.2% (6508256 kB) free!
[1548677871] SERVICE ALERT: db1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 8848 MB (17% inode=99%):
[1548677935] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.5% (6675312 kB) free!
[1548677982] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548678055] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.4% (7085164 kB) free!
[1548678102] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548678222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548678234] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548678834] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548678955] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548679075] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548679675] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548680022] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548680875] SERVICE ALERT: db1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1548680995] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548686395] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548686494] SERVICE FLAPPING ALERT: webservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548686515] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548687115] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548687236] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548687436] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548690222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548690342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548690462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548691435] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548691556] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548691676] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548691721] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47125 MB (20% inode=99%):
[1548691841] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47125 MB (20% inode=99%):
[1548691961] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47125 MB (20% inode=99%):
[1548692263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548692276] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548695124] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.2% (22847648 kB) free.
[1548695212] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.6% (15206680 kB) free.
[1548695354] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 93.5% (15191320 kB) free.
[1548695455] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 74.7% (36823856 kB) free.
[1548695471] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 77082 MB (34% inode=99%):
[1548696161] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48850 MB (21% inode=99%):
[1548696476] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548696596] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548696950] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 311 processes
[1548697558] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1548697695] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1548697698] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548697810] SERVICE ALERT: sigqscluster2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes