[1578681032] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.2% (29681284 kB) free.
[1578681080] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21599888 kB) free!
[1578681407] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168852 kB) free!
[1578681527] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179520 kB) free!
[1578681647] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196972 kB) free!
[1578681680] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1259180 kB) free!
[1578681775] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1578682280] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18167456 kB) free!
[1578682376] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1578682880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37783968 kB) free.
[1578682977] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1578683094] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184988 kB) free!
[1578683154] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187620 kB) free!
[1578683214] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (194704 kB) free!
[1578683274] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (184704 kB) free!
[1578683382] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1578683480] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13280360 kB) free!
[1578683574] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.2% (17455396 kB) free.
[1578683600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9363464 kB) free!
[1578683702] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/11 00:43:53 [error] 18366#0: *36900541 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:245/17, bytes from/to upstream:17/245
[1578683720] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5769088 kB) free!
[1578683817] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1578683823] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578683937] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1578684177] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1578684297] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1578684320] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20795936 kB) free!
[1578684417] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1578684920] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39442192 kB) free.
[1578685520] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16201456 kB) free!
[1578685617] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1578685641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12336320 kB) free!
[1578685761] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (8576976 kB) free!
[1578686217] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1578686338] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1578686361] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22731792 kB) free!
[1578686458] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1578686822] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/11 01:33:18 [error] 18365#0: *36916027 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:201/17, bytes from/to upstream:17/201
[1578686942] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578686961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31964000 kB) free.
[1578687542] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/11 01:39:33 [error] 18362#0: *36917940 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:214/17, bytes from/to upstream:17/214
[1578687561] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16873032 kB) free!
[1578687663] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578687681] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13100896 kB) free!
[1578687801] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (9328596 kB) free!
[1578688258] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1578688262] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/11 01:54:22 [error] 18368#0: *36922562 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:249/17, bytes from/to upstream:17/249
[1578688382] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578688401] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22727632 kB) free!
[1578688858] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1578689001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39042940 kB) free.
[1578689337] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1578689447] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (1933536 kB) free!
[1578689458] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1578689458] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1578689567] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239516 kB) free!
[1578689578] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1578689601] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15941704 kB) free!
[1578689687] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241928 kB) free!
[1578689721] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12094376 kB) free!
[1578689841] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (8516544 kB) free!
[1578690441] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19593988 kB) free!
[1578690779] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1578690782] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/11 02:35:43 [error] 18359#0: *36937824 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:210/17, bytes from/to upstream:17/210
[1578690887] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 81.4% (40103624 kB) free.
[1578690899] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1578690902] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578691019] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1578691041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32589756 kB) free.
[1578691641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11542324 kB) free!
[1578691761] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8164916 kB) free!
[1578691882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37130852 kB) free.
[1578691937] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.3% (43998216 kB) free.
[1578692448] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.9% (10993436 kB) free.
[1578692481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14756108 kB) free!
[1578692601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11134880 kB) free!
[1578692702] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/01/11 03:11:12 [error] 18370#0: *36950091 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578692704] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1325) < 203.212.70.170 [11/Jan/2020:03:12:32 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1578692721] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (7913940 kB) free!
[1578692761] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1696) < 2020/01/11 03:12:32 [error] 21682#0: *24188078 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578692822] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578692824] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578692881] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578693321] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17035128 kB) free!
[1578693921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27420648 kB) free.
[1578695121] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16930076 kB) free!
[1578695241] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13294872 kB) free!
[1578695361] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (9941132 kB) free!
[1578696419] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1578696561] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22505592 kB) free!
[1578697019] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1578697139] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1578697161] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27617216 kB) free.
[1578697259] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1578697761] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (573196 kB) free!
[1578697881] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402492 kB) free!
[1578698002] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29476452 kB) free.
[1578698601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4046380 kB) free!
[1578698721] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397260 kB) free!
[1578698841] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402340 kB) free!
[1578699049] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184676 kB) free!
[1578699169] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184672 kB) free!
[1578699289] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168672 kB) free!
[1578699660] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1578700041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26570024 kB) free.
[1578700260] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1578700380] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1578700478] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.2% (14619600 kB) free.
[1578700980] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1578701100] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1578701241] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14114104 kB) free!
[1578701361] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10110284 kB) free!
[1578701481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36379128 kB) free.
[1578702081] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15662584 kB) free!
[1578702201] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11695112 kB) free!
[1578702321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36236004 kB) free.
[1578702900] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578702921] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15192632 kB) free!
[1578703020] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1578703041] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11535256 kB) free!
[1578703161] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38206444 kB) free.
[1578703621] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578703741] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578703762] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19294056 kB) free!
[1578703861] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1578703882] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15634524 kB) free!
[1578704002] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38125160 kB) free.
[1578704602] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19415356 kB) free!
[1578704722] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15806476 kB) free!
[1578704842] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38047496 kB) free.
[1578705442] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19334768 kB) free!
[1578705562] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15652248 kB) free!
[1578705661] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1578705682] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37980924 kB) free.
[1578706282] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18934048 kB) free!
[1578706402] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15122444 kB) free!
[1578706522] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37726616 kB) free.
[1578706861] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1578706981] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1578707122] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18553372 kB) free!
[1578707242] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14855632 kB) free!
[1578707362] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37591424 kB) free.
[1578707581] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578707702] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578707822] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578707963] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17915720 kB) free!
[1578708083] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13964572 kB) free!
[1578708203] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36940452 kB) free.
[1578708803] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17603464 kB) free!
[1578708923] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13853700 kB) free!
[1578709024] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(869) < 103.255.146.164 [11/Jan/2020:07:44:35 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578709043] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36549984 kB) free.
[1578709081] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(869) < 2020/01/11 07:44:35 [error] 21673#0: *24236061 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578709082] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.7% (21732820 kB) free.
[1578709119] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40730392 kB) free.
[1578709145] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578709157] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.8% (22751088 kB) free.
[1578709196] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.3% (21892176 kB) free.
[1578709202] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578709623] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(53) < 2020/01/11 07:55:05 [error] 18367#0: *37044228 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578709643] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16761220 kB) free!
[1578709743] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578709744] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 180.179.218.150 [11/Jan/2020:07:52:38 +0530] TCP 502 0 0 0.000 "10.147.212.155:13000, backend_13000" "0, 0" "0, 0" "0.000, 0.000"
[1578709763] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13031372 kB) free!
[1578709801] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(468) < 2020/01/11 07:59:51 [error] 21681#0: *24238614 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578709864] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578709883] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35628324 kB) free.
[1578709921] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578710483] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13576152 kB) free!
[1578710603] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36526588 kB) free.
[1578710678] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (731948 kB) free!
[1578710798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (698008 kB) free!
[1578710822] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 334 processes
[1578710915] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1578710918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (645220 kB) free!
[1578710960] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.10271.2322> (10.147.212.158:37114 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1578711006] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (209844 kB) free!
[1578711035] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1578711080] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1578711126] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186944 kB) free!
[1578711157] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.2% (22988700 kB) free.
[1578711201] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.0% (21334332 kB) free.
[1578711203] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16030900 kB) free!
[1578711323] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12239148 kB) free!
[1578711422] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1578711443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36081240 kB) free.
[1578711805] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.0% (3435512 kB) free!
[1578711925] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.8% (3384936 kB) free!
[1578712043] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14707400 kB) free!
[1578712045] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3339216 kB) free!
[1578712164] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10766792 kB) free!
[1578712284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34671376 kB) free.
[1578712622] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1578712884] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14104248 kB) free!
[1578712921] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2020/01/11 08:48:38 [error] 21683#0: *24248684 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578713004] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10174400 kB) free!
[1578713042] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578713124] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33811568 kB) free.
[1578713222] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1578713245] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177840 kB) free!
[1578713343] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1578713463] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1578713724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11694500 kB) free!
[1578713844] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7357524 kB) free!
[1578713964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32629412 kB) free.
[1578714217] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1578714217] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.6% (21466876 kB) free.
[1578714564] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8698708 kB) free!
[1578714684] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4699716 kB) free!
[1578714804] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25816964 kB) free!
[1578714825] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.8% (2656124 kB) free!
[1578714945] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.3% (2519472 kB) free!
[1578715065] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2391248 kB) free!
[1578715404] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (9672004 kB) free!
[1578716004] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20192576 kB) free!
[1578716463] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1578716604] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407136 kB) free!
[1578717063] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1578717184] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1578717535] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 44.5% (58600452 kB) free.
[1578717597] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.1% (20378360 kB) free.
[1578718206] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (229592 kB) free!
[1578718235] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1578718326] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169096 kB) free!
[1578718355] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1578718368] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182136 kB) free!
[1578718383] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1578718446] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185332 kB) free!
[1578718489] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169288 kB) free!
[1578718503] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1578718609] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167104 kB) free!
[1578718623] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1578719327] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9043292 kB) free!
[1578719335] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21720716 kB) free!
[1578719447] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9042668 kB) free!
[1578719455] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17868312 kB) free!
[1578719567] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (9043288 kB) free!
[1578719575] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13897200 kB) free!
[1578719823] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578719943] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578720063] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1578720775] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28894392 kB) free.
[1578721305] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (2135828 kB) free!
[1578721375] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2684160 kB) free!
[1578721425] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1656272 kB) free!
[1578721495] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396568 kB) free!
[1578721545] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1174180 kB) free!
[1578721615] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (447052 kB) free!
[1578721832] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9256440 kB) free!
[1578721953] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.3% (8511804 kB) free!
[1578722073] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.9% (7839744 kB) free!
[1578722285] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.6% (4315740 kB) free!
[1578722337] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4907036 kB) free!
[1578722405] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.0% (4156036 kB) free!
[1578722457] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4433664 kB) free!
[1578722464] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1578722525] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.2% (3969828 kB) free!
[1578722567] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2879024 kB) free!
[1578722577] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.1% (3977476 kB) free!
[1578722673] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4512116 kB) free!
[1578722815] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23783212 kB) free!
[1578722993] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.3% (22541652 kB) free.
[1578723032] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.2% (16886448 kB) free.
[1578723064] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578723067] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40718416 kB) free.
[1578723112] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.6% (28891808 kB) free.
[1578723147] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.2% (21366708 kB) free.
[1578723177] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1673712 kB) free!
[1578723184] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1578723415] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395216 kB) free!
[1578723725] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (6488940 kB) free!
[1578723725] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2257924 kB) free!
[1578723737] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9520428 kB) free!
[1578723758] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176644 kB) free!
[1578723845] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (6086228 kB) free!
[1578723857] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9123304 kB) free!
[1578723878] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178488 kB) free!
[1578723965] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (5695016 kB) free!
[1578723977] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.7% (8720144 kB) free!
[1578723998] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179000 kB) free!
[1578724015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34031856 kB) free.
[1578724566] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.2% (32121144 kB) free.
[1578724615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9555852 kB) free!
[1578724735] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5427024 kB) free!
[1578724766] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.9% (2913552 kB) free!
[1578724855] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1470180 kB) free!
[1578724886] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.0% (2706580 kB) free!
[1578724984] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1578725006] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2499540 kB) free!
[1578725104] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578725224] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1578725396] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191100 kB) free!
[1578725455] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16674164 kB) free!
[1578725516] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172588 kB) free!
[1578725577] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.5% (23071328 kB) free.
[1578725606] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1484044 kB) free!
[1578725636] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (1111660 kB) free!
[1578725777] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4596428 kB) free!
[1578725825] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1578726055] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35446912 kB) free.
[1578726077] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9586252 kB) free!
[1578726197] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (7906796 kB) free!
[1578726317] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1578726317] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (7699756 kB) free!
[1578726655] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10212944 kB) free!
[1578726664] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(842) < 103.255.146.164 [11/Jan/2020:12:38:38 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578726762] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578726776] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (5999908 kB) free!
[1578726784] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578726841] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(842) < 2020/01/11 12:38:38 [error] 21681#0: *24288976 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578726882] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578726896] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1721356 kB) free!
[1578726961] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578727002] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578727025] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1578727145] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1578727496] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18343036 kB) free!
[1578727745] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1578727865] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 361 processes
[1578727985] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578728096] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410536 kB) free!
[1578729186] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1578729296] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32370884 kB) free.
[1578729786] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1578729896] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6881524 kB) free!
[1578729906] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1578729917] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241248 kB) free!
[1578730016] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2717992 kB) free!
[1578730026] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578730136] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397384 kB) free!
[1578730626] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1578730736] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15600856 kB) free!
[1578731226] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1578731336] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (6932752 kB) free!
[1578731636] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2866028 kB) free!
[1578731687] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9210380 kB) free!
[1578731807] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (9060852 kB) free!
[1578731827] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1578731927] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8902032 kB) free!
[1578731947] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1578731949] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (703600 kB) free!
[1578732067] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578732070] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (697532 kB) free!
[1578732190] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (704460 kB) free!
[1578732237] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1578732237] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180412 kB) free!
[1578732536] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34436860 kB) free.
[1578732582] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3166068 kB) free!
[1578732702] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3063696 kB) free!
[1578732823] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (2973464 kB) free!
[1578733136] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7648836 kB) free!
[1578733257] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3312096 kB) free!
[1578733267] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1578733377] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392748 kB) free!
[1578733598] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1578733977] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19480308 kB) free!
[1578734023] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 35.4% (5758684 kB) free.
[1578734435] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (3742540 kB) free!
[1578734555] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.8% (3544120 kB) free!
[1578734577] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403212 kB) free!
[1578734675] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (3319320 kB) free!
[1578735177] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15064656 kB) free!
[1578735184] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(677) < 103.255.146.164 [11/Jan/2020:15:01:05 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578735275] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.7% (16990236 kB) free.
[1578735304] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578735361] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(677) < 2020/01/11 15:01:05 [error] 21683#0: *24312618 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578735481] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578735777] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (9590668 kB) free!
[1578736127] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4466800 kB) free!
[1578736977] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37943780 kB) free.
[1578737023] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2583152 kB) free!
[1578737143] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2619112 kB) free!
[1578737183] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578737263] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2771104 kB) free!
[1578737303] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1578737577] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10996140 kB) free!
[1578737697] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6314320 kB) free!
[1578737817] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (1989296 kB) free!
[1578738317] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1578738417] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (17996320 kB) free!
[1578738428] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 15.6% (3823324 kB) free!
[1578738463] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1578738463] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 34.4% (5589408 kB) free.
[1578738481] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(97) < 2020/01/11 15:48:44 [error] 21682#0: *24241705 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:58597475/16068320, bytes from/to upstream:16068320/58597475
[1578738548] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 15.6% (3820508 kB) free!
[1578738601] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578738667] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1578738668] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 15.5% (3805624 kB) free!
[1578738787] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1578739017] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406108 kB) free!
[1578739063] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3200228 kB) free!
[1578739183] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3364176 kB) free.
[1578740217] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28197296 kB) free.
[1578740457] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1578740578] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1578740817] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (640752 kB) free!
[1578740937] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402772 kB) free!
[1578741057] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424648 kB) free!
[1578741657] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19758488 kB) free!
[1578742257] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (456900 kB) free!
[1578742857] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17435036 kB) free!
[1578743457] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (444212 kB) free!
[1578743637] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578744057] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14360300 kB) free!
[1578744187] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578744237] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4671300 kB) free!
[1578744838] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192436 kB) free!
[1578744838] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 617 processes
[1578744958] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 580 processes
[1578745078] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 456 processes
[1578745857] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (38027060 kB) free.
[1578746457] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10773492 kB) free!
[1578746577] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6403508 kB) free!
[1578746698] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2297112 kB) free!
[1578746983] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3033984 kB) free!
[1578747103] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 39.4% (6410180 kB) free.
[1578747298] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22699676 kB) free!
[1578747898] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400872 kB) free!
[1578748498] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18305388 kB) free!
[1578749503] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578749698] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12883348 kB) free!
[1578750298] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39010564 kB) free.
[1578750898] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10959544 kB) free!
[1578751018] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6158204 kB) free!
[1578751138] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1751916 kB) free!
[1578751738] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22322476 kB) free!
[1578752338] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389764 kB) free!
[1578752938] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19178252 kB) free!
[1578753538] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37740004 kB) free.
[1578754138] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15308340 kB) free!
[1578754258] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10739460 kB) free!
[1578754379] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6492648 kB) free!
[1578754979] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25826868 kB) free!
[1578755579] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2217596 kB) free!
[1578756179] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22779216 kB) free!
[1578756779] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41140160 kB) free.
[1578757379] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18087268 kB) free!
[1578757499] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (14029884 kB) free!
[1578757602] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1578757619] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (10212856 kB) free!
[1578758219] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25201784 kB) free!
[1578758819] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5815860 kB) free!
[1578759001] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(930) < 2020/01/11 21:35:14 [error] 21679#0: *24362625 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578759121] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578759305] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(930) < 103.255.146.164 [11/Jan/2020:21:35:14 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578759419] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20119836 kB) free!
[1578759425] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578760019] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36272052 kB) free.
[1578760025] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [11/Jan/2020:21:56:40 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 60.000, 0.000"
[1578760143] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(74) < 2020/01/11 21:58:39 [error] 18370#0: *37392459 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578760145] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578760263] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/11 21:59:23 [error] 18368#0: *37392348 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:62/0, bytes from/to upstream:0/62
[1578760321] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(715) < 2020/01/11 22:00:04 [error] 21672#0: *24367593 no live upstreams while connecting to upstream, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "backend_13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578760383] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578760441] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578760619] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12801196 kB) free!
[1578760739] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8651128 kB) free!
[1578760746] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(127) < 92.42.104.166 [11/Jan/2020:22:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1578760859] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41041532 kB) free.
[1578760866] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578761460] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16453108 kB) free!
[1578761580] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12345568 kB) free!
[1578761700] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8470144 kB) free!
[1578762300] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22825444 kB) free!
[1578762900] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40088096 kB) free.
[1578763500] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16703004 kB) free!
[1578763620] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12866076 kB) free!
[1578763740] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9266432 kB) free!
[1578764340] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21661028 kB) free!
[1578764940] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35502852 kB) free.
[1578765540] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14550668 kB) free!
[1578765660] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10961532 kB) free!
[1578765780] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40880036 kB) free.
[1578766380] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19618456 kB) free!
[1578766500] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15988336 kB) free!
[1578766621] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (12319832 kB) free!
[1578766918] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2163320 kB) free!
[1578767038] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (220872 kB) free!
[1578767158] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219824 kB) free!
[1578767221] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24504820 kB) free!