[1548441460] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548441580] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548441813] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.97, 9.28, 9.90
[1548442697] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 338 processes
[1548443898] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 334 processes
[1548445698] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 334 processes
[1548450880] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50214 MB (22% inode=99%):
[1548451480] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46118 MB (20% inode=99%):
[1548451600] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46118 MB (20% inode=99%):
[1548451720] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46118 MB (20% inode=99%):
[1548475121] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50209 MB (22% inode=99%):
[1548475793] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548475913] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548476033] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548476633] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548477833] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548477953] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548478073] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548478674] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1548478674] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548478781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1548478901] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1548479021] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548479874] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548479994] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548480114] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548481314] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548481421] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548484918] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 449 processes
[1548484978] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 450 processes
[1548484996] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1548487996] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1548491596] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548491954] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3244804 kB) free!
[1548492074] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3204788 kB) free!
[1548492114] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1548492194] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3161796 kB) free!
[1548492796] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1548492916] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1548493036] SERVICE FLAPPING ALERT: webservercluster2;Current Users;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1548493036] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1548495478] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 432 processes
[1548495538] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 433 processes
[1548496289] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 315 processes
[1548496394] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1363888 kB) free!
[1548496409] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1548497009] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1548497130] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1548497250] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1548500236] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548500836] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548500956] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548501076] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1548501450] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1548501570] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1548501794] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.8% (4359556 kB) free.
[1548502171] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1548502291] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1548502394] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1845712 kB) free!
[1548502514] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (1560640 kB) free!
[1548502634] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1548502634] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1545160 kB) free!
[1548504691] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1548504811] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1548504931] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1548505034] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 32.7% (5320748 kB) free.
[1548506132] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1548509732] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 311 processes
[1548509852] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 312 processes
[1548509972] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 312 processes
[1548512476] SERVICE FLAPPING ALERT: webservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548514635] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1548519434] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3205364 kB) free!
[1548519554] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3186008 kB) free!
[1548519572] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1548519674] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3166304 kB) free!
[1548520276] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548523413] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.13, 11.54, 10.55
[1548523533] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.18, 11.80, 10.76
[1548523653] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.67, 11.72, 10.86
[1548524853] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.99, 9.52, 9.88
[1548525076] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1548525196] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1548525316] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1548526516] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in