[1583605888] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36327936 kB) free.
[1583606347] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184196 kB) free!
[1583606467] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172200 kB) free!
[1583606488] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12526096 kB) free!
[1583606552] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181004 kB) free!
[1583606587] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179068 kB) free!
[1583606608] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8152484 kB) free!
[1583606712] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (205448 kB) free!
[1583606728] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4064004 kB) free!
[1583606739] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1583606832] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187548 kB) free!
[1583606952] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194176 kB) free!
[1583607328] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22730044 kB) free!
[1583607863] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (2093252 kB) free!
[1583607923] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173928 kB) free!
[1583607928] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2337632 kB) free!
[1583607988] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (193752 kB) free!
[1583608048] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (176680 kB) free!
[1583608348] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 73.4% (17999152 kB) free.
[1583608528] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (18992500 kB) free!
[1583609128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36844984 kB) free.
[1583609729] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13868724 kB) free!
[1583609849] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9613528 kB) free!
[1583609969] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (5522700 kB) free!
[1583610569] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22655956 kB) free!
[1583610752] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.5% (5552412 kB) free.
[1583611151] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233756 kB) free!
[1583611169] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2352020 kB) free!
[1583611271] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (221440 kB) free!
[1583611353] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199168 kB) free!
[1583611391] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234240 kB) free!
[1583611473] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190192 kB) free!
[1583611593] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165020 kB) free!
[1583611712] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1583611769] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17910848 kB) free!
[1583611837] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (410604 kB) free!
[1583611957] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382428 kB) free!
[1583612077] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399900 kB) free!
[1583612206] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 211526 MB (19% inode=99%):
[1583612326] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 211526 MB (19% inode=99%):
[1583612369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34495904 kB) free.
[1583612446] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 211526 MB (19% inode=99%):
[1583612588] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583612739] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583612969] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12786752 kB) free!
[1583613089] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8740848 kB) free!
[1583613192] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.3% (29734956 kB) free.
[1583613209] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5133312 kB) free!
[1583613809] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18492304 kB) free!
[1583614409] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33884620 kB) free.
[1583614494] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4639808 kB) free!
[1583614614] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (222604 kB) free!
[1583614734] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224524 kB) free!
[1583614836] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (221660 kB) free!
[1583614956] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244428 kB) free!
[1583615009] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12292828 kB) free!
[1583615076] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243156 kB) free!
[1583615129] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8476084 kB) free!
[1583615249] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4913336 kB) free!
[1583615334] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 79.3% (39100052 kB) free.
[1583615849] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17636808 kB) free!
[1583616276] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.8% (42795352 kB) free.
[1583616450] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31490932 kB) free.
[1583616798] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3222) < 203.212.70.170 [08/Mar/2020:03:02:32 +0530] TCP 502 0 0 120.001 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1583616920] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583616980] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(313) < 2020/03/08 03:01:32 [error] 12247#0: *67030157 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583617050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10644272 kB) free!
[1583617101] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583617152] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 51.1% (12521336 kB) free.
[1583617163] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(251) < 103.16.101.72 [08/Mar/2020:03:00:31 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583617170] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6929392 kB) free!
[1583617284] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583617290] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36920864 kB) free.
[1583617521] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583617641] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1583617890] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15076972 kB) free!
[1583618010] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11160888 kB) free!
[1583618130] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7449100 kB) free!
[1583618730] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19453736 kB) free!
[1583619330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32180248 kB) free.
[1583619487] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583619607] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583619930] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10864012 kB) free!
[1583620050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7149396 kB) free!
[1583620170] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35150124 kB) free.
[1583620770] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13576108 kB) free!
[1583620890] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9635492 kB) free!
[1583621010] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6156956 kB) free!
[1583621611] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16933612 kB) free!
[1583622211] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28387408 kB) free.
[1583622811] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9267896 kB) free!
[1583622931] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36714412 kB) free.
[1583622993] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583623477] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.7% (14145172 kB) free!
[1583623531] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15089148 kB) free!
[1583623651] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11193348 kB) free!
[1583623771] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36711252 kB) free.
[1583624352] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168328 kB) free!
[1583624371] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15673212 kB) free!
[1583624473] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167436 kB) free!
[1583624491] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11812240 kB) free!
[1583624593] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186968 kB) free!
[1583624611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28079848 kB) free.
[1583625211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17864100 kB) free!
[1583625331] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14217372 kB) free!
[1583625393] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.6% (12977536 kB) free.
[1583625451] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (10640540 kB) free!
[1583625788] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.7% (19541508 kB) free.
[1583626051] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19260180 kB) free!
[1583626651] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27869176 kB) free.
[1583627677] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394100 kB) free!
[1583627852] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15484216 kB) free!
[1583627972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11363812 kB) free!
[1583628092] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36435320 kB) free.
[1583628277] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.3% (18854744 kB) free!
[1583628646] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 232496 MB (21% inode=99%):
[1583628692] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7841312 kB) free!
[1583628812] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3513984 kB) free!
[1583628877] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1583628877] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379524 kB) free!
[1583628932] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36285004 kB) free.
[1583629477] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.5% (40147524 kB) free.
[1583629532] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9488516 kB) free!
[1583629652] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5129752 kB) free!
[1583629772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36093236 kB) free.
[1583630077] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.7% (14153428 kB) free!
[1583630197] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9656648 kB) free!
[1583630317] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5699636 kB) free!
[1583630372] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13361668 kB) free!
[1583630492] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9604436 kB) free!
[1583630612] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6049588 kB) free!
[1583630917] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.6% (40370796 kB) free.
[1583631212] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18377196 kB) free!
[1583631812] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27096124 kB) free.
[1583633012] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15428016 kB) free!
[1583633133] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37887672 kB) free.
[1583633421] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 48.1% (23697344 kB) free.
[1583633447] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 60.4% (14808732 kB) free.
[1583633478] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.6% (21464848 kB) free.
[1583633718] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(58) < 49.50.64.122 [08/Mar/2020:07:44:40 +0530] TCP 502 0 0 180.000 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1583633732] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18466212 kB) free!
[1583633838] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(9) < 172.105.59.173 [08/Mar/2020:07:46:40 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583633852] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14820356 kB) free!
[1583633958] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583633972] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37179428 kB) free.
[1583634572] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17824172 kB) free!
[1583634692] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13646096 kB) free!
[1583634812] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35647380 kB) free.
[1583635412] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14789052 kB) free!
[1583635532] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37430448 kB) free.
[1583635593] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175096 kB) free!
[1583635713] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174172 kB) free!
[1583635833] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193336 kB) free!
[1583636132] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18210624 kB) free!
[1583636252] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (14539820 kB) free!
[1583636372] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36405628 kB) free.
[1583636972] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16983124 kB) free!
[1583637092] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37211216 kB) free.
[1583637692] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19189908 kB) free!
[1583637812] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15441588 kB) free!
[1583637933] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37009432 kB) free.
[1583638533] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15988196 kB) free!
[1583638653] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11603076 kB) free!
[1583638773] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36365368 kB) free.
[1583638858] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.8% (21177920 kB) free.
[1583642317] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583643624] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18114024 kB) free!
[1583643744] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45517720 kB) free.
[1583644344] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20247100 kB) free!
[1583644464] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15890300 kB) free!
[1583644584] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (11651096 kB) free!
[1583645184] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13365840 kB) free!
[1583645784] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31719728 kB) free.
[1583645917] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.7% (23328488 kB) free!
[1583646037] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.8% (22113956 kB) free!
[1583646063] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4907860 kB) free!
[1583646157] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.4% (25604776 kB) free!
[1583646183] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.7% (3620724 kB) free!
[1583646262] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1583646303] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4177976 kB) free!
[1583646382] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1583646384] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7633624 kB) free!
[1583646502] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1583646504] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3489084 kB) free!
[1583646624] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36973180 kB) free.
[1583646757] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (27258584 kB) free.
[1583646962] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2930272 kB) free!
[1583647082] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.4% (2561808 kB) free!
[1583647202] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2223032 kB) free!
[1583647224] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10524024 kB) free!
[1583647345] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6087816 kB) free!
[1583647357] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.4% (21675984 kB) free!
[1583647465] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1841392 kB) free!
[1583647477] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.5% (24367836 kB) free!
[1583647503] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175444 kB) free!
[1583647597] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (27254164 kB) free.
[1583647992] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8509060 kB) free!
[1583648112] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7940456 kB) free!
[1583648197] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.6% (24570940 kB) free!
[1583648232] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.0% (7368108 kB) free!
[1583648317] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.0% (17103456 kB) free!
[1583648437] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.5% (20429560 kB) free!
[1583648482] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4814808 kB) free!
[1583648602] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.8% (4609348 kB) free!
[1583648665] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25669524 kB) free!
[1583648722] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.2% (4465484 kB) free!
[1583648832] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4587108 kB) free!
[1583649002] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (2884812 kB) free!
[1583649188] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (3764464 kB) free!
[1583649265] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (766628 kB) free!
[1583649308] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3576532 kB) free!
[1583649364] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (4271400 kB) free!
[1583649428] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.8% (3390392 kB) free!
[1583649433] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.5% (30312704 kB) free.
[1583649485] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3935004 kB) free!
[1583649502] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1583649603] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.0% (15498836 kB) free.
[1583649605] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.7% (3628336 kB) free!
[1583649865] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13993036 kB) free!
[1583649947] SERVICE ALERT: db1-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 470 processes
[1583650067] SERVICE ALERT: db1-vm2;CPU_Procs;OK;HARD;1;CPU OK: 467 processes
[1583650125] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.5% (20850772 kB) free.
[1583650205] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.0% (22700932 kB) free.
[1583650465] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (436004 kB) free!
[1583650522] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2030016 kB) free!
[1583650526] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 52.0% (68563644 kB) free.
[1583650556] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21662080 kB) free.
[1583650586] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.8% (14101520 kB) free.
[1583650628] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1900124 kB) free!
[1583650823] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8242636 kB) free!
[1583650943] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.6% (7205508 kB) free!
[1583651063] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (6863568 kB) free!
[1583652037] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1583652037] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (11115900 kB) free!
[1583652052] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4470532 kB) free!
[1583652172] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (4340012 kB) free!
[1583652292] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4229952 kB) free!
[1583652334] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11230844 kB) free!
[1583652454] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44511292 kB) free.
[1583652637] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.1% (23886888 kB) free!
[1583653054] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18193152 kB) free!
[1583653174] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13713148 kB) free!
[1583653294] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (9352028 kB) free!
[1583653463] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4061744 kB) free!
[1583653476] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9725112 kB) free!
[1583653493] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2180916 kB) free!
[1583653596] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9473792 kB) free!
[1583653716] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9276052 kB) free!
[1583653894] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23843524 kB) free!
[1583654332] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (3496532 kB) free!
[1583654452] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.6% (2873020 kB) free!
[1583654494] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40797700 kB) free.
[1583654572] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1971296 kB) free!
[1583655094] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15429848 kB) free!
[1583655214] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11258652 kB) free!
[1583655334] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (6799940 kB) free!
[1583655615] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1583655615] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.1% (20761264 kB) free.
[1583655649] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.6% (52142760 kB) free.
[1583655681] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.5% (39196060 kB) free.
[1583655705] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.6% (20497396 kB) free.
[1583655735] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.6% (20986816 kB) free.
[1583656641] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583656761] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583656881] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1583657165] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (1400244 kB) free!
[1583657285] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175748 kB) free!
[1583657316] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4631680 kB) free!
[1583657405] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174724 kB) free!
[1583657423] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164324 kB) free!
[1583657543] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179764 kB) free!
[1583657663] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166948 kB) free!
[1583658052] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19468820 kB) free!
[1583658172] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (14976284 kB) free!
[1583658292] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (10826888 kB) free!
[1583658892] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25269336 kB) free!
[1583659492] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40138128 kB) free.
[1583659618] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 338 processes
[1583659738] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 340 processes
[1583659858] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1583660092] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12908872 kB) free!
[1583660212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8203740 kB) free!
[1583660332] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42880060 kB) free.
[1583660358] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(185) < 49.50.67.175 [08/Mar/2020:15:00:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583660478] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583660932] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18160636 kB) free!
[1583661052] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13512744 kB) free!
[1583661160] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4895140 kB) free!
[1583661172] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9096716 kB) free!
[1583661280] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.9% (4873944 kB) free!
[1583661400] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.8% (4854612 kB) free!
[1583661681] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1583661772] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (21978128 kB) free!
[1583662372] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39791976 kB) free.
[1583662734] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9808744 kB) free!
[1583662854] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9717444 kB) free!
[1583662972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12759008 kB) free!
[1583662974] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (9663952 kB) free!
[1583663092] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8426120 kB) free!
[1583663213] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4409576 kB) free!
[1583663740] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1583663813] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15833688 kB) free!
[1583663860] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583663883] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 85.4% (20923452 kB) free.
[1583663925] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.4% (20815324 kB) free.
[1583663977] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.8% (44768348 kB) free.
[1583663980] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1583664037] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583664081] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583664143] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165280 kB) free!
[1583664201] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1583664263] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187612 kB) free!
[1583664383] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171872 kB) free!
[1583664413] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36586192 kB) free.
[1583664502] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 707 processes
[1583664580] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1583664623] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 707 processes
[1583664743] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 708 processes
[1583665013] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10786620 kB) free!
[1583665133] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6471484 kB) free!
[1583665180] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583665253] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2107296 kB) free!
[1583665300] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583665420] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583665853] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15192528 kB) free!
[1583666453] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27843940 kB) free.
[1583666620] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1583666620] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1583667053] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (388092 kB) free!
[1583667115] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4710784 kB) free!
[1583667173] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (410468 kB) free!
[1583667220] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1583667235] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4570756 kB) free!
[1583667293] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397072 kB) free!
[1583667340] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1583667355] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (4469628 kB) free!
[1583667383] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 25.6% (6264768 kB) free.
[1583667460] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583667683] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8511096 kB) free!
[1583667803] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8394452 kB) free!
[1583667923] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (8832252 kB) free!
[1583667983] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190088 kB) free!
[1583668103] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186820 kB) free!
[1583668223] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 42.1% (10330500 kB) free.
[1583668493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34652584 kB) free.
[1583668732] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583668823] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (363016 kB) free!
[1583668943] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (310264 kB) free!
[1583669063] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (274532 kB) free!
[1583669093] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3431180 kB) free!
[1583669213] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393740 kB) free!
[1583669260] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1583669334] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405844 kB) free!
[1583669437] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.2% (29238644 kB) free.
[1583669934] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14678608 kB) free!
[1583670534] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41008904 kB) free.
[1583670774] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4875272 kB) free!
[1583671134] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4126564 kB) free!
[1583671238] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.2% (24020008 kB) free!
[1583671254] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (380604 kB) free!
[1583671358] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.1% (27795588 kB) free.
[1583671374] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (421792 kB) free!
[1583671555] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2182244 kB) free!
[1583671869] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 88.4% (43572000 kB) free.
[1583671919] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45107320 kB) free.
[1583671955] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.5% (20703176 kB) free.
[1583671958] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (22715680 kB) free!
[1583672031] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1583672031] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.6% (20989172 kB) free.
[1583672078] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26483164 kB) free.
[1583673123] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11130036 kB) free!
[1583673243] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6819940 kB) free!
[1583673278] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (24766008 kB) free!
[1583673363] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2790088 kB) free!
[1583673398] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (27726168 kB) free.
[1583673998] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (22803756 kB) free!
[1583674118] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.5% (25756356 kB) free!
[1583674132] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4835404 kB) free!
[1583674238] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (28654940 kB) free.
[1583674252] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4756820 kB) free!
[1583674373] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4544440 kB) free!
[1583674563] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22385436 kB) free!
[1583674859] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 335 processes
[1583674943] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1583674979] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1583675164] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400952 kB) free!
[1583675438] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24405796 kB) free!
[1583675543] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1583675558] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (27246920 kB) free.
[1583675663] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1583675723] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4749336 kB) free!
[1583675764] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19423464 kB) free!
[1583675783] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1583676173] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1294392 kB) free!
[1583676364] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403552 kB) free!
[1583676383] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1583676383] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1583676758] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25602760 kB) free!
[1583676878] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (28314168 kB) free.
[1583676964] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16139736 kB) free!
[1583677564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37395956 kB) free.
[1583677971] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.7% (20407640 kB) free.
[1583678005] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.1% (36534356 kB) free.
[1583678164] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4628812 kB) free!
[1583678284] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (455604 kB) free!
[1583678404] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (447752 kB) free!
[1583678678] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26360104 kB) free!
[1583678798] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.0% (29052012 kB) free.
[1583679174] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163200 kB) free!
[1583679294] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189424 kB) free!
[1583679398] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.9% (23621244 kB) free!
[1583679414] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1583679414] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179692 kB) free!
[1583679460] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583679518] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.9% (24906160 kB) free!
[1583679604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30812904 kB) free.
[1583679638] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (27618340 kB) free.
[1583680204] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2654044 kB) free!
[1583680324] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406256 kB) free!
[1583680434] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1583680444] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38490748 kB) free.
[1583681044] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10740992 kB) free!
[1583681164] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6639940 kB) free!
[1583681285] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2812888 kB) free!
[1583681885] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16191668 kB) free!
[1583682485] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30628424 kB) free.
[1583683085] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7176348 kB) free!
[1583683205] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3482984 kB) free!
[1583683325] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36525312 kB) free.
[1583683925] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10834032 kB) free!
[1583684045] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6993632 kB) free!
[1583684165] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40011868 kB) free.
[1583684765] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14351564 kB) free!
[1583684885] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10240292 kB) free!
[1583685005] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6476040 kB) free!
[1583685293] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 70.1% (17298596 kB) free.
[1583685322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40827008 kB) free.
[1583687783] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583688324] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22082672 kB) free!
[1583688444] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17584388 kB) free!
[1583688564] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13539180 kB) free!
[1583689412] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9819972 kB) free!
[1583689532] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9818092 kB) free!
[1583689652] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9817968 kB) free!
[1583689764] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33169836 kB) free.
[1583690104] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583690964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20989108 kB) free!
[1583691084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17189328 kB) free!
[1583691204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45175044 kB) free.
[1583691804] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21538536 kB) free!
[1583691924] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.5% (17823900 kB) free!
[1583692044] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45162224 kB) free.