[1583001018] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.1% (28654792 kB) free.
[1583001075] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1583001128] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2020/03/01 00:01:51 [error] 12256#0: *62557557 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583001248] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583001388] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [01/Mar/2020:00:04:06 +0530] TCP 502 0 0 180.000 "10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1583001509] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583001675] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1583002103] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583002163] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23360092 kB) free!
[1583002283] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19337112 kB) free!
[1583002403] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15563632 kB) free!
[1583002875] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1583002995] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1583003003] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30386220 kB) free.
[1583003176] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188564 kB) free!
[1583003236] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183740 kB) free!
[1583003296] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.8% (17116052 kB) free.
[1583003595] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1583003603] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14457440 kB) free!
[1583003715] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1583003724] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42810460 kB) free.
[1583004323] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20099400 kB) free!
[1583004409] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.1% (13284356 kB) free!
[1583004443] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16156488 kB) free!
[1583004563] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12507020 kB) free!
[1583005009] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12382296 kB) free!
[1583005163] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25879132 kB) free!
[1583005515] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583005635] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1583005701] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1583005763] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41530240 kB) free.
[1583005793] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583006363] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19529612 kB) free!
[1583006483] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15778264 kB) free!
[1583006603] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (12130280 kB) free!
[1583007203] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (24046468 kB) free!
[1583007803] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37810296 kB) free.
[1583008009] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1583008009] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.0% (17200508 kB) free!
[1583008035] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583008155] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1583008403] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16628916 kB) free!
[1583008523] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.0% (46120776 kB) free.
[1583009123] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23698460 kB) free!
[1583009243] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19832000 kB) free!
[1583009364] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16222412 kB) free!
[1583009931] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220844 kB) free!
[1583009955] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 355 processes
[1583009964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27306532 kB) free.
[1583010051] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243220 kB) free!
[1583010075] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1583010112] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (241900 kB) free!
[1583010171] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232388 kB) free!
[1583010232] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244380 kB) free!
[1583010352] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218888 kB) free!
[1583010675] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1583010772] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 82.6% (40694952 kB) free.
[1583010795] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583010915] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583011030] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.9% (10025868 kB) free.
[1583011164] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19336892 kB) free!
[1583011284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (46041220 kB) free.
[1583011516] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1583011552] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.0% (42858812 kB) free.
[1583011884] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (26019028 kB) free!
[1583011987] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(261) < 182.18.144.246 [01/Mar/2020:03:00:37 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583012004] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.9% (22307240 kB) free!
[1583012048] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(319) < 2020/03/01 03:01:31 [error] 12257#0: *62611107 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583012109] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583012124] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18837304 kB) free!
[1583012169] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583012308] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2494) < 172.105.59.173 [01/Mar/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583012428] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583012716] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583012724] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27913888 kB) free.
[1583012836] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583012956] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1583013556] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1583013676] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1583013797] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1583014524] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12928344 kB) free!
[1583014644] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9469148 kB) free!
[1583014764] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34261488 kB) free.
[1583014996] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583015116] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1583015236] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1583015364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13547716 kB) free!
[1583015484] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10092816 kB) free!
[1583015605] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33268152 kB) free.
[1583015768] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(71) < 2020/03/01 04:01:49 [error] 12252#0: *62628161 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1583015889] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583016028] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(19) < 203.212.70.170 [01/Mar/2020:04:02:57 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1583016149] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583016204] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13377888 kB) free!
[1583016324] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35607288 kB) free.
[1583016436] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1583016924] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17438524 kB) free!
[1583017009] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1583017044] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13952560 kB) free!
[1583017164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34973204 kB) free.
[1583017764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12881756 kB) free!
[1583017884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8836592 kB) free!
[1583018004] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3613148 kB) free!
[1583018604] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14054376 kB) free!
[1583018613] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1583018733] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1583019430] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196616 kB) free!
[1583019551] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188420 kB) free!
[1583019671] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186996 kB) free!
[1583019804] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34721644 kB) free.
[1583020193] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 45.8% (11228800 kB) free.
[1583020404] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11603888 kB) free!
[1583020524] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7144248 kB) free!
[1583020644] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33739560 kB) free.
[1583020701] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 60.4% (14898432 kB) free.
[1583021245] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11553308 kB) free!
[1583021365] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8240436 kB) free!
[1583021485] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34250168 kB) free.
[1583022085] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15234372 kB) free!
[1583022205] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36285148 kB) free.
[1583022805] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17273032 kB) free!
[1583022925] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13355752 kB) free!
[1583023045] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37242620 kB) free.
[1583023248] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.3% (38110796 kB) free.
[1583023299] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.7% (21745000 kB) free.
[1583023645] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17986080 kB) free!
[1583023765] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14243132 kB) free!
[1583023885] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36160708 kB) free.
[1583024209] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.2% (33165584 kB) free.
[1583024485] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17210824 kB) free!
[1583024605] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13493940 kB) free!
[1583024725] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35639876 kB) free.
[1583024901] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (767392 kB) free!
[1583025021] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (295216 kB) free!
[1583025141] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (285612 kB) free!
[1583025325] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16422708 kB) free!
[1583025445] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36956640 kB) free.
[1583026046] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18726888 kB) free!
[1583026166] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15114152 kB) free!
[1583026286] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35543048 kB) free.
[1583026507] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22056012 kB) free.
[1583026886] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16018724 kB) free!
[1583027006] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36661992 kB) free.
[1583027236] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583027606] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18276764 kB) free!
[1583027726] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14626472 kB) free!
[1583027846] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35329804 kB) free.
[1583028446] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16255548 kB) free!
[1583028566] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36473992 kB) free.
[1583029166] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18817152 kB) free!
[1583029286] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (15069148 kB) free!
[1583029406] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35937084 kB) free.
[1583030006] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16295444 kB) free!
[1583030126] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36218148 kB) free.
[1583030511] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4905816 kB) free!
[1583030631] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.9% (4883256 kB) free!
[1583030727] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18719516 kB) free!
[1583030751] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 20.8% (5086132 kB) free.
[1583030847] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (15035072 kB) free!
[1583030967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36035412 kB) free.
[1583031148] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 49.50.64.122 [01/Mar/2020:08:22:19 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583031268] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(3) < 172.105.59.173 [01/Mar/2020:08:23:12 +0530] TCP 502 0 0 120.001 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1583031388] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583031488] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2020/03/01 08:20:29 [error] 12248#0: *62722067 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:3295/47820, bytes from/to upstream:47820/3295
[1583031567] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17150716 kB) free!
[1583031608] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583031687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (13243960 kB) free!
[1583031807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34406868 kB) free.
[1583031951] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4778184 kB) free!
[1583032071] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4751280 kB) free!
[1583032191] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.3% (4735180 kB) free!
[1583032407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13905060 kB) free!
[1583032527] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34290004 kB) free.
[1583033127] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15153912 kB) free!
[1583033247] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35481708 kB) free.
[1583033847] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15462772 kB) free!
[1583033967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11357480 kB) free!
[1583034087] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583034087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34082124 kB) free.
[1583034687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12771212 kB) free!
[1583034807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34009904 kB) free.
[1583034988] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 103.16.101.72 [01/Mar/2020:09:23:07 +0530] TCP 502 0 0 60.000 "10.147.212.143:13000, 10.147.212.142:13000, 10.147.212.154:13000, 10.147.212.155:13000" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.000, 0.000, 60.000"
[1583035108] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583035208] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/03/01 09:22:03 [error] 12253#0: *62752747 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583035328] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583035407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14748256 kB) free!
[1583035528] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10844504 kB) free!
[1583035648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33368376 kB) free.
[1583035791] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2268108 kB) free!
[1583035831] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1583035951] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1583036071] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1583036248] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12962868 kB) free!
[1583036368] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8843748 kB) free!
[1583036488] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31421576 kB) free.
[1583036528] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(48) < 2020/03/01 09:50:47 [error] 12252#0: *62766983 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583036648] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583037088] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9237612 kB) free!
[1583037208] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31709252 kB) free.
[1583037249] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2020/03/01 10:01:23 [error] 12244#0: *62772373 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583037369] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583037808] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10062992 kB) free!
[1583037928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5840728 kB) free!
[1583038048] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30204588 kB) free.
[1583038193] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4749352 kB) free!
[1583038313] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4594432 kB) free!
[1583038433] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4461620 kB) free!
[1583038648] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8116816 kB) free!
[1583038768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (4009412 kB) free!
[1583038888] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28932608 kB) free.
[1583039110] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.9% (4161300 kB) free!
[1583039230] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (4099236 kB) free!
[1583039350] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (3955444 kB) free!
[1583039488] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6033284 kB) free!
[1583039608] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1853984 kB) free!
[1583039633] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198320 kB) free!
[1583039729] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.7% (25960364 kB) free!
[1583039829] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1583039837] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1583039949] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1583039957] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1583040069] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1583040077] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1583040104] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176564 kB) free!
[1583040225] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188100 kB) free!
[1583040329] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396752 kB) free!
[1583040345] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203216 kB) free!
[1583040669] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1583040677] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1583040929] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18873776 kB) free!
[1583041150] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2362732 kB) free!
[1583041192] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.2% (21379924 kB) free.
[1583041277] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 355 processes
[1583041397] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 352 processes
[1583041517] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1583041529] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583041529] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32429736 kB) free.
[1583041852] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.9% (5853724 kB) free!
[1583041972] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217492 kB) free!
[1583042092] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215136 kB) free!
[1583042117] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583042117] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 368 processes
[1583042129] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8585808 kB) free!
[1583042249] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34798888 kB) free.
[1583042849] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10851072 kB) free!
[1583042969] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6572752 kB) free!
[1583043089] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2512516 kB) free!
[1583043362] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3553336 kB) free!
[1583043482] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.7% (3310892 kB) free!
[1583043599] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1583043602] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3040936 kB) free!
[1583043719] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1583043839] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1583044289] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19325856 kB) free!
[1583044569] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(216) < 2020/03/01 12:02:17 [error] 12254#0: *62831571 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583044627] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.4% (21809540 kB) free.
[1583044675] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39288564 kB) free.
[1583044689] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583044708] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 172.105.59.173 [01/Mar/2020:12:03:27 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1583044717] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.7% (21750772 kB) free.
[1583044770] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.8% (21520188 kB) free.
[1583044802] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.7% (15877556 kB) free.
[1583044828] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583044889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32130300 kB) free.
[1583045233] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.7% (3371348 kB) free!
[1583045353] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.1% (3235692 kB) free!
[1583045473] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1583045473] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.6% (3100844 kB) free!
[1583045489] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6578572 kB) free!
[1583045609] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2187064 kB) free!
[1583045640] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1583045729] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (11501080 kB) free!
[1583045974] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187168 kB) free!
[1583046073] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166728 kB) free!
[1583046094] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187120 kB) free!
[1583046214] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1583046214] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172820 kB) free!
[1583046929] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18734736 kB) free!
[1583047517] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1583047529] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31922748 kB) free.
[1583048129] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7227456 kB) free!
[1583048249] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3093196 kB) free!
[1583048281] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5783532 kB) free!
[1583048369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32682364 kB) free.
[1583048401] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.7% (5782856 kB) free!
[1583048521] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3871508 kB) free!
[1583048969] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8250396 kB) free!
[1583049089] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3947736 kB) free!
[1583049209] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33350240 kB) free.
[1583049809] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8301684 kB) free!
[1583049929] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3979764 kB) free!
[1583050049] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380816 kB) free!
[1583051249] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23198044 kB) free!
[1583051325] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166592 kB) free!
[1583051445] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172924 kB) free!
[1583051565] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172648 kB) free!
[1583051849] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (661852 kB) free!
[1583053049] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23223312 kB) free!
[1583053649] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415332 kB) free!
[1583054118] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1583054393] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.4% (4501416 kB) free!
[1583054513] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.9% (4390956 kB) free!
[1583054614] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 52.0% (12753128 kB) free.
[1583054633] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.5% (4287532 kB) free!
[1583054849] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21892000 kB) free!
[1583055028] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(183) < 92.42.104.166 [01/Mar/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583055148] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583055449] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382444 kB) free!
[1583056377] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1583056497] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1583056634] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.9% (20200000 kB) free.
[1583056649] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (24015592 kB) free!
[1583056669] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37267924 kB) free.
[1583056699] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 61.2% (30144852 kB) free.
[1583056728] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 53.6% (13145844 kB) free.
[1583056822] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (2817060 kB) free!
[1583056942] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (2820208 kB) free!
[1583057062] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (2820828 kB) free!
[1583057243] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.4% (2800852 kB) free!
[1583057363] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.5% (2601360 kB) free!
[1583057483] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1837120 kB) free!
[1583057671] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1583058214] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189668 kB) free!
[1583058233] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2180768 kB) free!
[1583058334] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177012 kB) free!
[1583058352] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9753992 kB) free!
[1583058454] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194928 kB) free!
[1583058472] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9697504 kB) free!
[1583058593] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (9635508 kB) free!
[1583058748] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(320) < 203.212.70.170 [01/Mar/2020:16:01:04 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583058868] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583059069] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19813600 kB) free!
[1583059090] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(88) < 2020/03/01 16:01:58 [error] 12244#0: *62958261 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583059190] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.2% (50322352 kB) free.
[1583059210] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583059472] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1583059517] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1583059592] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1583059637] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583059712] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1583059757] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583059789] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25256880 kB) free!
[1583059909] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.0% (21135476 kB) free!
[1583060029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.8% (52420404 kB) free.
[1583061557] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1583062157] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1583062277] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1583062429] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24465956 kB) free!
[1583062549] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20333832 kB) free!
[1583062669] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24545440 kB) free!
[1583062877] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1583062998] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583063118] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1583063869] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36713920 kB) free.
[1583064318] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1583064318] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1583064512] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1583064593] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4608596 kB) free!
[1583064918] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583065018] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.1% (21836704 kB) free.
[1583065038] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1583065045] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21786432 kB) free.
[1583065069] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24468188 kB) free!
[1583065072] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 29.5% (14518720 kB) free.
[1583065102] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.1% (22078380 kB) free.
[1583065158] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1583065189] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20285004 kB) free!
[1583065309] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.1% (48848192 kB) free.
[1583065711] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168788 kB) free!
[1583065758] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583065831] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178528 kB) free!
[1583065878] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1583065910] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25253644 kB) free!
[1583065951] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171756 kB) free!
[1583066030] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20530288 kB) free!
[1583066150] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16129212 kB) free!
[1583066478] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 356 processes
[1583066598] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1583067350] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38680960 kB) free.
[1583067798] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583067918] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583067950] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15717900 kB) free!
[1583068038] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583068070] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.3% (46539580 kB) free.
[1583068670] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21426632 kB) free!
[1583068790] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17176680 kB) free!
[1583068910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46753248 kB) free.
[1583069253] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583069510] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22565032 kB) free!
[1583069630] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18000180 kB) free!
[1583069750] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13981084 kB) free!
[1583070438] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1583070950] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33680400 kB) free.
[1583071039] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583071159] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1583071759] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583071879] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1583072151] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18573184 kB) free!
[1583072271] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13962400 kB) free!
[1583072391] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (9758320 kB) free!
[1583072991] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (18935824 kB) free!
[1583073079] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1583073199] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 382 processes
[1583073268] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(49) < 203.212.70.170 [01/Mar/2020:20:03:05 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1583073319] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 377 processes
[1583073388] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583073591] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31084728 kB) free.
[1583073610] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2020/03/01 20:01:51 [error] 12247#0: *63075083 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583073730] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583073919] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1583074372] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9582148 kB) free!
[1583074492] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9527272 kB) free!
[1583074612] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (9480620 kB) free!
[1583074702] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9258852 kB) free!
[1583074791] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19090676 kB) free!
[1583074822] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9055224 kB) free!
[1583074911] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14852388 kB) free!
[1583074942] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (8786588 kB) free!
[1583075031] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (11024272 kB) free!
[1583075119] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583075239] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1583075458] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38434044 kB) free.
[1583075484] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.5% (21451048 kB) free.
[1583075839] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1583075959] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1583076093] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (230284 kB) free!
[1583076213] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164804 kB) free!
[1583076333] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186988 kB) free!
[1583076559] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583076679] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 357 processes
[1583076799] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1583077261] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24388772 kB) free!
[1583077381] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.2% (20099728 kB) free!
[1583077501] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16073596 kB) free!
[1583078600] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1583078701] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35879964 kB) free.
[1583079301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14086560 kB) free!
[1583079421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41409136 kB) free.
[1583079678] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.6% (22215284 kB) free.
[1583080021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17860040 kB) free!
[1583080141] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (14062124 kB) free!
[1583080261] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40997252 kB) free.
[1583080861] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17636016 kB) free!
[1583080981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13793196 kB) free!
[1583081102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40310460 kB) free.
[1583081701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17687220 kB) free!
[1583081821] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39799792 kB) free.
[1583082421] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20864280 kB) free!
[1583082541] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16846840 kB) free!
[1583082661] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42950540 kB) free.
[1583083131] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4699484 kB) free!
[1583083251] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4644664 kB) free!
[1583083261] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20545304 kB) free!
[1583083371] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4585552 kB) free!
[1583083381] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (16676440 kB) free!
[1583083501] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42851764 kB) free.
[1583084101] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19447628 kB) free!
[1583084221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15552620 kB) free!
[1583084341] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40367924 kB) free.
[1583084941] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18382208 kB) free!
[1583085061] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42669168 kB) free.
[1583085454] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4840220 kB) free!
[1583085574] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4829420 kB) free!
[1583085661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21655868 kB) free!
[1583085694] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4672364 kB) free!
[1583085782] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17897868 kB) free!
[1583085902] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42586496 kB) free.
[1583085933] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1583086053] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1583086222] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (245936 kB) free!
[1583086342] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233528 kB) free!
[1583086462] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237948 kB) free!
[1583086502] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21236620 kB) free!
[1583086609] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.4% (20355712 kB) free!
[1583086622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17538764 kB) free!
[1583086729] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.2% (20068852 kB) free!
[1583086742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41435240 kB) free.
[1583086849] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.1% (19896360 kB) free!
[1583086884] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1583087342] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (19982172 kB) free!