[1558031737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558031858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558031978] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558032543] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3856556 kB) free!
[1558032578] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558033743] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.4% (9073972 kB) free!
[1558033778] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558033898] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558034018] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558035543] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3539504 kB) free!
[1558035818] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558036143] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5371584 kB) free!
[1558036305] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 314 processes
[1558036418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558036425] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 314 processes
[1558036538] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558036546] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 314 processes
[1558036574] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 21.09, 29.80, 18.39
[1558036658] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558036694] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.11, 20.41, 16.33
[1558036743] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4318300 kB) free!
[1558036814] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.42, 14.33, 14.59
[1558037343] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6303796 kB) free!
[1558038614] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.94, 1.89, 3.69
[1558038797] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 112.28, 26.84, 9.27
[1558038918] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 298.44, 126.32, 47.85
[1558039038] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 323.55, 192.88, 81.76
[1558039658] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1558039671] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 90.85, 169.77, 94.15
[1558039791] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.34, 113.63, 82.75
[1558039911] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.75, 76.06, 72.72
[1558040838] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.55, 1.45, 19.60
[1558040858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1558040978] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1558041578] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1558041698] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1558041711] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 0.54, 10.73
[1558042299] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558042419] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1558042638] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.84, 0.65, 3.31
[1558042911] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.14, 0.20, 3.12
[1558043019] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1558043139] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1558043739] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558043859] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1558044459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1558044579] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1558044698] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1558044818] SERVICE ALERT: sig2-vm1;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1558045143] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4443400 kB) free!
[1558045180] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1558045300] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1558045743] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6252840 kB) free!
[1558045900] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558046020] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1558046620] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1558046740] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1558047340] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1558047461] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1558048060] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1558048180] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1558048780] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1558048900] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1558049500] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558049620] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1558050220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1558050341] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1558050341] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1558050940] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1558051060] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1558051660] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1558051743] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3522056 kB) free!
[1558051780] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1558052343] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5584332 kB) free!
[1558052380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1558052500] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1558053100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558053221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1558053341] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558054796] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1558054916] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 9 users currently logged in
[1558055036] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1558055040] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (3134704 kB) free!
[1558055160] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3131176 kB) free!
[1558055280] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (3125368 kB) free!
[1558055407] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1914484 kB) free!
[1558055636] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1558056836] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1558058741] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558059341] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558059461] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558059581] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558059711] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 269.32, 171.97, 75.06
[1558059831] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 249.50, 200.40, 97.42
[1558059951] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.87, 134.18, 85.64
[1558060143] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3497856 kB) free!
[1558060181] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1558060181] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1558060744] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.0% (8880720 kB) free!
[1558060781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558060902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558061022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558061751] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 25.92, 15.00, 18.48
[1558061838] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 297.83, 185.00, 79.86
[1558061959] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 299.84, 223.10, 106.58
[1558062007] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (3083852 kB) free!
[1558062079] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.92, 156.18, 96.53
[1558062544] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3970644 kB) free!
[1558063080] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1612452 kB) free!
[1558063144] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.2% (9484808 kB) free!
[1558063551] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.29, 3.56
[1558063879] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 0.76, 14.25
[1558065007] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437736 kB) free!
[1558065079] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 0.61, 4.38
[1558065544] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4811616 kB) free!
[1558066546] SERVICE ALERT: sig3-vm2;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1558066744] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.9% (9295156 kB) free!
[1558068052] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.9% (23422740 kB) free.
[1558068103] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.5% (14712032 kB) free.
[1558068222] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558068544] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4868160 kB) free!
[1558068822] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558068942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558069062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558069663] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558069744] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.1% (8407600 kB) free!
[1558069811] SERVICE ALERT: webserver2;Root Partition;OK;HARD;1;DISK OK - free space: / 23350 MB (45.62% inode=99%):
[1558070263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558070383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558070503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558071103] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558071544] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2804300 kB) free!
[1558071703] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558071823] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558071943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558071978] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 45.7% (22521104 kB) free.
[1558075308] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.1% (2780892 kB) free!
[1558075428] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.8% (2568484 kB) free!
[1558075526] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46953 MB (20% inode=99%):
[1558075543] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558075548] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.5% (2352452 kB) free!
[1558075646] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46870 MB (20% inode=99%):
[1558075766] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46594 MB (20% inode=99%):
[1558076143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558076148] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1358160 kB) free!
[1558076247] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.58, 5.45, 5.01
[1558076263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558076367] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.81, 5.94, 5.24
[1558076384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558076487] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.25, 6.39, 5.49
[1558077449] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 540 processes
[1558077509] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 536 processes
[1558077569] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 530 processes
[1558077584] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558077966] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (1442748 kB) free!
[1558078086] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171684 kB) free!
[1558078184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558078206] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171512 kB) free!
[1558078304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558078424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558078608] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151912 kB) free!
[1558078662] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1558078728] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164476 kB) free!
[1558078782] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 345 processes
[1558078848] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161204 kB) free!
[1558078856] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (2409296 kB) free!
[1558078902] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 344 processes
[1558078976] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190048 kB) free!
[1558079096] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189172 kB) free!
[1558079624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558080224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558080345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558080465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558080615] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 91.53, 22.35, 8.67
[1558080687] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.52, 4.49, 4.75
[1558080702] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1558080735] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.97, 15.62, 7.86
[1558080793] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558080855] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.36, 11.02, 7.12
[1558080913] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558081033] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558081287] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.38, 5.87, 5.31
[1558081407] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.80, 6.50, 5.61
[1558081455] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.03, 3.20, 4.65
[1558081527] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4523704 kB) free!
[1558081527] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.56, 6.43, 5.69
[1558081633] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558081647] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4610816 kB) free!
[1558081665] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558081753] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558081767] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (4529724 kB) free!
[1558081874] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558082187] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558082265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558082385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558082505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558084182] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (1157812 kB) free!
[1558084302] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (1045948 kB) free!
[1558084422] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.6% (927660 kB) free!
[1558085022] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (282812 kB) free!
[1558085093] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2583204 kB) free!
[1558085213] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.4% (2338704 kB) free!
[1558085333] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.0% (2111824 kB) free!
[1558085367] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2165600 kB) free!
[1558085933] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1130512 kB) free!
[1558086296] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.7% (11775132 kB) free.
[1558087166] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24165 MB (10% inode=99%):
[1558087455] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.39, 32.01, 19.65
[1558087575] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.81, 21.97, 17.47
[1558087587] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9825780 kB) free!
[1558087695] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.80, 15.27, 15.56
[1558087707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 33.3% (16417280 kB) free.
[1558089296] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4889884 kB) free!
[1558089416] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4847384 kB) free!
[1558089495] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1558089495] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.76, 1.84, 3.76
[1558089536] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4786704 kB) free!
[1558089927] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.00, 4.31, 4.88
[1558090708] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8799640 kB) free!
[1558090828] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.7% (15128220 kB) free.
[1558091336] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1935660 kB) free!
[1558091896] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.46, 17.30, 9.24
[1558092016] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.93, 12.12, 8.32
[1558092136] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.82, 8.66, 7.51
[1558092736] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.97, 2.87, 4.87
[1558093228] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8515904 kB) free!
[1558093302] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 348 processes
[1558093305] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558093348] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.5% (15048680 kB) free.
[1558093423] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 348 processes
[1558093543] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 347 processes
[1558093766] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55482 MB (24% inode=99%):
[1558093905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558094026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558094146] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558094743] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558094748] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 40.7% (6623204 kB) free.
[1558095346] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558095748] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9422576 kB) free!
[1558095868] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.9% (15718040 kB) free.
[1558095946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558095951] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 269.02, 195.30, 91.72
[1558096066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558096071] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 43.43, 137.02, 83.24
[1558096186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558096191] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.11, 91.78, 73.18
[1558096206] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2686052 kB) free!
[1558096423] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 94.2% (6931328 kB) free.
[1558096733] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 79.6% (12938504 kB) free.
[1558096806] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2379256 kB) free!
[1558097148] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3252456 kB) free!
[1558097268] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3132604 kB) free!
[1558097387] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558097388] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (3026712 kB) free!
[1558097987] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558097991] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 0.59, 10.85
[1558098079] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 97.14, 182.66, 101.92
[1558098107] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558098199] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.32, 122.29, 89.59
[1558098227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558098319] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.18, 81.98, 78.78
[1558098536] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2691824 kB) free!
[1558098868] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8426016 kB) free!
[1558098988] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.1% (14827552 kB) free.
[1558099136] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1558099136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2465052 kB) free!
[1558099189] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1141568 kB) free!
[1558099191] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.33, 3.21
[1558099427] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558100027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558100119] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.37, 0.61, 11.70
[1558100147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558100268] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558101319] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.41, 0.44, 3.54
[1558101388] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8503300 kB) free!
[1558101468] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558101508] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.0% (14783704 kB) free.
[1558102068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558102188] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558102308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558102708] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6557096 kB) free!
[1558102829] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (4368144 kB) free!
[1558102908] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558102949] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11012728 kB) free.
[1558103508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558103549] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8926060 kB) free!
[1558103629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558103669] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7009560 kB) free!
[1558103749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558103789] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4986640 kB) free!
[1558104137] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558104389] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11741932 kB) free.
[1558104949] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558105549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558105589] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9433848 kB) free!
[1558105669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558105709] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7712088 kB) free!
[1558105789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558105829] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5814268 kB) free!
[1558106390] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558106429] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1558106429] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11691120 kB) free.
[1558106990] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558107110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558107230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558107630] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8383024 kB) free!
[1558107750] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6594212 kB) free!
[1558107870] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4291996 kB) free!
[1558108430] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558108470] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11385420 kB) free.
[1558109030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558109150] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558109270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558109670] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9063512 kB) free!
[1558109790] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7374936 kB) free!
[1558109910] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5270628 kB) free!
[1558110470] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558110510] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11433472 kB) free.
[1558110536] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558111070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558111110] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9795544 kB) free!
[1558111190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558111230] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8180132 kB) free!
[1558111310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558111350] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6549092 kB) free!
[1558112510] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558112550] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11467588 kB) free.
[1558112743] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 348 processes
[1558112863] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 350 processes
[1558112933] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (158012 kB) free!
[1558112983] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 350 processes
[1558113053] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (152900 kB) free!
[1558113110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558113173] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162028 kB) free!
[1558113231] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558113351] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558113583] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1558113736] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 291.94, 198.06, 89.95
[1558113750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8719160 kB) free!
[1558113856] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 292.73, 229.41, 114.55
[1558113870] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6923088 kB) free!
[1558113976] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 292.82, 250.43, 136.18
[1558113990] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4738600 kB) free!
[1558114374] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 44.3% (7199044 kB) free.
[1558114551] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558114590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10941096 kB) free.
[1558114783] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 350 processes
[1558114904] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558114974] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176824 kB) free!
[1558115094] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 80.9% (13153548 kB) free.
[1558115151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558115190] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9846900 kB) free!
[1558115271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558115310] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8262112 kB) free!
[1558115391] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558115430] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.5% (6651892 kB) free!
[1558115503] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 351 processes
[1558115623] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 350 processes
[1558115694] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162364 kB) free!
[1558115743] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 351 processes
[1558115814] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (188392 kB) free!
[1558115934] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (185564 kB) free!
[1558116343] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1558116343] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558116377] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.78, 1.70, 12.87
[1558116534] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1558116534] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.0% (14458640 kB) free.
[1558116592] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558116630] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11465024 kB) free.
[1558117192] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558117231] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9773820 kB) free!
[1558117312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558117351] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8226092 kB) free!
[1558117432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558117471] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.5% (6640492 kB) free!
[1558117577] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.58, 1.55, 4.63