[1561401570] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9877608 kB) free!
[1561401575] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561401624] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1561401691] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13099896 kB) free.
[1561401695] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561401753] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 40318 MB (78% inode=99%):
[1561401801] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1561401921] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1561402132] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1561402135] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1561402186] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.76, 8.72, 6.31
[1561402224] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.15, 0.52, 13.32
[1561402255] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1561402256] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1561402290] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9747688 kB) free!
[1561402295] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1561402295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561402314] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.21, 8.76, 6.63
[1561402410] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13105736 kB) free.
[1561402415] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561402425] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1561402440] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1561402440] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.66, 8.66, 6.86
[1561402549] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 331 processes
[1561403010] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9606604 kB) free!
[1561403015] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561403050] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1561403130] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13094076 kB) free.
[1561403131] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1561403135] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561403155] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1561403258] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1561403275] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 330 processes
[1561403425] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.17, 0.23, 3.83
[1561403462] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1561403535] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1561403582] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1561403650] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.80, 7.70, 7.69
[1561403655] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1561403730] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9514332 kB) free!
[1561403736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561403850] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13102852 kB) free.
[1561403856] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561404450] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9432892 kB) free!
[1561404456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561404570] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13102716 kB) free.
[1561404576] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561404850] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.23, 2.47, 3.89
[1561404915] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3865664 kB) free!
[1561404975] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 15.8% (3864688 kB) free!
[1561405035] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 15.8% (3868012 kB) free!
[1561405095] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 15.8% (3865016 kB) free!
[1561405170] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9360312 kB) free!
[1561405176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561405290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13093984 kB) free.
[1561405296] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561405855] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561405890] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9253796 kB) free!
[1561405896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561406011] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13103880 kB) free.
[1561406016] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561406610] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9119304 kB) free!
[1561406617] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561406730] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13093328 kB) free.
[1561406737] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561407330] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8981288 kB) free!
[1561407337] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561407450] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13099652 kB) free.
[1561407457] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561408050] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8855724 kB) free!
[1561408057] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561408170] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13099008 kB) free.
[1561408177] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561408459] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178092 kB) free!
[1561408579] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (172536 kB) free!
[1561408699] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166776 kB) free!
[1561408770] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8634660 kB) free!
[1561408777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561408890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13089384 kB) free.
[1561408897] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561409490] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8476640 kB) free!
[1561409497] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561409610] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13098404 kB) free.
[1561409618] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561410677] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 138.34, 45.38, 17.20
[1561410797] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 180.29, 92.20, 37.77
[1561410811] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9860900 kB) free!
[1561410817] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561410917] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 186.08, 123.52, 55.85
[1561410931] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13092984 kB) free.
[1561410937] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561411450] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.57, 38.92, 19.47
[1561411531] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9712132 kB) free!
[1561411537] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561411570] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6.25, 26.75, 17.37
[1561411651] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13100352 kB) free.
[1561411657] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561411690] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.84, 18.67, 15.55
[1561412251] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9601960 kB) free!
[1561412257] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561412371] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13102720 kB) free.
[1561412377] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561412717] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.09, 2.90, 18.78
[1561412971] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9449400 kB) free!
[1561412978] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561413091] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13082704 kB) free.
[1561413098] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561413490] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.44, 2.51, 4.34
[1561413691] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9149324 kB) free!
[1561413698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1561413811] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13088988 kB) free.
[1561413818] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561414411] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8857536 kB) free!
[1561414418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561414517] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.47, 1.64, 4.17
[1561414531] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13079268 kB) free.
[1561414538] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561415035] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4913396 kB) free!
[1561415117] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 103.05, 27.89, 11.85
[1561415132] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8521520 kB) free!
[1561415138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561415155] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4904108 kB) free!
[1561415237] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 195.74, 88.16, 35.88
[1561415252] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13086772 kB) free.
[1561415258] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561415275] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (4898628 kB) free!
[1561415357] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1561415357] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 207.84, 128.36, 56.95
[1561416452] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9708044 kB) free!
[1561416459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561416572] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13091344 kB) free.
[1561416579] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561417172] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9469540 kB) free!
[1561417179] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561417292] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13085860 kB) free.
[1561417299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561417757] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.55, 1.92, 12.43
[1561417892] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9305400 kB) free!
[1561417899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561418012] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13093212 kB) free.
[1561418019] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561418612] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9058216 kB) free!
[1561418619] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561418732] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13081168 kB) free.
[1561418739] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561418957] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.91, 1.66, 4.58
[1561419333] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8824592 kB) free!
[1561419340] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561419453] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13089096 kB) free.
[1561419460] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561420053] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8406084 kB) free!
[1561420060] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561420173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12314432 kB) free.
[1561420180] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561420300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561421373] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9661444 kB) free!
[1561421493] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13055392 kB) free.
[1561421500] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561422100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561422221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561422341] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561422661] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1561422781] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1561423103] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 56.1% (9125104 kB) free.
[1561423690] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561423893] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8349364 kB) free!
[1561424013] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12253692 kB) free.
[1561426855] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1561426975] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561427013] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9630212 kB) free!
[1561427133] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12862360 kB) free.
[1561427141] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1561427141] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561427741] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561427861] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561427981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561428475] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2456740 kB) free!
[1561429075] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4597164 kB) free!
[1561429181] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561429533] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9641408 kB) free!
[1561429653] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12866068 kB) free.
[1561429757] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561430253] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9609820 kB) free!
[1561430374] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12867400 kB) free.
[1561430890] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 179.43, 107.81, 47.01
[1561430895] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 25.9% (6346908 kB) free.
[1561430973] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9611276 kB) free!
[1561430981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561431010] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 181.23, 132.14, 63.30
[1561431093] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12856416 kB) free.
[1561431101] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561431130] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 181.50, 148.50, 77.64
[1561431693] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9546244 kB) free!
[1561431701] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561431813] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12865696 kB) free.
[1561431821] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561432413] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9535788 kB) free!
[1561432421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561432533] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12858140 kB) free.
[1561432541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561432930] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.62, 2.82, 13.30
[1561433133] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9484804 kB) free!
[1561433141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561433253] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12865152 kB) free.
[1561433261] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561433853] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9661180 kB) free!
[1561433862] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561433973] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12754872 kB) free.
[1561433982] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561434573] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9698248 kB) free!
[1561434582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561434694] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12853328 kB) free.
[1561434702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561434730] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.22, 2.70, 4.23
[1561435293] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9831884 kB) free!
[1561435302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561435413] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8286684 kB) free!
[1561435422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561435533] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12850860 kB) free.
[1561435542] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561435676] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2128140 kB) free!
[1561435825] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 85.03, 133.91, 76.73
[1561435945] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.60, 89.66, 67.45
[1561436065] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.76, 60.10, 59.31
[1561436133] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8121236 kB) free!
[1561436253] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12685536 kB) free.
[1561436276] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1561436276] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4401648 kB) free!
[1561436530] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 121.06, 73.32, 32.73
[1561436650] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 122.16, 89.54, 43.59
[1561436770] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1561436770] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 122.45, 100.48, 53.17
[1561437265] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.49, 1.48, 16.55
[1561437342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561437462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561437503] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (149240 kB) free!
[1561437582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561437623] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (162092 kB) free!
[1561437743] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163700 kB) free!
[1561437746] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (320256 kB) free!
[1561437866] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (319492 kB) free!
[1561437970] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.75, 5.05, 19.47
[1561437986] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180148 kB) free!
[1561438157] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 68.67, 31.83, 13.42
[1561438278] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 72.23, 45.39, 20.63
[1561438398] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 72.84, 54.55, 26.99
[1561438465] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.47, 0.51, 4.91
[1561438488] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5390 MB (10% inode=97%):
[1561439088] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6991 MB (13% inode=98%):
[1561439598] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.99, 4.15, 13.55
[1561440370] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.69, 2.68, 3.93
[1561440401] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 62.8% (10210892 kB) free.
[1561440476] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2263844 kB) free!
[1561440495] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4843616 kB) free!
[1561440555] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4845292 kB) free!
[1561440582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561440615] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.8% (4851408 kB) free!
[1561440676] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.8% (4856252 kB) free!
[1561440704] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (300108 kB) free!
[1561440781] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 4 crit, 0 warn out of 371 processes
[1561440824] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.3% (6537996 kB) free!
[1561440901] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 4 crit, 0 warn out of 369 processes
[1561440944] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.3% (6541272 kB) free!
[1561441021] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 367 processes
[1561441053] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8616144 kB) free!
[1561441173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12818112 kB) free.
[1561441398] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.65, 1.99, 3.74
[1561441621] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1561441773] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5642100 kB) free!
[1561441782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561441893] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12349684 kB) free.
[1561441902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561442022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561442267] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.95, 7.04, 5.96
[1561442387] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.99, 6.71, 5.97
[1561442507] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.88, 6.51, 6.00
[1561442745] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233340 kB) free!
[1561444422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561444545] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.3% (6552228 kB) free!
[1561444893] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8344688 kB) free!
[1561445013] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12310164 kB) free.
[1561445022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561445142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561445202] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.3% (2657296 kB) free!
[1561445262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561445322] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.1% (2462020 kB) free!
[1561445442] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.8% (2236288 kB) free!
[1561446042] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1174900 kB) free!
[1561446787] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 44.6% (7252752 kB) free.
[1561446857] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.0% (18518468 kB) free.
[1561446917] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.2% (21626124 kB) free.
[1561446982] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1561446982] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 92.0% (45352408 kB) free.
[1561447039] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1561447039] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.3% (14359716 kB) free.
[1561447062] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561447100] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.4% (21921564 kB) free.
[1561447204] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.4% (14863740 kB) free.
[1561447375] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561447398] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (153784 kB) free!
[1561447413] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9399252 kB) free!
[1561447495] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561447518] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171480 kB) free!
[1561447533] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12784748 kB) free.
[1561447639] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155772 kB) free!
[1561448069] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (234628 kB) free!
[1561448095] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561448129] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4345780 kB) free!
[1561448133] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8868324 kB) free!
[1561448189] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (204012 kB) free!
[1561448215] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561448221] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1561448249] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.3% (2519476 kB) free!
[1561448253] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12789832 kB) free.
[1561448309] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172928 kB) free!
[1561448341] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1561448369] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (615176 kB) free!
[1561448405] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (149052 kB) free!
[1561448525] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171100 kB) free!
[1561448645] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152512 kB) free!
[1561448815] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561448853] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8659688 kB) free!
[1561448862] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561448935] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561448941] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 2 warn out of 371 processes
[1561448973] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12790464 kB) free.
[1561448982] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561449062] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 371 processes
[1561449074] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1561449182] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 367 processes
[1561449194] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1561449314] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1561449573] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8746812 kB) free!
[1561449582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561449693] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12780740 kB) free.
[1561449702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561449782] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1561449782] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 371 processes
[1561449970] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1561450382] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 366 processes
[1561450893] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8616708 kB) free!
[1561450902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561450982] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1561451014] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12781932 kB) free.
[1561451023] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561451187] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244816 kB) free!
[1561451307] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (333908 kB) free!
[1561451427] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220336 kB) free!
[1561451613] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8184540 kB) free!
[1561451622] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561451733] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12774492 kB) free.
[1561451742] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561452049] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.8% (791432 kB) free!
[1561452169] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (669788 kB) free!
[1561452289] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (547532 kB) free!
[1561452342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561452462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561452506] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1561452582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561452798] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 240.04, 87.52, 33.66
[1561452919] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 283.64, 154.64, 64.79
[1561453039] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 289.20, 199.43, 92.10
[1561453307] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.11, 4.08, 4.73
[1561453743] SERVICE ALERT: sigqscluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 111001 MB (49% inode=99%):
[1561454114] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1561454133] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8199708 kB) free!
[1561454253] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12777268 kB) free.
[1561454383] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561454427] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 90.3% (44513904 kB) free.
[1561454665] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 217.29, 146.11, 65.62
[1561454785] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 218.05, 169.97, 84.12
[1561454843] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (3026596 kB) free!
[1561454905] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 80.40, 146.34, 86.77
[1561454963] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (2889172 kB) free!
[1561455083] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (2724904 kB) free!
[1561455289] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.5% (6799720 kB) free.
[1561455439] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.95, 3.77, 17.98
[1561455453] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9556056 kB) free!
[1561455573] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12780784 kB) free.
[1561455883] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.7% (18671740 kB) free.
[1561455959] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.2% (20155872 kB) free.
[1561456039] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.1% (14319148 kB) free.
[1561456102] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.7% (15060912 kB) free.
[1561456161] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 58.0% (9428184 kB) free.
[1561456174] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9024256 kB) free!
[1561456183] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561456224] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 35.0% (17269020 kB) free.
[1561456303] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561456705] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.72, 0.84, 12.94
[1561456761] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (154112 kB) free!
[1561456881] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (163336 kB) free!
[1561456903] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561456907] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.04, 5.06, 5.18
[1561457001] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1561457001] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169496 kB) free!
[1561457023] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561457027] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.74, 4.62, 5.00
[1561457089] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (216020 kB) free!
[1561457143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561457164] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193136 kB) free!
[1561457209] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184236 kB) free!
[1561457274] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1561457284] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171160 kB) free!
[1561457329] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173592 kB) free!
[1561457335] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561457394] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1561457404] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180148 kB) free!
[1561457455] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1561457514] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1561457575] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1561457628] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.50, 5.82, 5.50
[1561457652] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.41, 6.04, 5.02
[1561457748] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.97, 5.67, 5.47
[1561457773] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.98, 5.69, 5.02
[1561457839] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.32, 3.80, 4.88
[1561457868] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.34, 5.10, 5.28
[1561457893] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.59, 4.71, 4.76
[1561457905] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.41, 3.88
[1561458175] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561458468] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.45, 3.65, 4.42
[1561459093] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.09, 7.64, 6.21
[1561459213] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.23, 7.84, 6.45
[1561459333] SERVICE ALERT: webservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.45, 7.44, 6.48
[1561459830] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9394836 kB) free!
[1561459950] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13089468 kB) free.
[1561461175] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1561461295] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561461343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561461376] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 29.9% (7322444 kB) free.
[1561461750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9118404 kB) free!
[1561461870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13086176 kB) free.
[1561462333] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.13, 2.32, 4.33
[1561462382] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561462470] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9157196 kB) free!
[1561462590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13076204 kB) free.
[1561463143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561463190] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9124640 kB) free!
[1561463263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561463310] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13084096 kB) free.
[1561463427] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1561463839] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 39.77, 12.38, 5.48
[1561463863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561463911] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8908760 kB) free!
[1561463959] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 62.56, 30.07, 12.79
[1561463983] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561464031] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13072024 kB) free.
[1561464079] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 64.81, 41.73, 19.16
[1561464288] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5387 MB (10% inode=97%):
[1561464526] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.8% (18446544 kB) free.
[1561464583] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561464631] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8771232 kB) free!
[1561464650] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.7% (22233908 kB) free.
[1561464703] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561464715] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 68.1% (11066588 kB) free.
[1561464751] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.6% (12630640 kB) free.
[1561464888] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6950 MB (13% inode=98%):
[1561464889] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.1% (1034648 kB) free!
[1561465009] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.2% (895684 kB) free!
[1561465129] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.5% (774496 kB) free!
[1561465280] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1561465280] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.75, 4.92, 13.92
[1561465304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561465424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561465544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561465706] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (183140 kB) free!
[1561465729] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (123900 kB) free!
[1561465736] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1570292 kB) free!
[1561465827] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (187568 kB) free!
[1561465827] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (345008 kB) free!
[1561465856] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (1049948 kB) free!
[1561465880] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16.70, 92.17, 62.08
[1561465947] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173800 kB) free!
[1561465948] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233060 kB) free!
[1561465976] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191980 kB) free!
[1561466068] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235708 kB) free!
[1561466095] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1561466144] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561466215] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561466245] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (3142488 kB) free!
[1561466330] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1561466330] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.7% (6815824 kB) free.
[1561466365] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (2985880 kB) free!
[1561466455] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (4123740 kB) free!
[1561466485] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (2866800 kB) free!
[1561466551] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9159168 kB) free!
[1561466575] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (3316284 kB) free!
[1561466668] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 62.4% (30769868 kB) free.
[1561466671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12903936 kB) free.
[1561466695] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (2855792 kB) free!
[1561466744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561466770] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 35.18, 26.26, 13.71
[1561466864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561466890] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.52, 29.35, 16.37
[1561466984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561467010] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.36, 28.20, 17.63
[1561467080] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.08, 3.57, 18.52
[1561467115] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (179416 kB) free!
[1561467235] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (183964 kB) free!
[1561467295] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1561467295] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187744 kB) free!
[1561467355] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (925752 kB) free!
[1561467584] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561467610] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.84, 5.94, 10.39
[1561467685] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1455172 kB) free!
[1561467714] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1561468185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561468305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561468425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561468471] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9770744 kB) free!
[1561468591] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8174716 kB) free!
[1561468668] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.17, 6.57, 5.41
[1561468712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12273640 kB) free.
[1561468788] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.42, 5.46, 5.15
[1561468810] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.45, 2.64, 4.75
[1561468880] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.58, 1.98, 4.33
[1561468909] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.88, 4.65, 4.89
[1561469508] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.52, 5.48, 5.17
[1561469628] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.07, 4.99, 5.03
[1561469668] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (221436 kB) free!
[1561469748] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.15, 5.76, 5.30
[1561469788] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231712 kB) free!
[1561469908] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (252192 kB) free!
[1561469961] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.7% (18433176 kB) free.
[1561470010] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1561470010] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 93.2% (45944804 kB) free.
[1561470056] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.9% (22528880 kB) free.
[1561470106] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.3% (14352736 kB) free.
[1561470148] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.2% (14985976 kB) free.
[1561470209] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 70.0% (11372216 kB) free.
[1561470348] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.52, 3.69, 4.37
[1561470568] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.2% (3019776 kB) free!
[1561470688] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (1387216 kB) free!
[1561470748] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (2532868 kB) free!
[1561470808] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1561470808] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1374504 kB) free!
[1561470868] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (193436 kB) free!
[1561470988] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172340 kB) free!
[1561471111] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8594864 kB) free!
[1561471231] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12899880 kB) free.
[1561471865] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (2592980 kB) free!
[1561471985] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (1262916 kB) free!
[1561472105] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185392 kB) free!
[1561472148] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.11, 6.14, 5.54
[1561472268] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.57, 7.51, 6.12
[1561472388] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.66, 6.67, 5.98
[1561472415] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3225848 kB) free!
[1561472535] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (1686760 kB) free!
[1561472625] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561472655] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (245120 kB) free!
[1561473031] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8739360 kB) free!
[1561473151] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12512548 kB) free.
[1561473225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561473345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561473465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561473588] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1561473588] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.23, 4.20, 4.66
[1561474065] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561474188] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.03, 5.82, 5.10
[1561474308] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.02, 5.20, 4.96
[1561474615] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1561474665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561474735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1561474785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561474855] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1561474905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561475506] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561476106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561476226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561476346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561476946] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561477546] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561477666] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561477730] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561477916] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1561478018] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1561478266] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561478387] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561478507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561478551] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9021484 kB) free!
[1561478671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13072936 kB) free.
[1561479080] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561479707] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561479834] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.9% (17980648 kB) free.
[1561479903] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.5% (21929524 kB) free.
[1561480000] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 94.2% (46426656 kB) free.
[1561480082] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.7% (15067252 kB) free.
[1561480907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561481027] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561481035] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (1385784 kB) free!
[1561481155] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (1371304 kB) free!
[1561481275] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1251596 kB) free!
[1561481410] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 286.01, 119.90, 46.36
[1561481530] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 312.15, 184.87, 79.09
[1561481651] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 315.75, 228.35, 107.84
[1561481875] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 657 processes
[1561481996] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 657 processes
[1561482116] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 653 processes
[1561482228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561482271] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9148536 kB) free!
[1561482311] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (2077392 kB) free!
[1561482348] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561482391] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13063420 kB) free.
[1561482431] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (998308 kB) free!
[1561482551] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180856 kB) free!
[1561482716] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 325 processes
[1561482948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561482991] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9018412 kB) free!
[1561483068] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561483111] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13069396 kB) free.
[1561483668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561483711] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8871452 kB) free!
[1561483788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561483831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13059312 kB) free.
[1561483855] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1561483880] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 163.21, 93.77, 40.00
[1561483914] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1561483975] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1561484000] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 165.70, 117.72, 55.29
[1561484034] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1561484051] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.35, 2.54, 15.38
[1561484120] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 165.57, 133.56, 68.67
[1561484154] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1561484388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561484431] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8699164 kB) free!
[1561484508] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561484508] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561484551] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13067392 kB) free.
[1561485109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561485151] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9726924 kB) free!
[1561485229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561485272] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8489388 kB) free!
[1561485349] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561485366] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.7% (18671956 kB) free.
[1561485392] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13060212 kB) free.
[1561485408] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1561485443] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.1% (22579588 kB) free.
[1561485851] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.25, 2.31, 4.20
[1561485920] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.03, 2.34, 14.46
[1561485949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561485992] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9626764 kB) free!
[1561486069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561486105] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 68.02, 48.85, 22.82
[1561486106] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (4325824 kB) free!
[1561486112] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8366396 kB) free!
[1561486189] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561486225] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 10.96, 34.28, 20.71
[1561486226] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 28.0% (6863860 kB) free.
[1561486232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13060560 kB) free.
[1561486247] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 583 processes
[1561486249] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 583 processes
[1561486345] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.71, 23.03, 18.24
[1561486367] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 351 processes
[1561486571] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (256084 kB) free!
[1561486691] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (205008 kB) free!
[1561486789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561486811] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202768 kB) free!
[1561486832] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9556128 kB) free!
[1561486849] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 325 processes
[1561486909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561486952] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13069432 kB) free.
[1561487029] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561487249] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (3886792 kB) free!
[1561487369] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.6% (2851868 kB) free!