[1566239433] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566239471] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 22.8% (5595500 kB) free.
[1566239981] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177100 kB) free!
[1566240071] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200944 kB) free!
[1566240101] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164052 kB) free!
[1566240117] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.60, 2.86, 4.82
[1566240191] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192524 kB) free!
[1566240221] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200308 kB) free!
[1566240311] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174948 kB) free!
[1566240590] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8004124 kB) free!
[1566240633] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566240710] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11585880 kB) free.
[1566240753] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1566241525] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 32.1% (7922820 kB) free.
[1566241924] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9047388 kB) free!
[1566242044] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9047924 kB) free!
[1566242164] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (9047496 kB) free!
[1566242862] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566243560] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9807588 kB) free!
[1566243680] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227544 kB) free!
[1566243801] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (1702196 kB) free!
[1566244182] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.5% (19179096 kB) free!
[1566244302] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.0% (19750128 kB) free!
[1566244422] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.0% (19752240 kB) free!
[1566245164] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3349292 kB) free!
[1566247511] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1566248074] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 672 processes
[1566248125] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (300168 kB) free!
[1566248194] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 423 processes
[1566248246] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (209012 kB) free!
[1566248262] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 18.89, 197.46, 147.90
[1566248366] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (217948 kB) free!
[1566248382] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.97, 132.31, 130.03
[1566248502] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.74, 88.67, 114.32
[1566250302] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.44, 1.90, 17.24
[1566251021] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.5% (10665092 kB) free.
[1566251111] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 44.0% (10783508 kB) free.
[1566252102] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.38, 2.99
[1566259318] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 202.25, 51.01, 18.64
[1566259438] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 468.18, 202.79, 78.19
[1566259558] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 504.31, 304.46, 130.56
[1566261200] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 81.4% (40116376 kB) free.
[1566261222] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (26677268 kB) free.
[1566261822] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.4% (24241288 kB) free!
[1566261942] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.1% (23813132 kB) free!
[1566262063] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.7% (23325196 kB) free!
[1566262558] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.68, 2.69, 15.34
[1566262663] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 56.9% (75031560 kB) free.
[1566262766] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 54.2% (13377076 kB) free.
[1566264358] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.71, 2.69, 4.52
[1566265253] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 360 processes
[1566265373] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1566265973] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566266021] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.3% (2535568 kB) free!
[1566266093] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566266111] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (216732 kB) free!
[1566266141] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.3% (2518788 kB) free!
[1566266213] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566266231] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (218924 kB) free!
[1566266261] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2511452 kB) free!
[1566266351] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204384 kB) free!
[1566266861] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2439300 kB) free!
[1566272502] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 506.27, 271.17, 110.13
[1566272622] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 518.03, 353.43, 159.80
[1566272742] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 519.55, 408.50, 203.46
[1566273625] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566273745] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1566274346] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566274466] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566275142] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.36, 0.64, 19.11
[1566276942] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.55, 0.59, 3.33
[1566277860] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.6% (20989880 kB) free.
[1566277909] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40740260 kB) free.
[1566277975] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.0% (14632988 kB) free.
[1566278019] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.3% (21648960 kB) free.
[1566278366] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4859016 kB) free!
[1566278486] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4831308 kB) free!
[1566278581] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176624 kB) free!
[1566278606] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4804860 kB) free!
[1566278701] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (181960 kB) free!
[1566278821] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171420 kB) free!
[1566279040] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1566279115] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (225572 kB) free!
[1566279160] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1566279235] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225876 kB) free!
[1566279280] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1566279329] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 358.35, 101.19, 35.41
[1566279355] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237804 kB) free!
[1566279449] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 563.88, 264.87, 103.37
[1566279569] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 592.29, 374.61, 163.16
[1566280406] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200712 kB) free!
[1566281666] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566281786] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566281976] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1566282036] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1566282096] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1566282156] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1566282400] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.73, 6.42, 5.21
[1566282520] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.14, 5.78, 5.12
[1566282569] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.73, 0.84, 15.84
[1566282640] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.91, 5.44, 5.07
[1566283769] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.45, 0.60, 4.81
[1566284680] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1566287440] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.51, 4.35, 4.66
[1566288040] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.03, 6.68, 5.61
[1566288160] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.77, 7.55, 6.03
[1566288280] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.27, 6.37, 5.80
[1566290186] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566290306] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 403 processes
[1566290427] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1566291280] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1566291280] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.13, 3.33, 4.64
[1566292056] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1566292480] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1566292600] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1566292720] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1566292953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566293073] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1566293142] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 87.83, 19.87, 7.46
[1566293193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1566293262] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 499.48, 199.81, 74.91
[1566293382] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 555.36, 320.29, 134.20
[1566293680] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.69, 5.24, 5.07
[1566293793] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1566293800] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.77, 4.76, 4.92
[1566294206] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2703888 kB) free!
[1566294806] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2438116 kB) free!
[1566295000] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.79, 6.78, 5.93
[1566295056] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1566295120] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.87, 7.07, 6.12
[1566295240] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.38, 9.11, 6.96
[1566295656] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1566296382] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.84, 2.03, 14.35
[1566297582] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.02, 0.93, 4.57
[1566297823] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191716 kB) free!
[1566297943] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168224 kB) free!
[1566298063] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197612 kB) free!
[1566298264] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187320 kB) free!
[1566298385] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198284 kB) free!
[1566298505] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186876 kB) free!
[1566299565] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9875288 kB) free!
[1566299685] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9677704 kB) free!
[1566299805] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (9454220 kB) free!
[1566300156] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1566300414] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 607 processes
[1566300534] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 382 processes
[1566300569] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 577.45, 473.71, 245.19
[1566300689] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 165.81, 380.01, 240.74
[1566300809] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.92, 254.45, 211.61
[1566301056] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1566301558] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 361.07, 103.63, 37.46
[1566301678] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 566.88, 267.53, 105.55
[1566301798] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 594.51, 377.22, 165.38
[1566302103] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5274 MB (10% inode=97%):
[1566302320] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566302522] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 374 processes
[1566302703] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 8676 MB (16% inode=98%):
[1566302755] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 353 processes
[1566302805] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4036732 kB) free!
[1566303122] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 371 processes
[1566303209] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.55, 1.19, 17.18
[1566303356] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 349 processes
[1566303640] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1566303806] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.4% (21309184 kB) free.
[1566303905] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.5% (21703148 kB) free.
[1566303924] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 81.4% (19954412 kB) free.
[1566303954] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40408620 kB) free.
[1566304005] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.4% (47493636 kB) free.
[1566304021] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.5% (14229028 kB) free.
[1566304798] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.74, 2.90, 18.96
[1566305009] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 0.72, 3.06
[1566305821] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151696 kB) free!
[1566305941] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (152680 kB) free!
[1566306061] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155384 kB) free!
[1566306520] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566307056] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1566307198] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.78, 2.94, 4.09
[1566307956] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 54 users currently logged in
[1566308320] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1566308440] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.50, 4.46, 4.86
[1566309040] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.63, 5.95, 5.80
[1566309160] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.24, 5.69, 5.72
[1566309280] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.58, 5.26, 5.55
[1566309964] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9253128 kB) free!
[1566310084] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8478736 kB) free!
[1566310204] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (7684180 kB) free!
[1566310804] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3658404 kB) free!
[1566313542] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4730712 kB) free!
[1566313613] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1566313662] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4672960 kB) free!
[1566313782] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (4608852 kB) free!
[1566314426] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566314546] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1566316480] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.61, 3.73, 4.58
[1566318582] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2127932 kB) free!
[1566320000] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9715424 kB) free!
[1566320120] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9626848 kB) free!
[1566320240] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.4% (9537164 kB) free!
[1566322346] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566322466] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1566325043] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4149564 kB) free!
[1566325409] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.55, 9.07, 7.95
[1566325484] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.45, 8.75, 5.62
[1566325529] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.27, 6.16, 7.02
[1566325604] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.84, 8.75, 6.00
[1566325649] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 4.20, 6.20
[1566325724] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.20, 8.02, 6.09