[1544728807] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1544739473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1544739593] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1544761193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1544761313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544761433] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544761868] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (2962820 kB) free!
[1544765033] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544765153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544765273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544768945] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 77.64, 55.21, 23.07
[1544769007] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1544769065] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.80, 37.42, 20.45
[1544769128] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1544769185] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.56, 25.40, 18.10
[1544769248] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1544769785] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.02, 4.42, 10.04
[1544769848] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1544770985] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.77, 0.94, 3.44
[1544773449] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1544773569] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1544774169] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1544774289] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1544774409] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1544775473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544775594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544775714] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544776914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544777034] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544780409] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1544781234] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544781354] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544781474] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1544789874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544789994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544790114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544790714] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544797915] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544798035] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544798155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544798755] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1544798755] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1544800555] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544800675] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544800796] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544810396] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)