[1589654310] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (532844 kB) free!
[1589654431] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400124 kB) free!
[1589654533] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.46, 4.14, 4.61
[1589654551] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412016 kB) free!
[1589654750] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589655151] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17091620 kB) free!
[1589655751] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410712 kB) free!
[1589655922] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 1.3% (323632 kB) free!
[1589656351] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22326400 kB) free!
[1589656544] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (636524 kB) free!
[1589656630] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1589656667] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (525708 kB) free!
[1589656750] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1589656787] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (408332 kB) free!
[1589656951] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388860 kB) free!
[1589657111] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 37106 MB (16.51% inode=100%):
[1589657171] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 31728 MB (14.11% inode=100%):
[1589657231] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;3;DISK WARNING - free space: /home 27362 MB (12.17% inode=100%):
[1589657291] SERVICE ALERT: webserver2;HDD_Home status;WARNING;HARD;4;DISK WARNING - free space: /home 23268 MB (10.35% inode=100%):
[1589657388] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.5% (2033636 kB) free!
[1589657551] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (24052752 kB) free!
[1589657591] SERVICE ALERT: webserver2;HDD_Home status;CRITICAL;HARD;4;DISK CRITICAL - free space: /home 9484 MB (4.22% inode=100%):
[1589657891] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 60241 MB (26.80% inode=100%):
[1589658023] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 40.3% (9877672 kB) free.
[1589658151] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400512 kB) free!
[1589658751] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26189000 kB) free!
[1589659108] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230524 kB) free!
[1589659188] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (512524 kB) free!
[1589659228] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234924 kB) free!
[1589659304] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5711356 kB) free!
[1589659348] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228832 kB) free!
[1589659351] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407060 kB) free!
[1589659424] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385576 kB) free!
[1589659544] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400256 kB) free!
[1589659788] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (2057208 kB) free!
[1589659951] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27651788 kB) free.
[1589660423] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178300 kB) free!
[1589660483] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163756 kB) free!
[1589660543] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (196088 kB) free!
[1589660551] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1574032 kB) free!
[1589660603] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (177288 kB) free!
[1589660671] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389500 kB) free!
[1589660791] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439664 kB) free!
[1589661148] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.1% (8942368 kB) free!
[1589661391] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13441252 kB) free!
[1589661803] SERVICE FLAPPING ALERT: webserver2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1589661803] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 67.2% (16462088 kB) free.
[1589661992] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376468 kB) free!
[1589662334] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.68, 6.75, 5.90
[1589662454] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.70, 6.73, 6.00
[1589662574] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.62, 6.70, 6.07
[1589663396] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1703 processes
[1589663661] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.1% (44419892 kB) free.
[1589663792] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15094348 kB) free!
[1589663988] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3260980 kB) free.
[1589663996] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1694 processes
[1589664104] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.0% (44354680 kB) free.
[1589664203] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (741884 kB) free!
[1589664264] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (744156 kB) free!
[1589664324] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 3.0% (746196 kB) free!
[1589664384] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 3.0% (746724 kB) free!
[1589664392] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412400 kB) free!
[1589664597] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1701 processes
[1589664704] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(428) < 2020/05/17 03:01:32 [error] 22125#0: *88935695 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589664748] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.9% (44318680 kB) free.
[1589664824] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589664824] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(320) < 103.16.101.72 [17/May/2020:03:01:32 +0530] TCP 502 0 0 60.000 "10.147.212.38:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1589664945] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589664992] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15689980 kB) free!
[1589665142] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1701) < 203.212.70.170 [17/May/2020:03:01:52 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589665188] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (325096 kB) free!
[1589665262] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589665308] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3359888 kB) free.
[1589665592] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (413964 kB) free!
[1589665797] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1589665797] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1697 processes
[1589666192] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15758628 kB) free!
[1589666792] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403284 kB) free!
[1589667392] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21120704 kB) free!
[1589667984] SERVICE FLAPPING ALERT: webserver2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1589667992] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1211884 kB) free!
[1589668154] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1589668275] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1589668395] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1589668574] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 0.87, 1.73, 3.40
[1589668592] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18230108 kB) free!
[1589669192] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410820 kB) free!
[1589669792] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14761244 kB) free!
[1589670392] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31997668 kB) free.
[1589670795] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1589670993] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11112640 kB) free!
[1589671113] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7492476 kB) free!
[1589671221] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.2% (9858744 kB) free.
[1589671233] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3922584 kB) free!
[1589671821] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.1% (2961592 kB) free!
[1589671833] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18765152 kB) free!
[1589671908] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1589671941] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194816 kB) free!
[1589672061] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178352 kB) free!
[1589672433] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32467068 kB) free.
[1589673033] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13924280 kB) free!
[1589673153] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10234532 kB) free!
[1589673273] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6628196 kB) free!
[1589673401] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 67.0% (16423272 kB) free.
[1589673873] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20105684 kB) free!
[1589674197] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1703 processes
[1589674225] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 53.7% (13172156 kB) free.
[1589674473] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2166892 kB) free!
[1589675073] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19124736 kB) free!
[1589675397] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1700 processes
[1589675673] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32877264 kB) free.
[1589675774] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.52, 5.79, 5.29
[1589675894] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.16, 5.23, 5.16
[1589676014] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.30, 5.31, 5.19
[1589676195] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1589676273] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12580884 kB) free!
[1589676315] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 340 processes
[1589676393] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8777236 kB) free!
[1589676435] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1589676513] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5145472 kB) free!
[1589676597] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1707 processes
[1589677114] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18627676 kB) free!
[1589677197] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1696 processes
[1589677714] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31815208 kB) free.
[1589678144] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 65.0% (85722992 kB) free.
[1589678314] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12409636 kB) free!
[1589678414] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.57, 2.76, 3.98
[1589678434] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8658260 kB) free!
[1589678554] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32529220 kB) free.
[1589679154] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16350224 kB) free!
[1589679274] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12538444 kB) free!
[1589679394] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8865872 kB) free!
[1589679994] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20938404 kB) free!
[1589680594] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34976248 kB) free.
[1589681194] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14356888 kB) free!
[1589681314] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10515972 kB) free!
[1589681434] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (6701688 kB) free!
[1589682034] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19756396 kB) free!
[1589682634] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33685992 kB) free.
[1589683235] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11106556 kB) free!
[1589683355] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7283756 kB) free!
[1589683475] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3492148 kB) free!
[1589684075] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17830036 kB) free!
[1589684675] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32608864 kB) free.
[1589685275] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10610576 kB) free!
[1589685395] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6518404 kB) free!
[1589685515] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2405448 kB) free!
[1589686115] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15911248 kB) free!
[1589686284] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.9% (4387420 kB) free!
[1589686715] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32782204 kB) free.
[1589686814] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.09, 6.75, 5.38
[1589686934] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.83, 7.74, 5.93
[1589686974] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.12479.500> (10.147.212.158:57806 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1589687054] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.40, 8.30, 6.35
[1589687094] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589687315] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7761380 kB) free!
[1589687418] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 689 processes
[1589687435] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3783032 kB) free!
[1589687555] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405232 kB) free!
[1589688155] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16700092 kB) free!
[1589688224] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/17 09:27:40 [error] 22120#0: *88936836 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: rq-2-1gbps.server, request: "GET /favicon.ico HTTP/1.0", host: "rq-2-1gbps.server", referrer: "http://rq-2-1gbps.server/"
[1589688344] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589688597] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589688625] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (4003556 kB) free!
[1589688745] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3745608 kB) free!
[1589688755] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432648 kB) free!
[1589688865] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (3511692 kB) free!
[1589689035] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1589689465] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2196328 kB) free!
[1589689956] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31710124 kB) free.
[1589690508] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3123404 kB) free!
[1589690556] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5533776 kB) free!
[1589690628] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3533616 kB) free.
[1589690676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1285288 kB) free!
[1589690796] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383484 kB) free!
[1589690938] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589691058] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589691178] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589691396] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17847548 kB) free!
[1589691401] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178696 kB) free!
[1589691521] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189276 kB) free!
[1589691641] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176508 kB) free!
[1589691828] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (290248 kB) free!
[1589691948] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3063768 kB) free!
[1589691996] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387896 kB) free!
[1589692068] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3002640 kB) free!
[1589692378] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1589693150] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (3765948 kB) free!
[1589693270] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3128776 kB) free!
[1589693390] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.1% (2716900 kB) free!
[1589693990] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (786516 kB) free!
[1589696745] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.5% (6120980 kB) free!
[1589696865] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (5790592 kB) free!
[1589696941] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (4061956 kB) free!
[1589696985] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 15.7% (5491080 kB) free!
[1589697061] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.2% (3753600 kB) free!
[1589697181] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3326548 kB) free!
[1589697468] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (780196 kB) free!
[1589697585] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.0% (31864756 kB) free.
[1589697590] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16189584 kB) free.
[1589697662] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(245) < 103.255.146.164 [17/May/2020:12:01:49 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589697781] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1389552 kB) free!
[1589697782] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589698068] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (2205184 kB) free!
[1589698381] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.9% (22929184 kB) free.
[1589698982] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(702) < 103.255.146.164 [17/May/2020:12:31:20 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589699103] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589699268] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1589699268] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3755868 kB) free.
[1589700861] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9697812 kB) free!
[1589700981] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9447272 kB) free!
[1589700996] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14103164 kB) free!
[1589701068] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (2498164 kB) free!
[1589701101] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (9214020 kB) free!
[1589701188] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.5% (2360100 kB) free!
[1589701308] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (2190212 kB) free!
[1589701502] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(640) < 103.255.146.164 [17/May/2020:13:12:05 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589701596] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (442852 kB) free!
[1589701622] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589702196] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13365564 kB) free!
[1589702796] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446608 kB) free!
[1589703396] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17503880 kB) free!
[1589703704] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9719964 kB) free!
[1589703824] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9594488 kB) free!
[1589703944] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (9446220 kB) free!
[1589703996] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400384 kB) free!
[1589704308] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3314676 kB) free.
[1589704596] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16724132 kB) free!
[1589704701] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4208456 kB) free!
[1589704908] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3135516 kB) free!
[1589705028] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (3091604 kB) free!
[1589705054] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.25, 4.22, 4.86
[1589705148] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3073584 kB) free!
[1589705196] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416372 kB) free!
[1589705796] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19140340 kB) free!
[1589705822] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(373) < 103.255.146.164 [17/May/2020:14:18:19 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589705943] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589706226] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1589706347] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1589706397] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (458512 kB) free!
[1589706590] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3631380 kB) free!
[1589706710] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3426868 kB) free!
[1589706830] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3162204 kB) free!
[1589707430] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2099600 kB) free!
[1589707454] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.42, 5.84, 5.37
[1589707548] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1270244 kB) free!
[1589707574] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.01, 6.03, 5.51
[1589707694] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.75, 6.18, 5.61
[1589708031] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1589708031] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16188804 kB) free.
[1589708148] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (1851804 kB) free!
[1589708197] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23413784 kB) free!
[1589708240] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1589708342] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1357) < 172.105.57.57 [17/May/2020:15:08:59 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589708462] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(10) < 203.212.70.170 [17/May/2020:15:09:49 +0530] TCP 502 0 0 120.001 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1589708583] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589708744] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4801172 kB) free!
[1589708797] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388536 kB) free!
[1589708894] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.78, 3.73, 4.43
[1589709397] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32266476 kB) free.
[1589709782] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(146) < 103.255.146.164 [17/May/2020:15:29:41 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589709902] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589709997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (379720 kB) free!
[1589710040] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1589710117] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393384 kB) free!
[1589710160] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 340 processes
[1589710237] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393292 kB) free!
[1589710281] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1589710837] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23243312 kB) free!
[1589711148] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (887408 kB) free!
[1589711437] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408544 kB) free!
[1589711481] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589711748] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.9% (2256920 kB) free!
[1589711894] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.57, 7.65, 6.03
[1589712014] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 15.26, 9.60, 6.89
[1589712037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32524900 kB) free.
[1589712134] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1589712134] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.42, 10.30, 7.47
[1589712637] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (381708 kB) free!
[1589712757] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (411876 kB) free!
[1589712877] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374224 kB) free!
[1589712902] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(552) < 103.255.146.164 [17/May/2020:16:16:27 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589713022] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589713478] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24349988 kB) free!
[1589714078] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398688 kB) free!
[1589714678] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22863776 kB) free!
[1589715278] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389428 kB) free!
[1589715681] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1589715681] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1589715734] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.96, 3.86, 4.68
[1589715878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29543916 kB) free.
[1589716022] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(159) < 203.212.70.170 [17/May/2020:17:13:41 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1589716142] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589716478] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (414932 kB) free!
[1589716598] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (428256 kB) free!
[1589716718] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (445780 kB) free!
[1589716934] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.02, 5.86, 5.22
[1589717031] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4093704 kB) free!
[1589717054] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.86, 6.36, 5.49
[1589717148] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (183016 kB) free!
[1589717151] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.9% (3797924 kB) free!
[1589717174] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.69, 7.23, 5.92
[1589717272] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.5% (3486968 kB) free!
[1589717318] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14949036 kB) free!
[1589717481] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1589717602] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1589717722] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1589717748] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2623488 kB) free!
[1589717872] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2190036 kB) free!
[1589717918] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439100 kB) free!
[1589718518] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22776552 kB) free!
[1589718542] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(661) < 103.255.146.164 [17/May/2020:17:58:12 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589718663] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589718949] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1386700 kB) free!
[1589719072] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.7% (16108132 kB) free.
[1589719118] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415328 kB) free!
[1589719718] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25795440 kB) free!
[1589720122] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1589720149] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3421252 kB) free.
[1589720174] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.53, 4.16, 4.45
[1589720318] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407696 kB) free!
[1589720722] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1589720749] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3232268 kB) free!
[1589720869] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2891980 kB) free!
[1589720918] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30691568 kB) free.
[1589720989] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2846204 kB) free!
[1589721084] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (171728 kB) free!
[1589721384] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.5% (4290268 kB) free!
[1589721519] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (391868 kB) free!
[1589721639] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (398140 kB) free!
[1589721759] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393300 kB) free!
[1589721922] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1589721974] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.87, 5.67, 5.10
[1589722094] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.17, 5.11, 4.96
[1589722189] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.3% (3951628 kB) free.
[1589722359] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19091960 kB) free!
[1589722959] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416600 kB) free!
[1589723146] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 334 processes
[1589723266] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1589723389] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2599720 kB) free!
[1589723509] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3232900 kB) free!
[1589723559] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22949288 kB) free!
[1589723629] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3644368 kB) free.
[1589724159] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390232 kB) free!
[1589724229] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2118052 kB) free!
[1589724349] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2658692 kB) free!
[1589724469] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3287984 kB) free.
[1589724759] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25300336 kB) free!
[1589725069] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (1418184 kB) free!
[1589725190] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.8% (1923856 kB) free!
[1589725310] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (2375372 kB) free!
[1589725359] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386212 kB) free!
[1589725959] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25834324 kB) free!
[1589726559] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414912 kB) free!
[1589726894] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.84, 6.41, 5.47
[1589727014] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.85, 7.10, 5.82
[1589727134] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.47, 6.52, 5.77
[1589727159] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25563404 kB) free!
[1589727759] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (434116 kB) free!
[1589728359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26612892 kB) free.
[1589728935] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.23, 4.29, 4.76
[1589728959] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (382012 kB) free!
[1589729080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (436108 kB) free!
[1589729200] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (435620 kB) free!
[1589729510] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.3% (3948052 kB) free.
[1589729800] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15368904 kB) free!
[1589729872] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589730400] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416268 kB) free!
[1589731000] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16120988 kB) free!
[1589731600] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (457396 kB) free!
[1589732200] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15854260 kB) free!
[1589732722] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589732800] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (421652 kB) free!
[1589732998] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1709 processes
[1589733110] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2532080 kB) free!
[1589733230] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.0% (2433248 kB) free!
[1589733350] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2537080 kB) free!
[1589733400] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14299668 kB) free!
[1589733950] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.3% (4276892 kB) free.
[1589734000] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398068 kB) free!
[1589734600] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14167980 kB) free!
[1589735200] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (440804 kB) free!
[1589736350] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2977392 kB) free!
[1589736470] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (2746388 kB) free!
[1589736590] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2639792 kB) free!
[1589737190] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3987904 kB) free.
[1589738800] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31569196 kB) free.
[1589739135] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1589739400] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6259176 kB) free!
[1589739520] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (2027076 kB) free!
[1589739640] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403268 kB) free!
[1589739922] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes