[1581445820] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11816888 kB) free!
[1581446267] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2020/02/12 00:00:49 [error] 10222#0: *52417963 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581446276] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 49.50.64.122 [12/Feb/2020:00:00:54 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581446299] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184256 kB) free!
[1581446372] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.0% (29554080 kB) free.
[1581446387] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581446396] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581446409] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188932 kB) free!
[1581446419] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194752 kB) free!
[1581446530] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166068 kB) free!
[1581446539] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175820 kB) free!
[1581446650] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182648 kB) free!
[1581446670] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (209488 kB) free!
[1581446790] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (208444 kB) free!
[1581446910] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204960 kB) free!
[1581449050] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.6% (10009864 kB) free.
[1581449420] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.2% (13403564 kB) free!
[1581450193] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1581450250] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190504 kB) free!
[1581450313] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1581450370] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185776 kB) free!
[1581450491] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168100 kB) free!
[1581450710] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (481800 kB) free!
[1581453739] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.2% (12067132 kB) free.
[1581454659] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (218184 kB) free!
[1581454779] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242024 kB) free!
[1581454827] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 67.4% (33227384 kB) free.
[1581454899] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218088 kB) free!
[1581455500] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.7% (40762092 kB) free.
[1581457111] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(207) < 103.16.101.72 [12/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581457187] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(258) < 2020/02/12 03:00:22 [error] 10213#0: *52462207 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581457196] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2916) < 172.105.59.173 [12/Feb/2020:03:02:31 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1581457232] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581457307] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581457316] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581460907] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/12 04:07:38 [error] 10222#0: *52478124 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1581461027] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581463939] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182280 kB) free!
[1581464059] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (212520 kB) free!
[1581464179] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178584 kB) free!
[1581464891] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.6% (11256236 kB) free.
[1581464910] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.9% (20561560 kB) free.
[1581467113] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581467233] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581468620] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.6% (40335088 kB) free.
[1581470087] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.5% (39679244 kB) free.
[1581470179] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.1% (21597200 kB) free.
[1581471116] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(147) < 92.42.104.166 [12/Feb/2020:07:00:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581471236] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581473291] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (1801792 kB) free!
[1581473411] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (1788132 kB) free!
[1581473531] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1787412 kB) free!
[1581473753] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.9% (21683672 kB) free.
[1581474237] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 92.42.104.166 [12/Feb/2020:07:45:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581474357] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(148) < 92.42.104.166 [12/Feb/2020:07:55:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581474477] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581474744] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 411 processes
[1581474827] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(265) < 2020/02/12 08:03:45 [error] 10210#0: *52541319 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581474864] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 405 processes
[1581474948] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581474984] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 389 processes
[1581475077] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(651) < 172.105.59.173 [12/Feb/2020:08:04:36 +0530] TCP 502 0 0 127.014 "10.147.212.154:3001, 10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "7.012, 60.001, 60.001, 0.000"
[1581475197] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581476747] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/02/12 08:32:55 [error] 10212#0: *52552985 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581476867] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581486359] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.7% (3639652 kB) free!
[1581486479] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.3% (3523608 kB) free!
[1581486599] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (3431312 kB) free!
[1581486961] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1581487082] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1581487199] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1790984 kB) free!
[1581487284] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.4% (21562228 kB) free.
[1581488799] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.8% (4609552 kB) free!
[1581488914] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4473416 kB) free!
[1581488919] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.8% (4351664 kB) free!
[1581489034] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.1% (15526548 kB) free.
[1581489039] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.9% (4137576 kB) free!
[1581489169] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1581489289] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581489384] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1581489467] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2020/02/12 12:00:49 [error] 10217#0: *52636786 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581489587] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581489597] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [12/Feb/2020:12:02:06 +0530] TCP 502 0 0 67.015 "10.147.212.142:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 7.015, 60.000, 0.000"
[1581489717] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581490239] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2247796 kB) free!
[1581490588] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164648 kB) free!
[1581490709] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189284 kB) free!
[1581490829] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187812 kB) free!
[1581491089] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1581491210] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1581492385] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1581493288] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (2279788 kB) free!
[1581493408] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (1111852 kB) free!
[1581493528] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1805076 kB) free!
[1581493609] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1581493609] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.2% (21613584 kB) free.
[1581493664] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1581493664] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.7% (21648156 kB) free.
[1581493690] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21629504 kB) free.
[1581493729] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1581493918] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(478) < 103.255.146.164 [12/Feb/2020:13:19:38 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581494038] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581494688] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8954260 kB) free!
[1581494808] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8953812 kB) free!
[1581494928] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (8954820 kB) free!
[1581496300] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9100876 kB) free!
[1581496420] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8815168 kB) free!
[1581496540] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.4% (8551320 kB) free!
[1581496585] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1581496618] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4774048 kB) free!
[1581496670] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168056 kB) free!
[1581496710] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4457112 kB) free!
[1581496739] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4626264 kB) free!
[1581496746] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (1402932 kB) free!
[1581496790] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177988 kB) free!
[1581496831] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4385420 kB) free!
[1581496859] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (4493148 kB) free!
[1581496866] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (1403400 kB) free!
[1581496911] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168724 kB) free!
[1581496951] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4314480 kB) free!
[1581496986] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1400204 kB) free!
[1581497434] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4053148 kB) free!
[1581497554] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (3718192 kB) free!
[1581497674] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.0% (3378580 kB) free!
[1581497786] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1581497786] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581498059] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1581498059] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2263052 kB) free!
[1581498274] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1856884 kB) free!
[1581498340] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4129076 kB) free!
[1581498386] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1581498528] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3802552 kB) free!
[1581498751] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2388204 kB) free!
[1581499129] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1581499249] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1581499369] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581499772] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8269448 kB) free!
[1581499827] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9517392 kB) free!
[1581499893] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.9% (30029476 kB) free.
[1581499947] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9378556 kB) free!
[1581499960] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18260032 kB) free!
[1581500038] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(164) < 92.42.104.166 [12/Feb/2020:15:00:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581500067] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (9233856 kB) free!
[1581500075] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.3% (15552084 kB) free.
[1581500080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.0% (48756644 kB) free.
[1581500159] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581500890] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1581500900] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1581501119] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1581501279] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1581501359] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(750) < 103.255.146.164 [12/Feb/2020:15:25:36 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581501479] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581502538] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3193408 kB) free!
[1581502658] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (2950208 kB) free!
[1581502778] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.2% (5072980 kB) free.
[1581503546] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 669 processes
[1581503666] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 664 processes
[1581503787] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 363 processes
[1581503879] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(203) < 203.212.70.170 [12/Feb/2020:16:02:23 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1581503987] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(89) < 2020/02/12 16:01:07 [error] 10214#0: *52761797 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581503999] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581504107] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581504123] SERVICE ALERT: mapserver;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581504154] HOST ALERT: mapserver;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1581504243] SERVICE ALERT: mapserver;Total Processes;OK;HARD;1;PROCS OK: 213 processes
[1581504309] HOST ALERT: mapserver;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1581504386] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1581504386] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1581504541] SERVICE ALERT: mapserver;HDD_root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581504600] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(594) < 103.255.146.164 [12/Feb/2020:16:15:34 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581504662] SERVICE ALERT: mapserver;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581504671] SERVICE ALERT: mapserver;HDD_root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1581504678] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4783336 kB) free!
[1581504689] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.2% (18550420 kB) free.
[1581504692] HOST ALERT: mapserver;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1581504708] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(11) < 2020/02/12 16:19:57 [error] 19748#0: *52770588 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: sigqscluster2-1gbps.server, request: "GET /favicon.ico HTTP/1.1", host: "sigqscluster2-1gbps.server:8091", referrer: "http://sigqscluster2-1gbps.server:8091/"
[1581504720] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581504730] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 58.0% (28570768 kB) free.
[1581504756] HOST FLAPPING ALERT: mapserver;STARTED; Host appears to have started flapping (21.8% change > 20.0% threshold)
[1581504756] HOST ALERT: mapserver;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1581504760] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.4% (9569284 kB) free!
[1581504763] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1581504763] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 41.0% (20181832 kB) free.
[1581504782] SERVICE ALERT: mapserver;Memory;OK;HARD;1;OK - 89.4% (18229424 kB) free.
[1581504791] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 74.7% (18318512 kB) free.
[1581504791] SERVICE ALERT: mapserver;HDD_root status;OK;HARD;1;DISK OK - free space: / 210730 MB (97% inode=99%):
[1581504798] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.2% (4709176 kB) free!
[1581504802] HOST FLAPPING ALERT: mapserver;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1581504802] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1581504802] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1581504828] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581504851] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 68.4% (16757908 kB) free.
[1581504867] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4578676 kB) free!
[1581504889] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.2% (14169436 kB) free.
[1581504907] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 32.2% (15848136 kB) free.
[1581504918] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.9% (4637584 kB) free!
[1581505029] SERVICE ALERT: mapserver;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581505060] HOST ALERT: mapserver;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1581505076] HOST ALERT: mapserver;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 0.62 ms
[1581505159] SERVICE ALERT: mapserver;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1581505279] SERVICE ALERT: mapserver;CPU_Procs;OK;HARD;1;CPU OK: 215 processes
[1581505300] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (3084168 kB) free!
[1581505420] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194732 kB) free!
[1581505499] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21113976 kB) free!
[1581505540] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165000 kB) free!
[1581505559] SERVICE ALERT: mapserver;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581505619] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43396068 kB) free.
[1581505679] SERVICE ALERT: mapserver;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.01, 0.05
[1581506189] HOST ALERT: mapserver;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1581506253] HOST ALERT: mapserver;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.89 ms
[1581506579] HOST ALERT: mapserver;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1581506589] HOST ALERT: mapserver;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.83 ms
[1581506819] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22322724 kB) free!
[1581506940] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44187572 kB) free.
[1581507002] SERVICE ALERT: mapserver;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581507132] SERVICE ALERT: mapserver;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1581507216] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.4% (20096996 kB) free.
[1581507253] SERVICE ALERT: mapserver;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1581507480] SERVICE ALERT: mapserver;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581507539] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20864464 kB) free!
[1581507610] HOST FLAPPING ALERT: mapserver;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581507610] SERVICE ALERT: mapserver;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1581507659] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43275492 kB) free.
[1581507730] SERVICE ALERT: mapserver;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1581507918] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2257604 kB) free!
[1581508259] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20219332 kB) free!
[1581508379] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42349600 kB) free.
[1581508939] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8784832 kB) free!
[1581508979] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18769804 kB) free!
[1581508992] SERVICE ALERT: mapserver;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581509010] HOST ALERT: mapserver;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1581509023] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1623632 kB) free!
[1581509053] SERVICE ALERT: mapserver;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581509058] HOST ALERT: mapserver;UP;SOFT;1;PING OK - Packet loss = 60%, RTA = 0.59 ms
[1581509059] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (2910440 kB) free!
[1581509099] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (22725944 kB) free!
[1581509122] SERVICE ALERT: mapserver;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1581509143] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174400 kB) free!
[1581509173] SERVICE ALERT: mapserver;Total Processes;OK;HARD;1;PROCS OK: 212 processes
[1581509179] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2904384 kB) free!
[1581509219] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37312948 kB) free.
[1581509242] SERVICE ALERT: mapserver;Memory;OK;HARD;1;OK - 89.3% (18213316 kB) free.
[1581509264] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1581509264] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186220 kB) free!
[1581509378] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (2498744 kB) free!
[1581509498] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2599764 kB) free!
[1581509540] SERVICE ALERT: mapserver;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581509570] HOST ALERT: mapserver;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1581509601] SERVICE ALERT: mapserver;HDD_root status;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1581509618] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (2713856 kB) free!
[1581509661] SERVICE ALERT: mapserver;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1581509665] HOST FLAPPING ALERT: mapserver;STARTED; Host appears to have started flapping (20.5% change > 20.0% threshold)
[1581509665] HOST ALERT: mapserver;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1581509686] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.9% (22174676 kB) free.
[1581509714] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 59.4% (29249020 kB) free.
[1581509721] SERVICE ALERT: mapserver;HDD_root status;OK;HARD;1;DISK OK - free space: / 210730 MB (97% inode=99%):
[1581509750] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 75.7% (18566216 kB) free.
[1581509819] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14168592 kB) free!
[1581509826] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1581509939] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38124204 kB) free.
[1581510275] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4406088 kB) free!
[1581510395] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4159076 kB) free!
[1581510515] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.7% (3981256 kB) free!
[1581510539] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13736900 kB) free!
[1581510659] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37949048 kB) free.
[1581510818] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3916704 kB) free.
[1581511116] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.2% (15540336 kB) free.
[1581511259] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13058824 kB) free!
[1581511379] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35473832 kB) free.
[1581511493] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187776 kB) free!
[1581511613] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189228 kB) free!
[1581511734] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171308 kB) free!
[1581511979] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11677016 kB) free!
[1581512018] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2956904 kB) free!
[1581512100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.8% (47194984 kB) free.
[1581512138] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3025860 kB) free!
[1581512258] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3016980 kB) free!
[1581512699] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22493504 kB) free!
[1581512819] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.3% (47842980 kB) free.
[1581512858] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1581512858] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3693416 kB) free.
[1581513419] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21569972 kB) free!
[1581513539] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.9% (47357288 kB) free.
[1581514139] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21460736 kB) free!
[1581514259] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17590332 kB) free!
[1581514320] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(531) < 103.255.146.164 [12/Feb/2020:18:59:58 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581514379] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43610268 kB) free.
[1581514440] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581514979] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18566592 kB) free!
[1581515099] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44484712 kB) free.
[1581515699] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20956192 kB) free!
[1581515723] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9552936 kB) free!
[1581515819] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (16970508 kB) free!
[1581515843] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8427668 kB) free!
[1581515858] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2836592 kB) free!
[1581515939] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42169792 kB) free.
[1581515963] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (8368532 kB) free!
[1581515978] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2811404 kB) free!
[1581516098] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2766652 kB) free!
[1581516539] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18584108 kB) free!
[1581516659] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43207796 kB) free.
[1581517260] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19924348 kB) free!
[1581517380] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16282496 kB) free!
[1581517500] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41486384 kB) free.
[1581517898] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.6% (4644064 kB) free.
[1581518028] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(75) < 2020/02/12 20:01:09 [error] 19831#0: *52848207 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581518040] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 182.18.144.246 [12/Feb/2020:20:02:17 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581518100] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15990540 kB) free!
[1581518148] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581518160] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581518220] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41964192 kB) free.
[1581518600] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4640100 kB) free!
[1581518720] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4533400 kB) free!
[1581518820] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17681808 kB) free!
[1581518841] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4405296 kB) free!
[1581518940] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13489772 kB) free!
[1581519060] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40230932 kB) free.
[1581519660] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14883472 kB) free!
[1581519780] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40059312 kB) free.
[1581520135] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1581520163] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4785512 kB) free!
[1581520380] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16327840 kB) free!
[1581520500] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12206920 kB) free!
[1581520620] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39066048 kB) free.
[1581521220] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16989832 kB) free!
[1581521236] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581521241] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2270740 kB) free!
[1581521340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13115100 kB) free!
[1581521357] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581521461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36374440 kB) free.
[1581521539] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.2% (20295540 kB) free.
[1581521578] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 63.5% (31297448 kB) free.
[1581521624] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 72.4% (17744188 kB) free.
[1581522060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13941032 kB) free!
[1581522180] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38526880 kB) free.
[1581522746] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (646056 kB) free!
[1581522780] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17584836 kB) free!
[1581522866] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (631416 kB) free!
[1581522900] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13915656 kB) free!
[1581522986] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200976 kB) free!
[1581523020] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37334016 kB) free.
[1581523620] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15314408 kB) free!
[1581523740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38264092 kB) free.
[1581524340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15872852 kB) free!
[1581524460] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12315372 kB) free!
[1581524580] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36197952 kB) free.
[1581525169] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1581525180] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14917608 kB) free!
[1581525257] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.3% (22042752 kB) free.
[1581525289] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 334 processes
[1581525300] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37979540 kB) free.
[1581525409] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1581525900] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15793076 kB) free!
[1581526020] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (11994660 kB) free!
[1581526140] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34981536 kB) free.
[1581526741] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12669792 kB) free!
[1581526861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37477996 kB) free.
[1581527461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14001576 kB) free!
[1581527581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37464880 kB) free.
[1581528098] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1581528181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15863840 kB) free!
[1581528301] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12093504 kB) free!
[1581528421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35735876 kB) free.
[1581529021] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12672544 kB) free!
[1581529141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35641796 kB) free.
[1581529741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13538512 kB) free!
[1581529861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37085728 kB) free.
[1581530461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14261928 kB) free!
[1581530581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36964576 kB) free.
[1581531092] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244208 kB) free!
[1581531181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15558164 kB) free!
[1581531212] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234552 kB) free!
[1581531301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.2% (24024372 kB) free!
[1581531332] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241244 kB) free!
[1581531422] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33430104 kB) free.
[1581531620] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3680148 kB) free!
[1581531740] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407820 kB) free!
[1581531860] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387716 kB) free!
[1581532021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13312648 kB) free!
[1581532141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34958208 kB) free.