[1581618906] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5667056 kB) free!
[1581619026] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30367756 kB) free.
[1581619069] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2020/02/14 00:06:33 [error] 19841#0: *53416617 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1581619182] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.1% (29124180 kB) free.
[1581619189] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581619250] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188788 kB) free!
[1581619337] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182656 kB) free!
[1581619370] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199300 kB) free!
[1581619457] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179616 kB) free!
[1581619490] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195284 kB) free!
[1581619577] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185528 kB) free!
[1581619626] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8049104 kB) free!
[1581619746] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4401240 kB) free!
[1581619866] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27567928 kB) free.
[1581620466] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6179208 kB) free!
[1581620587] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30195056 kB) free.
[1581621187] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8336240 kB) free!
[1581621307] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4740940 kB) free!
[1581621427] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27345984 kB) free.
[1581622028] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6106768 kB) free!
[1581622148] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30015572 kB) free.
[1581622337] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 31.7% (7812864 kB) free.
[1581622748] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7478160 kB) free!
[1581622868] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29886708 kB) free.
[1581622937] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164216 kB) free!
[1581623057] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169748 kB) free!
[1581623177] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174012 kB) free!
[1581623468] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8971268 kB) free!
[1581623588] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.1% (23853092 kB) free!
[1581623708] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26282720 kB) free!
[1581624308] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5714048 kB) free!
[1581624532] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (199172 kB) free!
[1581626108] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15571124 kB) free!
[1581627536] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (2128796 kB) free!
[1581627596] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234780 kB) free!
[1581627656] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 73.3% (36123764 kB) free.
[1581627716] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227740 kB) free!
[1581627836] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232304 kB) free!
[1581627890] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.4% (10150272 kB) free.
[1581627908] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42833748 kB) free.
[1581628436] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.6% (41186912 kB) free.
[1581629401] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 103.255.146.164 [14/Feb/2020:03:00:01 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1581629522] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(2304) < 182.18.144.246 [14/Feb/2020:03:01:48 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581629554] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(195) < 103.16.101.72 [14/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581629642] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581629675] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581629989] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(245) < 2020/02/14 03:00:22 [error] 19831#0: *53462639 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581630109] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581631879] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581631999] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1581633242] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [14/Feb/2020:04:01:44 +0530] TCP 502 0 0 63.004 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 3.004, 60.000, 0.000"
[1581633362] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581636890] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190912 kB) free!
[1581637010] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193900 kB) free!
[1581637131] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194668 kB) free!
[1581637999] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1581638108] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1581638119] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1581638177] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.1% (20625908 kB) free.
[1581638177] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.1% (12868052 kB) free.
[1581641301] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 31.7% (41722132 kB) free.
[1581641994] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39302288 kB) free.
[1581642035] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.9% (22036804 kB) free.
[1581642143] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.9% (14125632 kB) free.
[1581642377] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (503984 kB) free!
[1581642497] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (498812 kB) free!
[1581642617] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (490272 kB) free!
[1581644335] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1581644455] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 363 processes
[1581644575] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1581644578] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195332 kB) free!
[1581644698] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201984 kB) free!
[1581644777] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (205452 kB) free!
[1581644820] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201824 kB) free!
[1581644897] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (218608 kB) free!
[1581645017] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (218408 kB) free!
[1581645176] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1581645962] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(725) < 203.212.70.170 [14/Feb/2020:07:35:07 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581646082] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581646309] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(228) < 2020/02/14 07:37:41 [error] 19835#0: *53540314 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1581646429] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581646474] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(170) < 103.16.101.72 [14/Feb/2020:07:37:14 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581646577] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1581646595] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581646682] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(206) < 103.16.101.72 [14/Feb/2020:07:41:32 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1581646697] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1581646802] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581646818] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1581647220] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 70.5% (17289872 kB) free.
[1581647418] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1581647629] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(44) < 2020/02/14 08:00:58 [error] 19839#0: *53549969 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581647750] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581647813] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22093752 kB) free.
[1581647850] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.8% (22251300 kB) free.
[1581648018] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1581648138] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1581648258] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1581654258] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1581654258] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1581655415] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 712 processes
[1581655535] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 712 processes
[1581655655] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 714 processes
[1581656216] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4849700 kB) free!
[1581656336] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4717564 kB) free!
[1581656456] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3760872 kB) free!
[1581656855] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1581657056] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2467708 kB) free!
[1581657858] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1581657978] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1581658099] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1581658508] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17884028 kB) free!
[1581658628] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38561496 kB) free.
[1581659228] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14377996 kB) free!
[1581659348] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34246856 kB) free.
[1581659435] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.6% (3338912 kB) free!
[1581659444] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.7% (2854184 kB) free!
[1581659556] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (1992332 kB) free!
[1581659564] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.8% (2432012 kB) free!
[1581659676] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (686416 kB) free!
[1581659684] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2003720 kB) free!
[1581659948] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8100660 kB) free!
[1581660068] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.8% (49815496 kB) free.
[1581660267] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1581660372] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.9% (21441628 kB) free.
[1581660387] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1581660410] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.5% (15607992 kB) free.
[1581660420] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.7% (3846116 kB) free!
[1581660459] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.6% (21962564 kB) free.
[1581660507] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 333 processes
[1581660541] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.6% (3584324 kB) free!
[1581660582] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8942776 kB) free!
[1581660597] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9826132 kB) free!
[1581660661] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.5% (3314216 kB) free!
[1581660668] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24810264 kB) free!
[1581660702] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8255988 kB) free!
[1581660717] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9825660 kB) free!
[1581660779] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 52.5% (69173936 kB) free.
[1581660822] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.3% (7552900 kB) free!
[1581660838] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9825800 kB) free!
[1581660843] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21645688 kB) free.
[1581660905] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4827904 kB) free!
[1581661025] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.4% (4296340 kB) free!
[1581661145] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 15.3% (3783356 kB) free!
[1581661422] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4125024 kB) free!
[1581661746] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1373864 kB) free!
[1581662023] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.2% (30183296 kB) free.
[1581662149] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(201) < 2020/02/14 12:02:22 [error] 19840#0: *53662231 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1581662269] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581662308] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1581662402] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 172.105.59.173 [14/Feb/2020:12:02:55 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581662452] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.6% (4797800 kB) free!
[1581662522] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581662572] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.5% (4771256 kB) free!
[1581662692] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.4% (4755580 kB) free!
[1581662908] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1581662946] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 91.2% (22514908 kB) free.
[1581663028] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1581663148] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1581663456] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581663577] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1581663836] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8659708 kB) free!
[1581663838] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2957036 kB) free!
[1581663956] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8330536 kB) free!
[1581663975] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167140 kB) free!
[1581664076] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.3% (8034980 kB) free!
[1581664094] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.4% (21569864 kB) free.
[1581664120] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.2% (22099856 kB) free.
[1581664160] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39288836 kB) free.
[1581664231] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.6% (43187760 kB) free.
[1581664348] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1581664348] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 334 processes
[1581664672] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (1471516 kB) free!
[1581664792] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (1278744 kB) free!
[1581664912] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (1063904 kB) free!
[1581665895] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184704 kB) free!
[1581665966] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (218964 kB) free!
[1581666015] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180776 kB) free!
[1581666050] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1581666086] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217776 kB) free!
[1581666135] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183100 kB) free!
[1581666170] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581666206] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1581666206] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240160 kB) free!
[1581666457] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4549056 kB) free!
[1581666499] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1581666546] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1581666546] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.3% (21061276 kB) free.
[1581666577] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4468140 kB) free!
[1581666697] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4363020 kB) free!
[1581667155] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177232 kB) free!
[1581667275] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172416 kB) free!
[1581667387] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23499784 kB) free!
[1581667395] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187444 kB) free!
[1581667507] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.3% (49169540 kB) free.
[1581667996] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 71.2% (17560528 kB) free.
[1581668107] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25661972 kB) free!
[1581668227] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.6% (21930412 kB) free!
[1581668347] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45030092 kB) free.
[1581668548] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1581668947] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (21988016 kB) free!
[1581669067] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.6% (25788412 kB) free!
[1581669188] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39617452 kB) free.
[1581669697] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2382800 kB) free!
[1581669787] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17697908 kB) free!
[1581669907] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44121560 kB) free.
[1581670015] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4452112 kB) free!
[1581670135] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4128448 kB) free!
[1581670255] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.1% (3828856 kB) free!
[1581670507] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20834700 kB) free!
[1581670627] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16846920 kB) free!
[1581670747] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42726176 kB) free.
[1581670996] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (2386004 kB) free!
[1581671116] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (2231804 kB) free!
[1581671220] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 920 processes
[1581671236] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2091052 kB) free!
[1581671340] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 919 processes
[1581671347] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20227000 kB) free!
[1581671455] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1539432 kB) free!
[1581671456] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9820648 kB) free!
[1581671460] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 322 processes
[1581671467] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16122460 kB) free!
[1581671576] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9698260 kB) free!
[1581671587] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (12258288 kB) free!
[1581671696] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9568088 kB) free!
[1581672187] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20303928 kB) free!
[1581672246] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.3% (4487844 kB) free!
[1581672366] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.7% (4342572 kB) free!
[1581672486] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.1% (4194356 kB) free!
[1581672655] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.5% (15819632 kB) free.
[1581672722] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(202) < 92.42.104.166 [14/Feb/2020:15:00:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581672748] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1581672842] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581672868] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 343 processes
[1581672988] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 339 processes
[1581673104] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21939308 kB) free.
[1581673136] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.4% (21676524 kB) free.
[1581673175] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39616672 kB) free.
[1581673199] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.0% (21474064 kB) free.
[1581673229] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.7% (21739548 kB) free.
[1581673263] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 85.0% (41912648 kB) free.
[1581673311] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 44.0% (57969484 kB) free.
[1581673588] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1581673805] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (1366952 kB) free!
[1581673925] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (1329292 kB) free!
[1581674042] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [14/Feb/2020:15:18:31 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.000, 0.000"
[1581674045] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (1089972 kB) free!
[1581674162] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581674188] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1581674307] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.8% (2884464 kB) free!
[1581674427] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.8% (2648576 kB) free!
[1581674547] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2460152 kB) free!
[1581674788] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1581674908] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 338 processes
[1581674977] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243168 kB) free!
[1581675029] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1581675097] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241916 kB) free!
[1581675115] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (21993004 kB) free!
[1581675217] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243136 kB) free!
[1581675236] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18441636 kB) free!
[1581675356] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1581675356] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15000504 kB) free!
[1581675629] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1581676176] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581676296] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1581676336] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.4% (19847924 kB) free.
[1581676380] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39532344 kB) free.
[1581676407] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.2% (21375360 kB) free.
[1581676416] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1581676438] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 59.5% (78454188 kB) free.
[1581676562] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(339) < 103.255.146.164 [14/Feb/2020:16:01:49 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581676669] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(150) < 2020/02/14 16:01:56 [error] 19838#0: *53790477 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581676682] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581676699] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1581676789] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581676948] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (322552 kB) free!
[1581677068] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191912 kB) free!
[1581677188] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189388 kB) free!
[1581677234] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21961008 kB) free.
[1581677258] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2978636 kB) free!
[1581677282] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(470) < 103.255.146.164 [14/Feb/2020:16:12:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581677378] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3268936 kB) free.
[1581677402] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581677586] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (2565340 kB) free!
[1581677647] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39757068 kB) free.
[1581678214] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.8% (3377012 kB) free!
[1581678246] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22191040 kB) free!
[1581678334] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (1989652 kB) free!
[1581678366] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18438736 kB) free!
[1581678437] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194460 kB) free!
[1581678454] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1581678454] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (541676 kB) free!
[1581678486] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.4% (51916556 kB) free.
[1581678557] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186644 kB) free!
[1581678677] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175064 kB) free!
[1581679086] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25366276 kB) free!
[1581679178] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2957224 kB) free!
[1581679206] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21458128 kB) free!
[1581679231] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9475996 kB) free!
[1581679298] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3335016 kB) free.
[1581679326] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17430896 kB) free!
[1581679351] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9218708 kB) free!
[1581679455] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.9% (5360832 kB) free!
[1581679472] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9032804 kB) free!
[1581679575] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.9% (5360780 kB) free!
[1581679695] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 26.1% (12850504 kB) free.
[1581679898] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (1417500 kB) free!
[1581679926] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28299588 kB) free.
[1581680018] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1904324 kB) free!
[1581680138] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (2547460 kB) free!
[1581680338] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21915192 kB) free.
[1581680374] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.0% (22054024 kB) free.
[1581680738] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3700772 kB) free.
[1581680895] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9779516 kB) free!
[1581681002] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(607) < 203.212.70.170 [14/Feb/2020:17:19:54 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581681015] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9780736 kB) free!
[1581681029] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1581681055] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4328620 kB) free!
[1581681123] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581681127] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17790912 kB) free!
[1581681135] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (9780592 kB) free!
[1581681175] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4129584 kB) free!
[1581681247] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13775572 kB) free!
[1581681295] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.3% (3882792 kB) free!
[1581681367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43579772 kB) free.
[1581681395] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(182) < 103.16.101.72 [14/Feb/2020:17:21:35 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581681516] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581681589] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(259) < 2020/02/14 17:22:26 [error] 19840#0: *53834100 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581681709] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581681872] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4804044 kB) free!
[1581681938] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1244064 kB) free!
[1581681967] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16905244 kB) free!
[1581682058] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.0% (1788616 kB) free!
[1581682087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12586008 kB) free!
[1581682178] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (2390400 kB) free!
[1581682207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31970488 kB) free.
[1581682495] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1581682495] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1912996 kB) free!
[1581682519] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1581682639] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581682778] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1581682778] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (501104 kB) free!
[1581682807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5660368 kB) free!
[1581682927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1611924 kB) free!
[1581683047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35383028 kB) free.
[1581683378] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.3% (4595532 kB) free.
[1581683647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9295556 kB) free!
[1581683695] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.4% (15812852 kB) free.
[1581683767] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4764052 kB) free!
[1581683887] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (672076 kB) free!
[1581684487] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17546744 kB) free!
[1581684578] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.6% (2049336 kB) free!
[1581684578] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4806464 kB) free!
[1581684698] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.5% (2351224 kB) free!
[1581684699] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4638888 kB) free!
[1581684747] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (216904 kB) free!
[1581684818] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2937832 kB) free!
[1581684819] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4504776 kB) free!
[1581684867] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 69.9% (11363392 kB) free.
[1581685087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35713000 kB) free.
[1581685099] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1581685148] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166392 kB) free!
[1581685219] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1581685268] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174528 kB) free!
[1581685322] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(431) < 103.255.146.164 [14/Feb/2020:18:31:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581685335] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4466184 kB) free!
[1581685339] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1581685388] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183032 kB) free!
[1581685418] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3869668 kB) free.
[1581685419] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163292 kB) free!
[1581685442] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581685687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8650124 kB) free!
[1581685745] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.7% (3638584 kB) free!
[1581685808] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4227240 kB) free!
[1581685865] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.2% (3507140 kB) free!
[1581685928] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380032 kB) free!
[1581685985] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (3382784 kB) free!
[1581686018] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (2696592 kB) free!
[1581686138] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3107128 kB) free!
[1581686258] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3582420 kB) free.
[1581686528] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13668116 kB) free!
[1581686570] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1581686691] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1581687128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32113196 kB) free.
[1581687185] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1997204 kB) free!
[1581687629] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581687637] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4797384 kB) free!
[1581687728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6210900 kB) free!
[1581687749] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1581687757] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.3% (4720408 kB) free!
[1581687848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2424832 kB) free!
[1581687869] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1581687869] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1581687877] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.0% (4655408 kB) free!
[1581687968] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395904 kB) free!
[1581688469] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1581688589] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1581688658] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (2361204 kB) free!
[1581688709] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1581688778] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.7% (2547140 kB) free!
[1581688898] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (2665732 kB) free!
[1581689168] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31743448 kB) free.
[1581689498] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3649512 kB) free.
[1581689768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5761980 kB) free!
[1581689888] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1364796 kB) free!
[1581689909] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1581690008] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39155164 kB) free.
[1581690029] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1581690278] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2349348 kB) free!
[1581690608] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15003732 kB) free!
[1581690709] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(49) < 2020/02/14 20:01:17 [error] 19836#0: *53908202 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581690728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9572748 kB) free!
[1581690829] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581690848] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4565180 kB) free!
[1581691298] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (2421728 kB) free!
[1581691418] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.5% (2515660 kB) free!
[1581691449] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (19979836 kB) free!
[1581691538] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2627104 kB) free!
[1581691829] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1581691949] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1581692049] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37335004 kB) free.
[1581692649] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11059580 kB) free!
[1581692738] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3847736 kB) free.
[1581692769] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (7032320 kB) free!
[1581692889] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3287264 kB) free!
[1581693135] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1581693149] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1581693269] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1581693489] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18018508 kB) free!
[1581693672] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46263 MB (20% inode=99%):
[1581693792] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46263 MB (20% inode=99%):
[1581693912] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46263 MB (20% inode=99%):
[1581694089] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34327552 kB) free.
[1581694495] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581694689] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7881308 kB) free!
[1581694809] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4203832 kB) free!
[1581694929] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (742388 kB) free!
[1581695069] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1581695189] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1581695309] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1581695529] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15623084 kB) free!
[1581695910] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581696030] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1581696129] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30747580 kB) free.
[1581696150] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1581696729] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7986528 kB) free!
[1581696750] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1581696819] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581696849] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4366972 kB) free!
[1581696969] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36514448 kB) free.
[1581697569] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12280428 kB) free!
[1581697690] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8417628 kB) free!
[1581697810] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4772420 kB) free!
[1581697985] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581698410] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16130732 kB) free!
[1581699010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30468176 kB) free.
[1581699150] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1581699610] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6574680 kB) free!
[1581699730] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2815220 kB) free!
[1581699751] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1581699850] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35914128 kB) free.
[1581699871] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1581699991] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1581699991] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1581700450] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10310984 kB) free!
[1581700570] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6307288 kB) free!
[1581700591] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1581700690] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2495584 kB) free!
[1581701191] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1581701290] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13678872 kB) free!
[1581701311] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1581701539] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1581701911] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1581702032] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 342 processes
[1581702152] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1581702490] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5354832 kB) free!
[1581702752] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1581702872] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1581702992] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1581703090] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16043584 kB) free!
[1581703538] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581703690] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28475052 kB) free.
[1581704023] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237988 kB) free!
[1581704143] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236992 kB) free!
[1581704192] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1581704264] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238928 kB) free!
[1581704291] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7011092 kB) free!
[1581704301] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5641776 kB) free!
[1581704411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35199624 kB) free.
[1581704421] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389700 kB) free!
[1581704541] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414632 kB) free!
[1581704783] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 222355 MB (20% inode=99%):
[1581704792] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1581704903] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 214929 MB (20% inode=99%):
[1581704912] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes