[1549652215] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.98, 4.95, 4.98
[1549652815] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.91, 5.08, 5.03
[1549652935] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.03, 5.03, 5.02
[1549653055] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.62, 4.95, 5.00
[1549653659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549653779] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549653899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549654032] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9571048 kB) free!
[1549654152] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11704488 kB) free.
[1549654499] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549654855] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.45, 5.20, 5.07
[1549654975] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.44, 5.34, 5.14
[1549655095] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.35, 5.33, 5.16
[1549655099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549655219] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549655819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549655939] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549656895] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.85, 4.87, 5.00
[1549657152] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9617912 kB) free!
[1549657272] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11701360 kB) free.
[1549658095] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.88, 5.09, 5.03
[1549658215] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.67, 4.98, 5.00
[1549659072] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9798040 kB) free!
[1549659192] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11691248 kB) free.
[1549659539] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549659659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549659780] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549660380] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1549660380] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549661215] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.90, 5.10, 5.07
[1549661335] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.58, 5.17, 5.09
[1549661455] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1549661455] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.52, 5.27, 5.14
[1549663855] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.88, 4.90, 4.97
[1549665180] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549665300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549665420] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549668056] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.32, 5.14, 5.05
[1549668176] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.41, 5.17, 5.07
[1549668296] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.88, 5.04, 5.03
[1549668896] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.04, 5.01, 5.00
[1549670820] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549670940] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549671060] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549671896] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.08, 5.11, 5.03
[1549672016] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.98, 5.06, 5.02
[1549672136] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.89, 4.99, 5.00
[1549672260] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549678136] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.60, 5.01, 5.07
[1549678256] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.14, 5.05, 5.07
[1549678376] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.97, 4.97, 5.04
[1549678860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549678980] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549679100] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549683300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549683420] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549683540] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549684376] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1549685143] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.2% (15148412 kB) free.
[1549685216] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 92.1% (14971680 kB) free.
[1549685340] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549685940] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549686060] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549686180] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549686776] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.01, 4.97, 4.99
[1549687976] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.88, 5.03, 5.01
[1549688096] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.50, 5.14, 5.05
[1549688216] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.14, 5.11, 5.05
[1549688334] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (1867304 kB) free!
[1549688454] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (1838576 kB) free!
[1549688574] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1813080 kB) free!
[1549690380] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1549690380] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549690980] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549691100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549691221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549692421] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549693021] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549693141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549693261] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1549693374] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.5% (22916680 kB) free.
[1549693861] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549693981] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549694582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549694702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549694822] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549695416] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.01, 4.89, 4.96
[1549695555] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.82, 9.81, 5.95
[1549695675] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.85, 7.08, 5.41
[1549695795] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.08, 5.39, 5.00
[1549696616] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.71, 5.07, 5.08
[1549696622] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549696736] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.16, 5.05, 5.07
[1549696856] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1549696856] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.59, 5.17, 5.10
[1549697222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549697342] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549697943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549698063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549698183] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549698783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549698903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549699256] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.34, 4.79, 4.93
[1549699503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549699623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549699743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549699995] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 57.16, 16.49, 6.68
[1549700115] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.83, 11.48, 6.04
[1549700235] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.40, 8.13, 5.47
[1549700835] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.52, 2.46, 3.63
[1549701056] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.31, 5.07, 5.01
[1549701176] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.28, 5.19, 5.06
[1549701296] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.00, 5.12, 5.05
[1549702893] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1549702953] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 51 users currently logged in
[1549703013] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 51 users currently logged in
[1549703073] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 51 users currently logged in
[1549703943] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549704435] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.73, 10.43, 9.74
[1549704543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549704556] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.49, 7.48, 8.74
[1549704664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1549704676] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.46, 5.50, 7.87
[1549704784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549705276] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1549705276] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.85, 2.16, 4.87
[1549705773] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1549706373] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1549706433] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 51 users currently logged in
[1549706493] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 51 users currently logged in
[1549706553] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 51 users currently logged in
[1549706853] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1549706853] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 49 users currently logged in
[1549707676] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.86, 7.99, 8.39
[1549707796] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.48, 5.81, 7.54
[1549707916] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.47, 4.41, 6.82
[1549708516] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.30, 1.82, 4.26
[1549708985] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549709585] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549709696] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.17, 4.66, 4.88
[1549709705] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549709825] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549710425] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549711025] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549711145] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549711496] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.19, 5.06, 5.02
[1549711616] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.30, 5.12, 5.04
[1549711737] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.56, 4.95, 5.00
[1549712336] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.31, 5.07, 5.04
[1549712346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549712456] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.98, 5.06, 5.04
[1549712466] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549712553] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549712576] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.35, 5.13, 5.06
[1549712586] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549712716] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.48, 18.45, 9.18
[1549712836] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.22, 12.88, 8.27
[1549712956] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.89, 9.20, 7.48
[1549713453] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 52 users currently logged in
[1549713513] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 52 users currently logged in
[1549713556] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.09, 2.84, 4.76
[1549713573] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1549713633] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 53 users currently logged in
[1549718356] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 21.97, 25.72, 12.14
[1549718476] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.28, 17.70, 10.84
[1549718586] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1549718596] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.04, 12.40, 9.73
[1549718765] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 332 processes
[1549718885] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 332 processes
[1549718922] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1549719005] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 331 processes
[1549719012] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 317 processes
[1549719042] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 327 processes
[1549719132] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 317 processes
[1549719163] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 328 processes
[1549719176] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.78, 4.93, 4.99
[1549719186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549719252] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 317 processes
[1549719306] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1549719796] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.28, 1.63, 3.78
[1549719907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1549720027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549720147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549722033] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 45 users currently logged in
[1549722176] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.08, 5.14, 5.05
[1549722296] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.92, 5.10, 5.04
[1549722416] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.86, 5.10, 5.05
[1549723616] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.52, 4.88, 5.00
[1549724948] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549725792] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9796640 kB) free!
[1549725912] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11350888 kB) free.
[1549726016] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.84, 5.00, 5.01
[1549726136] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.36, 5.14, 5.06
[1549726257] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.84, 5.02, 5.03
[1549726512] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9479720 kB) free!
[1549726632] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10180096 kB) free.
[1549726748] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549726857] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.94, 4.88, 4.95
[1549726868] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549726988] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549727232] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9641008 kB) free!
[1549727352] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10267680 kB) free.
[1549727588] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549727709] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549728057] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.17, 5.23, 5.15
[1549728177] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.32, 5.30, 5.19
[1549728297] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.85, 5.16, 5.15
[1549728308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549728428] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549728552] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8289648 kB) free!
[1549728672] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.4% (11555128 kB) free.
[1549729028] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549729148] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549729268] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549729272] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9786684 kB) free!
[1549729392] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9836680 kB) free!
[1549729513] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.2% (9456556 kB) free!
[1549730596] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1549730743] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.9% (2739252 kB) free!
[1549730863] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.6% (2541156 kB) free!
[1549730983] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.5% (2361824 kB) free!
[1549731583] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1451784 kB) free!
[1549731706] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7785208 kB) free!
[1549731826] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10825604 kB) free.
[1549733026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8119092 kB) free!
[1549733146] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11015772 kB) free.
[1549734068] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549734313] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4319664 kB) free!
[1549734346] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7820416 kB) free!
[1549734424] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3095896 kB) free!
[1549734466] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10805108 kB) free.
[1549734544] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.0% (3095112 kB) free!
[1549734664] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.0% (3091880 kB) free!
[1549734668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549734788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549734913] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.8% (6805576 kB) free!
[1549735066] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7838596 kB) free!
[1549735186] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10473040 kB) free.
[1549735389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549735509] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549735513] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1549735513] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4392036 kB) free!
[1549735608] SERVICE ALERT: sig6;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1549735629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549735728] SERVICE ALERT: sig6;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1549735786] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6797776 kB) free!
[1549735848] SERVICE ALERT: sig6;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1549735907] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9731060 kB) free!
[1549736027] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.6% (9167720 kB) free!
[1549736464] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1392584 kB) free!
[1549736537] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 390 processes
[1549736657] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 390 processes
[1549736777] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 391 processes
[1549736829] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes