[1590345277] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1033180 kB) free!
[1590345397] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (1431392 kB) free!
[1590345517] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (1826516 kB) free!
[1590346673] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29261248 kB) free.
[1590347273] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3891628 kB) free!
[1590347317] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1149428 kB) free!
[1590347393] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374548 kB) free!
[1590347500] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1590347513] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407828 kB) free!
[1590347917] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (2381768 kB) free!
[1590348113] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14350532 kB) free!
[1590348517] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (592408 kB) free!
[1590348713] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408908 kB) free!
[1590348825] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 2.1% (505184 kB) free!
[1590349117] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.6% (2054124 kB) free!
[1590349717] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3875376 kB) free.
[1590349914] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29010752 kB) free.
[1590350025] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 76.4% (18715604 kB) free.
[1590350435] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219280 kB) free!
[1590350514] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4589992 kB) free!
[1590350555] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236572 kB) free!
[1590350626] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9942092 kB) free!
[1590350634] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (851284 kB) free!
[1590350675] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236660 kB) free!
[1590350746] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (412000 kB) free!
[1590350754] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432296 kB) free!
[1590350866] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391356 kB) free!
[1590351354] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13407432 kB) free!
[1590351954] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (456276 kB) free!
[1590352475] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.6% (9177048 kB) free!
[1590353154] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24359996 kB) free!
[1590353754] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403664 kB) free!
[1590354354] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23859208 kB) free!
[1590354764] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (2786192 kB) free!
[1590354954] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1148200 kB) free!
[1590355313] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (45235624 kB) free.
[1590355357] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.97, 4.23, 4.89
[1590355462] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 93.4% (46018344 kB) free.
[1590355554] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16976516 kB) free!
[1590355900] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1590355919] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(582) < 2020/05/25 03:01:42 [error] 22123#0: *90592153 no live upstreams while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590356005] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2748) < 203.212.70.170 [25/May/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590356035] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(474) < 49.50.86.152 [25/May/2020:03:01:42 +0530] TCP 502 0 0 60.001 "10.147.212.37:4000, backend3" "0, 0" "0, 0" "60.001, 0.000"
[1590356039] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590356075] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.8% (44272700 kB) free.
[1590356125] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590356154] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402596 kB) free!
[1590356156] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590357354] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28627076 kB) free.
[1590357955] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6037476 kB) free!
[1590358075] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2603456 kB) free!
[1590358195] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407180 kB) free!
[1590358379] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1553560 kB) free!
[1590358795] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16161248 kB) free!
[1590359395] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31322336 kB) free.
[1590359995] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12953224 kB) free!
[1590360115] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9539676 kB) free!
[1590360235] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34792188 kB) free.
[1590360835] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17042020 kB) free!
[1590360955] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13578644 kB) free!
[1590361075] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (10168988 kB) free!
[1590361117] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590361675] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20613748 kB) free!
[1590362275] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31970356 kB) free.
[1590362793] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.3% (9639416 kB) free.
[1590362875] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13664328 kB) free!
[1590362995] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10127104 kB) free!
[1590363116] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34752976 kB) free.
[1590363164] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178344 kB) free!
[1590363393] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170636 kB) free!
[1590363513] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191280 kB) free!
[1590363633] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177520 kB) free!
[1590363715] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16502180 kB) free!
[1590363835] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12936188 kB) free!
[1590363955] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (9365008 kB) free!
[1590364555] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19253528 kB) free!
[1590364922] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.0% (10365308 kB) free.
[1590364963] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.7% (18809284 kB) free.
[1590364964] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 49.2% (12067460 kB) free.
[1590365155] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30002928 kB) free.
[1590365755] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10366020 kB) free!
[1590365875] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6584840 kB) free!
[1590365995] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33852628 kB) free.
[1590366595] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12055904 kB) free!
[1590366715] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7660020 kB) free!
[1590366835] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3386388 kB) free!
[1590367435] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14884072 kB) free!
[1590368035] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27040388 kB) free.
[1590368635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7317484 kB) free!
[1590368755] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33644628 kB) free.
[1590369122] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.8% (3889248 kB) free!
[1590369242] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177156 kB) free!
[1590369356] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12633804 kB) free!
[1590369362] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202472 kB) free!
[1590369476] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8681216 kB) free!
[1590369596] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4875088 kB) free!
[1590370196] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16943616 kB) free!
[1590370796] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27346712 kB) free.
[1590370964] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4854128 kB) free!
[1590371084] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4827440 kB) free!
[1590371204] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (4792396 kB) free!
[1590371266] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 73.8% (97307964 kB) free.
[1590371396] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7714048 kB) free!
[1590371516] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32990920 kB) free.
[1590372116] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12897752 kB) free!
[1590372236] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8872524 kB) free!
[1590372356] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24585596 kB) free!
[1590374156] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3799088 kB) free!
[1590374360] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590374480] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590374501] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1590374563] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171088 kB) free!
[1590374601] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590374621] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590374683] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191320 kB) free!
[1590374741] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590374756] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13573548 kB) free!
[1590374803] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179112 kB) free!
[1590375201] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1590375341] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1590375957] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3292740 kB) free!
[1590376004] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437084 kB) free!
[1590376557] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13809148 kB) free!
[1590377157] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27106464 kB) free.
[1590377557] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.25, 6.47, 5.26
[1590377677] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.40, 7.09, 5.63
[1590377741] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1590377757] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2641124 kB) free!
[1590377797] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.10, 7.04, 5.79
[1590377861] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590377877] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (439248 kB) free!
[1590377894] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 66.6% (16320092 kB) free.
[1590377981] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590377997] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16758472 kB) free!
[1590378240] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(31) < 2020/05/25 09:07:48 [error] 22130#0: *90633843 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590378360] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590378597] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (8571012 kB) free!
[1590379197] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23919972 kB) free!
[1590379597] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.21, 3.82, 4.62
[1590379797] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1539008 kB) free!
[1590380382] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1590380382] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590380397] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14925876 kB) free!
[1590380982] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1590380997] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29343436 kB) free.
[1590381102] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 354 processes
[1590381222] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 354 processes
[1590381597] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3711860 kB) free!
[1590381717] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394428 kB) free!
[1590381837] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396304 kB) free!
[1590382597] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.85, 6.16, 5.33
[1590382717] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.33, 5.52, 5.19
[1590382837] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.14, 5.55, 5.23
[1590383037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28859156 kB) free.
[1590383638] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3529392 kB) free!
[1590383758] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (390572 kB) free!
[1590383878] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399060 kB) free!
[1590384037] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1590384037] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.82, 4.56, 4.98
[1590384638] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.29, 6.32, 5.64
[1590384758] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.52, 6.38, 5.75
[1590384878] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.87, 6.01, 5.68
[1590385078] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27701740 kB) free.
[1590385678] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1500300 kB) free!
[1590385694] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.6% (3822192 kB) free!
[1590385798] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (449940 kB) free!
[1590385814] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.9% (3648644 kB) free!
[1590385819] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3256192 kB) free!
[1590385918] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410864 kB) free!
[1590385935] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.1% (3466016 kB) free!
[1590385939] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.5% (3024404 kB) free!
[1590386022] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1590386059] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.5% (2816712 kB) free!
[1590386126] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(646) < 103.255.146.164 [25/May/2020:11:20:06 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1590386246] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590386659] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1608292 kB) free!
[1590387118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29735012 kB) free.
[1590387135] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1592292 kB) free!
[1590387361] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/25 11:37:52 [error] 22123#0: *90658789 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: _, request: "GET /favicon.ico HTTP/1.1", host: "148qservercluster1.com", referrer: "http://148qservercluster1.com/"
[1590387481] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590387718] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3021932 kB) free!
[1590387838] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397224 kB) free!
[1590387958] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423928 kB) free!
[1590388422] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 357 processes
[1590388460] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16065904 kB) free.
[1590389159] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25263212 kB) free!
[1590389317] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3041084 kB) free!
[1590389437] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4057008 kB) free.
[1590389678] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.62, 2.47, 4.32
[1590389759] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397068 kB) free!
[1590389881] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/25 12:27:21 [error] 22135#0: *90669967 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590390001] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590390179] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (1712724 kB) free!
[1590390359] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16197148 kB) free!
[1590390959] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (429868 kB) free!
[1590391559] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 332 processes
[1590391559] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15224916 kB) free!
[1590391679] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 4 warn out of 333 processes
[1590391799] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 329 processes
[1590392159] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (475268 kB) free!
[1590392399] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1590392678] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.25, 9.60, 6.68
[1590392798] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 26.04, 14.94, 8.92
[1590392918] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.97, 14.20, 9.39
[1590393359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34059568 kB) free.
[1590393959] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3682452 kB) free!
[1590394079] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397880 kB) free!
[1590394199] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411964 kB) free!
[1590394237] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (706708 kB) free!
[1590394357] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (708520 kB) free!
[1590394422] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590394477] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (874868 kB) free!
[1590394800] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15505076 kB) free!
[1590394913] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9807472 kB) free!
[1590395033] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9547916 kB) free!
[1590395077] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (3252396 kB) free!
[1590395153] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9267248 kB) free!
[1590395400] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (427344 kB) free!
[1590395660] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.5% (2575464 kB) free!
[1590395677] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3745780 kB) free.
[1590395780] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (2076636 kB) free!
[1590395900] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1615868 kB) free!
[1590396000] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13847712 kB) free!
[1590396296] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < webmachine error: path="/api/queues/%2F/FT5"
[1590396416] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590396600] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423492 kB) free!
[1590396601] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/25 14:13:57 [error] 22130#0: *90699673 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/ HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/", host: "altsmsc.tesync.net"
[1590396721] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590396877] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (2471480 kB) free!
[1590396997] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.1% (2294420 kB) free!
[1590397117] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1590397117] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.9% (2093696 kB) free!
[1590397553] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4371012 kB) free!
[1590398662] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8906880 kB) free!
[1590398782] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (8689104 kB) free!
[1590398902] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (8504300 kB) free!
[1590398917] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3470788 kB) free.
[1590399121] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/25 14:52:20 [error] 22135#0: *90710580 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590399201] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1590399242] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590399321] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1590399441] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1590399446] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(126) < 185.255.8.58 [25/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590399501] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.3% (16226084 kB) free.
[1590399517] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3161540 kB) free!
[1590399566] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590399637] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3105588 kB) free!
[1590399757] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3063388 kB) free!
[1590399841] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/05/25 15:10:37 [error] 22128#0: *90715110 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590399961] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590400335] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 57.7% (14152964 kB) free.
[1590400422] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1590400641] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1590400803] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(8) < webmachine error: path="/api/exchanges/%2f/NST1_EX"
[1590400924] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590400957] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1227396 kB) free!
[1590401400] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27973528 kB) free.
[1590401902] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4702948 kB) free!
[1590401918] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1590402000] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (504448 kB) free!
[1590402120] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (439048 kB) free!
[1590402157] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (1946380 kB) free!
[1590402241] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (421416 kB) free!
[1590402518] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.08, 3.48, 4.59
[1590402757] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (771252 kB) free!
[1590403119] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.69, 6.82, 5.89
[1590403239] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.75, 6.34, 5.82
[1590403357] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3077432 kB) free!
[1590403359] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.35, 5.99, 5.75
[1590405157] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (706216 kB) free!
[1590405361] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(39) < 2020/05/25 16:37:18 [error] 22121#0: *90745370 connect() failed (111: Connection refused) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590405481] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590405757] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (2331184 kB) free!
[1590405759] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.87, 3.06, 4.00
[1590405841] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22488628 kB) free!
[1590406441] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (426492 kB) free!
[1590406935] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.6% (4302248 kB) free!
[1590407041] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16529576 kB) free!
[1590407055] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.2% (4218084 kB) free!
[1590407175] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.9% (4140800 kB) free!
[1590407641] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (429248 kB) free!
[1590408157] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1347704 kB) free!
[1590408241] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19107804 kB) free!
[1590408757] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2592260 kB) free!
[1590408841] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417492 kB) free!
[1590409101] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.6% (3941268 kB) free!
[1590409221] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.9% (3797908 kB) free!
[1590409341] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (3666324 kB) free!
[1590409441] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18828032 kB) free!
[1590409575] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2090560 kB) free!
[1590409941] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.9% (16148360 kB) free.
[1590410041] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (464936 kB) free!
[1590410641] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13597288 kB) free!
[1590410881] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/25 18:09:33 [error] 22135#0: *90772422 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590411001] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590411242] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409288 kB) free!
[1590411757] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (305708 kB) free!
[1590412357] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (2717160 kB) free!
[1590412957] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (746860 kB) free!
[1590413557] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (2823112 kB) free!
[1590413559] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.03, 17.04, 10.45
[1590413679] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 15.06, 15.87, 10.80
[1590413799] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.67, 14.25, 10.83
[1590414157] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3802992 kB) free.
[1590415357] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.0% (1946456 kB) free!
[1590415477] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2595820 kB) free!
[1590415597] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3200252 kB) free!
[1590416198] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1630584 kB) free!
[1590416798] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (2896496 kB) free!
[1590416799] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.72, 2.38, 3.88
[1590417242] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33027824 kB) free.
[1590417842] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4350376 kB) free!
[1590417962] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (412908 kB) free!
[1590418082] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (437872 kB) free!
[1590418598] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (293176 kB) free!
[1590418682] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16925440 kB) free!
[1590419198] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (2688456 kB) free!
[1590419283] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (455276 kB) free!
[1590419798] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1400440 kB) free!
[1590420398] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2612724 kB) free!
[1590420602] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(19) < 2020/05/25 20:57:31 [error] 22132#0: *90807926 connect() failed (111: Connection refused) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590420722] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590421683] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32696364 kB) free.
[1590422198] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.4% (4294036 kB) free.
[1590422283] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7211592 kB) free!
[1590422403] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2824140 kB) free!
[1590422523] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417228 kB) free!
[1590422798] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.7% (2396432 kB) free!
[1590422918] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2598260 kB) free!
[1590422999] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1590423038] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (2802796 kB) free!
[1590423119] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1590423123] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14675280 kB) free!
[1590423638] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.5% (4140804 kB) free.
[1590423719] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1590423723] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400964 kB) free!
[1590423839] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1590424838] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.2% (1976268 kB) free!
[1590424923] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29764256 kB) free.
[1590424958] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2096936 kB) free!
[1590425078] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.8% (2248496 kB) free!
[1590425523] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3180848 kB) free!
[1590425643] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (383760 kB) free!
[1590425763] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404692 kB) free!
[1590426878] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3287764 kB) free.
[1590426963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31840044 kB) free.
[1590427563] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7617180 kB) free!
[1590427683] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3428740 kB) free!
[1590427803] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394592 kB) free!
[1590428403] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15474420 kB) free!
[1590428678] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2843412 kB) free!
[1590428798] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (2776728 kB) free!
[1590428918] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (2801216 kB) free!
[1590429003] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (440796 kB) free!
[1590430118] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3644716 kB) free.
[1590430203] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26588652 kB) free.
[1590430803] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4577232 kB) free!
[1590430923] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (810848 kB) free!
[1590431043] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33023836 kB) free.
[1590431318] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2131764 kB) free!