[1545110278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1545110307] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 375 processes
[1545110398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1545110518] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1545110585] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.03, 16.23, 12.52
[1545110705] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.54, 10.99, 11.04
[1545110825] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.62, 7.54, 9.77
[1545110907] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 2 warn out of 326 processes
[1545111186] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1545111306] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1545112025] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.98, 0.72, 3.06
[1545112318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1545112439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1545112559] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1545113759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1545113879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1545113999] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1545114599] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1545119962] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8281592 kB) free!
[1545119999] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1545120082] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9973900 kB) free.
[1545120120] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1545122519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1545122639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1545122759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1545123610] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1545123731] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 343 processes
[1545123851] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1545123960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1545124080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1545124162] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.3% (1669268 kB) free!
[1545124200] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1545124268] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1629260 kB) free!
[1545124800] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1545124800] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1545125362] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1623316 kB) free!
[1545132852] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1545135600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1545135720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1545135840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1545135852] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1545135972] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 343 processes
[1545136093] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1545136693] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 342 processes
[1545137500] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4218100 kB) free!
[1545137620] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 21.6% (5284360 kB) free.
[1545137641] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1545139420] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4221504 kB) free!
[1545139540] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (4222896 kB) free!
[1545139660] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (4209780 kB) free!
[1545140364] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10232 MB (19% inode=98%):
[1545140484] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 10232 MB (19% inode=98%):
[1545140604] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 12060 MB (23% inode=98%):
[1545141841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1545141962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1545142082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1545143893] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1545144493] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 345 processes
[1545144613] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1545144734] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1545144734] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1545145934] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1545148082] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1545156734] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1545157506] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1545157627] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes