[1596911777] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596911898] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199756 kB) free!
[1596912018] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163940 kB) free!
[1596912271] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181076 kB) free!
[1596912392] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.6% (19027624 kB) free.
[1596912618] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 83.6% (20483352 kB) free.
[1596914939] SERVICE ALERT: pservercluster2;Open-Files;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596914981] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915017] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915020] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915028] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915031] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915065] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915069] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915111] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915147] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915150] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915158] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915161] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915185] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915185] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1596915190] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596915231] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596915277] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915280] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915288] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915291] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915307] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915315] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915315] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596915427] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 417.72, 362.49, 150.31
[1596915435] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 331 processes
[1596915738] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 68.14, 47.60, 44.31
[1596915790] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 8256 open files (0% of max 2424156)
[1596915877] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 161725 MB (71% inode=99%):
[1596915880] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596915888] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1596915888] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.6% (21479852 kB) free.
[1596915896] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.21% system=3.64% iowait=0.00% idle=95.15%
[1596915915] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1596916200] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916224] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916233] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916271] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916318] SERVICE ALERT: pservercluster1;Current Users;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596916320] SERVICE ALERT: pservercluster1;Total Processes;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596916335] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393380 kB) free!
[1596916338] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 46.18, 43.98, 43.53
[1596916344] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916344] SERVICE ALERT: pservercluster1;zombie_procs;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596916348] SERVICE ALERT: pservercluster1;CPU Load;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596916353] SERVICE ALERT: pservercluster1;Sensors Status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596916366] SERVICE ALERT: pservercluster1;CPU-STATISTICS;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596916391] SERVICE ALERT: pservercluster1;HDD_Home status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596916392] SERVICE ALERT: pservercluster1;HDD_Root status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596916431] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596916440] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916443] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916451] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916464] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916468] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916469] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916478] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916486] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916491] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916512] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916516] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916522] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (958540 kB) free!
[1596916568] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916576] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916589] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916593] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916606] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916608] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916611] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916616] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916637] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916642] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200464 kB) free!
[1596916701] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916728] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916731] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916736] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916762] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198432 kB) free!
[1596916851] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916853] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596916926] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3716976 kB) free!
[1596917046] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3713992 kB) free!
[1596917166] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3711424 kB) free!
[1596917362] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2732728 kB) free!
[1596917539] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1596917539] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 46.66, 43.89, 45.02
[1596917962] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 57.9% (14194420 kB) free.
[1596918562] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172896 kB) free!
[1596918682] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 52.8% (12930404 kB) free.
[1596918840] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 313 processes
[1596918864] SERVICE ALERT: pservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596918878] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1596918916] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 172709 MB (78% inode=99%):
[1596918966] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 63.5% (31306680 kB) free.
[1596918968] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596918989] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 9792 open files (0% of max 2440681)
[1596918991] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185452 kB) free!
[1596918993] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.29, 2.79, 2.89
[1596919022] SERVICE ALERT: pservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.12% system=5.67% iowait=5.54% idle=87.67%
[1596919037] SERVICE ALERT: pservercluster1;HDD_Root status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596919111] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.3% (19685564 kB) free.
[1596919349] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 103.255.146.164 [09/Aug/2020:02:03:40 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1596919445] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919469] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919469] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596919483] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919521] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919568] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919570] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919589] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919593] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919594] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919608] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919627] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919641] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919642] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919690] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919693] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919713] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919714] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919719] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919733] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919752] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919766] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919818] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919838] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919844] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919877] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596919882] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190352 kB) free!
[1596919939] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 43.28, 39.65, 42.08
[1596920002] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168800 kB) free!
[1596920122] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1596920122] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195452 kB) free!
[1596920539] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 61.37, 46.02, 44.20
[1596920642] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 75.2% (37069748 kB) free.
[1596920723] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.8% (22994356 kB) free.
[1596922339] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 34.30, 36.34, 39.70
[1596924139] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 61.65, 43.79, 42.00
[1596924739] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 47.22, 43.96, 42.79
[1596927289] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596928690] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 312 processes
[1596928701] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596928714] SERVICE ALERT: pservercluster1;zombie_procs;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596928733] SERVICE ALERT: pservercluster1;Sensors Status;UNKNOWN;HARD;3;NRPE: Unable to read output
[1596928736] SERVICE ALERT: pservercluster1;CPU_Procs;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596928766] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 172555 MB (78% inode=99%):
[1596928844] SERVICE ALERT: pservercluster1;Open-Files;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596928851] SERVICE ALERT: pservercluster1;Memory;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596928853] SERVICE ALERT: pservercluster1;Messages Log Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596928877] SERVICE ALERT: pservercluster1;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596929242] SERVICE ALERT: pservercluster1;HDD_Root status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596929290] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596929301] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596929314] SERVICE FLAPPING ALERT: pservercluster1;zombie_procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1596929314] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596929333] SERVICE FLAPPING ALERT: pservercluster1;Sensors Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1596929333] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596929336] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596929366] SERVICE ALERT: pservercluster1;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596929410] SERVICE ALERT: pservercluster1;Total Processes;UNKNOWN;SOFT;2;Could not open pipe: /usr/bin/ps
[1596929418] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596929438] SERVICE ALERT: pservercluster1;CPU Load;UNKNOWN;HARD;3;Error opening /usr/bin/uptime
[1596929444] SERVICE FLAPPING ALERT: pservercluster1;Open-Files;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1596929444] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596929451] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596929453] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596929477] SERVICE FLAPPING ALERT: pservercluster1;CPU-STATISTICS;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1596929477] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596929486] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 172563 MB (78% inode=99%):
[1596929530] SERVICE FLAPPING ALERT: pservercluster1;Total Processes;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1596929530] SERVICE ALERT: pservercluster1;Total Processes;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596929847] SERVICE FLAPPING ALERT: pservercluster1;HDD_Root status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1596929847] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596929933] SERVICE ALERT: pservercluster1;Sensors Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596929936] SERVICE ALERT: pservercluster1;CPU_Procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596930018] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596930038] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1596930038] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596930077] SERVICE ALERT: pservercluster1;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596930086] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596930130] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596930138] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1596930206] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;2;NRPE: Call to popen() failed
[1596930263] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1596930263] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596930326] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 172541 MB (78% inode=99%):
[1596930501] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596930514] SERVICE ALERT: pservercluster1;zombie_procs;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596930863] SERVICE ALERT: pservercluster1;Current Users;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596931101] SERVICE FLAPPING ALERT: pservercluster1;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1596931101] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596931114] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596931133] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596931136] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1596931136] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596931463] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596931701] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596931714] SERVICE ALERT: pservercluster1;zombie_procs;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596931733] SERVICE ALERT: pservercluster1;Sensors Status;UNKNOWN;HARD;3;NRPE: Unable to read output
[1596931736] SERVICE ALERT: pservercluster1;CPU_Procs;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596931838] SERVICE ALERT: pservercluster1;CPU Load;UNKNOWN;HARD;3;Error opening /usr/bin/uptime
[1596931877] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596931930] SERVICE ALERT: pservercluster1;Total Processes;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1596932123] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596932445] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 11936 open files (0% of max 2440681)
[1596932451] SERVICE ALERT: pservercluster1;Memory;UNKNOWN;HARD;3;NRPE: Unable to read output
[1596932453] SERVICE ALERT: pservercluster1;Messages Log Status;WARNING;HARD;3;NRPE: Unable to read output
[1596932477] SERVICE ALERT: pservercluster1;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596932530] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;NRPE: Call to popen() failed
[1596932663] SERVICE ALERT: pservercluster1;Current Users;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596932726] SERVICE ALERT: pservercluster1;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596932851] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596932906] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596932920] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596932936] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596932938] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596932976] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596933043] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596933050] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596933053] SERVICE FLAPPING ALERT: pservercluster1;Messages Log Status;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1596933053] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596933056] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1596933056] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596933082] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596933175] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596933268] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596933300] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596933454] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37466 MB (73% inode=99%):
[1596933506] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596933520] SERVICE ALERT: pservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596933536] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1596933538] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1596933576] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 172461 MB (78% inode=99%):
[1596933656] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (3782552 kB) free!
[1596933687] SERVICE ALERT: pservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=0.88% system=10.61% iowait=0.00% idle=88.51%
[1596933730] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 318 processes
[1596933868] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596933900] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 10144 open files (0% of max 2440681)
[1596934254] SERVICE ALERT: pservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596934256] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.6% (20393920 kB) free.
[1596934335] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.5% (13849596 kB) free!
[1596934339] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 56.09, 45.70, 41.97
[1596934935] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 37.4% (49354456 kB) free.
[1596934939] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 31.49, 36.42, 39.26
[1596935539] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 68.38, 55.85, 46.46
[1596937843] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.92, 1.05, 314.46
[1596938124] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197824 kB) free!
[1596938244] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164156 kB) free!
[1596938364] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171916 kB) free!
[1596938964] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 79.7% (19525532 kB) free.
[1596939564] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180204 kB) free!
[1596939684] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165228 kB) free!
[1596939804] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178000 kB) free!
[1596940404] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1596940404] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.7% (22967136 kB) free.
[1596941689] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15214.75, 7014.72, 2705.06
[1596941809] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16031.95, 10040.46, 4336.39
[1596941854] SERVICE FLAPPING ALERT: pservercluster1;HDD_Root status;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1596941929] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16141.98, 12064.94, 5769.87
[1596944530] SERVICE FLAPPING ALERT: pservercluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1596944907] SERVICE FLAPPING ALERT: pservercluster1;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596944920] SERVICE FLAPPING ALERT: pservercluster1;zombie_procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596944929] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.30, 5.72, 746.61
[1596944936] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596944938] SERVICE FLAPPING ALERT: pservercluster1;Sensors Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596945044] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1596945054] SERVICE FLAPPING ALERT: pservercluster1;Messages Log Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1596945092] SERVICE FLAPPING ALERT: pservercluster1;CPU-STATISTICS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596945268] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596945300] SERVICE FLAPPING ALERT: pservercluster1;Open-Files;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596945427] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14174.96, 5534.93, 2039.41
[1596945529] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.84, 1.65, 391.72
[1596945547] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15910.85, 9057.70, 3754.17
[1596945656] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596945667] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16147.57, 11493.82, 5319.59
[1596946256] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4462668 kB) free!
[1596946376] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (1690416 kB) free!
[1596946496] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181820 kB) free!
[1596946939] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596947097] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 32.5% (8019572 kB) free.
[1596948297] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180224 kB) free!
[1596948417] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 65.7% (16204108 kB) free.
[1596948667] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.88, 4.22, 630.16
[1596949017] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.3% (2789972 kB) free!
[1596949137] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167488 kB) free!
[1596949258] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.4% (18608192 kB) free.
[1596949857] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (4236476 kB) free!
[1596949977] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (1512892 kB) free!
[1596950097] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1299216 kB) free!
[1596950697] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1596950697] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.1% (20518684 kB) free.
[1596951804] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596953929] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15996.56, 13241.25, 6911.62
[1596954049] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15998.61, 14153.70, 8013.45
[1596954169] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596956443] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15634.28, 8026.91, 3203.41
[1596956563] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16089.16, 10718.00, 4774.41
[1596956683] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16150.14, 12518.54, 6154.88
[1596960283] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.25, 2.30, 331.75
[1596962097] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596963504] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1596963625] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1596963745] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596963745] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1596963865] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 326 processes
[1596966269] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 103.255.146.164 [09/Aug/2020:15:13:17 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1596966390] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596970867] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16151.00, 13874.35, 7551.56
[1596970987] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16152.82, 14628.24, 8594.45
[1596971107] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 13139.27, 14475.27, 9296.30
[1596973345] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1596973465] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596973508] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.35, 5.93, 701.96
[1596973585] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1596974108] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.66, 1.94, 368.38
[1596978898] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.7% (6442108 kB) free!
[1596979018] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.6% (6408404 kB) free!
[1596979138] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.5% (6379900 kB) free!
[1596979585] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1596979706] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1596987539] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 44.5% (14571592 kB) free.
[1596988442] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9796120 kB) free!
[1596988562] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9759140 kB) free!
[1596988683] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.8% (9737356 kB) free!
[1596990308] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16100.60, 11827.42, 5588.05
[1596990428] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16119.70, 13248.63, 6865.37
[1596990548] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16122.15, 14199.62, 7987.83
[1596990620] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1377 processes
[1596991939] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 38.82, 45.35, 45.00
[1596993548] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.01, 2.78, 525.13
[1596993739] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 53.56, 46.35, 45.62
[1596996139] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 47.25, 41.37, 43.45
[1596997339] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1596997339] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 51.27, 48.52, 45.77