[1547404246] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547404514] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547404634] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547404693] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 333 processes
[1547404754] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547405212] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.26, 4.22, 4.71
[1547406211] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547407643] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 14.54, 12.43, 10.87
[1547407763] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.01, 11.68, 10.78
[1547407883] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.73, 11.43, 10.80
[1547413846] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1547428124] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 419 processes
[1547428184] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 418 processes
[1547436683] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.40, 8.52, 9.43
[1547437246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547437366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547437393] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9691592 kB) free!
[1547437486] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547437513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11725944 kB) free.
[1547438086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547438206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547438212] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.46, 8.63, 6.20
[1547438235] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5528 MB (10% inode=97%):
[1547438326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547438332] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.20, 9.24, 6.72
[1547438354] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547438452] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.06, 9.48, 7.11
[1547438474] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547438588] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.59, 24.89, 14.86
[1547438594] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547438708] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.22, 17.13, 13.23
[1547438718] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 68.8% (16864160 kB) free.
[1547438829] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.79, 12.06, 11.85
[1547438835] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 8091 MB (15% inode=98%):
[1547439083] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.34, 11.04, 10.36
[1547439195] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547439204] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.60, 11.46, 10.58
[1547439313] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8799608 kB) free!
[1547439314] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547439324] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.10, 11.94, 10.86
[1547439433] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11286280 kB) free.
[1547439760] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46791 MB (20% inode=99%):
[1547439880] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46791 MB (20% inode=99%):
[1547439914] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547440000] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46791 MB (20% inode=99%):
[1547440029] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.66, 2.00, 4.60
[1547440033] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9001320 kB) free!
[1547440034] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547440153] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7980280 kB) free!
[1547440273] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10674452 kB) free.
[1547440634] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547440654] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1463000 kB) free!
[1547440693] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 5 warn out of 335 processes
[1547440754] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547440873] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7584820 kB) free!
[1547440874] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547440993] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10276416 kB) free.
[1547441926] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547442074] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547442194] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7331792 kB) free!
[1547442194] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547442314] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9885524 kB) free.
[1547442314] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547443126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547443246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547443367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547443484] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 442 processes
[1547443518] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3998268 kB) free!
[1547443544] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 443 processes
[1547443604] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 2 crit, 0 warn out of 446 processes
[1547443638] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.7% (3600888 kB) free!
[1547443664] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 2 crit, 0 warn out of 446 processes
[1547443758] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (3132964 kB) free!
[1547444114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6919128 kB) free!
[1547444114] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547444234] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9382156 kB) free!
[1547444354] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7824860 kB) free!
[1547444358] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181896 kB) free!
[1547444864] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 2 warn out of 445 processes
[1547444958] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1547444958] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 39.7% (9729628 kB) free.
[1547445164] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 1 warn out of 450 processes
[1547445167] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1547445167] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1547445314] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547445434] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547445554] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547445648] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.1% (4966308 kB) free!
[1547445663] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4906128 kB) free!
[1547445723] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4910732 kB) free!
[1547445767] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547445783] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4915864 kB) free.
[1547445887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547446000] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 71806 MB (31% inode=99%):
[1547446007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547446083] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4912020 kB) free!
[1547446143] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4912148 kB) free!
[1547446203] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4915212 kB) free.
[1547446248] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234456 kB) free!
[1547446364] SERVICE FLAPPING ALERT: webserver2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1547446364] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 2 warn out of 449 processes
[1547446503] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4901496 kB) free!
[1547446563] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4820260 kB) free!
[1547446623] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.7% (4827776 kB) free!
[1547446664] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 1 warn out of 457 processes
[1547446676] SERVICE ALERT: sig2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2945048 kB) free!
[1547446683] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.7% (4823100 kB) free!
[1547446754] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1547446754] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547446758] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (1507064 kB) free!
[1547446796] SERVICE ALERT: sig2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2873120 kB) free!
[1547446879] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (1039612 kB) free!
[1547446916] SERVICE ALERT: sig2;Memory;WARNING;HARD;3;WARNING - 17.2% (2801748 kB) free!
[1547446999] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (559064 kB) free!
[1547447354] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547447448] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.9% (5846448 kB) free!
[1547447475] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547447808] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547447808] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 10.2% (1649856 kB) free!
[1547447939] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.9% (2750972 kB) free!
[1547448048] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1547448048] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3979708 kB) free!
[1547448074] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547448194] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547448408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547448408] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (284900 kB) free!
[1547448464] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 2 warn out of 454 processes
[1547448528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547448539] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1607892 kB) free!
[1547448590] SERVICE ALERT: db1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10177 MB (19% inode=99%):
[1547448648] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547448710] SERVICE ALERT: db1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 9977 MB (19% inode=99%):
[1547448794] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547448830] SERVICE ALERT: db1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 9977 MB (19% inode=99%):
[1547448914] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547449034] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547449093] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 10 warn out of 342 processes
[1547449139] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 24.8% (4036596 kB) free.
[1547449248] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.7% (6746276 kB) free!
[1547449316] SERVICE ALERT: sig2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1553036 kB) free!
[1547449693] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 6 warn out of 338 processes
[1547449739] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3126712 kB) free!
[1547449848] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3158724 kB) free!
[1547449859] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2987516 kB) free!
[1547449980] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1547449980] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.4% (2828736 kB) free!
[1547450234] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547450448] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.0% (7385260 kB) free!
[1547450448] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1547450834] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547450954] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547450987] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 7 warn out of 311 processes
[1547451048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547451074] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547451168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547451289] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547451430] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.75, 11.44, 6.32
[1547451550] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.01, 8.21, 5.76
[1547451648] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4170668 kB) free!
[1547451670] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.32, 6.20, 5.32
[1547451799] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 25.5% (6250580 kB) free.
[1547452399] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4712604 kB) free!
[1547452489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547452519] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4474852 kB) free!
[1547452639] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4287932 kB) free!
[1547452664] SERVICE FLAPPING ALERT: webserver2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1547452693] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 10 warn out of 341 processes
[1547452754] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4720584 kB) free!
[1547452848] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.7% (8731940 kB) free!
[1547452870] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.80, 2.61, 3.83
[1547452874] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547453089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547453209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547453329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547453354] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.0% (5413264 kB) free!
[1547453474] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547453594] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547453714] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547453839] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187708 kB) free!
[1547453864] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 455 processes
[1547454439] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.5% (5033272 kB) free.
[1547454648] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (1694572 kB) free!
[1547454670] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.72, 7.08, 5.68
[1547454790] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.00, 5.32, 5.20
[1547454910] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.13, 5.07, 5.13
[1547455039] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (4127932 kB) free!
[1547455129] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547455159] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3934248 kB) free!
[1547455248] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.6% (8170372 kB) free!
[1547455279] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (3715528 kB) free!
[1547455510] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1547455510] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.06, 3.12, 4.16
[1547455514] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547455730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547455754] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1547455754] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3993728 kB) free!
[1547455850] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1547456115] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547456235] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547456354] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6413188 kB) free!
[1547456355] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547456449] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3847900 kB) free!
[1547456450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547456479] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1985156 kB) free!
[1547456570] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547457113] SERVICE ALERT: sig2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547457116] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1547457180] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1547457233] SERVICE ALERT: sig2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1547457649] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.5% (8121792 kB) free!
[1547457679] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 44.4% (10879560 kB) free.
[1547457716] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 93.9% (15253880 kB) free.
[1547457770] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547457890] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547457910] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.01, 11.50, 7.98
[1547458010] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547458030] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.86, 8.81, 7.39
[1547458150] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.43, 6.52, 6.72
[1547458611] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547458755] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547458849] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4741392 kB) free!
[1547459355] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547459475] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547459596] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547459950] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.89, 3.25, 4.39
[1547459954] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4940224 kB) free!
[1547460049] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.3% (8053840 kB) free!
[1547460411] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547460531] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547460651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547460679] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190084 kB) free!
[1547460799] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (1409244 kB) free!
[1547460919] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186428 kB) free!
[1547461154] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5049476 kB) free!
[1547461380] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1616144 kB) free!
[1547461396] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547461754] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3990076 kB) free!
[1547461849] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2440692 kB) free!
[1547461851] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547461996] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547462116] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547462236] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547462350] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 20.84, 7.11, 4.92
[1547462354] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6202548 kB) free!
[1547462451] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1547462470] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.00, 5.56, 4.62
[1547462572] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547462692] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1547462836] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547462955] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2979288 kB) free!
[1547462956] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547463049] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.6% (8204120 kB) free!
[1547463077] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547463292] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547463357] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 314 processes
[1547463412] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1547463477] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1547463532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1547463626] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547463629] HOST ALERT: sig5;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.57)
[1547463678] SERVICE ALERT: sig5;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig5' was out of bounds)
[1547463684] SERVICE ALERT: sig5;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig5' was out of bounds)
[1547463747] SERVICE ALERT: sig5;CPU_Procs;WARNING;HARD;1;CPU WARNING: 1 warn out of 299 processes
[1547463756] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1547463798] SERVICE ALERT: sig5;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1547463804] SERVICE ALERT: sig5;Total Processes;OK;HARD;1;PROCS OK: 301 processes
[1547464249] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4220648 kB) free!
[1547464346] SERVICE ALERT: sig5;CPU_Procs;OK;HARD;1;CPU OK: 286 processes
[1547464732] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547464755] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.7% (5781332 kB) free!
[1547465332] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1547465355] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3235320 kB) free!
[1547465452] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1547465477] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547465882] HOST ALERT: sig5;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1547465903] SERVICE ALERT: sig5;HDD_Home status;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547465904] SERVICE ALERT: sig5;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig5' was out of bounds)
[1547465904] SERVICE ALERT: sig5;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig5' was out of bounds)
[1547465907] SERVICE ALERT: sig5;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig5' was out of bounds)
[1547465907] SERVICE ALERT: sig5;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig5' was out of bounds)
[1547465955] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4993224 kB) free!
[1547466009] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1547466023] SERVICE ALERT: sig5;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 221436 MB (98% inode=99%):
[1547466024] SERVICE ALERT: sig5;Sensors Status;OK;HARD;1;SENSORS OK
[1547466025] SERVICE ALERT: sig5;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547466027] SERVICE ALERT: sig5;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38811 MB (75% inode=98%):
[1547466028] SERVICE ALERT: sig5;Memory;OK;HARD;1;OK - 94.4% (15336044 kB) free.
[1547466049] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.1% (8409720 kB) free!
[1547466053] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547466077] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547466173] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547466197] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547466317] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547466917] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547467038] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547467158] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547467373] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1547467380] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1547467493] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1547467758] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547467849] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3418488 kB) free!
[1547467878] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547467998] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547468987] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 6 warn out of 314 processes
[1547469049] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.7% (8220988 kB) free!
[1547469587] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 313 processes
[1547469894] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547469919] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1547470014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547470134] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547470270] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1547470849] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (1094776 kB) free!
[1547470870] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 22.12, 57.04, 29.86
[1547470990] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.76, 38.84, 26.49
[1547470999] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547471110] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.23, 26.62, 23.51
[1547471599] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547471711] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.39, 5.31, 13.25
[1547471719] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547471839] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547471988] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 7 warn out of 312 processes
[1547472588] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1547472588] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 312 processes
[1547472649] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.5% (8111704 kB) free!
[1547472911] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.57, 1.82, 5.00
[1547473511] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 27.03, 19.17, 10.52
[1547473519] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (2577888 kB) free!
[1547473631] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.88, 13.57, 9.51
[1547473635] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 3970 MB (7% inode=96%):
[1547473734] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547473751] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1547473751] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.25, 9.66, 8.57
[1547474119] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1693336 kB) free!
[1547474449] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3385868 kB) free!
[1547474839] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547474934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547474951] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.70, 1.94, 3.74
[1547475054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547475174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547475919] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (2947032 kB) free!
[1547476039] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547476159] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547476250] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.8% (8290300 kB) free!
[1547476279] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547476519] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1547476519] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2284308 kB) free!
[1547476879] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547476999] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547477119] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547477356] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547478050] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3945036 kB) free!
[1547478175] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547478775] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547478895] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547479015] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547479250] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.9% (8826360 kB) free!
[1547479519] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547480119] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547480239] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547480359] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547480450] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2772620 kB) free!
[1547481050] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.6% (9153304 kB) free!
[1547481650] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4475092 kB) free!
[1547482250] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.9% (9333124 kB) free!
[1547482850] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4042988 kB) free!
[1547483450] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.7% (8223708 kB) free!
[1547483988] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547484542] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547484544] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547484544] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1547484548] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2' was out of bounds)
[1547484565] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig2' was out of bounds)
[1547484650] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4073708 kB) free!
[1547484662] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38563 MB (75% inode=98%):
[1547484664] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 332 processes
[1547484666] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1547484668] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547484686] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547485064] SERVICE ALERT: sig2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547485067] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1547485071] SERVICE ALERT: sig2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig2' was out of bounds)
[1547485131] HOST FLAPPING ALERT: sig2;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1547485131] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1547485159] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547485184] SERVICE ALERT: sig2;Sensors Status;OK;HARD;1;SENSORS OK
[1547485191] SERVICE ALERT: sig2;CPU Load;OK;HARD;1;OK - load average: 0.15, 0.12, 0.05
[1547485250] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.3% (6568524 kB) free!
[1547485615] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547485751] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1547485885] HOST FLAPPING ALERT: sig2;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1547486359] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547486450] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10162228 kB) free.
[1547486479] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547486815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547486935] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1547487650] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.0% (6383716 kB) free!
[1547487679] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547487770] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10223748 kB) free.
[1547487799] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547487919] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547490319] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.4% (5007380 kB) free.