[1585074693] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20066420 kB) free!
[1585074846] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2992120 kB) free!
[1585074966] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3915280 kB) free.
[1585075293] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385432 kB) free!
[1585075566] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.5% (2522184 kB) free!
[1585075686] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3231500 kB) free!
[1585075806] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3984048 kB) free.
[1585075893] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14559484 kB) free!
[1585076406] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.6% (2373680 kB) free!
[1585076478] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1585076494] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409940 kB) free!
[1585076526] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3007456 kB) free!
[1585076646] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3749252 kB) free.
[1585077097] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44706 MB (19.89% inode=100%):
[1585077157] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44509 MB (19.80% inode=100%):
[1585077217] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;3;DISK WARNING - free space: /home 40683 MB (18.10% inode=100%):
[1585077246] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.7% (1893936 kB) free!
[1585077277] SERVICE ALERT: webserver2;HDD_Home status;WARNING;HARD;4;DISK WARNING - free space: /home 39285 MB (17.48% inode=100%):
[1585077366] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.3% (2328248 kB) free!
[1585077470] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 73.7% (18072096 kB) free.
[1585077486] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (2910892 kB) free!
[1585077577] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 83022 MB (36.94% inode=100%):
[1585077679] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4412936 kB) free!
[1585078086] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.7% (4497752 kB) free.
[1585078279] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164588 kB) free!
[1585078686] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.3% (1838300 kB) free!
[1585078806] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2094792 kB) free!
[1585078926] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.5% (2350328 kB) free!
[1585080103] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224988 kB) free!
[1585080126] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3305232 kB) free.
[1585080135] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11160384 kB) free!
[1585080223] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233532 kB) free!
[1585080255] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2647184 kB) free!
[1585080343] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213276 kB) free!
[1585080373] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < operation none caused a connection exception frame_error: "Malformed UTF-8 in shortstr"
[1585080375] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3577928 kB) free!
[1585080494] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1585081006] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.91, 5.91, 5.15
[1585081126] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.23, 5.74, 5.19
[1585081246] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.78, 5.80, 5.28
[1585081294] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28638928 kB) free.
[1585081385] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1585081505] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1585081894] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3276988 kB) free!
[1585082014] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (390588 kB) free!
[1585082134] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415320 kB) free!
[1585082186] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 724.19, 1942.37, 1038.68
[1585082199] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(52) < 203.212.70.170 [25/Mar/2020:02:02:53 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1585082306] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 99.32, 1300.34, 912.98
[1585082319] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585082426] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.84, 870.97, 802.60
[1585082446] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.47, 3.86, 4.61
[1585082743] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.8% (8291276 kB) free!
[1585083026] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.00, 118.70, 421.63
[1585083601] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1585083722] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1585083839] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1585083842] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1585083943] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 26.2% (12900596 kB) free.
[1585083959] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585083976] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.2% (16046572 kB) free!
[1585084063] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 40.0% (19689188 kB) free.
[1585084079] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1585084246] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.49, 6.22, 5.34
[1585084278] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 993 processes
[1585084366] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.98, 6.12, 5.40
[1585084486] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.18, 5.79, 5.37
[1585084535] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24958124 kB) free!
[1585084735] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.13132.4453> (10.147.212.192:34980 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1585084855] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1585085135] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3794424 kB) free!
[1585085187] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.2% (45934068 kB) free.
[1585085253] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 16.2% (3960280 kB) free!
[1585085280] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1585085293] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.31654.515> (10.147.212.158:56848 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1585085328] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 85.9% (21052868 kB) free.
[1585085414] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1585085443] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1585085499] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 23.3% (5699656 kB) free.
[1585085563] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1585085642] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1585085683] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1585085686] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1585085686] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.21, 4.49, 4.75
[1585085735] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17760892 kB) free!
[1585085872] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 714 processes
[1585085919] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(384) < 172.105.59.173 [25/Mar/2020:03:08:38 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1585085992] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 713 processes
[1585086039] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(260) < 203.212.70.170 [25/Mar/2020:03:09:57 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1585086078] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 2229 processes
[1585086099] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (2444700 kB) free!
[1585086112] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 714 processes
[1585086160] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(13) < 203.212.70.170 [25/Mar/2020:03:11:57 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1585086199] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 2228 processes
[1585086220] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (2451088 kB) free!
[1585086242] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 343 processes
[1585086286] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.21, 5.14, 5.03
[1585086308] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 37.1% (9094476 kB) free.
[1585086319] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 2222 processes
[1585086335] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382972 kB) free!
[1585086340] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2444808 kB) free!
[1585086362] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 341 processes
[1585086406] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.78, 4.70, 4.88
[1585086482] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 341 processes
[1585086760] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585086883] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1585087082] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1585087082] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 342 processes
[1585087535] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29593752 kB) free.
[1585088135] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6489172 kB) free!
[1585088255] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3026528 kB) free!
[1585088375] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417808 kB) free!
[1585089575] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27220644 kB) free.
[1585090082] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1585090175] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6541300 kB) free!
[1585090295] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3199804 kB) free!
[1585090415] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400552 kB) free!
[1585090926] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585091616] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23253896 kB) free!
[1585092216] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (4998648 kB) free!
[1585092816] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14801000 kB) free!
[1585092909] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169940 kB) free!
[1585093029] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187144 kB) free!
[1585093149] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188656 kB) free!
[1585094016] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385064 kB) free!
[1585094384] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 62.7% (15377940 kB) free.
[1585094479] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.7% (11278184 kB) free.
[1585094616] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17702628 kB) free!
[1585094740] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 48.6% (11906844 kB) free.
[1585095216] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34516968 kB) free.
[1585095816] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8170236 kB) free!
[1585095936] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4231752 kB) free!
[1585096056] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424876 kB) free!
[1585096576] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408020 kB) free!
[1585096606] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585097256] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25147988 kB) free!
[1585097395] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 372 processes
[1585097515] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1585097776] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 67.2% (88585828 kB) free.
[1585097856] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2652136 kB) free!
[1585099056] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23912844 kB) free!
[1585099082] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1585099657] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1807792 kB) free!
[1585100257] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20001856 kB) free!
[1585100857] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29926784 kB) free.
[1585101457] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10279528 kB) free!
[1585101577] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34627716 kB) free.
[1585102177] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12763100 kB) free!
[1585102297] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8518288 kB) free!
[1585102370] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 12.7% (3118812 kB) free!
[1585102417] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22434112 kB) free!
[1585102430] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 12.7% (3118744 kB) free!
[1585102490] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 12.7% (3118512 kB) free!
[1585102550] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 12.7% (3121556 kB) free!
[1585104050] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 6.8% (1662928 kB) free!
[1585104217] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33422384 kB) free.
[1585104817] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11480544 kB) free!
[1585104937] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7313820 kB) free!
[1585105057] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22704380 kB) free!
[1585105657] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12936000 kB) free!
[1585105965] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(115) < 92.42.104.166 [25/Mar/2020:08:35:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585106086] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585106206] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.15, 6.80, 6.00
[1585106258] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22228732 kB) free!
[1585106326] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.42, 6.38, 5.94
[1585106446] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.99, 5.83, 5.79
[1585106479] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188040 kB) free!
[1585106599] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171140 kB) free!
[1585106719] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196008 kB) free!
[1585106858] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32319500 kB) free.
[1585107458] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8008012 kB) free!
[1585107578] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3436888 kB) free!
[1585107698] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20836080 kB) free!
[1585107715] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 365 processes
[1585107835] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1585108298] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8770448 kB) free!
[1585108898] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21671660 kB) free!
[1585109140] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4672072 kB) free!
[1585109260] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4458772 kB) free!
[1585109380] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (4213496 kB) free!
[1585109498] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391896 kB) free!
[1585110580] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (991220 kB) free!
[1585110698] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27967940 kB) free.
[1585111246] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 24.64, 23.71, 20.13
[1585111298] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4057364 kB) free!
[1585111418] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393736 kB) free!
[1585111538] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399348 kB) free!
[1585111729] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.7% (3608784 kB) free!
[1585111849] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 13.3% (3257968 kB) free!
[1585111970] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 11.8% (2883916 kB) free!
[1585112384] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4871624 kB) free!
[1585112504] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4736400 kB) free!
[1585112570] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (1073848 kB) free!
[1585112624] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4521952 kB) free!
[1585112738] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27979668 kB) free.
[1585113338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (386972 kB) free!
[1585113459] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394056 kB) free!
[1585113579] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406984 kB) free!
[1585113885] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(87) < 103.255.146.164 [25/Mar/2020:10:44:55 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585114005] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585114153] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.6% (3269204 kB) free!
[1585114179] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17036644 kB) free!
[1585114246] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.94, 12.77, 19.22
[1585114273] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3078460 kB) free!
[1585114393] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (2852232 kB) free!
[1585114424] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1775180 kB) free!
[1585114779] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399512 kB) free!
[1585114994] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (921084 kB) free!
[1585115979] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18028124 kB) free!
[1585116406] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(37) < 103.255.146.164 [25/Mar/2020:11:29:48 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585116526] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585116579] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396848 kB) free!
[1585116985] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3860348 kB) free!
[1585117105] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3592592 kB) free!
[1585117126] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(40) < 103.255.146.164 [25/Mar/2020:11:41:31 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585117179] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20279644 kB) free!
[1585117225] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3340240 kB) free!
[1585117246] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585117779] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408024 kB) free!
[1585117825] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2365512 kB) free!
[1585117994] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.7% (15645432 kB) free.
[1585118446] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(43) < 103.255.146.164 [25/Mar/2020:12:01:00 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585118567] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585119025] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.8% (22917280 kB) free.
[1585119766] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(31) < 103.255.146.164 [25/Mar/2020:12:31:01 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585119886] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585120179] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26145292 kB) free!
[1585120486] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(15) < 103.255.146.164 [25/Mar/2020:12:41:38 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585120607] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585120779] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2674620 kB) free!
[1585121187] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8682144 kB) free!
[1585121307] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8396356 kB) free!
[1585121379] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24690988 kB) free!
[1585121427] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (8511040 kB) free!
[1585121806] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(31) < 103.255.146.164 [25/Mar/2020:13:00:59 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585121926] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585121980] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2108240 kB) free!
[1585122580] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19587288 kB) free!
[1585123126] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(41) < 103.255.146.164 [25/Mar/2020:13:27:15 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585123180] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393280 kB) free!
[1585123227] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4857600 kB) free!
[1585123247] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585123780] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16248888 kB) free!
[1585124380] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401008 kB) free!
[1585124446] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(24) < 103.255.146.164 [25/Mar/2020:13:50:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585124566] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585124980] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13263052 kB) free!
[1585125453] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.7% (2383096 kB) free!
[1585125580] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (440060 kB) free!
[1585125766] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(24) < 103.255.146.164 [25/Mar/2020:14:11:24 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585125887] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585127263] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9715052 kB) free!
[1585127383] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9645560 kB) free!
[1585127503] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9573396 kB) free!
[1585127594] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4110084 kB) free!
[1585127686] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(38) < 103.255.146.164 [25/Mar/2020:14:41:46 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585127714] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (3999756 kB) free!
[1585127806] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585127834] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.3% (3874568 kB) free!
[1585127980] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27986836 kB) free.
[1585128580] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3832228 kB) free!
[1585128700] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395688 kB) free!
[1585128821] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (462300 kB) free!
[1585129034] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.5% (15606240 kB) free.
[1585129421] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14333532 kB) free!
[1585129606] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(24) < 103.255.146.164 [25/Mar/2020:15:13:42 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585129726] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585130021] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (455940 kB) free!
[1585130621] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14617616 kB) free!
[1585131221] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403560 kB) free!
[1585131527] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(38) < 103.255.146.164 [25/Mar/2020:15:47:56 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585131647] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585132421] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16796556 kB) free!
[1585132846] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.15, 3.84, 4.52
[1585133021] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390548 kB) free!
[1585133447] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(28) < 103.255.146.164 [25/Mar/2020:16:11:06 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585133567] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585133621] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19314576 kB) free!
[1585134221] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396764 kB) free!
[1585134646] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.17, 7.77, 6.55
[1585134766] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.92, 7.44, 6.57
[1585134821] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19563232 kB) free!
[1585134886] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.57, 6.89, 6.49
[1585135367] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(30) < 103.255.146.164 [25/Mar/2020:16:52:14 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585135421] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396280 kB) free!
[1585135487] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585135903] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4915828 kB) free!
[1585136021] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17876380 kB) free!
[1585136621] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (430336 kB) free!
[1585136688] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(37) < 103.255.146.164 [25/Mar/2020:17:11:53 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585136808] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585137126] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.0% (2444368 kB) free!
[1585137221] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15017076 kB) free!
[1585137246] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2882836 kB) free!
[1585137287] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.81, 4.46, 4.66
[1585137366] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3360316 kB) free.
[1585137821] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385064 kB) free!
[1585138008] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(39) < 103.255.146.164 [25/Mar/2020:17:36:48 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585138128] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(6) < 103.255.146.164 [25/Mar/2020:17:36:49 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585138248] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585138422] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14399352 kB) free!
[1585139022] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419868 kB) free!
[1585140048] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(57) < 103.255.146.164 [25/Mar/2020:18:00:59 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585140168] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585140768] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(24) < 103.255.146.164 [25/Mar/2020:18:15:58 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585140887] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.96, 7.27, 5.93
[1585140889] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585141007] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.89, 7.22, 6.08
[1585141127] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1585141127] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.09, 6.80, 6.06
[1585141422] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33738436 kB) free.
[1585141488] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(34) < 103.255.146.164 [25/Mar/2020:18:31:00 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585141608] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585142022] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4316340 kB) free!
[1585142142] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (418540 kB) free!
[1585142208] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(30) < 103.255.146.164 [25/Mar/2020:18:42:43 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585142262] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (434412 kB) free!
[1585142328] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585142862] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17552996 kB) free!
[1585143366] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2849864 kB) free!
[1585143462] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (441652 kB) free!
[1585143486] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (2957560 kB) free!
[1585143606] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (3082464 kB) free!
[1585144062] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15850648 kB) free!
[1585144206] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3508116 kB) free.
[1585144662] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395572 kB) free!
[1585144727] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.67, 4.27, 4.72
[1585145327] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.27, 6.41, 5.62
[1585145447] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.78, 6.98, 5.93
[1585145567] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.74, 6.49, 5.88
[1585147367] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.93, 4.41, 4.77
[1585148262] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30829388 kB) free.
[1585148862] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6638568 kB) free!
[1585148982] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2279040 kB) free!
[1585149102] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409180 kB) free!
[1585149702] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15595960 kB) free!
[1585150128] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 92.42.104.166 [25/Mar/2020:20:50:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585150248] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585150302] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417372 kB) free!
[1585152702] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27256052 kB) free.
[1585153302] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2810656 kB) free!
[1585153422] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396256 kB) free!
[1585153542] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393872 kB) free!
[1585153556] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4476.16, 1119.84, 375.73
[1585153676] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 9351.93, 4095.24, 1556.30
[1585153736] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 9899.61, 5297.41, 2145.22
[1585153796] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10009.57, 6086.04, 2593.67
[1585153848] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(15) < 203.212.70.170 [25/Mar/2020:21:59:30 +0530] TCP 502 0 0 240.003 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 60.001, 60.001"
[1585153856] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 10121.31, 6905.21, 3116.63
[1585153968] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585153976] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2619.48, 5565.58, 3120.48
[1585154742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28894084 kB) free.
[1585155342] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4923248 kB) free!
[1585155462] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (739036 kB) free!
[1585155582] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401908 kB) free!
[1585155596] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.05, 30.41, 610.87
[1585155776] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.88, 15.19, 450.42
[1585156196] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.29, 5.18, 320.70
[1585156782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28570996 kB) free.
[1585157382] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5829460 kB) free!
[1585157502] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1988000 kB) free!
[1585157567] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1585157622] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (11398372 kB) free!
[1585158822] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24959796 kB) free!
[1585159422] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4021744 kB) free!
[1585160022] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15577960 kB) free!
[1585160622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29214092 kB) free.