[1563820411] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563820449] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6360724 kB) free!
[1563820531] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563820569] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11041620 kB) free.
[1563820817] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168972 kB) free!
[1563820882] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563820937] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196508 kB) free!
[1563820987] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1754080 kB) free!
[1563821057] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190772 kB) free!
[1563821107] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1738836 kB) free!
[1563821131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563821169] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6153572 kB) free!
[1563821228] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1730528 kB) free!
[1563821251] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563821289] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11047504 kB) free.
[1563821512] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'rq-2' was out of bounds)
[1563821512] SERVICE ALERT: rq-2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'rq-2' was out of bounds)
[1563821512] SERVICE ALERT: rq-2;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'rq-2' was out of bounds)
[1563821515] HOST ALERT: rq-2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.204)
[1563821521] SERVICE ALERT: rq-2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'rq-2' was out of bounds)
[1563821541] HOST ALERT: rq-1;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1563821595] SERVICE ALERT: rq-1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'rq-1' was out of bounds)
[1563821595] SERVICE ALERT: rq-1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'rq-1' was out of bounds)
[1563821595] SERVICE ALERT: rq-1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'rq-1' was out of bounds)
[1563821635] SERVICE ALERT: rq-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'rq-2' was out of bounds)
[1563821635] SERVICE ALERT: rq-2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'rq-2' was out of bounds)
[1563821635] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'rq-2' was out of bounds)
[1563821666] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'rq-1' was out of bounds)
[1563821666] SERVICE ALERT: rq-1;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'rq-1' was out of bounds)
[1563821705] HOST ALERT: rq-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1563821715] SERVICE ALERT: rq-1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 401278 MB (99% inode=99%):
[1563821715] SERVICE ALERT: rq-1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 44641 MB (87% inode=99%):
[1563821715] SERVICE ALERT: rq-1;CPU Load;OK;HARD;1;OK - load average: 0.43, 0.17, 0.06
[1563821719] HOST ALERT: rq-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1563821786] SERVICE ALERT: rq-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /rq1vns1 1068147 MB (99% inode=99%):
[1563821786] SERVICE ALERT: rq-1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1563821851] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563821889] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5954972 kB) free!
[1563821971] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563822009] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11038240 kB) free.
[1563822090] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563822232] SERVICE ALERT: rq-2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1563822232] SERVICE FLAPPING ALERT: rq-2;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1563822232] SERVICE ALERT: rq-2;CPU Load;OK;HARD;1;OK - load average: 0.17, 0.08, 0.06
[1563822232] SERVICE FLAPPING ALERT: rq-2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1563822232] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 98.7% (130163888 kB) free.
[1563822241] SERVICE ALERT: rq-2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 44638 MB (87% inode=99%):
[1563822386] SERVICE ALERT: rq-1;HDD_VNS1 status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'rq-1' was out of bounds)
[1563822386] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'rq-1' was out of bounds)
[1563822516] SERVICE ALERT: rq-1;HDD_VNS1 status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1563822516] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1563822571] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563822586] SERVICE ALERT: redis-2;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'redis-2' was out of bounds)
[1563822639] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;HARD;3;(Return code of 255 for service 'CPU_Procs' on host 'rq-1' was out of bounds)
[1563822639] SERVICE ALERT: rq-1;HDD_VNS1 status;CRITICAL;HARD;3;(Return code of 255 for service 'HDD_VNS1 status' on host 'rq-1' was out of bounds)
[1563822642] SERVICE ALERT: rq-1;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'rq-1' was out of bounds)
[1563822642] SERVICE ALERT: rq-1;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'rq-1' was out of bounds)
[1563822645] HOST ALERT: rq-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.203)
[1563822651] SERVICE ALERT: redis-2;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-2' was out of bounds)
[1563822690] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'redis-2' was out of bounds)
[1563822692] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563822709] HOST ALERT: rq-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1563822716] SERVICE ALERT: redis-2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1563822724] SERVICE ALERT: redis-2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1563822727] HOST ALERT: redis-2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.202)
[1563822758] SERVICE ALERT: redis-2;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'redis-2' was out of bounds)
[1563822762] SERVICE ALERT: rq-1;Memory;OK;HARD;1;OK - 98.8% (130185112 kB) free.
[1563822762] SERVICE ALERT: rq-1;Total Processes;OK;HARD;1;PROCS OK: 351 processes
[1563822774] SERVICE ALERT: redis-2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-2' was out of bounds)
[1563822813] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.3% (7462552 kB) free!
[1563822813] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'redis-2' was out of bounds)
[1563822829] SERVICE ALERT: redis-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'redis-2' was out of bounds)
[1563822839] SERVICE ALERT: redis-2;HDD_Home status;CRITICAL;HARD;2;(Return code of 255 for service 'HDD_Home status' on host 'redis-2' was out of bounds)
[1563822847] SERVICE ALERT: redis-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'redis-2' was out of bounds)
[1563822847] SERVICE ALERT: redis-2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-2' was out of bounds)
[1563822933] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (266512 kB) free!
[1563822963] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1563823053] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (256368 kB) free!
[1563823056] SERVICE ALERT: redis-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'redis-2' was out of bounds)
[1563823122] SERVICE ALERT: redis-2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'redis-2' was out of bounds)
[1563823209] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7466416 kB) free!
[1563823239] SERVICE FLAPPING ALERT: rq-1;HDD_VNS1 status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563823239] SERVICE ALERT: rq-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /rq1vns1 1068147 MB (99% inode=99%):
[1563823239] SERVICE FLAPPING ALERT: rq-1;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563823239] SERVICE ALERT: rq-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1563823291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563823295] HOST ALERT: redis-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1563823329] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11049892 kB) free.
[1563823371] SERVICE ALERT: redis-2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45449 MB (88% inode=99%):
[1563823410] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 79.6% (104953820 kB) free.
[1563823411] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563823436] SERVICE ALERT: redis-2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 344326 MB (92% inode=99%):
[1563823444] SERVICE ALERT: redis-2;Sensors Status;OK;HARD;1;SENSORS OK
[1563823473] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1563823476] HOST ALERT: redis-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.201)
[1563823478] SERVICE ALERT: redis-2;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 1069171 MB (99% inode=99%):
[1563823490] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4856904 kB) free!
[1563823491] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'redis-1' was out of bounds)
[1563823514] SERVICE ALERT: redis-1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-1' was out of bounds)
[1563823551] SERVICE ALERT: redis-1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-1' was out of bounds)
[1563823567] SERVICE ALERT: redis-2;Total Processes;OK;HARD;1;PROCS OK: 375 processes
[1563823596] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'redis-1' was out of bounds)
[1563823599] SERVICE ALERT: redis-1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'redis-1' was out of bounds)
[1563823684] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1563823776] SERVICE ALERT: redis-2;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.74, 0.45
[1563823785] SERVICE ALERT: redis-1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'redis-1' was out of bounds)
[1563823842] SERVICE ALERT: redis-2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1563823863] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (187868 kB) free!
[1563823905] SERVICE ALERT: redis-1;CPU Load;OK;HARD;1;OK - load average: 1.43, 0.55, 0.20
[1563823909] HOST ALERT: redis-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1563823929] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7343228 kB) free!
[1563824011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563824049] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11048452 kB) free.
[1563824131] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563824194] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1563824212] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1563824234] SERVICE ALERT: redis-1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45410 MB (88% inode=99%):
[1563824253] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 96.3% (126963688 kB) free.
[1563824271] SERVICE ALERT: redis-1;Sensors Status;OK;HARD;1;SENSORS OK
[1563824320] SERVICE ALERT: redis-1;Total Processes;OK;HARD;1;PROCS OK: 396 processes
[1563824482] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178436 kB) free!
[1563824602] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165276 kB) free!
[1563824649] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7109116 kB) free!
[1563824722] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191660 kB) free!
[1563824769] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11040936 kB) free.
[1563825331] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563825369] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7259452 kB) free!
[1563825451] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563825489] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11048844 kB) free.
[1563825572] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563825639] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1563825692] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.67, 17.47, 10.98
[1563825760] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1563825812] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.41, 12.71, 10.02
[1563825828] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (2756096 kB) free!
[1563825880] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563825915] SERVICE ALERT: rq-1;CPU Load;WARNING;SOFT;1;WARNING - load average: 17.06, 11.22, 6.62
[1563825932] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.68, 9.40, 9.14
[1563825948] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (2750820 kB) free!
[1563826035] SERVICE ALERT: rq-1;CPU Load;WARNING;SOFT;2;WARNING - load average: 18.86, 13.58, 8.03
[1563826068] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (2754996 kB) free!
[1563826089] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7600684 kB) free!
[1563826155] SERVICE ALERT: rq-1;CPU Load;WARNING;HARD;3;WARNING - load average: 21.15, 16.11, 9.63
[1563826171] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563826209] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11044704 kB) free.
[1563826291] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563826490] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 22.8% (5584924 kB) free.
[1563826628] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2611304 kB) free!
[1563826809] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7734392 kB) free!
[1563826891] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563826929] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6062492 kB) free!
[1563827011] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563827049] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11050248 kB) free.
[1563827132] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.79, 2.99, 4.62
[1563827611] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563827649] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6055276 kB) free!
[1563827690] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189648 kB) free!
[1563827719] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.30, 7.10, 8.90
[1563827722] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 44.1% (10800376 kB) free.
[1563827731] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563827769] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11040236 kB) free.
[1563827810] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176804 kB) free!
[1563827874] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.08, 4.07, 4.95
[1563827930] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163980 kB) free!
[1563827955] SERVICE FLAPPING ALERT: rq-1;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1563827955] SERVICE ALERT: rq-1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.34, 22.38, 20.46
[1563828322] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190052 kB) free!
[1563828331] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563828369] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6025380 kB) free!
[1563828428] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2447484 kB) free!
[1563828442] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194600 kB) free!
[1563828452] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563828490] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11046236 kB) free.
[1563828562] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190556 kB) free!
[1563829051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563829089] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6431300 kB) free!
[1563829171] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563829209] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11041792 kB) free.
[1563829228] SERVICE FLAPPING ALERT: rq-2;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563829239] SERVICE FLAPPING ALERT: rq-2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563829242] SERVICE ALERT: rq-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1563829362] SERVICE ALERT: rq-2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1563829482] SERVICE ALERT: rq-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1563829771] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563829809] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6751628 kB) free!
[1563829891] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563829929] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11049056 kB) free.
[1563829962] SERVICE FLAPPING ALERT: rq-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563829962] SERVICE FLAPPING ALERT: rq-1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563830491] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563830529] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7093076 kB) free!
[1563830611] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563830649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11049184 kB) free.
[1563830883] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 378 processes
[1563831003] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1563831211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563831249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7335336 kB) free!
[1563831332] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563831369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11039840 kB) free.
[1563831452] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563831562] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1563831562] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.7% (17576720 kB) free.
[1563831717] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176860 kB) free!
[1563831837] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.1% (9762520 kB) free.
[1563831882] SERVICE ALERT: rq-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1563831969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7493776 kB) free!
[1563832052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563832068] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1226432 kB) free!
[1563832089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11045256 kB) free.
[1563832162] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172848 kB) free!
[1563832172] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563832282] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163436 kB) free!
[1563832402] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173760 kB) free!
[1563832432] SERVICE ALERT: rq-2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.97, 6.74, 5.77
[1563832437] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167456 kB) free!
[1563832552] SERVICE ALERT: rq-2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.75, 6.48, 5.81
[1563832557] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.7% (14258444 kB) free.
[1563832653] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5621008 kB) free!
[1563832673] SERVICE ALERT: rq-2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.68, 5.84, 5.65
[1563832689] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7694500 kB) free!
[1563832772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563832773] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5625672 kB) free!
[1563832810] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (6026656 kB) free!
[1563832892] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563832893] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5624312 kB) free!
[1563832930] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11047212 kB) free.
[1563833273] SERVICE ALERT: rq-2;CPU Load;OK;HARD;1;OK - load average: 2.17, 2.80, 4.16
[1563833492] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563833530] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6392692 kB) free!
[1563833612] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563833632] SERVICE FLAPPING ALERT: rq-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563833650] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11046524 kB) free.
[1563834213] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563834250] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6802068 kB) free!
[1563834333] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563834370] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11035592 kB) free.
[1563834428] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2527196 kB) free!
[1563834555] SERVICE ALERT: rq-1;CPU Load;WARNING;HARD;3;WARNING - load average: 18.78, 18.07, 18.93
[1563834603] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 375 processes
[1563834639] SERVICE FLAPPING ALERT: rq-1;HDD_VNS1 status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563834724] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1563834933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563834970] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7130112 kB) free!
[1563835028] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1563835028] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2450032 kB) free!
[1563835053] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563835090] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11045500 kB) free.
[1563835403] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 45.0% (11018460 kB) free.
[1563835480] SERVICE FLAPPING ALERT: rq-1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1563835653] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563835690] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7424240 kB) free!
[1563835773] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563835810] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11036456 kB) free.
[1563835893] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563836003] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168328 kB) free!
[1563836123] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184564 kB) free!
[1563836243] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180012 kB) free!
[1563836410] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7683484 kB) free!
[1563836493] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563836531] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5985320 kB) free!
[1563836614] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563836651] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11045620 kB) free.
[1563837213] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563837251] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6299120 kB) free!
[1563837333] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563837371] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11040936 kB) free.
[1563837933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563837971] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6559048 kB) free!
[1563838053] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563838091] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11044540 kB) free.
[1563838653] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563838691] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6812868 kB) free!
[1563838730] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563838773] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563838811] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11037820 kB) free.
[1563839373] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563839411] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7052152 kB) free!
[1563839494] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563839531] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11044332 kB) free.
[1563840093] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563840131] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7085660 kB) free!
[1563840213] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563840251] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11038556 kB) free.
[1563840333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563840555] SERVICE FLAPPING ALERT: rq-1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563840852] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9836588 kB) free!
[1563840972] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8489028 kB) free!
[1563841092] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7084392 kB) free!
[1563841692] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11011328 kB) free.
[1563842493] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.4% (19003160 kB) free!
[1563842892] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8604000 kB) free!
[1563843012] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7268480 kB) free!
[1563843093] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (11480800 kB) free!
[1563843132] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10985156 kB) free.
[1563844375] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 500 processes
[1563844435] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 501 processes
[1563844673] SERVICE FLAPPING ALERT: rq-2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563845133] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563845493] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 60.5% (79781048 kB) free.
[1563845532] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6635376 kB) free!
[1563845653] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10967228 kB) free.
[1563846252] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6913560 kB) free!
[1563846372] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10960072 kB) free.
[1563846428] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563846933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563846972] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7130356 kB) free!
[1563847053] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563847092] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10964740 kB) free.
[1563847173] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563847643] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563847692] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7285088 kB) free!
[1563847774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563847812] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10956920 kB) free.
[1563847894] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563848412] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7387996 kB) free!
[1563848494] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563848523] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 376 processes
[1563848532] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10963312 kB) free.
[1563848614] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563848643] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1563849132] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7537056 kB) free!
[1563849214] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563849243] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 376 processes
[1563849252] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 18.5% (9109324 kB) free!
[1563849334] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563849363] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1563849372] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10959092 kB) free.
[1563849934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563849973] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6090204 kB) free!
[1563850054] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563850093] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10964316 kB) free.
[1563850655] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563850693] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6367392 kB) free!
[1563850775] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563850813] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10957376 kB) free.
[1563851375] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563851413] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6140512 kB) free!
[1563851495] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563851533] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10961316 kB) free.
[1563852095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563852133] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6097472 kB) free!
[1563852215] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563852253] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10965780 kB) free.
[1563852815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563852853] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6501212 kB) free!
[1563852935] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563852963] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 375 processes
[1563852973] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10956380 kB) free.
[1563853083] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1563853536] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563853573] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7810360 kB) free!
[1563853656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563853683] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 374 processes
[1563853693] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6870100 kB) free!
[1563853776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563853803] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1563853814] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10948400 kB) free.
[1563853931] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216696 kB) free!
[1563854051] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (222864 kB) free!
[1563854172] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227092 kB) free!
[1563854403] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 379 processes
[1563854413] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7475600 kB) free!
[1563854523] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1563854533] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10962192 kB) free.
[1563855133] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7507920 kB) free!
[1563855253] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10868060 kB) free.
[1563855853] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7644924 kB) free!
[1563855973] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5909488 kB) free!
[1563856093] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10959460 kB) free.
[1563856176] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563856693] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5801296 kB) free!
[1563856813] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10949372 kB) free.
[1563857376] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563857413] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5961388 kB) free!
[1563857496] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563857533] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10960072 kB) free.
[1563858096] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563858134] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5933036 kB) free!
[1563858216] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563858254] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10950812 kB) free.
[1563858816] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563858854] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6214112 kB) free!
[1563858936] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563858974] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10939368 kB) free.
[1563859536] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563859574] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6564688 kB) free!
[1563859656] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563859694] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10947320 kB) free.
[1563859958] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 13.6% (2205976 kB) free!
[1563860078] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 11.4% (1852384 kB) free!
[1563860198] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1452488 kB) free!
[1563860257] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563860294] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6554100 kB) free!
[1563860377] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563860414] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10934508 kB) free.
[1563860523] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 375 processes
[1563860643] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1563860835] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1563860955] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1563860977] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563861009] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.85, 5.61, 5.04
[1563861014] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6389916 kB) free!
[1563861075] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1563861097] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563861129] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.08, 5.77, 5.15
[1563861134] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10948088 kB) free.
[1563861249] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.54, 4.92, 4.92
[1563861317] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45817 MB (20% inode=99%):
[1563861437] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45721 MB (20% inode=99%):
[1563861557] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45097 MB (20% inode=99%):
[1563861644] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (335020 kB) free!
[1563861675] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1563861697] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563861734] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6423736 kB) free!
[1563861765] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (131016 kB) free!
[1563861817] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563861854] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10943648 kB) free.
[1563861885] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (130528 kB) free!
[1563862417] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563862455] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6509612 kB) free!
[1563862458] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.1% (9656932 kB) free.
[1563862537] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563862575] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10935132 kB) free.
[1563863049] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.23, 6.72, 5.90
[1563863138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563863169] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.75, 7.24, 6.19
[1563863175] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6393636 kB) free!
[1563863258] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563863289] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.31, 7.62, 6.45
[1563863295] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10938072 kB) free.
[1563863685] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.3% (6715360 kB) free.
[1563863858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563863895] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6732624 kB) free!
[1563863978] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563864015] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10929948 kB) free.
[1563864578] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563864615] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6291168 kB) free!
[1563864698] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563864735] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10936664 kB) free.
[1563864843] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 375 processes
[1563864901] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 78.44, 426.71, 284.99
[1563864963] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1563865021] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.34, 285.77, 250.52
[1563865141] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.47, 188.42, 219.10
[1563865298] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563865418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563865539] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563866058] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188904 kB) free!
[1563866163] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 375 processes
[1563866179] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181904 kB) free!
[1563866283] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 379 processes
[1563866299] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174148 kB) free!
[1563867339] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563867541] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.94, 1.36, 17.57
[1563867556] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.7% (14250304 kB) free.
[1563867582] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 62.2% (15256260 kB) free.
[1563867613] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.8% (21271088 kB) free.
[1563867647] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.9% (20566896 kB) free.
[1563867673] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.6% (15218460 kB) free.
[1563867711] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.1% (37999288 kB) free.
[1563867735] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7336100 kB) free!
[1563867763] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.2% (22246448 kB) free.
[1563867801] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.2% (20873724 kB) free.
[1563867837] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 34.7% (17087840 kB) free.
[1563868161] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (158236 kB) free!
[1563868277] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (1581092 kB) free!
[1563868281] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (170188 kB) free!
[1563868319] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (2307700 kB) free!
[1563868368] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.3% (3037500 kB) free!
[1563868397] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (182632 kB) free!
[1563868401] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (195648 kB) free!
[1563868439] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242588 kB) free!
[1563868485] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.1% (741584 kB) free!
[1563868489] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194256 kB) free!
[1563868518] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175084 kB) free!
[1563868539] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563868559] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (253672 kB) free!
[1563868605] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (259752 kB) free!
[1563868609] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1563868609] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189652 kB) free!
[1563868659] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563868712] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 369 processes
[1563868725] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (130652 kB) free!
[1563868832] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1563869341] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 1.42, 4.00
[1563869357] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 98185 MB (43% inode=99%):
[1563869608] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3559444 kB) free!
[1563869718] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1563869718] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 58.9% (9575144 kB) free.
[1563869728] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186036 kB) free!
[1563869848] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180920 kB) free!
[1563869860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563869980] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563870047] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182272 kB) free!
[1563870100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563870167] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180492 kB) free!
[1563870287] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164420 kB) free!
[1563870357] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (1633456 kB) free!
[1563870477] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (1223704 kB) free!
[1563870597] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (822380 kB) free!
[1563870837] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8711572 kB) free!
[1563870957] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7018684 kB) free!
[1563871077] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11835484 kB) free.
[1563872118] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.0% (2440960 kB) free!
[1563872238] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 12.6% (2046564 kB) free!
[1563872307] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.12, 5.18, 5.08
[1563872358] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 10.4% (1685904 kB) free!
[1563872427] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.26, 5.53, 5.21
[1563872547] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.05, 5.27, 5.15
[1563872877] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7258036 kB) free!
[1563872959] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159352 kB) free!
[1563872997] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12173368 kB) free.
[1563873100] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1563873100] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563873147] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 4.55, 4.41, 4.75
[1563873597] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7062360 kB) free!
[1563873718] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12169756 kB) free.
[1563873903] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.4% (13717244 kB) free.
[1563873936] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22052488 kB) free.
[1563873966] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.0% (20586396 kB) free.
[1563873998] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.4% (6722676 kB) free.
[1563874033] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 62.5% (30799232 kB) free.
[1563874065] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.9% (15102132 kB) free.
[1563874097] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.0% (14784984 kB) free.
[1563874127] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.1% (20372688 kB) free.
[1563874300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563874420] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563874570] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177288 kB) free!
[1563874690] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165156 kB) free!
[1563874728] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (2136788 kB) free!
[1563874810] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1563874810] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187076 kB) free!
[1563874848] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (2070284 kB) free!
[1563874917] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8534920 kB) free!
[1563874947] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.66, 5.76, 5.29
[1563874968] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1563874968] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2006260 kB) free!
[1563875037] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12171996 kB) free.
[1563875067] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.94, 6.54, 5.63
[1563875139] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197636 kB) free!
[1563875187] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.01, 6.75, 5.83
[1563875259] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174144 kB) free!
[1563875269] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (170820 kB) free!
[1563875301] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173400 kB) free!
[1563875380] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190628 kB) free!
[1563875389] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164140 kB) free!
[1563875421] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (150000 kB) free!
[1563875510] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (184712 kB) free!
[1563875541] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1563875541] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163664 kB) free!
[1563875620] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563875637] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8438584 kB) free!
[1563875740] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563875757] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12163856 kB) free.
[1563876340] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563876357] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8351688 kB) free!
[1563876460] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563876477] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12168944 kB) free.
[1563876710] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 45.6% (7408984 kB) free.
[1563877061] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563877077] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8312348 kB) free!
[1563877181] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563877197] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12169716 kB) free.
[1563877507] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 10.9% (1779136 kB) free!
[1563877587] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1563877587] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 4.02, 3.93, 4.79
[1563877627] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (1404264 kB) free!
[1563877637] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8427756 kB) free!
[1563877747] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1092072 kB) free!
[1563877757] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.6% (6682788 kB) free!
[1563877781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563877797] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8151428 kB) free!
[1563877877] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4828520 kB) free!
[1563877901] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563877917] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12167388 kB) free.
[1563878347] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1563878347] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.2% (1978616 kB) free!
[1563878501] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563878510] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2832672 kB) free!
[1563878518] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7951724 kB) free!
[1563878621] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563878630] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.7% (2550184 kB) free!
[1563878638] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12159280 kB) free.
[1563878750] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.9% (2253872 kB) free!
[1563878947] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (858812 kB) free!
[1563879221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563879238] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7670704 kB) free!
[1563879341] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563879350] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.7% (14736640 kB) free.
[1563879358] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12151328 kB) free.
[1563879835] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 586 processes
[1563879895] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 589 processes
[1563879942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563879958] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7924776 kB) free!
[1563880005] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (1245976 kB) free!
[1563880062] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563880078] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12108952 kB) free.
[1563880125] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (1049164 kB) free!
[1563880245] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.3% (828632 kB) free!
[1563880387] SERVICE ALERT: sig2-vm2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1563880390] HOST ALERT: sig2-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.180)
[1563880480] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig2-vm2' was out of bounds)
[1563880500] SERVICE ALERT: sig2-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2-vm2' was out of bounds)
[1563880510] SERVICE ALERT: sig2-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig2-vm2' was out of bounds)
[1563880549] SERVICE ALERT: sig2-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig2-vm2' was out of bounds)
[1563880598] SERVICE ALERT: sig2-vm2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig2-vm2' was out of bounds)
[1563880662] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563880663] SERVICE ALERT: sig2-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig2-vm2' was out of bounds)
[1563880678] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1594204 kB) free!
[1563880689] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.46, 4.79, 4.85
[1563880728] SERVICE ALERT: sig2-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig2-vm2' was out of bounds)
[1563880782] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563880798] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12120900 kB) free.
[1563880845] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (132240 kB) free!
[1563880952] HOST ALERT: sig2-vm2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.180)
[1563881289] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.41, 6.15, 5.48
[1563881409] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.33, 5.55, 5.35
[1563881529] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.17, 4.79, 5.10
[1563881982] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563882102] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563882129] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.77, 4.26, 4.71
[1563882183] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1269.61, 810.29, 353.68
[1563882303] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1276.61, 964.94, 465.71
[1563882423] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1277.54, 1068.42, 564.16
[1563882598] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8722784 kB) free!
[1563882703] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563882719] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7045672 kB) free!
[1563882769] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (2549652 kB) free!
[1563882823] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563882827] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.9% (14774428 kB) free.
[1563882839] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12117300 kB) free.
[1563882872] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21870724 kB) free.
[1563882916] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.3% (21395868 kB) free.
[1563883012] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1563883012] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.3% (6786760 kB) free.
[1563883062] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.1% (37998600 kB) free.
[1563883104] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 85.3% (13860180 kB) free.
[1563883369] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2229896 kB) free!
[1563883423] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563883427] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166464 kB) free!
[1563883439] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7030376 kB) free!
[1563883543] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1563883543] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563883547] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (160688 kB) free!
[1563883559] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12115416 kB) free.
[1563883667] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (187236 kB) free!
[1563884121] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183248 kB) free!
[1563884143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563884159] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7010092 kB) free!
[1563884241] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167940 kB) free!
[1563884263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563884279] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12120388 kB) free.
[1563884361] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164904 kB) free!
[1563884863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563884879] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6987728 kB) free!
[1563884883] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 382 processes
[1563884983] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563884999] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12124260 kB) free.
[1563885004] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 381 processes
[1563885456] HOST ALERT: sig2-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1563885469] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8659560 kB) free!
[1563885584] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563885589] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (8260896 kB) free!
[1563885599] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6917632 kB) free!
[1563885704] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563885709] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (7232996 kB) free!
[1563885719] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12121756 kB) free.
[1563886023] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 0.65, 13.80
[1563886304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563886309] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1563886309] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1301096 kB) free!
[1563886319] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7271892 kB) free!
[1563886424] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563886439] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12120948 kB) free.
[1563886478] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4416896 kB) free!
[1563886598] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (4387872 kB) free!
[1563886718] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4359468 kB) free!
[1563886987] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1563886989] HOST ALERT: sig2-vm1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.45)
[1563886989] SERVICE ALERT: sig2-vm1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig2-vm1' was out of bounds)
[1563887002] SERVICE ALERT: sig2-vm1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2-vm1' was out of bounds)
[1563887024] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563887040] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7459292 kB) free!
[1563887046] SERVICE ALERT: sig2-vm1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig2-vm1' was out of bounds)
[1563887067] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig2-vm1' was out of bounds)
[1563887110] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig2-vm1' was out of bounds)
[1563887137] SERVICE ALERT: sig2-vm1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig2-vm1' was out of bounds)
[1563887144] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563887150] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.2% (2638260 kB) free!
[1563887160] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12110792 kB) free.
[1563887161] SERVICE ALERT: sig2-vm1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig2-vm1' was out of bounds)
[1563887187] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1563887223] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.88, 1.12, 4.46
[1563887255] HOST ALERT: sig2-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.180)
[1563887270] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.4% (2501164 kB) free!
[1563887318] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182496 kB) free!
[1563887390] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.6% (2371880 kB) free!
[1563887554] HOST ALERT: sig2-vm1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.45)
[1563887697] HOST ALERT: sig2-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1563887744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563887760] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7421892 kB) free!
[1563887864] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563887880] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12116008 kB) free.
[1563887941] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.82, 7.28, 5.21
[1563888061] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.39, 6.28, 5.10
[1563888183] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.92, 5.66, 5.01
[1563888464] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563888480] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7116908 kB) free!
[1563888585] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563888590] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168620 kB) free!
[1563888600] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12118120 kB) free.
[1563888783] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.19, 4.75, 4.99
[1563889048] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46473 MB (20% inode=99%):
[1563889168] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46428 MB (20% inode=99%):
[1563889184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563889200] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6878260 kB) free!
[1563889288] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46428 MB (20% inode=99%):
[1563889304] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563889320] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12115012 kB) free.
[1563889499] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172540 kB) free!
[1563889619] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171312 kB) free!
[1563889739] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189716 kB) free!
[1563889875] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1563889904] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563889996] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1563890024] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563890116] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1563890308] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1563890520] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8516892 kB) free!
[1563890583] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.05, 5.84, 5.45
[1563890624] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1563890640] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12119524 kB) free.
[1563890703] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.68, 5.45, 5.35
[1563890744] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563890823] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.06, 4.91, 5.17
[1563891241] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8323348 kB) free!
[1563891361] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12110744 kB) free.
[1563891423] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1563891423] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.55, 4.76, 4.86
[1563891945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563891961] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8192404 kB) free!
[1563892017] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1563892023] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.91, 6.03, 5.58
[1563892065] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563892081] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12115176 kB) free.
[1563892143] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.18, 6.09, 5.66
[1563892263] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.78, 5.94, 5.66
[1563892665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563892681] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8143668 kB) free!
[1563892739] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 62.9% (15404708 kB) free.
[1563892785] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563892801] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12111624 kB) free.
[1563892888] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 62231 MB (27% inode=99%):
[1563893116] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1563893211] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1563893236] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1563893331] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1563893357] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1563893385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563893401] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7865996 kB) free!
[1563893505] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563893521] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12113356 kB) free.
[1563893931] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1563894051] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1563894105] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563894121] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7608576 kB) free!
[1563894225] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563894241] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12110420 kB) free.
[1563894557] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1563894561] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563894769] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563895067] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563895157] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1563895277] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1563895425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563895441] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7076400 kB) free!
[1563895545] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563895562] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12101432 kB) free.
[1563895877] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1563895997] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1563896117] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1563896145] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563896161] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7063600 kB) free!
[1563896265] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563896281] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12109664 kB) free.
[1563896718] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1563896838] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1563896865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563896881] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6932912 kB) free!
[1563896940] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191756 kB) free!
[1563896958] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1563896985] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563897001] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12103160 kB) free.
[1563897060] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198976 kB) free!
[1563897180] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188100 kB) free!
[1563897585] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563897706] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563897710] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563897780] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1563897780] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 52.4% (12837356 kB) free.
[1563898158] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1563898158] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1563898201] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8417760 kB) free!
[1563898321] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12107444 kB) free.
[1563898718] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563898758] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1563898878] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1563898905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563898921] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8584456 kB) free!
[1563898980] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170204 kB) free!
[1563898998] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1563899025] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563899041] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12109816 kB) free.
[1563899100] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 43.3% (10605524 kB) free.
[1563899145] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563899390] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563899641] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8765264 kB) free!
[1563899745] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563899761] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7034320 kB) free!
[1563899865] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563899882] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12107724 kB) free.
[1563900465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563900481] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7095268 kB) free!
[1563900586] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563900592] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1563900601] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12099316 kB) free.
[1563901185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563901305] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563901801] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8492264 kB) free!
[1563901905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563901921] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12105628 kB) free.
[1563902025] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563902521] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8458416 kB) free!
[1563902641] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12107332 kB) free.
[1563903225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563903241] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8365828 kB) free!
[1563903345] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563903361] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12098288 kB) free.
[1563903663] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563903946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563903961] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8255748 kB) free!
[1563904066] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563904082] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12107132 kB) free.
[1563904666] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563904681] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8407212 kB) free!
[1563904786] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563904801] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12107652 kB) free.
[1563905082] SERVICE ALERT: rq-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1563905386] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563905401] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8477152 kB) free!
[1563905463] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 1.93, 4.80
[1563905506] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563905521] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12107132 kB) free.
[1563906106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563906121] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8525488 kB) free!
[1563906226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563906241] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12104428 kB) free.
[1563906347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes