[1576434974] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(82) < 2019/12/16 00:04:24 [error] 20756#0: *17390077 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576435053] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1576435095] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576435173] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1576435397] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (214804 kB) free!
[1576436995] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.6% (3579696 kB) free!
[1576437115] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3577676 kB) free!
[1576437235] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (3567484 kB) free!
[1576437797] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.5% (12453612 kB) free.
[1576438397] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (4055816 kB) free!
[1576438517] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188804 kB) free!
[1576438637] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163904 kB) free!
[1576439237] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1576439237] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.2% (3245840 kB) free!
[1576439450] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576439571] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576439838] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166388 kB) free!
[1576440058] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9080820 kB) free!
[1576440178] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (7869040 kB) free!
[1576440298] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (7871280 kB) free!
[1576444038] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 55.5% (13690324 kB) free.
[1576447094] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(55) < 2019/12/16 03:28:05 [error] 20756#0: *17422165 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576447214] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(28) < 2019/12/16 03:30:14 [error] 20752#0: *17422351 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576447334] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576447935] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/12/16 03:35:56 [error] 20756#0: *17423120 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576448055] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576448570] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1576448691] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576451838] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1576456436] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195840 kB) free!
[1576456638] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198196 kB) free!
[1576456758] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199720 kB) free!
[1576456878] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (214156 kB) free!
[1576457036] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.4% (20195300 kB) free.
[1576457330] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22656972 kB) free.
[1576459455] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(35) < 2019/12/16 06:48:44 [error] 20756#0: *17453700 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576459575] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576460175] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(58) < 2019/12/16 06:59:30 [error] 20756#0: *17455063 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576460295] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(25) < 2019/12/16 07:07:11 [error] 20756#0: *17456139 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576460416] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576466542] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1576466662] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1576471815] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2019/12/16 10:19:52 [error] 20757#0: *17486623 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576471935] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576472174] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1013 processes
[1576472295] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1015 processes
[1576472415] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 369 processes
[1576472536] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(68) < 2019/12/16 10:27:59 [error] 20755#0: *17487596 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576472656] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576473087] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1576473208] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1576473256] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/12/16 10:38:48 [error] 20759#0: *17489195 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576473328] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1576473376] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576473536] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191484 kB) free!
[1576473656] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186144 kB) free!
[1576473776] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171192 kB) free!
[1576474462] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 372 processes
[1576474582] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1576474818] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22320880 kB) free.
[1576475098] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4135824 kB) free!
[1576476836] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.5% (3320012 kB) free!
[1576476956] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.2% (3000568 kB) free!
[1576476983] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 371 processes
[1576477076] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (2823072 kB) free!
[1576477103] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1576477226] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.1% (3733028 kB) free!
[1576477346] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.9% (3671052 kB) free!
[1576477466] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.6% (3596464 kB) free!
[1576477676] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2050084 kB) free!
[1576478066] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1576478066] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2193392 kB) free!
[1576478096] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.9% (22275328 kB) free.
[1576478140] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.7% (22622944 kB) free.
[1576478190] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.1% (39001388 kB) free.
[1576478303] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 382 processes
[1576478423] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1576479871] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(44) < 169.38.76.212 [16/Dec/2019:12:32:12 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576479976] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(67) < 2019/12/16 12:32:12 [error] 20758#0: *17508000 no live upstreams while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "backend_13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576479991] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576480097] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576481090] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576481210] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576482345] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (2413576 kB) free!
[1576482465] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (2308876 kB) free!
[1576482585] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2184680 kB) free!
[1576484486] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (4288180 kB) free!
[1576484606] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4185024 kB) free!
[1576484679] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9532032 kB) free!
[1576484726] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (4064164 kB) free!
[1576484799] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9280224 kB) free!
[1576484919] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9050112 kB) free!
[1576485308] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177944 kB) free!
[1576485428] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189524 kB) free!
[1576485548] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175812 kB) free!
[1576485926] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2387524 kB) free!
[1576486096] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2019/12/16 14:09:48 [error] 20760#0: *17524416 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576486216] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576487319] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4540308 kB) free!
[1576487578] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (3681108 kB) free!
[1576487698] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3511192 kB) free!
[1576487818] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (3323480 kB) free!
[1576488395] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.9% (22428888 kB) free.
[1576488418] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2215256 kB) free!
[1576488521] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 54.2% (13284404 kB) free.
[1576488551] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16122108 kB) free.
[1576488580] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 35.2% (17331824 kB) free.
[1576488617] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.6% (22687124 kB) free.
[1576490068] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4828272 kB) free!
[1576490188] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4746576 kB) free!
[1576490308] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4647996 kB) free!
[1576490423] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1576490543] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1576490886] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (256444 kB) free!
[1576491006] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (262560 kB) free!
[1576491126] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (252680 kB) free!
[1576491395] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177060 kB) free!
[1576491515] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166808 kB) free!
[1576491635] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197636 kB) free!
[1576492216] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(109) < 2019/12/16 15:59:56 [error] 20756#0: *17542101 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576492337] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576493308] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2350888 kB) free!
[1576493847] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9774684 kB) free!
[1576493967] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9698064 kB) free!
[1576494087] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (9611432 kB) free!
[1576494396] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2382300 kB) free!
[1576494516] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (2383524 kB) free!
[1576494636] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2383032 kB) free!
[1576495188] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (573016 kB) free!
[1576495308] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (376412 kB) free!
[1576495429] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231664 kB) free!
[1576496420] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1576496540] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1576496660] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1576497736] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2019/12/16 17:32:06 [error] 20756#0: *17557519 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576497856] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576498917] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.0% (22219348 kB) free.
[1576498970] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.6% (38747352 kB) free.
[1576499009] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 27.5% (13532992 kB) free.
[1576499052] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 72.9% (17858952 kB) free.
[1576499090] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.6% (14894124 kB) free.
[1576500527] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4545860 kB) free!
[1576500648] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4442380 kB) free!
[1576500687] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4624020 kB) free!
[1576500768] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (4335892 kB) free!
[1576501327] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (1998500 kB) free!
[1576501447] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (1277140 kB) free!
[1576501567] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1203652 kB) free!
[1576501825] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4894876 kB) free!
[1576501945] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (4760684 kB) free!
[1576502065] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4614368 kB) free!
[1576502656] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(59) < 2019/12/16 18:49:48 [error] 20755#0: *17570978 no live upstreams while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "backend_13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576502776] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576502791] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(37) < 103.209.99.7 [16/Dec/2019:18:49:48 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576502912] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576503168] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1759692 kB) free!
[1576504465] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2370196 kB) free!
[1576504858] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24938836 kB) free!
[1576504978] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43430032 kB) free.
[1576505172] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.4% (22555788 kB) free.
[1576505210] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 79.2% (19409368 kB) free.
[1576505246] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 29.0% (14292620 kB) free.
[1576505282] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.9% (21549132 kB) free.
[1576508166] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 565 processes
[1576508176] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4873796 kB) free!
[1576508286] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 566 processes
[1576508296] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4435004 kB) free!
[1576508406] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 573 processes
[1576508416] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4358864 kB) free!
[1576508777] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2019/12/16 20:28:40 [error] 20757#0: *17587039 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576508897] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576509606] SERVICE ALERT: webservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 610 processes
[1576510816] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1760768 kB) free!
[1576512488] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.6% (22367148 kB) free.
[1576514897] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(99) < 2019/12/16 22:17:58 [error] 20758#0: *17604247 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576515017] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576520898] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)