[1553538824] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553539093] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8290264 kB) free!
[1553539214] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.4% (13011488 kB) free.
[1553539357] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5942048 kB) free!
[1553539425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553539478] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (4636100 kB) free!
[1553539545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553539598] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1553539598] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (2592568 kB) free!
[1553539665] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553539814] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.6% (7667460 kB) free!
[1553539934] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.6% (12601856 kB) free.
[1553540198] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6897896 kB) free!
[1553540265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553540385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553540505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553540510] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44320 MB (19% inode=99%):
[1553540534] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8207608 kB) free!
[1553540630] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55749 MB (24% inode=99%):
[1553540654] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13163156 kB) free.
[1553540798] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11707124 kB) free.
[1553541254] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8099676 kB) free!
[1553541374] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13163396 kB) free.
[1553541398] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3595336 kB) free!
[1553541518] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11719852 kB) free.
[1553541705] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1553542118] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6278796 kB) free!
[1553542238] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.1% (4989308 kB) free!
[1553542306] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553542358] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3254420 kB) free!
[1553542426] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553542958] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7379484 kB) free!
[1553543026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553543146] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553543174] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9245596 kB) free!
[1553543266] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553543294] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13163788 kB) free.
[1553543558] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11709200 kB) free.
[1553544159] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4497916 kB) free!
[1553544230] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44341 MB (19% inode=99%):
[1553544279] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698600 kB) free.
[1553544350] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55739 MB (24% inode=99%):
[1553544466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553544879] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6828460 kB) free!
[1553544999] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5519080 kB) free!
[1553545066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553545119] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4154908 kB) free!
[1553545187] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553545307] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553545694] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8653652 kB) free!
[1553545719] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.6% (8157380 kB) free!
[1553545814] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13167584 kB) free.
[1553545907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553546027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553546147] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1553546319] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698492 kB) free.
[1553546747] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553546867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553546919] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5075508 kB) free!
[1553546987] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553547039] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3376868 kB) free!
[1553547159] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11708152 kB) free.
[1553547759] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6023776 kB) free!
[1553547879] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (4720512 kB) free!
[1553548000] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2776632 kB) free!
[1553548187] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553548600] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (7018196 kB) free!
[1553548787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553548814] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9497516 kB) free!
[1553548907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553548934] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13176184 kB) free.
[1553549027] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553549200] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11700432 kB) free.
[1553549627] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553549747] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553549800] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3970704 kB) free!
[1553549867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553549920] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11695680 kB) free.
[1553550520] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6484808 kB) free!
[1553550640] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5185936 kB) free!
[1553550760] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3580236 kB) free!
[1553551067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1553551335] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9853840 kB) free!
[1553551360] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7553752 kB) free!
[1553551455] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.7% (7230848 kB) free!
[1553551575] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.8% (13188848 kB) free.
[1553551667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553551787] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553551960] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11686012 kB) free.
[1553552387] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553552507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553552560] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4688964 kB) free!
[1553552627] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553552681] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (2663648 kB) free!
[1553552801] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11688452 kB) free.
[1553553401] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5617536 kB) free!
[1553553521] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4300056 kB) free!
[1553553641] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11677400 kB) free.
[1553553828] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553553975] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9041372 kB) free!
[1553554095] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.8% (13194676 kB) free.
[1553554241] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6601660 kB) free!
[1553554361] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.8% (5315416 kB) free!
[1553554428] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553554481] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3809512 kB) free!
[1553554548] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1553554550] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47132 MB (20% inode=99%):
[1553554668] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553554670] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55736 MB (24% inode=99%):
[1553555081] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.8% (7778164 kB) free!
[1553555268] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553555388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553555508] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1553555681] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11681004 kB) free.
[1553556108] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553556229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553556281] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4909640 kB) free!
[1553556349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1553556402] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (3096912 kB) free!
[1553556522] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11689484 kB) free.
[1553557122] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5954376 kB) free!
[1553557242] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (4648188 kB) free!
[1553557362] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2662348 kB) free!
[1553557549] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1553557962] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (7011988 kB) free!
[1553558149] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1553558269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553558270] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45975 MB (20% inode=99%):
[1553558389] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1553558390] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55735 MB (24% inode=99%):
[1553558562] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11690180 kB) free.
[1553558989] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553559110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553559162] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4416316 kB) free!
[1553559230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553559282] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11678784 kB) free.
[1553559882] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6857180 kB) free!
[1553560002] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5586708 kB) free!
[1553560123] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4315144 kB) free!
[1553560430] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1553560723] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.2% (8459536 kB) free!
[1553561030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553561150] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553561270] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553561323] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11678892 kB) free.
[1553561870] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553561923] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5395356 kB) free!
[1553561991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553562043] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (4016244 kB) free!
[1553562111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553562163] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11686020 kB) free.
[1553562763] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6543000 kB) free!
[1553562883] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5282148 kB) free!
[1553563003] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3764040 kB) free!
[1553563311] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553563603] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7842508 kB) free!
[1553563911] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553564031] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553564203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11683048 kB) free.
[1553564631] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553564752] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553564803] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4969736 kB) free!
[1553564872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553564896] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9435840 kB) free!
[1553564924] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (3232384 kB) free!
[1553565016] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13143384 kB) free.
[1553565044] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11669928 kB) free.
[1553565644] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6065848 kB) free!
[1553565764] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4773144 kB) free!
[1553565884] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2890548 kB) free!
[1553566072] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553566484] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7048616 kB) free!
[1553566672] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553566792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553566816] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8420568 kB) free!
[1553566912] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553566936] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.5% (13075336 kB) free.
[1553567536] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8271036 kB) free!
[1553567656] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.4% (12999860 kB) free.
[1553567684] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10369176 kB) free.
[1553568256] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7955572 kB) free!
[1553568376] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.6% (12605312 kB) free.
[1553568884] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3356488 kB) free!
[1553569004] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11492644 kB) free.
[1553569576] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8678332 kB) free!
[1553569604] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (2994068 kB) free!
[1553569696] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.3% (12979516 kB) free.
[1553569725] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11657704 kB) free.
[1553570296] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8084700 kB) free!
[1553570324] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4068524 kB) free!
[1553570417] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.7% (12673976 kB) free.
[1553570444] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11652940 kB) free.
[1553571016] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.6% (7198124 kB) free!
[1553571044] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3801772 kB) free!
[1553571136] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11940940 kB) free.
[1553571164] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11645852 kB) free.
[1553571764] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3334408 kB) free!
[1553571884] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11533612 kB) free.
[1553572336] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8162208 kB) free!
[1553572456] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.8% (12705576 kB) free.
[1553572484] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3787168 kB) free!
[1553572604] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11639608 kB) free.
[1553573056] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.0% (6923452 kB) free!
[1553573176] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.1% (11876692 kB) free.
[1553573204] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3771684 kB) free!
[1553573324] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11645632 kB) free.
[1553573925] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4067972 kB) free!
[1553574045] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11628380 kB) free.
[1553574376] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8220224 kB) free!
[1553574496] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.9% (12753752 kB) free.
[1553574645] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4223088 kB) free!
[1553574765] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11636516 kB) free.
[1553575096] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.7% (7221264 kB) free!
[1553575216] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.4% (12039120 kB) free.
[1553575365] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4312552 kB) free!
[1553575485] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11461084 kB) free.
[1553576082] SERVICE ALERT: pservercluster1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46149 MB (20% inode=99%):
[1553576085] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4220428 kB) free!
[1553576202] SERVICE ALERT: pservercluster1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46149 MB (20% inode=99%):
[1553576205] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11613976 kB) free.
[1553576322] SERVICE ALERT: pservercluster1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46149 MB (20% inode=99%):
[1553576416] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8777304 kB) free!
[1553576536] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.9% (5846764 kB) free!
[1553576656] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.1% (10912164 kB) free.
[1553576805] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4165980 kB) free!
[1553576925] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (2154432 kB) free!
[1553577045] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10858924 kB) free.
[1553577256] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.2% (5992688 kB) free!
[1553577356] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 19.5% (1433316 kB) free!
[1553577376] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11107812 kB) free.
[1553577476] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 19.4% (1426892 kB) free!
[1553577596] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 19.3% (1419088 kB) free!
[1553577645] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2280680 kB) free!
[1553577713] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553577766] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10993316 kB) free.
[1553577976] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.0% (5422708 kB) free!
[1553578033] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.9% (2908044 kB) free!
[1553578096] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10639640 kB) free.
[1553578153] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2791732 kB) free!
[1553578273] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.6% (2702092 kB) free!
[1553578313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553578365] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (2124716 kB) free!
[1553578433] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1553578485] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10901916 kB) free.
[1553578696] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.4% (5144020 kB) free!
[1553578816] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10375328 kB) free.
[1553579033] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553579085] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2276508 kB) free!
[1553579153] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1553579205] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10826360 kB) free.
[1553579416] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.2% (5535824 kB) free!
[1553579536] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10522132 kB) free.
[1553579753] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1553579874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553579994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553580137] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.3% (5592784 kB) free!
[1553580257] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10451804 kB) free.
[1553580405] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4060124 kB) free!
[1553580525] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1553580525] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11565952 kB) free.
[1553580857] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.2% (5504268 kB) free!
[1553580977] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10509472 kB) free.
[1553581125] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4567692 kB) free!
[1553581245] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (2789712 kB) free!
[1553581365] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11098468 kB) free.
[1553581965] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2609780 kB) free!
[1553582086] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10893512 kB) free.
[1553582177] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.3% (7071528 kB) free!
[1553582297] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11424100 kB) free.
[1553582897] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.1% (6949208 kB) free!
[1553583017] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.1% (11403660 kB) free.
[1553583283] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1553583285] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4870900 kB) free!
[1553583403] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1553583405] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3539988 kB) free!
[1553583523] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1553583525] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11550960 kB) free.
[1553583617] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.5% (6657864 kB) free!
[1553583673] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.1% (14810360 kB) free.
[1553583737] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11338160 kB) free.
[1553584123] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1553584125] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3564252 kB) free!
[1553584190] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43338 MB (19% inode=99%):
[1553584245] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11536412 kB) free.
[1553584310] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54607 MB (24% inode=99%):
[1553584337] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.3% (6567508 kB) free!
[1553584457] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.9% (11297224 kB) free.
[1553584845] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3128508 kB) free!
[1553584965] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11327588 kB) free.
[1553585057] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.3% (6085856 kB) free!
[1553585178] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.3% (10977532 kB) free.
[1553585565] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2380440 kB) free!
[1553585685] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10697796 kB) free.
[1553585777] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.2% (6022156 kB) free!
[1553585897] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.2% (10927352 kB) free.
[1553586497] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.2% (5499652 kB) free!
[1553586617] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10542060 kB) free.
[1553586711] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46470 MB (20% inode=99%):
[1553586831] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46462 MB (20% inode=99%):
[1553586885] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3515144 kB) free!
[1553586951] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55614 MB (24% inode=99%):
[1553587005] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11526364 kB) free.
[1553587217] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.0% (5434832 kB) free!
[1553587337] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10491160 kB) free.
[1553587605] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2876716 kB) free!
[1553587725] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11124644 kB) free.
[1553587937] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.0% (5413884 kB) free!
[1553588057] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10485628 kB) free.
[1553588325] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2476756 kB) free!
[1553588394] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553588445] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10841032 kB) free.
[1553588657] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.6% (5216648 kB) free!
[1553588777] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10187312 kB) free.
[1553588994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553589115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553589235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553589377] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.0% (5900856 kB) free!
[1553589497] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.4% (11016656 kB) free.
[1553589645] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3997768 kB) free!
[1553589765] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11513328 kB) free.
[1553590097] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.8% (7273760 kB) free!
[1553590217] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.3% (6039712 kB) free!
[1553590338] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9918468 kB) free.
[1553590365] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3686464 kB) free!
[1553590485] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11518560 kB) free.
[1553590937] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.9% (5381880 kB) free!
[1553591057] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10373180 kB) free.
[1553591086] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3363120 kB) free!
[1553591206] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11500840 kB) free.
[1553591657] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.1% (5946484 kB) free!
[1553591777] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10859188 kB) free.
[1553591806] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3008136 kB) free!
[1553591926] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11282284 kB) free.
[1553592377] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.8% (6288092 kB) free!
[1553592497] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.9% (10806620 kB) free.
[1553592526] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2664848 kB) free!
[1553592646] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10968920 kB) free.
[1553593097] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.2% (6031552 kB) free!
[1553593217] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.8% (10749112 kB) free.
[1553593246] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2270672 kB) free!
[1553593366] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10626924 kB) free.
[1553593817] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.2% (5994576 kB) free!
[1553593944] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.7% (10679708 kB) free.
[1553594544] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.0% (5899476 kB) free!
[1553594566] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3934344 kB) free!
[1553594664] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10624460 kB) free.
[1553594686] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11489892 kB) free.
[1553595264] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.9% (5871228 kB) free!
[1553595287] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4248868 kB) free!
[1553595384] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.5% (10575396 kB) free.
[1553595407] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (2348712 kB) free!
[1553595527] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10773824 kB) free.
[1553595984] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6188648 kB) free!
[1553596104] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10519280 kB) free.
[1553596127] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2694716 kB) free!
[1553596247] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11304784 kB) free.
[1553596281] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4903340 kB) free!
[1553596401] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4901092 kB) free!
[1553596521] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (4902336 kB) free!
[1553596704] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.7% (6742532 kB) free!
[1553596824] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4352516 kB) free!
[1553596847] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3087068 kB) free!
[1553596944] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.5% (9608672 kB) free!
[1553596967] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11337848 kB) free.
[1553597544] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4597648 kB) free!
[1553597567] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2925528 kB) free!
[1553597635] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553597687] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11212768 kB) free.
[1553598144] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.0% (7414112 kB) free!
[1553598235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553598287] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2455696 kB) free!
[1553598356] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1553598407] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10805676 kB) free.
[1553598744] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1553598744] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10041292 kB) free.
[1553598955] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553599075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553599195] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553599345] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.0% (5905272 kB) free!
[1553599465] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10370640 kB) free.
[1553599608] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4662252 kB) free!
[1553599728] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2875292 kB) free!
[1553599795] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553599848] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11114000 kB) free.
[1553600065] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.8% (6286444 kB) free!
[1553600185] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10345116 kB) free.
[1553600395] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553600448] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4048048 kB) free!
[1553600515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553600568] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11449656 kB) free.
[1553600635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553600785] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.2% (6515232 kB) free!
[1553600905] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10296888 kB) free.
[1553601168] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3627292 kB) free!
[1553601288] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11436048 kB) free.
[1553601505] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.6% (6694224 kB) free!
[1553601625] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4185072 kB) free!
[1553601745] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.1% (9421224 kB) free!
[1553601888] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3095288 kB) free!
[1553602008] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11221580 kB) free.
[1553602345] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4435364 kB) free!
[1553602608] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2372688 kB) free!
[1553602728] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10689804 kB) free.
[1553602945] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.2% (7005596 kB) free!
[1553603721] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 25.2% (6183740 kB) free.
[1553603928] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4393620 kB) free!
[1553604048] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (2507968 kB) free!
[1553604145] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4420732 kB) free!
[1553604168] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10984776 kB) free.
[1553604745] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.8% (7298624 kB) free!
[1553604768] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2512648 kB) free!
[1553604835] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553604888] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10783660 kB) free.
[1553605435] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553605555] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553605676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553605945] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4477520 kB) free!
[1553606088] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3809312 kB) free!
[1553606151] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46760 MB (20% inode=99%):
[1553606208] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11421000 kB) free.
[1553606271] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54951 MB (24% inode=99%):
[1553606545] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.6% (7190604 kB) free!
[1553606808] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3443908 kB) free!
[1553606871] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43180 MB (19% inode=99%):
[1553606928] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11410656 kB) free.
[1553606991] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54449 MB (24% inode=99%):
[1553607529] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3108520 kB) free!
[1553607649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11253420 kB) free.
[1553607745] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4678428 kB) free!
[1553608249] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2653088 kB) free!
[1553608345] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.7% (7262356 kB) free!
[1553608369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11060420 kB) free.
[1553608969] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2813320 kB) free!
[1553609089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11105184 kB) free.
[1553609689] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2489244 kB) free!
[1553609809] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11031044 kB) free.
[1553610409] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2917932 kB) free!
[1553610477] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553610529] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11132096 kB) free.
[1553611077] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553611129] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2870064 kB) free!
[1553611197] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553611249] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11080336 kB) free.
[1553611797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553611832] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6986 MB (13% inode=98%):
[1553611850] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3400648 kB) free!
[1553611917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553611970] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11382956 kB) free.
[1553612037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553612545] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3942548 kB) free!
[1553612570] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3175000 kB) free!
[1553612637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553612690] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11371292 kB) free.
[1553613146] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.6% (7195632 kB) free!
[1553613237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553613290] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2758624 kB) free!
[1553613358] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1553613410] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11029004 kB) free.
[1553613943] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (1472612 kB) free!
[1553613957] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1553614010] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2440992 kB) free!
[1553614063] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 20.0% (1472248 kB) free!
[1553614077] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1553614130] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10792072 kB) free.
[1553614183] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 20.0% (1471868 kB) free!
[1553614677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553614730] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3993176 kB) free!
[1553614797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553614850] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10490004 kB) free.
[1553614917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553616051] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3678412 kB) free!
[1553616171] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11361980 kB) free.
[1553616771] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5692728 kB) free!
[1553616891] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (4380948 kB) free!
[1553617011] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2388804 kB) free!
[1553617317] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553617611] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6707448 kB) free!
[1553617917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553618037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1553618211] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11361184 kB) free.
[1553618637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553618756] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 322 processes
[1553618757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553618811] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3811632 kB) free!
[1553618877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553618931] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11364580 kB) free.
[1553619039] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 338 processes
[1553619531] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6278712 kB) free!
[1553619651] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.1% (4972836 kB) free!
[1553619746] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4740156 kB) free!
[1553619772] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1553619772] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3484184 kB) free!
[1553620078] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553620346] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.7% (7242148 kB) free!
[1553620372] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7488616 kB) free!
[1553620678] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553620798] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553620972] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11362740 kB) free.
[1553621398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553621518] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553621547] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4721044 kB) free!
[1553621572] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4493316 kB) free!
[1553621638] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553621692] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (2619492 kB) free!
[1553621812] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11369464 kB) free.
[1553622147] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.4% (7097656 kB) free!
[1553622412] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5475408 kB) free!
[1553622532] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4181492 kB) free!
[1553622652] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11363824 kB) free.
[1553622838] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553623252] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6504156 kB) free!
[1553623347] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4743996 kB) free!
[1553623372] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5210712 kB) free!
[1553623439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553623493] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3843748 kB) free!
[1553623559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553623679] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553623947] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.2% (7965400 kB) free!
[1553624093] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7834028 kB) free!
[1553624279] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553624399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553624519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553624547] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9976752 kB) free.
[1553624693] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11359668 kB) free.