[1583865264] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20380512 kB) free!
[1583865384] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (17051080 kB) free!
[1583865504] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583865504] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38498096 kB) free.
[1583866104] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17832792 kB) free!
[1583866224] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39473812 kB) free.
[1583866828] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18217292 kB) free!
[1583866948] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39442740 kB) free.
[1583867021] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583867081] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181868 kB) free!
[1583867141] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (178032 kB) free!
[1583867201] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (177756 kB) free!
[1583867502] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 64.6% (15829168 kB) free.
[1583867548] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18556976 kB) free!
[1583867668] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39743176 kB) free.
[1583868268] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18854076 kB) free!
[1583868388] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39929460 kB) free.
[1583868988] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18972640 kB) free!
[1583869108] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39599088 kB) free.
[1583869283] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4896264 kB) free!
[1583869883] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197176 kB) free!
[1583870308] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20902816 kB) free!
[1583870355] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231972 kB) free!
[1583870428] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40883336 kB) free.
[1583870475] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243000 kB) free!
[1583870595] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239920 kB) free!
[1583871028] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19794092 kB) free!
[1583871148] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40190424 kB) free.
[1583871748] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19583752 kB) free!
[1583871869] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38757236 kB) free.
[1583872395] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 47.8% (23576672 kB) free.
[1583872403] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(78) < 103.16.101.72 [11/Mar/2020:02:01:14 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1583872468] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9420008 kB) free!
[1583872523] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583872588] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37402596 kB) free.
[1583873188] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23932324 kB) free!
[1583873308] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36881952 kB) free.
[1583873908] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10819176 kB) free!
[1583873977] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 11.1% (5449272 kB) free!
[1583874028] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37535640 kB) free.
[1583874097] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (248448 kB) free!
[1583874217] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222760 kB) free!
[1583874268] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 60.3% (29728144 kB) free.
[1583874628] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12039464 kB) free!
[1583874748] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37547052 kB) free.
[1583874817] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.5% (43610408 kB) free.
[1583875245] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.6% (9962752 kB) free.
[1583875822] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(327) < 2020/03/11 03:00:21 [error] 12256#0: *68448560 no live upstreams while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583875942] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583875948] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18832984 kB) free!
[1583876068] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42766020 kB) free.
[1583876123] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3670) < 172.105.59.173 [11/Mar/2020:03:01:41 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1583876243] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583876249] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(272) < 49.50.86.152 [11/Mar/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583876371] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583876668] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23322128 kB) free!
[1583876788] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42188732 kB) free.
[1583877989] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24200164 kB) free!
[1583878109] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41887260 kB) free.
[1583883046] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183072 kB) free!
[1583883166] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174384 kB) free!
[1583883286] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189984 kB) free!
[1583883683] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.4% (14656260 kB) free.
[1583884084] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.8% (19549160 kB) free.
[1583884190] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 366 processes
[1583884310] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1583886886] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 223451 MB (20% inode=99%):
[1583887006] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 222868 MB (20% inode=99%):
[1583887126] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 222367 MB (20% inode=99%):
[1583887324] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (377748 kB) free!
[1583887444] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7859132 kB) free!
[1583887474] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 93.0% (122538632 kB) free.
[1583887521] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.4% (39119440 kB) free.
[1583887557] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.8% (20789400 kB) free.
[1583887726] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 259920 MB (24% inode=99%):
[1583890883] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4778656 kB) free!
[1583891003] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4809456 kB) free!
[1583891123] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4802952 kB) free!
[1583892323] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (540920 kB) free!
[1583899294] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.0% (4176884 kB) free!
[1583899347] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1583899414] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.3% (4006748 kB) free!
[1583899467] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 706 processes
[1583899534] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.6% (3818064 kB) free!
[1583899587] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 706 processes
[1583900134] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2425468 kB) free!
[1583900309] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19237292 kB) free!
[1583900429] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40286368 kB) free.
[1583901029] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14818684 kB) free!
[1583901144] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/03/11 09:58:07 [error] 12254#0: *68614623 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /favicon.ico HTTP/1.1", host: "qservercluster1.com:8091"
[1583901149] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35832692 kB) free.
[1583901264] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583902349] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5734504 kB) free!
[1583902469] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34144508 kB) free.
[1583902619] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1583902739] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1583903069] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8167816 kB) free!
[1583903187] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 663 processes
[1583903190] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3651540 kB) free!
[1583903310] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28828984 kB) free.
[1583903787] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 715 processes
[1583903907] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 715 processes
[1583903910] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5259424 kB) free!
[1583904027] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 715 processes
[1583904030] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29464388 kB) free.
[1583904299] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4265668 kB) free!
[1583904407] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21616696 kB) free.
[1583904419] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.6% (3956096 kB) free!
[1583904420] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.8% (21180192 kB) free.
[1583904539] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.9% (3579200 kB) free!
[1583904539] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1583904630] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4783176 kB) free!
[1583904659] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1583904750] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.5% (25761056 kB) free!
[1583904870] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19789196 kB) free!
[1583905139] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1675736 kB) free!
[1583905470] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36734504 kB) free.
[1583905613] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197804 kB) free!
[1583905643] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1583905684] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189056 kB) free!
[1583905733] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190476 kB) free!
[1583905763] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1583905804] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206328 kB) free!
[1583905853] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1583905853] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181648 kB) free!
[1583905924] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187380 kB) free!
[1583906070] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12816440 kB) free!
[1583906165] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185652 kB) free!
[1583906190] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37148200 kB) free.
[1583906285] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185500 kB) free!
[1583906405] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189200 kB) free!
[1583906790] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12914484 kB) free!
[1583906910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35584064 kB) free.
[1583907027] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1583907027] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1583907195] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8086488 kB) free!
[1583907315] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7617504 kB) free!
[1583907424] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (4242728 kB) free!
[1583907435] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.4% (7108640 kB) free!
[1583907510] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11801936 kB) free!
[1583907544] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4116096 kB) free!
[1583907584] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1583907630] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7533104 kB) free!
[1583907653] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 71.4% (17511276 kB) free.
[1583907664] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.1% (3962728 kB) free!
[1583907704] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583907751] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33648980 kB) free.
[1583907824] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1583907845] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 12.9% (3195812 kB) free!
[1583907965] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 11.3% (2785844 kB) free!
[1583908035] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4739624 kB) free!
[1583908085] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2376940 kB) free!
[1583908227] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 712 processes
[1583908347] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 709 processes
[1583908350] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7817064 kB) free!
[1583908467] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 709 processes
[1583908470] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3455328 kB) free!
[1583908570] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 36.36, 38.21, 31.68
[1583908590] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29603620 kB) free.
[1583908636] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.2% (30166196 kB) free.
[1583908690] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 30.00, 35.23, 31.38
[1583908740] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.5% (15383324 kB) free.
[1583908811] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 24.49, 31.97, 30.69
[1583908864] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (251824 kB) free!
[1583909069] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9092512 kB) free!
[1583909124] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9055412 kB) free!
[1583909189] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (8891140 kB) free!
[1583909190] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4823620 kB) free!
[1583909244] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9030892 kB) free!
[1583909286] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.1% (22737120 kB) free.
[1583909309] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (8689908 kB) free!
[1583909310] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (625688 kB) free!
[1583909364] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (9032364 kB) free!
[1583909411] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 23.27, 25.29, 27.79
[1583909430] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26906444 kB) free.
[1583909843] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(156) < 92.42.104.166 [11/Mar/2020:12:20:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583909963] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583910030] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (1100548 kB) free!
[1583910150] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394744 kB) free!
[1583910217] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8582232 kB) free!
[1583910225] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1583910270] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27953824 kB) free.
[1583910337] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8388412 kB) free!
[1583910457] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.3% (8012608 kB) free!
[1583910870] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1757736 kB) free!
[1583910990] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30462624 kB) free.
[1583911590] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5782648 kB) free!
[1583911657] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4367920 kB) free!
[1583911709] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4141872 kB) free!
[1583911710] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1329420 kB) free!
[1583911812] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 96.50, 40.20, 31.08
[1583911830] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30356500 kB) free.
[1583911932] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 35.31, 36.22, 30.77
[1583912052] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 25.98, 33.32, 30.45
[1583912365] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220612 kB) free!
[1583912431] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3613740 kB) free!
[1583912524] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 27.9% (6844016 kB) free.
[1583912551] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395576 kB) free!
[1583912644] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1583912644] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.4% (19841668 kB) free.
[1583912652] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1583912652] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 25.11, 25.45, 27.54
[1583912671] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31200060 kB) free.
[1583912704] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.7% (43240152 kB) free.
[1583912758] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.1% (38505248 kB) free.
[1583912815] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.6% (45150712 kB) free.
[1583912860] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.2% (20871636 kB) free.
[1583913253] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187664 kB) free!
[1583913271] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4480124 kB) free!
[1583913362] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236800 kB) free!
[1583913373] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191980 kB) free!
[1583913391] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (388816 kB) free!
[1583913482] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216460 kB) free!
[1583913493] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192892 kB) free!
[1583913511] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29100680 kB) free.
[1583913602] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1583913602] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238552 kB) free!
[1583914111] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2606840 kB) free!
[1583914231] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404560 kB) free!
[1583914324] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4689316 kB) free!
[1583914351] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33487928 kB) free.
[1583914444] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4591116 kB) free!
[1583914564] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4491004 kB) free!
[1583914951] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8222608 kB) free!
[1583915067] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1583915071] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4116976 kB) free!
[1583915191] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30425340 kB) free.
[1583915340] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4272712 kB) free!
[1583915460] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (3938028 kB) free!
[1583915580] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (3453752 kB) free!
[1583915667] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1583915788] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1583915791] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3816636 kB) free!
[1583915865] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (1954628 kB) free!
[1583915911] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405484 kB) free!
[1583915985] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1742632 kB) free!
[1583916032] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34038916 kB) free.
[1583916105] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1499400 kB) free!
[1583916180] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1583916180] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1224156 kB) free!
[1583916364] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1583916364] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2123368 kB) free!
[1583916631] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8688424 kB) free!
[1583916751] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3965164 kB) free!
[1583916844] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583916845] SERVICE ALERT: rq-2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583916874] HOST ALERT: rq-2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1583916881] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583916896] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583916899] HOST ALERT: redis-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.201)
[1583916911] HOST ALERT: rq-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.203)
[1583916929] SERVICE ALERT: mrq4;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583916934] HOST ALERT: redis-2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.202)
[1583916948] HOST ALERT: mrq4;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1583916967] SERVICE ALERT: rq-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'rq-2' was out of bounds)
[1583916967] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'rq-2' was out of bounds)
[1583916988] SERVICE ALERT: mrq4;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'mrq4' was out of bounds)
[1583917019] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'redis-1' was out of bounds)
[1583917052] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'mrq4' was out of bounds)
[1583917085] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'rq-2' was out of bounds)
[1583917107] SERVICE ALERT: rq-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'rq-2' was out of bounds)
[1583917128] SERVICE ALERT: redis-1;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'redis-1' was out of bounds)
[1583917156] HOST ALERT: mrq4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1583917164] HOST ALERT: rq-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1583917164] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(36583) < 203.212.70.170 [11/Mar/2020:14:29:23 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583917170] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1270) < 103.16.101.72 [11/Mar/2020:14:29:28 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583917187] HOST ALERT: redis-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1583917190] HOST ALERT: rq-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1583917205] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1583917215] HOST ALERT: redis-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1583917227] SERVICE ALERT: rq-2;Total Processes;OK;HARD;1;PROCS OK: 345 processes
[1583917248] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 317588 MB (29% inode=99%):
[1583917284] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(3495) < 103.255.146.164 [11/Mar/2020:14:31:24 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583917291] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(516) < 103.16.101.72 [11/Mar/2020:14:31:30 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583917307] SERVICE ALERT: mrq4;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1583917405] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(28631) < 103.13.97.241 [11/Mar/2020:14:33:24 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583917411] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;HARD;3;(2185) < 103.16.101.72 [11/Mar/2020:14:33:31 +0530] TCP 502 0 0 0.002 "10.147.212.38:4000, 10.147.212.61:4000, 10.147.212.49:4000, 10.147.212.37:4000" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 0.000, 0.001"
[1583917427] SERVICE ALERT: mrq4;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1583917453] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 195.42, 112.65, 58.02
[1583917464] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7236) < 2020/03/11 14:34:24 [error] 12250#0: *68789886 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583917481] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1583917481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 59.7% (78644652 kB) free.
[1583917547] SERVICE ALERT: mrq4;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1583917564] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (4187492 kB) free!
[1583917564] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 98.6% (129956564 kB) free.
[1583917565] SERVICE ALERT: rq-2;CPU Load;OK;HARD;1;OK - load average: 0.03, 0.09, 0.07
[1583917573] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 198.46, 141.40, 75.19
[1583917584] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(667) < 2020/03/11 14:36:05 [error] 12245#0: *68791273 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583917586] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 89.32, 65.61, 31.64
[1583917616] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 72.9% (96071908 kB) free.
[1583917649] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 97.2% (31823412 kB) free.
[1583917693] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.30, 114.49, 74.04
[1583917704] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583917706] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 92.71, 74.67, 39.07
[1583917708] SERVICE ALERT: mrq4;Total Processes;OK;HARD;1;PROCS OK: 219 processes
[1583917827] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 85.10, 78.12, 44.68
[1583917853] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583918025] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 326 processes
[1583918081] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18707960 kB) free!
[1583918135] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.17895.3791> (10.147.212.181:54769 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1583918145] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1583918187] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 707 processes
[1583918201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14704436 kB) free!
[1583918255] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583918265] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1583918307] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 709 processes
[1583918321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32597712 kB) free.
[1583918400] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.17397.2> (10.147.212.158:44520 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1583918428] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 707 processes
[1583918520] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583918605] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583918612] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583918893] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 16.82, 19.71, 34.34
[1583918921] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9457640 kB) free!
[1583919041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38215960 kB) free.
[1583919181] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.3% (15552908 kB) free.
[1583919206] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(230) < 92.42.104.166 [11/Mar/2020:15:00:22 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583919326] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583919364] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1904620 kB) free!
[1583919486] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (4047548 kB) free!
[1583919494] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 19.63, 21.54, 28.56
[1583919606] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3597308 kB) free!
[1583919641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (386664 kB) free!
[1583919726] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (3130380 kB) free!
[1583919761] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374696 kB) free!
[1583919881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37202192 kB) free.
[1583920326] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 91.8% (22663048 kB) free.
[1583920481] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9093716 kB) free!
[1583920601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5100932 kB) free!
[1583920721] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43722104 kB) free.
[1583920859] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1583920980] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1583921196] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.83, 6.06, 5.05
[1583921316] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.23, 5.81, 5.08
[1583921322] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10563748 kB) free!
[1583921436] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.44, 5.40, 5.02
[1583921442] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6402348 kB) free!
[1583921562] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2652000 kB) free!
[1583922036] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.97, 4.24, 4.55
[1583922762] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21238184 kB) free!
[1583923065] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1583923093] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583923362] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40765780 kB) free.
[1583923665] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1583923853] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.7% (4587668 kB) free!
[1583923962] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11474324 kB) free!
[1583923973] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.6% (4548080 kB) free!
[1583924070] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3035108 kB) free!
[1583924082] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7499296 kB) free!
[1583924093] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.7% (4340752 kB) free!
[1583924190] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3526508 kB) free.
[1583924202] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31705152 kB) free.
[1583924265] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1583924265] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 333 processes
[1583924402] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583924802] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11527404 kB) free!
[1583924922] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7477656 kB) free!
[1583925042] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3288776 kB) free!
[1583925861] SERVICE ALERT: qservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46078 MB (20% inode=99%):
[1583925893] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2230004 kB) free!
[1583925982] SERVICE ALERT: qservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45997 MB (20% inode=99%):
[1583926102] SERVICE ALERT: qservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45906 MB (20% inode=99%):
[1583926242] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20002956 kB) free!
[1583926843] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36036620 kB) free.
[1583927428] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1583927443] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6106720 kB) free!
[1583927563] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1888856 kB) free!
[1583927683] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6128708 kB) free!
[1583927791] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (2737004 kB) free!
[1583927911] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3151072 kB) free!
[1583928031] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3571440 kB) free.
[1583928181] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1583928781] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.6% (3288084 kB) free!
[1583928883] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20976800 kB) free!
[1583928901] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3056720 kB) free!
[1583929021] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (2840788 kB) free!
[1583929483] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34101632 kB) free.
[1583929564] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583929621] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1580156 kB) free!
[1583930083] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5985152 kB) free!
[1583930203] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1699028 kB) free!
[1583930221] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1583930221] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.5% (15609048 kB) free.
[1583930323] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400732 kB) free!
[1583930431] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3118448 kB) free!
[1583930551] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3557292 kB) free.
[1583930826] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227932 kB) free!
[1583930894] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583930946] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244796 kB) free!
[1583931066] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237576 kB) free!
[1583931314] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (222528 kB) free!
[1583931434] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235308 kB) free!
[1583931438] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46238 MB (20% inode=99%):
[1583931523] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22121756 kB) free!
[1583931554] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224528 kB) free!
[1583931558] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46206 MB (20% inode=99%):
[1583931678] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46169 MB (20% inode=99%):
[1583932123] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374800 kB) free!
[1583932228] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1583932236] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.53, 5.74, 5.40
[1583932266] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 22.6% (11115512 kB) free.
[1583932278] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 89979 MB (40% inode=99%):
[1583932351] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2946572 kB) free!
[1583932356] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.03, 5.53, 5.37
[1583932471] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3061756 kB) free!
[1583932476] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.95, 5.28, 5.29
[1583932591] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3272180 kB) free.
[1583932657] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1583932693] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 23.9% (5906396 kB) free.
[1583932702] SERVICE ALERT: qservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 97150 MB (43% inode=99%):
[1583932723] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.3% (51807364 kB) free.
[1583932754] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 72.7% (35816532 kB) free.
[1583932778] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583932789] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.3% (21401788 kB) free.
[1583932801] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.9% (35453384 kB) free.
[1583932811] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.7% (20762312 kB) free.
[1583932821] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.7% (21242360 kB) free.
[1583933076] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1583933076] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.45, 4.55, 4.90
[1583933293] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170188 kB) free!
[1583933404] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (250408 kB) free!
[1583933413] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189732 kB) free!
[1583933524] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231868 kB) free!
[1583933527] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(4) < {writer,send_failed,{error,timeout}}
[1583933533] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184512 kB) free!
[1583933644] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (251076 kB) free!
[1583933648] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583933676] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.41, 6.30, 5.33
[1583933797] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.52, 6.36, 5.47
[1583933917] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.49, 6.69, 5.69
[1583933977] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(19) < {writer,send_failed,{error,timeout}}
[1583934097] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583934120] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1583934241] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(3) < {writer,send_failed,{error,timeout}}
[1583934361] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;3;(1) < {writer,send_failed,{error,timeout}}
[1583935214] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.8% (3132116 kB) free!
[1583935334] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (3020680 kB) free!
[1583935455] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2857636 kB) free!
[1583935665] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583935723] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6169708 kB) free!
[1583935828] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 714 processes
[1583935843] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2318280 kB) free!
[1583935948] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 713 processes
[1583935963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37319560 kB) free.
[1583935979] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1583936055] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (523400 kB) free!
[1583936068] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 711 processes
[1583936099] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1583936162] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583936563] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7683984 kB) free!
[1583936683] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3598164 kB) free!
[1583936803] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (483416 kB) free!
[1583937027] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.93, 1.16, 17.88
[1583937403] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14253504 kB) free!
[1583937421] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3648300 kB) free!
[1583937517] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 0.06, 0.32, 2.82
[1583937541] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.1% (3401712 kB) free!
[1583937661] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3172048 kB) free!
[1583938003] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30646000 kB) free.
[1583938261] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1883468 kB) free!
[1583938603] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3999540 kB) free!
[1583938723] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399288 kB) free!
[1583938827] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.41, 0.51, 3.02
[1583938843] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38461552 kB) free.
[1583939444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10840872 kB) free!
[1583939467] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9194376 kB) free!
[1583939564] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6761984 kB) free!
[1583939587] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9068384 kB) free!
[1583939684] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2802664 kB) free!
[1583939707] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (8945716 kB) free!
[1583940284] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14507228 kB) free!
[1583940661] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.6% (15635752 kB) free.
[1583940884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27107732 kB) free.
[1583941484] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2495548 kB) free!
[1583941604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37080608 kB) free.
[1583941918] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.5% (20108444 kB) free.
[1583941950] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 42.3% (20829516 kB) free.
[1583941997] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1583941997] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.9% (21310516 kB) free.
[1583942204] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (9000080 kB) free!
[1583942324] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4845652 kB) free!
[1583942444] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (876392 kB) free!
[1583942525] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185112 kB) free!
[1583942550] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (250540 kB) free!
[1583942645] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178092 kB) free!
[1583942670] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (278384 kB) free!
[1583942765] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1583942765] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196016 kB) free!
[1583942790] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1583942790] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235240 kB) free!
[1583943644] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27171624 kB) free.
[1583944244] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3982972 kB) free!
[1583944364] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36057940 kB) free.
[1583944964] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11232264 kB) free!
[1583945085] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7496068 kB) free!
[1583945205] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3989344 kB) free!
[1583945805] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14393572 kB) free!
[1583945917] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1583946803] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583947005] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (4979616 kB) free!
[1583947605] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14149600 kB) free!
[1583948107] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4893356 kB) free!
[1583948805] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35252636 kB) free.
[1583949405] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12514568 kB) free!
[1583949465] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1583949525] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8816788 kB) free!
[1583949645] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5355048 kB) free!
[1583950245] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13408628 kB) free!
[1583950261] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583950526] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 203.212.70.170 [11/Mar/2020:23:40:23 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1583950646] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found