[1588445266] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35844340 kB) free.
[1588445866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13021772 kB) free!
[1588445986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9325764 kB) free!
[1588446106] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 10.4% (2559148 kB) free!
[1588446106] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35808176 kB) free.
[1588446166] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190472 kB) free!
[1588446226] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (188996 kB) free!
[1588446286] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (164916 kB) free!
[1588446706] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13493804 kB) free!
[1588446826] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9727192 kB) free!
[1588446946] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35711244 kB) free.
[1588447105] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1144552 kB) free!
[1588447186] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 80.6% (19747496 kB) free.
[1588447225] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (475024 kB) free!
[1588447345] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1457204 kB) free!
[1588447420] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2417448 kB) free!
[1588447546] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13387272 kB) free!
[1588447666] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9499540 kB) free!
[1588447786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35661080 kB) free.
[1588448386] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14166568 kB) free!
[1588448506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10467120 kB) free!
[1588448627] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35621492 kB) free.
[1588449226] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14752380 kB) free!
[1588449346] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11105636 kB) free!
[1588449424] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 204465 MB (19% inode=99%):
[1588449466] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35672624 kB) free.
[1588449537] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11823088 kB) free!
[1588449544] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 194434 MB (18% inode=99%):
[1588449616] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233356 kB) free!
[1588449657] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1333300 kB) free!
[1588449664] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 184463 MB (17% inode=99%):
[1588449736] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233656 kB) free!
[1588449777] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (454156 kB) free!
[1588449856] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212944 kB) free!
[1588450066] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14213476 kB) free!
[1588450186] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10502248 kB) free!
[1588450306] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35589260 kB) free.
[1588450906] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12784452 kB) free!
[1588451026] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9164348 kB) free!
[1588451146] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33338408 kB) free.
[1588451656] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.8% (6785676 kB) free!
[1588451746] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11742228 kB) free!
[1588451866] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.6% (23211340 kB) free!
[1588451986] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32280948 kB) free.
[1588452586] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11530872 kB) free!
[1588452706] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35497048 kB) free.
[1588453307] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13645340 kB) free!
[1588453346] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (3180928 kB) free!
[1588453427] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10165372 kB) free!
[1588453547] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32666280 kB) free.
[1588453810] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.9% (43323068 kB) free.
[1588454147] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11588468 kB) free!
[1588454225] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.3% (44484364 kB) free.
[1588454267] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33314640 kB) free.
[1588454867] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13569392 kB) free!
[1588454987] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35488788 kB) free.
[1588455139] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2293) < 172.105.59.173 [03/May/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1588455257] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.3% (44519252 kB) free.
[1588455260] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(3) < 172.105.59.173 [03/May/2020:03:02:32 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1588455327] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(497) < 2020/05/03 03:01:32 [error] 22126#0: *86650059 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588455380] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588455439] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(378) < 182.18.144.246 [03/May/2020:03:00:25 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1588455447] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588455515] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1588455559] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588455587] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14541964 kB) free!
[1588455707] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35476212 kB) free.
[1588456307] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15863920 kB) free!
[1588456427] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12210756 kB) free!
[1588456547] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34494228 kB) free.
[1588457147] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13846688 kB) free!
[1588457267] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35385632 kB) free.
[1588457867] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15117224 kB) free!
[1588457892] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.6% (10196872 kB) free.
[1588457988] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35373796 kB) free.
[1588458146] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3522532 kB) free.
[1588458587] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16390252 kB) free!
[1588458707] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34502432 kB) free.
[1588459092] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (208240 kB) free!
[1588459212] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (208520 kB) free!
[1588459307] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16889136 kB) free!
[1588459332] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (251544 kB) free!
[1588459427] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1588459427] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34125856 kB) free.
[1588460027] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11722656 kB) free!
[1588460147] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5726436 kB) free!
[1588460267] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32573516 kB) free.
[1588460532] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.1% (9577624 kB) free.
[1588460867] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9181552 kB) free!
[1588460987] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33899664 kB) free.
[1588461587] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10612132 kB) free!
[1588461707] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34307876 kB) free.
[1588462307] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11533388 kB) free!
[1588462332] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (257596 kB) free!
[1588462427] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34109460 kB) free.
[1588462452] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (231776 kB) free!
[1588462572] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1588462572] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177920 kB) free!
[1588463027] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12288156 kB) free!
[1588463147] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35010080 kB) free.
[1588463666] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.5% (18254900 kB) free.
[1588463748] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15956144 kB) free!
[1588463752] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 53.2% (13127316 kB) free.
[1588463868] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34204880 kB) free.
[1588463880] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 52.0% (12748376 kB) free.
[1588465068] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19913132 kB) free!
[1588465188] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37043916 kB) free.
[1588465346] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1588465788] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18336104 kB) free!
[1588465908] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36276252 kB) free.
[1588467108] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17703560 kB) free!
[1588467228] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36575808 kB) free.
[1588467828] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15935208 kB) free!
[1588467948] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35277356 kB) free.
[1588467952] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4564648 kB) free!
[1588468072] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (897292 kB) free!
[1588468115] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1588468192] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (309136 kB) free!
[1588468235] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1588468264] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 233678 MB (21% inode=99%):
[1588468355] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1588468379] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 55.9% (73731652 kB) free.
[1588469148] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17574180 kB) free!
[1588469268] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35798052 kB) free.
[1588469869] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14754928 kB) free!
[1588469989] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34375212 kB) free.
[1588471189] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16090604 kB) free!
[1588471309] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34909004 kB) free.
[1588471909] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14196408 kB) free!
[1588472029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33136748 kB) free.
[1588472772] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588473229] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13488912 kB) free!
[1588473349] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33434164 kB) free.
[1588473481] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175668 kB) free!
[1588473601] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187720 kB) free!
[1588473714] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.57, 6.85, 5.07
[1588473721] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184288 kB) free!
[1588473834] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.00, 6.48, 5.15
[1588473949] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10475204 kB) free!
[1588473954] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.94, 5.86, 5.08
[1588474069] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32116364 kB) free.
[1588474669] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8987920 kB) free!
[1588474789] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30836896 kB) free.
[1588474946] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3138760 kB) free!
[1588474955] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1588475066] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (1559160 kB) free!
[1588475186] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (2802640 kB) free!
[1588475389] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7760568 kB) free!
[1588475509] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29929524 kB) free.
[1588475521] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 36.7% (8984796 kB) free.
[1588475573] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3976160 kB) free!
[1588475693] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.8% (3782032 kB) free!
[1588475813] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.0% (3595760 kB) free!
[1588476110] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8626364 kB) free!
[1588476230] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29080764 kB) free.
[1588476354] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.92, 3.77, 4.80
[1588476830] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7082836 kB) free!
[1588476950] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.6% (25853716 kB) free!
[1588476954] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.23, 8.66, 6.47
[1588477014] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.9% (15924224 kB) free.
[1588477070] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16515536 kB) free!
[1588477075] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.97, 9.81, 7.15
[1588477195] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.35, 9.10, 7.24
[1588477321] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (1227928 kB) free!
[1588477441] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168312 kB) free!
[1588477561] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187988 kB) free!
[1588477670] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33031016 kB) free.
[1588478270] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6449420 kB) free!
[1588478390] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32479132 kB) free.
[1588478990] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9496612 kB) free!
[1588479110] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.7% (20753568 kB) free!
[1588479230] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26786156 kB) free.
[1588479595] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1588479595] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.96, 3.93, 4.35
[1588480430] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25428156 kB) free!
[1588480550] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37927760 kB) free.
[1588480561] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1588480561] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3329628 kB) free!
[1588480795] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.24, 5.55, 5.04
[1588480915] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.20, 6.44, 5.42
[1588481035] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.07, 6.54, 5.58
[1588481150] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15189128 kB) free!
[1588481162] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174108 kB) free!
[1588481270] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38204128 kB) free.
[1588481786] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (385396 kB) free!
[1588481871] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13957596 kB) free!
[1588481991] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38570872 kB) free.
[1588482386] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2527756 kB) free!
[1588482591] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13289236 kB) free!
[1588482711] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37094480 kB) free.
[1588483311] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13449736 kB) free!
[1588483431] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9060360 kB) free!
[1588483551] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35411276 kB) free.
[1588483614] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3668712 kB) free!
[1588483734] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3125460 kB) free!
[1588483854] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.8% (2652272 kB) free!
[1588484151] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8814524 kB) free!
[1588484271] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33827036 kB) free.
[1588484454] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1588484454] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (177448 kB) free!
[1588484871] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9712668 kB) free!
[1588484991] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5362148 kB) free!
[1588485111] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32253720 kB) free.
[1588485508] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.26224.42> (10.147.212.147:46504 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588485628] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588485711] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6333308 kB) free!
[1588485831] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31108596 kB) free.
[1588485866] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4565472 kB) free!
[1588485986] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1588485986] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1423608 kB) free!
[1588485986] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4464724 kB) free!
[1588486106] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4374840 kB) free!
[1588486431] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7666504 kB) free!
[1588486552] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3231132 kB) free!
[1588486672] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29123224 kB) free.
[1588487186] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (2106676 kB) free!
[1588487272] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4859532 kB) free!
[1588487392] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (640984 kB) free!
[1588487455] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.7% (16089732 kB) free.
[1588487512] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26573408 kB) free.
[1588488112] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (2985908 kB) free!
[1588488232] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28589180 kB) free.
[1588488832] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6089984 kB) free!
[1588488835] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.68, 3.13, 4.58
[1588488952] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2051648 kB) free!
[1588489072] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27514352 kB) free.
[1588489106] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2380924 kB) free!
[1588489580] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(669) < 103.255.146.164 [03/May/2020:12:32:25 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588489586] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (598356 kB) free!
[1588489672] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3977040 kB) free!
[1588489701] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588489792] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (398600 kB) free!
[1588489912] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25751164 kB) free!
[1588490512] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1588490512] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2416404 kB) free!
[1588490900] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(330) < 103.255.146.164 [03/May/2020:12:54:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588491020] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588491386] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.9% (2102580 kB) free!
[1588491713] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17684500 kB) free!
[1588491835] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.67, 6.41, 5.04
[1588491955] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.58, 7.62, 5.65
[1588491986] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (657208 kB) free!
[1588492075] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.28, 7.34, 5.78
[1588492220] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(504) < 103.255.146.164 [03/May/2020:13:19:50 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588492313] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28635468 kB) free.
[1588492340] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588492563] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588492913] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2711488 kB) free!
[1588493033] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (765104 kB) free!
[1588493153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30280304 kB) free.
[1588493225] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9075880 kB) free!
[1588493345] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8664484 kB) free!
[1588493465] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.7% (8236596 kB) free!
[1588493753] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2770156 kB) free!
[1588493787] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3484916 kB) free.
[1588493852] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 615 processes
[1588493873] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (418296 kB) free!
[1588493972] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 359 processes
[1588493993] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17793820 kB) free!
[1588494140] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(177) < 103.255.146.164 [03/May/2020:13:43:24 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588494260] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588494387] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3205068 kB) free!
[1588494507] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3151536 kB) free!
[1588494593] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7407628 kB) free!
[1588494627] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (2870924 kB) free!
[1588494655] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 10.3% (2316208 kB) free!
[1588494775] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1603020 kB) free!
[1588494895] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (1222068 kB) free!
[1588495193] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21720448 kB) free!
[1588495265] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4403968 kB) free!
[1588495793] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28235828 kB) free.
[1588495827] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1111400 kB) free!
[1588496393] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5544080 kB) free!
[1588496427] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.5% (2355880 kB) free!
[1588496513] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33346104 kB) free.
[1588497113] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8747920 kB) free!
[1588497233] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4806168 kB) free!
[1588497354] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24648448 kB) free!
[1588497954] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (11574696 kB) free!
[1588498227] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (1480612 kB) free!
[1588498460] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(117) < 185.255.8.58 [03/May/2020:15:00:03 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1588498496] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16256860 kB) free.
[1588498554] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20715000 kB) free!
[1588498580] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588498827] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2934900 kB) free!
[1588499154] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31016424 kB) free.
[1588499754] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7962040 kB) free!
[1588499781] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(778) < 103.255.146.164 [03/May/2020:15:18:41 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588499874] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4079988 kB) free!
[1588499875] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.10, 2.64, 3.90
[1588499901] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588499994] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31670456 kB) free.
[1588500594] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8864068 kB) free!
[1588500714] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4798104 kB) free!
[1588500834] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31500856 kB) free.
[1588501101] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(173) < 103.255.146.164 [03/May/2020:15:45:39 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588501124] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1588501221] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588501244] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 351 processes
[1588501364] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1588501434] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8972740 kB) free!
[1588501554] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4862408 kB) free!
[1588501674] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31250764 kB) free.
[1588501827] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (914784 kB) free!
[1588501887] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1588502007] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1588502274] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5301824 kB) free!
[1588502275] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.87, 6.63, 5.34
[1588502394] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (637000 kB) free!
[1588502395] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.12, 7.30, 5.75
[1588502514] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7449324 kB) free!
[1588502515] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.13, 7.37, 5.95
[1588502804] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1588502924] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1588503164] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1588503621] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(723) < 103.255.146.164 [03/May/2020:16:26:05 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588503741] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588504227] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3173904 kB) free!
[1588504315] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.27, 3.71, 4.65
[1588504420] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 82.9% (20324956 kB) free.
[1588504663] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 26.4% (6478520 kB) free.
[1588505514] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24871208 kB) free!
[1588506027] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1396540 kB) free!
[1588506114] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392804 kB) free!
[1588506627] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.2% (2313680 kB) free!
[1588506715] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.05, 5.00, 5.08
[1588506836] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.87, 5.40, 5.22
[1588506896] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1588506956] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.28, 5.87, 5.40
[1588507314] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27968144 kB) free.
[1588507810] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9638856 kB) free!
[1588507827] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151692 kB) free!
[1588507914] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (858028 kB) free!
[1588507930] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9514732 kB) free!
[1588508034] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393576 kB) free!
[1588508050] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (9395328 kB) free!
[1588508096] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.6% (3055980 kB) free!
[1588508154] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41313636 kB) free.
[1588508217] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.3% (2756700 kB) free!
[1588508337] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.7% (2405100 kB) free!
[1588508754] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5566104 kB) free!
[1588508874] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1263876 kB) free!
[1588508937] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (765520 kB) free!
[1588508994] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398596 kB) free!
[1588509027] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (2479460 kB) free!
[1588509141] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(761) < 103.255.146.164 [03/May/2020:17:58:03 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588509262] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588509537] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1588509537] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16118168 kB) free.
[1588509627] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3283020 kB) free.
[1588510195] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34578492 kB) free.
[1588510461] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(815) < 103.255.146.164 [03/May/2020:18:23:40 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588510556] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.93, 4.60, 4.89
[1588510581] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588510795] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7619424 kB) free!
[1588510827] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1176712 kB) free!
[1588510915] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38653688 kB) free.
[1588510947] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (1385312 kB) free!
[1588511068] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1682744 kB) free!
[1588511156] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.15, 5.41, 5.16
[1588511276] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.22, 5.38, 5.19
[1588511396] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.62, 5.45, 5.24
[1588511515] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13291912 kB) free!
[1588511635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8430580 kB) free!
[1588511755] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3867364 kB) free!
[1588512868] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (440992 kB) free!
[1588512955] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25350720 kB) free!
[1588512981] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(794) < 103.255.146.164 [03/May/2020:18:57:18 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588513101] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588513450] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4492272 kB) free!
[1588513555] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379400 kB) free!
[1588514068] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.1% (2452004 kB) free!
[1588514396] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.23, 3.50, 4.78
[1588514668] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.0% (3898760 kB) free.
[1588514755] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26737400 kB) free.
[1588515268] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.2% (2139260 kB) free!
[1588515355] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1171060 kB) free!
[1588515388] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.9% (2255932 kB) free!
[1588515475] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38513236 kB) free.
[1588515508] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.1% (2460136 kB) free!
[1588516075] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6922892 kB) free!
[1588516108] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3494756 kB) free.
[1588516195] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1820836 kB) free!
[1588516315] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392460 kB) free!
[1588516420] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4782236 kB) free!
[1588516540] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.0% (4646896 kB) free!
[1588516660] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.3% (4479140 kB) free!
[1588517308] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3257388 kB) free!
[1588517396] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.98, 6.55, 5.51
[1588517428] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1160824 kB) free!
[1588517515] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25335948 kB) free!
[1588517516] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.47, 5.40, 5.22
[1588517548] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1361992 kB) free!
[1588517636] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.66, 4.71, 4.98
[1588518115] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (463976 kB) free!
[1588518148] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (1987680 kB) free!
[1588518460] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2434716 kB) free!
[1588519137] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.1% (3849160 kB) free!
[1588519257] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3656308 kB) free!
[1588519315] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17271096 kB) free!
[1588519377] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.5% (3480308 kB) free!
[1588519915] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29937940 kB) free.
[1588519977] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16114812 kB) free.
[1588520515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1217156 kB) free!
[1588520548] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.4% (4133928 kB) free.
[1588520635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394172 kB) free!
[1588520755] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37145024 kB) free.
[1588521148] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2015264 kB) free!
[1588521268] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.1% (1962104 kB) free!
[1588521355] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11413440 kB) free!
[1588521388] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3766608 kB) free.
[1588521475] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6917788 kB) free!
[1588521595] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2887088 kB) free!
[1588522795] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (25469768 kB) free!
[1588523395] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2535760 kB) free!
[1588523636] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1588523788] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3084676 kB) free!
[1588523908] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3002240 kB) free!
[1588524028] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2819716 kB) free!
[1588524596] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23576544 kB) free!
[1588525196] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1209540 kB) free!
[1588525229] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3517492 kB) free.
[1588525829] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3155824 kB) free!
[1588525949] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2914132 kB) free!
[1588526069] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2693664 kB) free!
[1588526396] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21139076 kB) free!
[1588526669] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3359896 kB) free.
[1588526996] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34041124 kB) free.
[1588527596] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9512304 kB) free!
[1588527716] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5931576 kB) free!
[1588527836] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34563432 kB) free.
[1588527869] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.0% (2434524 kB) free!
[1588527989] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.7% (2219804 kB) free!
[1588528109] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.8% (1910952 kB) free!
[1588528436] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11132836 kB) free!
[1588528556] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7522948 kB) free!
[1588528676] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34452488 kB) free.
[1588529276] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11195792 kB) free!
[1588529396] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7543712 kB) free!
[1588529516] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34405028 kB) free.
[1588530116] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11770132 kB) free!
[1588530237] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8187076 kB) free!
[1588530357] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33824440 kB) free.