[1585594113] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13780512 kB) free!
[1585595407] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185004 kB) free!
[1585595467] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173784 kB) free!
[1585595527] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (177600 kB) free!
[1585595587] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (192668 kB) free!
[1585595887] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.3% (17470240 kB) free.
[1585596513] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28260608 kB) free.
[1585597113] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8864024 kB) free!
[1585597219] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585597233] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27351440 kB) free.
[1585598433] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9360760 kB) free!
[1585598553] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.9% (26203116 kB) free!
[1585598673] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21712736 kB) free!
[1585599566] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.7% (10702648 kB) free.
[1585600499] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(110) < 92.42.104.166 [31/Mar/2020:01:55:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585600620] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585601073] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26456748 kB) free.
[1585601366] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186280 kB) free!
[1585601486] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182612 kB) free!
[1585601606] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176888 kB) free!
[1585602346] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 71.4% (35189164 kB) free.
[1585602638] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.2% (44468920 kB) free.
[1585603844] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(336) < 2020/03/31 03:00:22 [error] 22135#0: *79130155 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585603964] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(24) < 2020/03/31 03:01:32 [error] 22131#0: *79130458 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585604070] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(283) < 182.18.144.246 [31/Mar/2020:03:01:32 +0530] TCP 502 0 0 60.000 "10.147.212.61:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1585604084] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585604191] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585604219] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1986) < 203.212.70.170 [31/Mar/2020:03:02:32 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1585604339] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585604657] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4951924 kB) free!
[1585607673] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23436312 kB) free!
[1585607793] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17310368 kB) free!
[1585607914] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12612792 kB) free!
[1585609114] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13434112 kB) free!
[1585609714] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12529820 kB) free!
[1585610314] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16198832 kB) free!
[1585611514] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29932816 kB) free.
[1585612406] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.6% (20877800 kB) free.
[1585612423] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.0% (20095492 kB) free.
[1585612552] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.6% (20004284 kB) free.
[1585613314] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25107804 kB) free!
[1585613434] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.1% (21255512 kB) free!
[1585613554] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17467976 kB) free!
[1585614154] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29821760 kB) free.
[1585615354] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26349032 kB) free!
[1585615475] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (22661448 kB) free!
[1585615595] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19070608 kB) free!
[1585616795] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29548724 kB) free.
[1585617995] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26104740 kB) free!
[1585618115] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.0% (22433188 kB) free!
[1585618235] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18796688 kB) free!
[1585619435] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28992404 kB) free.
[1585620635] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24608200 kB) free!
[1585620755] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.0% (21126400 kB) free!
[1585620875] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17576260 kB) free!
[1585621475] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29069208 kB) free.
[1585622676] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24895888 kB) free!
[1585622796] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20873208 kB) free!
[1585622916] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16763688 kB) free!
[1585624116] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27227480 kB) free.
[1585624629] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 326 processes
[1585624716] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25672128 kB) free!
[1585624749] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1585624836] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.6% (21885808 kB) free!
[1585624956] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18012272 kB) free!
[1585625693] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.53, 6.23, 5.24
[1585625813] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.36, 5.90, 5.25
[1585625933] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.62, 6.36, 5.48
[1585626156] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12525976 kB) free!
[1585627356] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24849572 kB) free!
[1585628023] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4853620 kB) free!
[1585628143] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4761728 kB) free!
[1585628255] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 18.3% (4479800 kB) free!
[1585628263] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4663516 kB) free!
[1585628376] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 18.3% (4478876 kB) free!
[1585628496] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 18.3% (4476104 kB) free!
[1585628933] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.74, 4.62, 4.96
[1585629533] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.44, 4.87, 5.23
[1585629653] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.96, 4.52, 5.05
[1585629773] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.03, 4.86, 5.12
[1585630552] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189228 kB) free!
[1585630576] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1585630672] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169192 kB) free!
[1585630696] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1585630792] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187204 kB) free!
[1585630806] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1585630816] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1585630926] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1585630956] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12740880 kB) free!
[1585631046] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1585631263] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2418480 kB) free!
[1585631416] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1585632247] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585632295] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 69.7% (17071932 kB) free.
[1585632367] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585632487] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1585633063] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.9% (2675032 kB) free!
[1585633149] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1585633270] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1585633663] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1585633663] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2202404 kB) free!
[1585635156] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19462076 kB) free!
[1585635443] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.2% (2731560 kB) free!
[1585635563] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.2% (2509388 kB) free!
[1585635683] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.2% (2282948 kB) free!
[1585636087] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1585636283] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 67.9% (15248680 kB) free.
[1585637007] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4776656 kB) free!
[1585637069] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4598472 kB) free!
[1585637127] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4661292 kB) free!
[1585637189] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (4357256 kB) free!
[1585637247] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4545400 kB) free!
[1585637309] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.8% (22915260 kB) free.
[1585637392] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 39.0% (9561816 kB) free.
[1585638773] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.19, 4.04, 4.90
[1585639373] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.27, 4.89, 5.08
[1585639493] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.04, 4.63, 4.96
[1585639647] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2299136 kB) free!
[1585640093] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.82, 6.33, 5.68
[1585640213] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.07, 5.89, 5.60
[1585640333] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1585640333] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.82, 5.12, 5.35
[1585641069] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1585641189] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1585642389] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1585642509] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1585642629] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 334 processes
[1585642687] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585642792] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185736 kB) free!
[1585642807] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1585642816] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1585642912] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (216804 kB) free!
[1585642927] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1585642936] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1585642956] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1585643032] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166920 kB) free!
[1585643094] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179280 kB) free!
[1585643214] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187092 kB) free!
[1585643229] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1585643334] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192216 kB) free!
[1585645063] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585645283] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3308584 kB) free!
[1585645403] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.2% (2730188 kB) free!
[1585645523] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.3% (2305352 kB) free!
[1585645956] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (602948 kB) free!
[1585646123] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1585646123] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (318528 kB) free!
[1585646909] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (4094852 kB) free!
[1585647030] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3942656 kB) free!
[1585647150] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 15.2% (3749812 kB) free!
[1585647324] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.4% (15356484 kB) free.
[1585647540] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(125) < 92.42.104.166 [31/Mar/2020:15:00:13 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585647660] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585647756] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14243368 kB) free!
[1585648350] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.9% (22926632 kB) free.
[1585650020] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (2914228 kB) free!
[1585650140] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2982172 kB) free!
[1585650260] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (2975992 kB) free!
[1585650860] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 32.2% (5239724 kB) free.
[1585650946] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9658536 kB) free!
[1585651066] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9545656 kB) free!
[1585651133] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585651135] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 41.0% (10039592 kB) free.
[1585651186] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (9442176 kB) free!
[1585651327] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1585651356] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26656148 kB) free.
[1585651447] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1585651485] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2020/03/31 16:08:53 [error] 22123#0: *79391441 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585651567] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1585651605] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585651838] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9524780 kB) free!
[1585651854] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1585651956] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5414708 kB) free!
[1585651958] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9408408 kB) free!
[1585651974] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1585652076] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.0% (25089512 kB) free!
[1585652078] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9300536 kB) free!
[1585652094] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1585652167] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1585652196] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1585652196] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20786100 kB) free!
[1585653996] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5397144 kB) free!
[1585654567] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1585654688] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1585654808] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1585655032] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (3678652 kB) free!
[1585655632] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (213516 kB) free!
[1585655788] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 7852.32, 4965.15, 2155.10
[1585655829] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1585655908] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2665.80, 4392.72, 2318.10
[1585655949] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1585656008] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1585656008] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1585656028] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 361.38, 2939.99, 2037.29
[1585656396] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14039428 kB) free!
[1585656860] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3106952 kB) free!
[1585656879] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3783264 kB) free!
[1585656924] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4423960 kB) free!
[1585656980] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.7% (4344264 kB) free.
[1585657005] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/03/31 17:44:42 [error] 22135#0: *79426758 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/ HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/", host: "altsmsc.tesync.net", referrer: "http://spanel.sms24hours.com:6080/MIS/index.jsp"
[1585657044] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (3855316 kB) free!
[1585657125] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585657136] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.3% (3508984 kB) free!
[1585657165] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (3314404 kB) free!
[1585657186] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4301384 kB) free!
[1585657228] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.54, 55.16, 561.02
[1585657256] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 12.5% (3055712 kB) free!
[1585657376] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 11.1% (2721008 kB) free!
[1585657596] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394256 kB) free!
[1585657725] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2020/03/31 17:56:13 [error] 22131#0: *79431359 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/ HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/", host: "altsmsc.tesync.net", referrer: "http://spanel.sms24hours.com:6080/MIS/index.jsp"
[1585657765] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2046564 kB) free!
[1585657845] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585657976] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2106240 kB) free!
[1585658365] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.1% (15281312 kB) free.
[1585658445] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/03/31 18:08:37 [error] 22133#0: *79436408 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/ HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/", host: "altsmsc.tesync.net", referrer: "http://spanel.sms24hours.com:6080/MIS/index.jsp"
[1585658496] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.9% (2433076 kB) free!
[1585658565] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585659166] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/03/31 18:22:37 [error] 22132#0: *79441808 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1585659232] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.5% (3056084 kB) free!
[1585659286] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585659396] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (18939200 kB) free!
[1585659886] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/03/31 18:33:55 [error] 22135#0: *79446061 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1585660006] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585660432] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1585660432] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1561688 kB) free!
[1585660596] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405112 kB) free!
[1585660606] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2020/03/31 18:46:03 [error] 22134#0: *79450561 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/dashboard"
[1585660726] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/03/31 18:48:40 [error] 22132#0: *79451481 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/ HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/", host: "altsmsc.tesync.net"
[1585660846] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585661139] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 14 users currently logged in
[1585661259] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 13 users currently logged in
[1585661379] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 13 users currently logged in
[1585661447] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2020/03/31 18:59:50 [error] 22135#0: *79455071 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net"
[1585661567] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/03/31 19:01:24 [error] 22129#0: *79455572 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net"
[1585661687] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2020/03/31 19:04:46 [error] 22123#0: *79456737 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net"
[1585661694] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1585662287] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585662380] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3139504 kB) free!
[1585662501] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3353884 kB) free.
[1585662661] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(14023) < 103.255.146.164 [31/Mar/2020:19:18:55 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585662781] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585662887] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/03/31 19:23:08 [error] 22132#0: *79462692 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/ HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/", host: "altsmsc.tesync.net", referrer: "http://spanel.sms24hours.com:6080/MIS/index.jsp"
[1585663007] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585663100] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.6% (2380244 kB) free!
[1585663220] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2536132 kB) free!
[1585663340] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (2718176 kB) free!
[1585663733] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 24.25, 45.82, 31.40
[1585663940] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.6% (4161576 kB) free.
[1585664196] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33984520 kB) free.
[1585664365] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3587344 kB) free!
[1585664485] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (3011576 kB) free!
[1585664540] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1042852 kB) free!
[1585664605] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.3% (2531676 kB) free!
[1585664660] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (1447056 kB) free!
[1585664780] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (1816896 kB) free!
[1585664796] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6177264 kB) free!
[1585664808] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/31 19:50:29 [error] 22123#0: *79472318 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1585664916] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2218640 kB) free!
[1585664928] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585664934] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.95, 13.53, 18.57
[1585665036] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31976168 kB) free.
[1585665205] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1081832 kB) free!
[1585665380] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1585665380] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3398248 kB) free.
[1585665636] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10468176 kB) free!
[1585665757] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6576652 kB) free!
[1585665877] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31303812 kB) free.
[1585665980] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (1632300 kB) free!
[1585666100] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (1353376 kB) free!
[1585666128] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/31 20:17:20 [error] 22133#0: *79480782 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1585666220] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.2% (4580084 kB) free.
[1585666248] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585666477] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3949096 kB) free!
[1585666597] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391700 kB) free!
[1585666717] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (18927540 kB) free!
[1585666749] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1585666808] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585666869] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1585666950] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 10.3% (2534956 kB) free!
[1585667070] SERVICE ALERT: db1-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (1840272 kB) free!
[1585667190] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1227912 kB) free!
[1585667317] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4887944 kB) free!
[1585667334] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.24, 27.06, 21.54
[1585667417] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.4% (24203700 kB) free!
[1585667420] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (1386240 kB) free!
[1585667537] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.7% (37892300 kB) free.
[1585667540] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.8% (2087616 kB) free!
[1585667660] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (2659224 kB) free!
[1585667917] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17235992 kB) free!
[1585667934] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1585667934] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 16.34, 15.71, 17.59
[1585668049] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/31 20:47:20 [error] 22135#0: *79490168 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1585668169] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585668260] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (875932 kB) free!
[1585668517] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403604 kB) free!
[1585668737] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (25133024 kB) free!
[1585668805] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.8% (15457676 kB) free.
[1585668857] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.0% (33013272 kB) free.
[1585668861] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3172596 kB) free!
[1585669461] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1360196 kB) free!
[1585669717] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19618856 kB) free!
[1585670061] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3806408 kB) free.
[1585670190] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.3% (22792236 kB) free.
[1585670317] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448220 kB) free!
[1585670379] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1585670658] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.4% (24198144 kB) free!
[1585670778] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.3% (33335216 kB) free.
[1585671232] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585671517] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25206244 kB) free!
[1585671534] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.12, 21.95, 17.99
[1585671861] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (2260612 kB) free!
[1585671981] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2597060 kB) free!
[1585672101] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3456476 kB) free.
[1585672117] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390756 kB) free!
[1585672134] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.24, 18.33, 18.49
[1585672214] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.28556.714> (10.147.212.158:35436 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1585672334] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1585672717] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14646340 kB) free!
[1585673918] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (531712 kB) free!
[1585674518] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20600568 kB) free!
[1585675101] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3237932 kB) free!
[1585675118] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1279116 kB) free!
[1585675134] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.16, 4.14, 4.90
[1585675221] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.3% (4115308 kB) free.
[1585676318] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21778280 kB) free!
[1585676335] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.22, 6.08, 5.40
[1585676455] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.13, 6.03, 5.46
[1585676575] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.86, 5.95, 5.50
[1585676918] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1848596 kB) free!
[1585677175] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.88, 3.71, 4.60
[1585677205] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1585677621] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3026204 kB) free!
[1585677741] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.6% (3836860 kB) free.
[1585678118] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25219272 kB) free!
[1585678718] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5469044 kB) free!
[1585678941] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2849476 kB) free!
[1585679061] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3760128 kB) free.
[1585679318] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16050288 kB) free!