[1584815998] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1134532 kB) free!
[1584816118] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.4% (1687192 kB) free!
[1584816144] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584816169] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188344 kB) free!
[1584816189] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196636 kB) free!
[1584816238] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (2370168 kB) free!
[1584816290] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184580 kB) free!
[1584816309] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188716 kB) free!
[1584816410] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177300 kB) free!
[1584816429] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201908 kB) free!
[1584816838] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1584816838] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.9% (4042620 kB) free.
[1584816961] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(43) < 203.212.70.170 [22/Mar/2020:00:24:00 +0530] TCP 502 0 0 0.001 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "0.001, 0.000"
[1584817081] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584817439] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (2675180 kB) free!
[1584817537] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (190116 kB) free!
[1584817559] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3266400 kB) free.
[1584817682] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(69) < 172.105.57.57 [22/Mar/2020:00:34:50 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584817802] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584818403] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1465) < 203.212.70.170 [22/Mar/2020:00:48:57 +0530] TCP 502 0 0 183.006 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "3.006, 60.000, 60.000, 60.000"
[1584818523] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584818676] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2020/03/22 00:50:48 [error] 20713#0: *75164079 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.144.246, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:593/17, bytes from/to upstream:17/593
[1584818759] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (939356 kB) free!
[1584818796] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584818879] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1155564 kB) free!
[1584818971] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.1% (4682204 kB) free!
[1584818974] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40248 MB (17.90% inode=100%):
[1584818999] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1625044 kB) free!
[1584819016] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584819034] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40248 MB (17.90% inode=100%):
[1584819091] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.8% (4595912 kB) free!
[1584819094] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 95108 MB (42.31% inode=100%):
[1584819107] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584819211] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.4% (4513652 kB) free!
[1584819337] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 80.5% (19727340 kB) free.
[1584819343] SERVICE FLAPPING ALERT: qservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1584819599] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.4% (5427160 kB) free.
[1584819672] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584819724] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(530) < 182.18.144.246 [22/Mar/2020:01:08:12 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584819844] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584820719] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34184716 kB) free.
[1584820859] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (241624 kB) free!
[1584820979] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214056 kB) free!
[1584821099] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (252600 kB) free!
[1584821319] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3313404 kB) free!
[1584821439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (381320 kB) free!
[1584821530] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8378172 kB) free!
[1584821559] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387900 kB) free!
[1584821645] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(68) < 203.212.70.170 [22/Mar/2020:01:40:37 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584821650] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (416780 kB) free!
[1584821684] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.58, 6.28, 5.26
[1584821765] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584821770] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3625268 kB) free!
[1584821796] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2020/03/22 01:38:42 [error] 20705#0: *75182062 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584821804] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.20, 6.35, 5.39
[1584821917] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584821924] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.64, 6.01, 5.38
[1584822759] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16409076 kB) free!
[1584822811] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2430672 kB) free!
[1584822899] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 28.1% (13857500 kB) free.
[1584822965] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(842) < 172.105.59.173 [22/Mar/2020:02:06:02 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584823086] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(79) < 203.212.70.170 [22/Mar/2020:02:08:04 +0530] TCP 502 0 0 120.001 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1584823124] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.80, 4.09, 4.64
[1584823199] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (1601200 kB) free!
[1584823206] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584823320] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (1277376 kB) free!
[1584823359] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399068 kB) free!
[1584823440] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1112856 kB) free!
[1584823942] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8296676 kB) free!
[1584824036] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584824040] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 36.3% (5897464 kB) free.
[1584824062] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2877312 kB) free!
[1584824182] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240516 kB) free!
[1584824324] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.46, 5.77, 5.20
[1584824406] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9378) < 172.105.57.57 [22/Mar/2020:02:26:37 +0530] TCP 502 0 0 120.002 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1584824416] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 19 users currently logged in
[1584824444] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.22, 6.24, 5.44
[1584824527] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584824564] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.57, 6.71, 5.71
[1584824770] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.4% (15092336 kB) free!
[1584825759] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31327332 kB) free.
[1584825943] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 78.9% (38897772 kB) free.
[1584825970] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (9873672 kB) free!
[1584826359] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (385736 kB) free!
[1584826479] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374512 kB) free!
[1584826570] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1584826570] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.6% (13984020 kB) free!
[1584826582] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 73.5% (36205136 kB) free.
[1584826599] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399140 kB) free!
[1584826926] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1738) < 49.50.86.152 [22/Mar/2020:03:08:52 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584827046] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584828371] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6542660 kB) free!
[1584828764] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1584828764] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.08, 4.35, 4.93
[1584828846] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(543) < 203.212.70.170 [22/Mar/2020:03:44:04 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584828967] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(3) < 203.212.70.170 [22/Mar/2020:03:45:54 +0530] TCP 502 0 0 166.109 "10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "31.069, 15.040, 60.000, 60.000"
[1584828971] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.0% (18514440 kB) free!
[1584828999] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27321356 kB) free.
[1584829010] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.9% (9767932 kB) free.
[1584829087] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(4) < 172.105.59.173 [22/Mar/2020:03:47:05 +0530] TCP 502 0 0 180.001 "10.147.212.154:3001, 10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 60.000, 0.000"
[1584829116] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/03/22 03:43:43 [error] 20709#0: *75234673 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584829236] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584829571] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (10270580 kB) free!
[1584829599] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (385828 kB) free!
[1584829681] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584829687] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584829719] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385228 kB) free!
[1584829802] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584829839] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391980 kB) free!
[1584830171] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.3% (21508136 kB) free!
[1584830887] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(693) < 203.212.70.170 [22/Mar/2020:04:15:17 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584831007] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584831558] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44891 MB (19% inode=99%):
[1584831678] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44891 MB (19% inode=99%):
[1584831798] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44891 MB (19% inode=99%):
[1584832239] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29171376 kB) free.
[1584832807] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(823) < 172.105.59.173 [22/Mar/2020:04:47:32 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584832839] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (394132 kB) free!
[1584832927] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584832959] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405804 kB) free!
[1584833079] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398688 kB) free!
[1584833598] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51602 MB (22% inode=99%):
[1584833810] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179228 kB) free!
[1584833930] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178940 kB) free!
[1584834050] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195628 kB) free!
[1584835237] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(19) < 2020/03/22 05:24:44 [error] 20711#0: *75275990 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1584835328] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6033) < 172.105.59.173 [22/Mar/2020:05:27:19 +0530] TCP 502 0 0 180.001 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 60.000, 0.000"
[1584835357] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584835440] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584835448] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584835479] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20856796 kB) free!
[1584835634] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 51.1% (12531840 kB) free.
[1584835872] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 31.8% (7840144 kB) free.
[1584836079] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380352 kB) free!
[1584836564] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1584837371] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (9982248 kB) free!
[1584837673] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.9% (3675572 kB) free!
[1584837793] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196576 kB) free!
[1584837849] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(919) < 172.105.59.173 [22/Mar/2020:06:07:06 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1584837913] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186276 kB) free!
[1584837969] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584838571] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 44.9% (59123612 kB) free.
[1584839679] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40996532 kB) free.
[1584839769] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(966) < 182.18.144.246 [22/Mar/2020:06:45:58 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584839771] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.5% (13869656 kB) free!
[1584839889] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(18) < 103.13.97.241 [22/Mar/2020:06:47:44 +0530] TCP 502 0 0 120.001 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1584839892] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10470580 kB) free!
[1584840010] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(8) < 172.105.59.173 [22/Mar/2020:06:48:34 +0530] TCP 502 0 0 240.000 "10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 60.000"
[1584840012] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 38.4% (50630124 kB) free.
[1584840198] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43286 MB (19% inode=99%):
[1584840279] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3889384 kB) free!
[1584840318] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43286 MB (19% inode=99%):
[1584840399] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384228 kB) free!
[1584840438] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43286 MB (19% inode=99%):
[1584840519] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392980 kB) free!
[1584840610] SERVICE FLAPPING ALERT: primary-server;Nginx Access Log Status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584840610] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584840612] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25010556 kB) free!
[1584840732] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.7% (20727120 kB) free!
[1584840852] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.4% (16343068 kB) free!
[1584841453] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 39.2% (51727280 kB) free.
[1584842410] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2337) < 203.212.70.170 [22/Mar/2020:07:29:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584842530] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(27) < 49.50.86.152 [22/Mar/2020:07:31:37 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584842651] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584842919] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26353400 kB) free!
[1584843250] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(61) < 172.105.59.173 [22/Mar/2020:07:43:29 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584843370] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(3) < 172.105.59.173 [22/Mar/2020:07:44:19 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584843490] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584843519] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392468 kB) free!
[1584844090] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [22/Mar/2020:07:54:57 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584844210] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584844983] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1584845103] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1584845410] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2838) < 172.105.59.173 [22/Mar/2020:08:14:28 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1584845416] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 24 users currently logged in
[1584845530] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584845536] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 24 users currently logged in
[1584845656] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 24 users currently logged in
[1584845834] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (327348 kB) free!
[1584845919] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29948204 kB) free.
[1584845954] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194476 kB) free!
[1584846074] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175700 kB) free!
[1584846129] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.1% (21345556 kB) free.
[1584846152] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 57.7% (14235332 kB) free.
[1584846178] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 74.7% (18309596 kB) free.
[1584846199] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 30.5% (15009044 kB) free.
[1584846213] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1584846220] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.8% (17595560 kB) free.
[1584846236] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584846256] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 26 users currently logged in
[1584846333] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1584846356] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1584846438] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 79920 MB (35% inode=99%):
[1584846453] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1584846476] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1584846519] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3353960 kB) free!
[1584846639] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397848 kB) free!
[1584846734] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192060 kB) free!
[1584846759] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385544 kB) free!
[1584846760] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188700 kB) free!
[1584846802] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8683416 kB) free!
[1584846821] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172464 kB) free!
[1584846854] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193848 kB) free!
[1584846880] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166356 kB) free!
[1584846922] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8392252 kB) free!
[1584846941] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192592 kB) free!
[1584846975] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 79.0% (19359320 kB) free.
[1584847000] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189636 kB) free!
[1584847042] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.3% (8008276 kB) free!
[1584847061] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1584847061] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186188 kB) free!
[1584847302] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 71.09, 38.79, 18.09
[1584847351] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 26 users currently logged in
[1584847359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29738988 kB) free.
[1584847389] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1489 processes
[1584847422] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 26.77, 37.10, 20.30
[1584847542] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 56.48, 38.24, 22.37
[1584847559] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1584847574] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194960 kB) free!
[1584847601] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1584847601] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.1% (21489904 kB) free.
[1584847677] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1584847694] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184316 kB) free!
[1584847797] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21678560 kB) free.
[1584847930] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1557) < 103.13.97.241 [22/Mar/2020:08:55:29 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1584847959] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3161080 kB) free!
[1584848050] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584848079] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (448192 kB) free!
[1584848137] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 11.0% (2697668 kB) free!
[1584848142] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 22.87, 17.73, 18.30
[1584848197] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 11.0% (2690640 kB) free!
[1584848199] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446368 kB) free!
[1584848261] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.4% (19457328 kB) free.
[1584848267] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584848277] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1584848328] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 25.0% (6138696 kB) free.
[1584848346] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (206276 kB) free!
[1584848397] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1584848466] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (229412 kB) free!
[1584848517] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1584848586] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (220984 kB) free!
[1584848742] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 19.49, 48.51, 41.47
[1584848799] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31109368 kB) free.
[1584848801] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169560 kB) free!
[1584848846] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (535008 kB) free!
[1584848862] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 310.14, 117.07, 65.33
[1584848921] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182820 kB) free!
[1584848983] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 528.08, 263.31, 125.23
[1584849041] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191976 kB) free!
[1584849227] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3392400 kB) free!
[1584849287] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1801988 kB) free!
[1584849301] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.2% (21513792 kB) free.
[1584849330] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 58.7% (14390220 kB) free.
[1584849347] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.9% (222208 kB) free!
[1584849350] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1584849350] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 51.7% (25480444 kB) free.
[1584849400] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3032264 kB) free!
[1584849407] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (199744 kB) free!
[1584849520] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386352 kB) free!
[1584849640] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403180 kB) free!
[1584849707] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 18.2% (4449524 kB) free!
[1584849764] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.54, 6.93, 5.56
[1584849884] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.86, 6.16, 5.45
[1584849911] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179204 kB) free!
[1584849952] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7372672 kB) free!
[1584850004] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.24, 5.25, 5.21
[1584850007] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 8.0% (1957108 kB) free!
[1584850031] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179028 kB) free!
[1584850072] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.9% (6343052 kB) free!
[1584850151] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177784 kB) free!
[1584850159] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584850192] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (5660812 kB) free!
[1584850286] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 198070 MB (88% inode=99%):
[1584850450] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(86) < 203.212.70.170 [22/Mar/2020:09:39:29 +0530] TCP 502 0 0 151.063 "10.147.212.154:3001, 10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "31.063, 60.000, 60.000, 0.000"
[1584850570] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584850792] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (2416740 kB) free!
[1584850801] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1584850921] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584851983] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1584851983] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 35.06, 29.63, 37.55
[1584852040] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23048136 kB) free!
[1584852117] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1584852117] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584852189] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1501 processes
[1584852309] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1502 processes
[1584852404] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.87, 4.02, 4.46
[1584852430] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1503 processes
[1584852583] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 8.31, 9.81, 22.67
[1584852640] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388956 kB) free!
[1584852853] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584852970] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(87) < 203.212.70.170 [22/Mar/2020:10:20:36 +0530] TCP 502 0 0 211.082 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "31.082, 60.000, 60.000, 60.000"
[1584853030] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1499 processes
[1584853090] SERVICE FLAPPING ALERT: primary-server;Nginx Access Log Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584853090] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584854890] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [22/Mar/2020:10:57:45 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584855010] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1297) < 172.105.59.173 [22/Mar/2020:10:59:15 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584855040] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22790828 kB) free!
[1584855130] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(1) < 172.105.59.173 [22/Mar/2020:11:00:15 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1584855304] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1584855424] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1584855640] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379348 kB) free!
[1584855731] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584855751] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 31 users currently logged in
[1584855871] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 31 users currently logged in
[1584855991] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 31 users currently logged in
[1584856241] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31562444 kB) free.
[1584856841] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3540360 kB) free!
[1584856961] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (413848 kB) free!
[1584857081] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396224 kB) free!
[1584857182] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9150500 kB) free!
[1584857302] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8937160 kB) free!
[1584857422] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.7% (8747048 kB) free!
[1584857531] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2558) < 172.105.59.173 [22/Mar/2020:11:41:50 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584857651] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(76) < 203.212.70.170 [22/Mar/2020:11:43:09 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584857681] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34497336 kB) free.
[1584857772] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584858117] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1584858237] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 340 processes
[1584858281] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7584860 kB) free!
[1584858357] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1584858371] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(119) < 172.105.59.173 [22/Mar/2020:11:49:51 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584858401] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2488448 kB) free!
[1584858491] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584858521] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396672 kB) free!
[1584858958] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1584859061] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584859091] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(38) < 103.255.146.164 [22/Mar/2020:12:01:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584859151] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584859211] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584859822] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4707392 kB) free!
[1584859957] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/03/22 12:22:18 [error] 20712#0: *75455380 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584860077] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(24) < 2020/03/22 12:23:55 [error] 20705#0: *75455769 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584860132] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4595108 kB) free!
[1584860197] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584860252] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4419408 kB) free!
[1584860372] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4307376 kB) free!
[1584860396] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.3% (4240420 kB) free!
[1584860411] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1897) < 103.255.146.164 [22/Mar/2020:12:26:26 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584860516] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.8% (4114544 kB) free!
[1584860531] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584860584] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.9% (20958248 kB) free.
[1584860609] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14198936 kB) free!
[1584860632] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 36.5% (18005404 kB) free.
[1584860636] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.8% (4106980 kB) free!
[1584860659] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 41.5% (20459028 kB) free.
[1584860758] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584860797] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(859) < 2020/03/22 12:35:57 [error] 20703#0: *75461955 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1584860918] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584861129] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(404) < 103.16.101.72 [22/Mar/2020:12:34:28 +0530] TCP 502 0 0 60.001 "10.147.212.49:4000, 10.147.212.37:4000, 10.147.212.61:4000, backend3" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 60.000, 0.000"
[1584861131] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2557) < 103.255.146.164 [22/Mar/2020:12:42:01 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584861191] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194772 kB) free!
[1584861212] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28214152 kB) free.
[1584861241] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.8% (6815420 kB) free!
[1584861249] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584861252] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584861311] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182268 kB) free!
[1584861361] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.2% (6002960 kB) free!
[1584861431] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187836 kB) free!
[1584861481] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (5759272 kB) free!
[1584861812] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7728192 kB) free!
[1584861851] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(779) < 49.50.86.152 [22/Mar/2020:12:53:58 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584861932] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3585020 kB) free!
[1584861971] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(59) < 203.212.70.170 [22/Mar/2020:12:55:49 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584862052] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379592 kB) free!
[1584862081] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4867136 kB) free!
[1584862091] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584862380] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (4577100 kB) free!
[1584862500] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4408240 kB) free!
[1584862620] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (4275680 kB) free!
[1584862772] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1689872 kB) free!
[1584863036] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1929164 kB) free!
[1584863291] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(628) < 172.105.59.173 [22/Mar/2020:13:17:30 +0530] TCP 502 0 0 240.000 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 60.000"
[1584863411] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584863852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31992612 kB) free.
[1584863984] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584864452] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7540288 kB) free!
[1584864572] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3130460 kB) free!
[1584864693] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386056 kB) free!
[1584864958] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 70.0% (34488172 kB) free.
[1584864991] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.2% (21858024 kB) free.
[1584865020] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2393032 kB) free!
[1584865023] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23684284 kB) free!
[1584865031] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27962532 kB) free.
[1584865211] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(916) < 182.18.144.246 [22/Mar/2020:13:50:10 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584865331] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(962) < 203.212.70.170 [22/Mar/2020:13:52:00 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584865452] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(18) < 49.50.86.152 [22/Mar/2020:13:52:51 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1584866052] SERVICE FLAPPING ALERT: primary-server;Nginx Access Log Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1584866052] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584866231] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7060916 kB) free!
[1584866351] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (2973360 kB) free!
[1584866471] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384696 kB) free!
[1584866652] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(33) < 103.255.146.164 [22/Mar/2020:14:14:12 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584866761] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (241388 kB) free!
[1584866772] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(3) < 103.255.146.164 [22/Mar/2020:14:14:12 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584866791] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 30 users currently logged in
[1584866881] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (556800 kB) free!
[1584866892] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584867001] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (526928 kB) free!
[1584867391] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 31 users currently logged in
[1584867493] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1592) < 203.212.70.170 [22/Mar/2020:14:28:12 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1584867511] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 31 users currently logged in
[1584867613] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(14) < 203.212.70.170 [22/Mar/2020:14:28:20 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584867631] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 31 users currently logged in
[1584867733] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584868061] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4869288 kB) free!
[1584868181] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4801668 kB) free!
[1584868272] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29426032 kB) free.
[1584868301] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (4718328 kB) free!
[1584868872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (1999436 kB) free!
[1584868992] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400928 kB) free!
[1584869112] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (373556 kB) free!
[1584869533] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1801) < 203.212.70.170 [22/Mar/2020:15:00:42 +0530] TCP 502 0 0 180.000 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1584869653] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584869712] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17011920 kB) free!
[1584870253] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1130) < 172.105.59.173 [22/Mar/2020:15:10:24 +0530] TCP 502 0 0 180.001 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 60.000, 0.000"
[1584870312] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402820 kB) free!
[1584870373] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584870631] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 30 users currently logged in
[1584870958] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584870974] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(37) < 103.255.146.164 [22/Mar/2020:15:18:22 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584871094] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584871831] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 31 users currently logged in
[1584871951] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 30 users currently logged in
[1584872712] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38932920 kB) free.
[1584873101] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1266100 kB) free!
[1584873312] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (399444 kB) free!
[1584873343] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9732352 kB) free!
[1584873432] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396928 kB) free!
[1584873463] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9642016 kB) free!
[1584873494] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(49) < 103.255.146.164 [22/Mar/2020:16:01:03 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584873552] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398720 kB) free!
[1584873583] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9544184 kB) free!
[1584873614] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584874214] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1288) < 172.105.59.173 [22/Mar/2020:16:19:55 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1584874255] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 788 processes
[1584874328] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.5029.262> (10.147.212.158:37954 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1584874335] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(15) < 203.212.70.170 [22/Mar/2020:16:20:54 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1584874375] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 337 processes
[1584874448] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584874455] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584874717] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/03/22 16:18:42 [error] 20703#0: *75553017 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584874838] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584875952] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39410156 kB) free.
[1584876002] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584876255] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(90) < 203.212.70.170 [22/Mar/2020:16:53:20 +0530] TCP 502 0 0 151.076 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.154:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 31.076, 60.000, 60.000"
[1584876375] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(71) < 172.105.59.173 [22/Mar/2020:16:54:20 +0530] TCP 502 0 0 211.075 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "31.071, 60.002, 60.002, 60.000"
[1584876495] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584876552] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12366308 kB) free!
[1584876672] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7543052 kB) free!
[1584876792] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2934040 kB) free!
[1584877695] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(69) < 103.255.146.164 [22/Mar/2020:17:15:26 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584877816] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584878415] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2645) < 172.105.59.173 [22/Mar/2020:17:28:40 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1584878535] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584879135] SERVICE FLAPPING ALERT: primary-server;Nginx Access Log Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584879192] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22347668 kB) free!
[1584879736] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(80) < 103.255.146.164 [22/Mar/2020:17:43:50 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584879792] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405940 kB) free!
[1584879856] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584880456] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1878) < 172.105.59.173 [22/Mar/2020:18:04:04 +0530] TCP 502 0 0 120.002 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1584880576] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584880784] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8.33, 50.33, 37.18
[1584880904] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.25, 35.94, 33.50
[1584881024] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.32, 26.78, 30.46
[1584881177] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 172.105.57.57 [22/Mar/2020:18:15:15 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584881297] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584881624] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.74, 8.99, 19.01
[1584882192] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (14008792 kB) free!
[1584882498] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(985) < 203.212.70.170 [22/Mar/2020:18:37:56 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584882618] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 203.212.70.170 [22/Mar/2020:18:39:13 +0530] TCP 502 0 0 60.001 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1584882637] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(55) < 2020/03/22 18:38:01 [error] 20713#0: *75602710 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1584882738] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(1) < 203.212.70.170 [22/Mar/2020:18:40:20 +0530] TCP 502 0 0 180.003 "10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.154:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 60.000, 60.001, 60.001"
[1584882757] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584882792] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401160 kB) free!
[1584883183] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4909220 kB) free!
[1584883392] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32327916 kB) free.
[1584883602] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.0% (4669040 kB) free!
[1584883723] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.7% (4579516 kB) free!
[1584883843] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.3% (4486148 kB) free!
[1584883939] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584883992] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1178616 kB) free!
[1584884112] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (376420 kB) free!
[1584884232] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401296 kB) free!
[1584884539] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3572) < 203.212.70.170 [22/Mar/2020:19:11:49 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584884659] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(2) < 172.105.59.173 [22/Mar/2020:19:12:39 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584884779] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584884832] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33372520 kB) free.
[1584885432] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2466516 kB) free!
[1584885552] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396072 kB) free!
[1584885672] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380964 kB) free!
[1584886272] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33504216 kB) free.
[1584886843] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1998032 kB) free!
[1584886872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3300940 kB) free!
[1584886992] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374592 kB) free!
[1584887112] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392804 kB) free!
[1584887179] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(325) < 182.18.144.246 [22/Mar/2020:19:48:59 +0530] TCP 502 0 0 120.002 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1584887300] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584887552] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 31 users currently logged in
[1584887672] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 31 users currently logged in
[1584887712] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34251128 kB) free.
[1584887792] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 30 users currently logged in
[1584888312] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2871080 kB) free!
[1584888432] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405836 kB) free!
[1584888553] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (426372 kB) free!
[1584889024] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1584889099] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3685) < 203.212.70.170 [22/Mar/2020:20:26:39 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584889144] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1584889153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32743644 kB) free.
[1584889220] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584889357] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(380) < 2020/03/22 20:25:40 [error] 20704#0: *75642524 upstream timed out (110: Connection timed out) while connecting to upstream, client: 182.18.144.246, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584889450] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(260) < 103.16.101.72 [22/Mar/2020:20:25:39 +0530] TCP 502 0 0 60.000 "10.147.212.49:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1584889477] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584889571] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584889753] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2780432 kB) free!
[1584889873] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406296 kB) free!
[1584889993] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411948 kB) free!
[1584890593] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26377904 kB) free!
[1584891020] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(19) < 203.212.70.170 [22/Mar/2020:21:00:16 +0530] TCP 502 0 0 60.002 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 0.001, 60.000"
[1584891140] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(2747) < 203.212.70.170 [22/Mar/2020:21:02:18 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584891193] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (7199512 kB) free!
[1584891260] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(1198) < 172.105.59.173 [22/Mar/2020:21:04:15 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584891277] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(598) < 2020/03/22 21:04:05 [error] 20705#0: *75657704 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584891371] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(445) < 103.16.101.72 [22/Mar/2020:21:02:54 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1584891398] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584891492] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584892030] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1501 processes
[1584892151] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1495 processes
[1584892461] SERVICE FLAPPING ALERT: primary-server;Nginx Access Log Status;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1584892461] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584893061] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 172.105.57.57 [22/Mar/2020:21:34:21 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584893181] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1541) < 49.50.86.152 [22/Mar/2020:21:35:22 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584893302] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(41) < 203.212.70.170 [22/Mar/2020:21:38:21 +0530] TCP 502 0 0 180.001 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 60.000, 0.000"
[1584893593] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20500196 kB) free!
[1584894193] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400088 kB) free!
[1584894793] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32877680 kB) free.
[1584895102] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584895393] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7594160 kB) free!
[1584895513] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2954496 kB) free!
[1584895633] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375840 kB) free!
[1584895703] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3139) < 203.212.70.170 [22/Mar/2020:22:12:54 +0530] TCP 502 0 0 180.001 "10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 60.000, 0.000"
[1584895823] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584896234] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27433544 kB) free.
[1584896834] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (406908 kB) free!
[1584896954] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382824 kB) free!
[1584897024] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(13) < 203.212.70.170 [22/Mar/2020:22:40:07 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1584897074] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394904 kB) free!
[1584897144] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(99) < 172.105.59.173 [22/Mar/2020:22:41:07 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584897264] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584897674] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28357212 kB) free.
[1584898274] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (386620 kB) free!
[1584898394] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408292 kB) free!
[1584898514] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407832 kB) free!
[1584899114] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23916264 kB) free!
[1584899714] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383480 kB) free!
[1584900314] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28014172 kB) free.
[1584900864] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1862) < 49.50.86.152 [22/Mar/2020:23:41:20 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1584900914] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1800040 kB) free!
[1584900984] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584901034] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406936 kB) free!
[1584901154] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395820 kB) free!
[1584901754] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21422064 kB) free!