[1581359505] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2020/02/11 00:01:00 [error] 10220#0: *51876273 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581359625] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581359732] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.9% (29537348 kB) free.
[1581362077] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.0% (9874968 kB) free.
[1581363159] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4347628 kB) free!
[1581363499] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26474348 kB) free.
[1581363877] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190296 kB) free!
[1581363997] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171024 kB) free!
[1581364099] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25777832 kB) free!
[1581364117] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176816 kB) free!
[1581364219] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26466536 kB) free.
[1581366390] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.7% (10222780 kB) free.
[1581366825] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/11 01:58:17 [error] 10220#0: *51902697 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1581366945] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581367819] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (26002732 kB) free!
[1581367939] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26432516 kB) free.
[1581368467] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229020 kB) free!
[1581368539] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26422760 kB) free!
[1581368587] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.8% (41799368 kB) free.
[1581368659] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 20.0% (26420636 kB) free!
[1581368779] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 20.0% (26420832 kB) free!
[1581369380] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1581369380] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (27926512 kB) free.
[1581370469] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(189) < 103.16.101.72 [11/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581370545] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(231) < 2020/02/11 03:00:22 [error] 10215#0: *51918027 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581370590] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581370665] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581370675] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2688) < 203.212.70.170 [11/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581370795] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581377790] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179656 kB) free!
[1581377911] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 31.7% (7766364 kB) free.
[1581378510] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190740 kB) free!
[1581378517] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.9% (11326368 kB) free.
[1581378630] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186696 kB) free!
[1581378721] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.6% (20732136 kB) free.
[1581378750] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183716 kB) free!
[1581380780] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581386989] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39721672 kB) free.
[1581387021] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.0% (22057276 kB) free.
[1581387046] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.1% (14161572 kB) free.
[1581388795] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(68) < 203.212.70.170 [11/Feb/2020:08:03:09 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581388915] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581389449] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (3212700 kB) free!
[1581389570] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (3210072 kB) free!
[1581389690] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (2927348 kB) free!
[1581391915] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 49.50.86.152 [11/Feb/2020:08:57:37 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581392035] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581392343] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.0% (14145820 kB) free.
[1581395317] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4579088 kB) free!
[1581395437] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4441196 kB) free!
[1581395557] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4145596 kB) free!
[1581396158] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1730584 kB) free!
[1581399288] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.6% (21606816 kB) free.
[1581400728] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 337 processes
[1581400848] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1581400968] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 337 processes
[1581401568] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1581402772] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 721 processes
[1581402835] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(86) < 49.50.86.152 [11/Feb/2020:12:01:58 +0530] TCP 502 0 0 66.014 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "3.006, 3.008, 60.000, 0.000"
[1581402892] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 718 processes
[1581402955] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581402989] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(34) < 103.16.101.72 [11/Feb/2020:12:01:14 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581403012] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 326 processes
[1581403066] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(295) < 2020/02/11 12:02:08 [error] 10219#0: *52091320 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581403110] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581403165] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1581403186] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581403197] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9822152 kB) free!
[1581403285] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1581403317] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (8043556 kB) free!
[1581403405] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1581403437] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (8523256 kB) free!
[1581404005] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 694 processes
[1581404280] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39701212 kB) free.
[1581404386] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/02/11 12:19:50 [error] 10222#0: *52102010 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /favicon.ico HTTP/1.1", host: "qservercluster1.com:8091"
[1581404506] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581404699] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170328 kB) free!
[1581404756] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(697) < 103.255.146.164 [11/Feb/2020:12:29:37 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1581404819] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191500 kB) free!
[1581404876] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581404940] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189204 kB) free!
[1581405168] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1581405288] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1581405878] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21790432 kB) free.
[1581406321] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.9% (4399304 kB) free!
[1581406441] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.5% (4287552 kB) free!
[1581406473] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21626508 kB) free.
[1581407112] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1581407232] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1581409490] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (4370076 kB) free!
[1581409610] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4244976 kB) free!
[1581409730] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4159268 kB) free!
[1581410876] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(712) < 103.255.146.164 [11/Feb/2020:14:12:46 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581410930] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1581410930] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (748760 kB) free!
[1581410996] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581411628] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4708720 kB) free!
[1581411748] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 21.3% (5219176 kB) free.
[1581411788] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9008232 kB) free!
[1581411888] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1581411908] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8845648 kB) free!
[1581412008] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1581412013] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21953264 kB) free.
[1581412028] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.7% (8704040 kB) free!
[1581412048] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.7% (44182488 kB) free.
[1581412196] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 92.42.104.166 [11/Feb/2020:14:32:34 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581412316] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581412348] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (4495588 kB) free!
[1581412468] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (4370608 kB) free!
[1581412588] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (4259664 kB) free!
[1581412833] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4115072 kB) free!
[1581412953] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (3977596 kB) free!
[1581413073] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.2% (3851100 kB) free!
[1581413516] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(168) < 92.42.104.166 [11/Feb/2020:15:00:14 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581413636] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581413674] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.6% (15619464 kB) free.
[1581414988] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2186268 kB) free!
[1581416214] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1581416334] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1581416454] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1581416522] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4855472 kB) free!
[1581416642] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4801272 kB) free!
[1581416678] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (1095812 kB) free!
[1581416762] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4752000 kB) free!
[1581416798] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 46.7% (11440268 kB) free.
[1581417054] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1581417106] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2020/02/11 16:00:42 [error] 10216#0: *52217319 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581417226] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581417362] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1786812 kB) free!
[1581417654] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1581417775] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1581417895] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1581417998] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169080 kB) free!
[1581418018] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169380 kB) free!
[1581418118] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190660 kB) free!
[1581418138] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199764 kB) free!
[1581418238] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183868 kB) free!
[1581418258] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204428 kB) free!
[1581418688] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9411644 kB) free!
[1581418808] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8981816 kB) free!
[1581418823] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1581418929] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (8339396 kB) free!
[1581418944] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1581419808] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1581419928] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1581421596] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.6% (21459196 kB) free.
[1581421624] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.2% (21511572 kB) free.
[1581421662] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.5% (39670116 kB) free.
[1581421696] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.1% (21846668 kB) free.
[1581421737] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.7% (21993932 kB) free.
[1581421937] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1581422057] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1581422095] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1581422095] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1581422177] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1581422828] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (1969812 kB) free!
[1581422832] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1581422948] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190436 kB) free!
[1581422952] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581423068] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185156 kB) free!
[1581423295] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1581423415] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1581423535] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1581423874] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3965420 kB) free!
[1581423994] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.9% (3798092 kB) free!
[1581424114] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.1% (3618640 kB) free!
[1581424137] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4809052 kB) free!
[1581424257] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4756552 kB) free!
[1581424377] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1581424377] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (4179560 kB) free!
[1581424714] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1581424714] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.5% (15605168 kB) free.
[1581424840] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22741292 kB) free!
[1581424960] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40415124 kB) free.
[1581425160] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1581425187] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9819120 kB) free!
[1581425265] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 7.5% (1845892 kB) free!
[1581425280] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1581425307] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (9753188 kB) free!
[1581425385] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 7.4% (1822068 kB) free!
[1581425427] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (9673280 kB) free!
[1581425505] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 7.4% (1803884 kB) free!
[1581426480] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1581426598] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.1% (21490452 kB) free.
[1581426600] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1581426645] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 89.7% (21989624 kB) free.
[1581426720] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1581426977] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1581427128] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1581427248] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1581427320] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1581427368] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1581427440] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1581427560] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1581427977] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (215028 kB) free!
[1581428027] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(81) < 2020/02/11 19:01:53 [error] 10220#0: *52324496 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581428036] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [11/Feb/2020:19:01:47 +0530] TCP 502 0 0 67.017 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "7.016, 60.001, 0.000"
[1581428147] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581428156] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581428160] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1581430135] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1581431747] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(88) < 2020/02/11 20:01:49 [error] 10210#0: *52348785 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581431756] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 172.105.59.173 [11/Feb/2020:20:00:48 +0530] TCP 502 0 0 0.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "0.000, 0.000"
[1581431868] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581431876] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581433209] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1581436115] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581436377] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1581437667] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.8% (21765964 kB) free.
[1581439669] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9802116 kB) free!
[1581439789] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9803516 kB) free!
[1581439910] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9803068 kB) free!
[1581442227] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4938056 kB) free!
[1581442752] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581442872] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1581442992] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1581443593] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581444332] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8454324 kB) free!
[1581444452] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (1490260 kB) free!
[1581444572] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240116 kB) free!
[1581445580] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12376196 kB) free!
[1581445700] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12394180 kB) free!