[1552501832] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.6% (5725044 kB) free!
[1552501832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552501952] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1552501952] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4030796 kB) free!
[1552501952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552502552] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7460556 kB) free!
[1552502762] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.2% (13420184 kB) free.
[1552503152] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552503152] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11679360 kB) free.
[1552503962] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8116752 kB) free!
[1552504082] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 15.5% (7632196 kB) free!
[1552504202] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.4% (7088180 kB) free!
[1552504352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552504352] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7107608 kB) free!
[1552504472] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552504472] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.8% (5806616 kB) free!
[1552504592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552504592] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4173456 kB) free!
[1552504802] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229852 kB) free!
[1552505192] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.5% (7615652 kB) free!
[1552505714] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29752 MB (14% inode=99%):
[1552505792] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552505792] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11677648 kB) free.
[1552506392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552506393] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3735372 kB) free!
[1552506513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11677912 kB) free.
[1552506513] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552507112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552507113] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6012632 kB) free!
[1552507232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552507233] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4500988 kB) free!
[1552507352] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552507353] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11678716 kB) free.
[1552507952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552507953] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6519712 kB) free!
[1552508072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552508073] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5228692 kB) free!
[1552508192] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552508193] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11672976 kB) free.
[1552508792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552508793] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7060764 kB) free!
[1552508913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552508913] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5766184 kB) free!
[1552509033] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552509033] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4095004 kB) free!
[1552509633] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7526600 kB) free!
[1552510233] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552510234] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11671896 kB) free.
[1552510833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552510953] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552511434] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7159452 kB) free!
[1552511553] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552511554] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5852832 kB) free!
[1552511673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552511674] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4246004 kB) free!
[1552511794] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552512274] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7606140 kB) free!
[1552512393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552512513] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552512633] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552512874] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11676880 kB) free.
[1552513233] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552513353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552513473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552513474] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3846532 kB) free!
[1552513594] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11670928 kB) free.
[1552514194] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6408904 kB) free!
[1552514314] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5168824 kB) free!
[1552514434] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11670820 kB) free.
[1552514674] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552514714] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21449 MB (10% inode=99%):
[1552515035] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7075564 kB) free!
[1552515155] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.8% (5815624 kB) free!
[1552515274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552515275] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4193740 kB) free!
[1552515394] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552515875] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.5% (7652568 kB) free!
[1552515994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552516114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552516234] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552516475] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11673676 kB) free.
[1552516834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552516954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552517074] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552517075] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3821860 kB) free!
[1552517195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11667716 kB) free.
[1552517795] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6113488 kB) free!
[1552517915] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4699084 kB) free!
[1552518035] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11668416 kB) free.
[1552518274] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552518635] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6682696 kB) free!
[1552518756] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.0% (5439408 kB) free!
[1552518874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552518876] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3598644 kB) free!
[1552518994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552519202] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.0% (8863432 kB) free!
[1552519476] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7266200 kB) free!
[1552519514] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 33262 MB (15% inode=99%):
[1552519594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552519714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552519834] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552520076] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11672048 kB) free.
[1552520434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552520555] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552520675] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552521276] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7178660 kB) free!
[1552521396] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5897296 kB) free!
[1552521516] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4362040 kB) free!
[1552521602] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4751384 kB) free!
[1552522116] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7750212 kB) free!
[1552522475] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552522716] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11669608 kB) free.
[1552523075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552523195] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552523315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552523316] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4126796 kB) free!
[1552523436] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11664812 kB) free.
[1552524037] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6301368 kB) free!
[1552524157] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.2% (5012648 kB) free!
[1552524277] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11667564 kB) free.
[1552524516] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552524877] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6802568 kB) free!
[1552524997] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5556712 kB) free!
[1552525116] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552525117] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3793072 kB) free!
[1552525236] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552525717] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7377324 kB) free!
[1552525836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552525956] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552526076] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552526317] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11202864 kB) free.
[1552526676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552526796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552526917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552526917] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (5994896 kB) free!
[1552527037] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (4118008 kB) free!
[1552527157] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (2406372 kB) free!
[1552527228] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55772 MB (24% inode=99%):
[1552527517] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552527758] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7429636 kB) free!
[1552528117] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552528237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552528357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552528358] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (255420 kB) free!
[1552529557] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552529558] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.0% (8398408 kB) free!
[1552530002] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.2% (8977444 kB) free!
[1552530758] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552530758] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (712604 kB) free!
[1552530878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552530998] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552531598] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552532198] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552532318] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552532402] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1303276 kB) free!
[1552532558] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6027236 kB) free!
[1552532918] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552533038] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552533914] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22392 MB (10% inode=99%):
[1552534202] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.5% (8108072 kB) free!
[1552534238] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552534358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552534478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552534514] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 33406 MB (16% inode=99%):
[1552536002] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1552536002] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4631832 kB) free!
[1552536158] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1173460 kB) free!
[1552536278] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552537373] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46924 MB (20% inode=99%):
[1552537426] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 23.88, 15.70, 7.06
[1552537479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552537493] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46923 MB (20% inode=99%):
[1552537546] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.83, 11.08, 6.41
[1552537599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552537613] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46902 MB (20% inode=99%):
[1552537666] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.70, 7.84, 5.79
[1552537719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552538266] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.24, 2.21, 3.69
[1552539314] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22311 MB (10% inode=99%):
[1552539436] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.7% (2231544 kB) free!
[1552539519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552539758] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5068484 kB) free!
[1552539851] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.93, 5.76, 5.06
[1552539971] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.74, 5.73, 5.13
[1552540037] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1583856 kB) free!
[1552540092] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.80, 5.73, 5.20
[1552540119] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552540239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552540359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552541892] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.00, 4.36, 4.76
[1552542759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552543358] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (315368 kB) free!
[1552543359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552543479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552543599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552545492] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.56, 5.39, 5.10
[1552545612] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.27, 5.39, 5.14
[1552545732] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.03, 5.33, 5.15
[1552545999] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552546203] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5106140 kB) free!
[1552546332] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1552546332] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.05, 3.69, 4.43
[1552546514] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25041 MB (12% inode=99%):
[1552546599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552546666] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 46.1% (22724776 kB) free.
[1552546711] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.0% (12804692 kB) free.
[1552546719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552546763] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22118368 kB) free.
[1552546808] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.4% (23141520 kB) free.
[1552546836] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 331 processes
[1552546839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552546850] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 95.8% (7044796 kB) free.
[1552546896] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 95.9% (7054856 kB) free.
[1552546902] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1552547035] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.1% (22563084 kB) free.
[1552547044] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.6% (7029340 kB) free.
[1552547082] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.5% (7027220 kB) free.
[1552547120] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 96.1% (7069088 kB) free.
[1552547154] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.3% (15003620 kB) free.
[1552547196] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.6% (14891536 kB) free.
[1552547436] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 3 warn out of 337 processes
[1552547502] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 323 processes
[1552547622] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 1 warn out of 323 processes
[1552547742] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 322 processes
[1552547875] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5514232 kB) free!
[1552547911] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7398100 kB) free!
[1552547995] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.8% (13712548 kB) free.
[1552548031] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.3% (6554544 kB) free!
[1552548151] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.3% (5557528 kB) free!
[1552548314] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17358 MB (8% inode=99%):
[1552548636] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 340 processes
[1552548752] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212608 kB) free!
[1552549239] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552549795] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6040076 kB) free!
[1552549915] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.6% (14092160 kB) free.
[1552550439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552550559] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552551115] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6363900 kB) free!
[1552551235] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.5% (14025340 kB) free.
[1552551759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552551879] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552552435] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5870212 kB) free!
[1552552555] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.9% (13728216 kB) free.
[1552553079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552553199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552553319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552553436] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1552553436] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 336 processes
[1552553755] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5543524 kB) free!
[1552553875] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13362936 kB) free.
[1552554475] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1552555119] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552555675] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6373980 kB) free!
[1552555796] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.3% (13930900 kB) free.
[1552556320] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552556440] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552556979] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 261 processes
[1552556995] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6143476 kB) free!
[1552557040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552557099] SERVICE ALERT: sig3-vm2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 264 processes
[1552557115] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.5% (14024120 kB) free.
[1552557160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552557219] SERVICE ALERT: sig3-vm2;CPU_Procs;OK;HARD;1;CPU OK: 271 processes
[1552557280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552557732] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1552558315] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5907288 kB) free!
[1552558435] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.9% (13752640 kB) free.
[1552559080] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552559635] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5503412 kB) free!
[1552559755] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13370364 kB) free.
[1552560152] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1552560280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552560400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552560520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552560636] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1552560956] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5219124 kB) free!
[1552561076] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13207788 kB) free.
[1552562163] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (242356 kB) free!
[1552562276] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5628372 kB) free!
[1552562283] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164032 kB) free!
[1552562304] SERVICE ALERT: webservercluster1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: / 10338 MB (20% inode=98%):
[1552562320] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552562396] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13168740 kB) free.
[1552562403] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155680 kB) free!
[1552562424] SERVICE ALERT: webservercluster1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: / 10337 MB (20% inode=98%):
[1552562544] SERVICE ALERT: webservercluster1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: / 10337 MB (20% inode=98%):
[1552562920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552563041] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552563161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552563596] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5307276 kB) free!
[1552563716] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13105852 kB) free.
[1552564916] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5182092 kB) free!
[1552564961] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552565036] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13148056 kB) free.
[1552565561] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552565681] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552565801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552566237] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5366060 kB) free!
[1552566357] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13372688 kB) free.
[1552567557] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5184720 kB) free!
[1552567601] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552567677] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13153092 kB) free.
[1552568201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552568321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552568441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552568877] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5193800 kB) free!
[1552568997] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13096200 kB) free.
[1552570197] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5220152 kB) free!
[1552570241] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552570317] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13050376 kB) free.
[1552570666] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.47, 6.52, 5.01
[1552570786] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.88, 5.71, 4.92
[1552570842] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552570962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552571082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552571517] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5141060 kB) free!
[1552571637] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (12988676 kB) free.
[1552572837] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.7% (5267992 kB) free!
[1552572882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552572957] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12912992 kB) free.
[1552573482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552573602] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552573723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552574157] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4661864 kB) free!
[1552574277] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12698348 kB) free.
[1552575477] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4379236 kB) free!
[1552575523] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552575597] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12575536 kB) free.
[1552575612] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4836516 kB) free!
[1552575732] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4804164 kB) free!
[1552575853] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (4787884 kB) free!
[1552576123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552576243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552576363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552576453] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.8% (5106608 kB) free.
[1552576797] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5080484 kB) free!
[1552576917] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12686504 kB) free.
[1552578117] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5100868 kB) free!
[1552578163] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552578237] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12640860 kB) free.
[1552578763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552578883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552579003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552579437] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5072224 kB) free!
[1552579557] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12574616 kB) free.
[1552580213] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24721 MB (10% inode=99%):
[1552580757] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5353164 kB) free!
[1552580803] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552580878] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13369000 kB) free.
[1552581403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552581404] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (944064 kB) free!
[1552581413] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24739 MB (11% inode=99%):
[1552581524] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552581524] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 41.1% (6672944 kB) free.
[1552581571] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (211456 kB) free!
[1552581691] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (208680 kB) free!
[1552581811] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (208528 kB) free!
[1552581853] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4814616 kB) free!
[1552581973] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4773180 kB) free!
[1552582013] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24690 MB (10% inode=99%):
[1552582077] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5675536 kB) free!
[1552582093] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 26.5% (6486820 kB) free.
[1552582123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552582197] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13227660 kB) free.
[1552582243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552582363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552582952] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.2% (13387800 kB) free.
[1552583397] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5709880 kB) free!
[1552583517] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13344460 kB) free.
[1552583563] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552584117] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7958612 kB) free!
[1552584152] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.3% (8042956 kB) free!
[1552584163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552584237] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6682248 kB) free!
[1552584272] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.4% (7090268 kB) free!
[1552584283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552584357] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13349568 kB) free.
[1552584392] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.1% (5950468 kB) free!
[1552584404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552584957] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8500688 kB) free!
[1552584993] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1784852 kB) free!
[1552585003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552585078] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7209352 kB) free!
[1552585123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552585198] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5450244 kB) free!
[1552585243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552585843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552585963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552586083] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552586358] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1552586398] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13339952 kB) free.
[1552586478] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1552586500] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (159480 kB) free!
[1552586598] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1552586620] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155384 kB) free!
[1552586740] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155616 kB) free!
[1552587284] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552587598] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8714884 kB) free!
[1552587718] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7429960 kB) free!
[1552587838] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.9% (5853660 kB) free!
[1552587884] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552588004] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552588124] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157644 kB) free!