[1563733851] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (11182464 kB) free!
[1563733880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563733920] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11243328 kB) free.
[1563733972] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (18368540 kB) free.
[1563734101] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6458 MB (12% inode=97%):
[1563734139] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (510592 kB) free!
[1563734259] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187484 kB) free!
[1563734379] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173348 kB) free!
[1563734480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563734488] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4906632 kB) free!
[1563734519] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7204816 kB) free!
[1563734600] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563734608] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4897320 kB) free!
[1563734639] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11245352 kB) free.
[1563734712] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1168868 kB) free!
[1563734728] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4884824 kB) free!
[1563735171] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12149732 kB) free!
[1563735200] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563735239] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6914556 kB) free!
[1563735291] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (18992124 kB) free.
[1563735321] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563735359] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11234352 kB) free.
[1563735920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563735959] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6667444 kB) free!
[1563736040] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563736079] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11242364 kB) free.
[1563736491] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (12453288 kB) free!
[1563736611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (19093392 kB) free.
[1563736640] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563736679] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6721744 kB) free!
[1563736760] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563736799] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11234256 kB) free.
[1563737360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563737399] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6470596 kB) free!
[1563737480] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1563737519] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11232264 kB) free.
[1563737811] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (11802132 kB) free!
[1563737931] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (18250076 kB) free.
[1563738081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563738119] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6265956 kB) free!
[1563738201] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563738240] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11234220 kB) free.
[1563738801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563738839] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6569888 kB) free!
[1563738921] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563738959] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11225768 kB) free.
[1563739521] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563739559] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6835052 kB) free!
[1563739641] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563739651] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.62, 2.49, 13.76
[1563739679] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11230408 kB) free.
[1563739731] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (11924568 kB) free!
[1563739851] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (18413724 kB) free.
[1563740241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563740279] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7057796 kB) free!
[1563740361] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563740399] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11234688 kB) free.
[1563740962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563740999] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7317028 kB) free!
[1563741082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563741119] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11233412 kB) free.
[1563741451] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.80, 2.53, 4.12
[1563741682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563741719] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7456632 kB) free!
[1563741802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563741839] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11235768 kB) free.
[1563741922] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563742439] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7682800 kB) free!
[1563742522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563742560] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11223944 kB) free.
[1563742642] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563742980] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 12953.07, 5365.86, 2007.41
[1563743100] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14150.70, 8334.33, 3502.42
[1563743159] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7834676 kB) free!
[1563743220] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14312.52, 10320.70, 4816.20
[1563743242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563743279] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6148660 kB) free!
[1563743362] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563743399] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11234256 kB) free.
[1563743963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563743979] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (3004004 kB) free!
[1563743999] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6552196 kB) free!
[1563744051] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (12010136 kB) free!
[1563744083] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563744119] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11231040 kB) free.
[1563744171] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (16921900 kB) free.
[1563744683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563744719] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6840692 kB) free!
[1563744771] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (9956036 kB) free!
[1563744803] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563744839] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11233156 kB) free.
[1563744891] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (16071196 kB) free.
[1563745403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563745439] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6943852 kB) free!
[1563745523] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563745559] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11222140 kB) free.
[1563746123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563746159] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6902356 kB) free!
[1563746243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563746279] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11231428 kB) free.
[1563746631] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4785564 kB) free!
[1563746691] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (11453316 kB) free!
[1563746811] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (16971984 kB) free.
[1563746844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563746879] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6966988 kB) free!
[1563746964] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563747000] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11231364 kB) free.
[1563747411] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12252928 kB) free!
[1563747531] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (16987620 kB) free.
[1563747564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563747599] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6710940 kB) free!
[1563747684] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563747719] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11234120 kB) free.
[1563748284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563748319] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6789080 kB) free!
[1563748404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563748439] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11231288 kB) free.
[1563748503] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8559.68, 2396.70, 819.88
[1563748623] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 13479.98, 6317.79, 2447.96
[1563748743] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14142.86, 8941.09, 3878.48
[1563749004] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563749039] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6798944 kB) free!
[1563749124] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563749159] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11223560 kB) free.
[1563749220] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 0.47, 17.31
[1563749725] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563749728] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.3% (5475188 kB) free.
[1563749759] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6838680 kB) free!
[1563749845] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563749879] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11231504 kB) free.
[1563750420] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.38, 0.36, 5.00
[1563750445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563750479] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6898452 kB) free!
[1563750565] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563750599] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11233756 kB) free.
[1563750774] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 517 processes
[1563750834] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 520 processes
[1563751165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563751200] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6934756 kB) free!
[1563751285] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563751320] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11232344 kB) free.
[1563751731] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (12892808 kB) free!
[1563751851] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (17094732 kB) free.
[1563751885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563751920] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6983180 kB) free!
[1563752005] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563752040] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11108200 kB) free.
[1563752451] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11642104 kB) free!
[1563752571] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (17053676 kB) free.
[1563752606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563752640] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9823980 kB) free!
[1563752726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563752760] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.3% (8517516 kB) free!
[1563752846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563752880] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7201432 kB) free!
[1563752979] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2462496 kB) free!
[1563753446] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563753480] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11205216 kB) free.
[1563754046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563754166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563754286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563754371] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12852612 kB) free!
[1563754491] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (9851196 kB) free!
[1563754611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (16467364 kB) free.
[1563754680] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8547688 kB) free!
[1563754743] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.15, 0.23, 19.92
[1563754800] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7221388 kB) free!
[1563754886] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563754886] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563754920] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11180024 kB) free.
[1563755211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (10784212 kB) free!
[1563755331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (16828352 kB) free.
[1563755486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563755607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563755727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563755931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (11930124 kB) free!
[1563756052] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (9939052 kB) free!
[1563756120] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9013256 kB) free!
[1563756172] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (16470968 kB) free.
[1563756240] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7726316 kB) free!
[1563756360] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5986244 kB) free!
[1563756543] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.31, 3.15
[1563756928] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4875772 kB) free!
[1563757048] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4867724 kB) free!
[1563757168] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4860208 kB) free!
[1563757734] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 523 processes
[1563757795] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 520 processes
[1563757972] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (10419020 kB) free!
[1563758092] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (16434880 kB) free.
[1563758160] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1563758160] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10516876 kB) free.
[1563758727] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563759292] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (10604892 kB) free!
[1563759412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (16640028 kB) free.
[1563759928] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563759960] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6743272 kB) free!
[1563760048] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563760080] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10974212 kB) free.
[1563761212] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11607912 kB) free!
[1563761332] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (17301840 kB) free.
[1563761932] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (9971296 kB) free!
[1563762052] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (15955508 kB) free.
[1563762448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563762568] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563763080] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7495608 kB) free!
[1563763168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563763200] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11152828 kB) free.
[1563763253] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11807064 kB) free!
[1563763289] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563763373] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (17208352 kB) free.
[1563763800] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7614716 kB) free!
[1563763888] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563763920] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11153668 kB) free.
[1563763973] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (12411776 kB) free!
[1563764008] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563764093] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (17151780 kB) free.
[1563764437] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (149120 kB) free!
[1563764520] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7685220 kB) free!
[1563764557] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (170016 kB) free!
[1563764608] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563764640] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11150948 kB) free.
[1563764677] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173056 kB) free!
[1563764693] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (12443496 kB) free!
[1563764728] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563764813] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (9299044 kB) free!
[1563764933] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (15938644 kB) free.
[1563765240] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7879404 kB) free!
[1563765277] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 70.7% (11489280 kB) free.
[1563765328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563765360] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6211932 kB) free!
[1563765448] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563765476] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4853872 kB) free!
[1563765481] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11150380 kB) free.
[1563765533] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10919416 kB) free!
[1563765596] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4823504 kB) free!
[1563765653] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (17011456 kB) free.
[1563765716] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (4792236 kB) free!
[1563766048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563766080] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6491220 kB) free!
[1563766169] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563766200] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11150360 kB) free.
[1563766253] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12357208 kB) free!
[1563766373] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (9378364 kB) free!
[1563766493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (16525440 kB) free.
[1563766768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563766779] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.3% (10440288 kB) free.
[1563766800] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6189108 kB) free!
[1563766888] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563766920] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11147556 kB) free.
[1563767093] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (9280688 kB) free!
[1563767213] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (16380288 kB) free.
[1563767488] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563767520] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7510024 kB) free!
[1563767608] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563767640] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (6029292 kB) free!
[1563767728] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563767760] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11133688 kB) free.
[1563767813] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (8518468 kB) free!
[1563767888] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.68, 7.18, 5.67
[1563767933] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (15491296 kB) free.
[1563768008] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.34, 7.74, 6.05
[1563768128] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.57, 7.19, 6.05
[1563768328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563768360] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7326228 kB) free!
[1563768448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563768480] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11147856 kB) free.
[1563768534] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (8095992 kB) free!
[1563768569] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563768654] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (14879444 kB) free.
[1563768728] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.95, 4.11, 4.96
[1563768851] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.4% (14855616 kB) free.
[1563768895] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 82.7% (20265516 kB) free.
[1563768939] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39742596 kB) free.
[1563768983] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 79.7% (19542392 kB) free.
[1563769080] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7647552 kB) free!
[1563769168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563769200] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11147532 kB) free.
[1563769254] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (6964916 kB) free!
[1563769288] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1563769288] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563769374] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (14236824 kB) free.
[1563769801] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7771328 kB) free!
[1563769888] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563769921] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5978444 kB) free!
[1563769928] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.14, 5.35, 5.13
[1563769974] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (8383620 kB) free!
[1563770008] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563770009] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (2436320 kB) free!
[1563770041] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563770041] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11144176 kB) free.
[1563770048] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.99, 5.41, 5.19
[1563770094] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (13638716 kB) free.
[1563770129] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (269616 kB) free!
[1563770168] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.37, 5.26, 5.15
[1563770249] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (236888 kB) free!
[1563770608] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563770641] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6133900 kB) free!
[1563770728] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563770761] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11135104 kB) free.
[1563770849] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 56.1% (13756212 kB) free.
[1563771294] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (8146888 kB) free!
[1563771328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563771361] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6628516 kB) free!
[1563771388] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4887236 kB) free!
[1563771414] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (14258648 kB) free.
[1563771449] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563771481] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11140068 kB) free.
[1563771508] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4756468 kB) free!
[1563771596] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 17 users currently logged in
[1563771628] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (4606924 kB) free!
[1563771878] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (3094024 kB) free!
[1563771900] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4867104 kB) free!
[1563771998] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (2867412 kB) free!
[1563772014] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (7878040 kB) free!
[1563772020] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.8% (4841192 kB) free!
[1563772118] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.3% (2641744 kB) free!
[1563772134] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (13740412 kB) free.
[1563772140] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.7% (4826432 kB) free!
[1563772620] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 11601.44, 9354.85, 4759.34
[1563772649] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4539600 kB) free!
[1563772718] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1563772718] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1292328 kB) free!
[1563772734] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (6914848 kB) free!
[1563772740] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3047.82, 7369.80, 4630.23
[1563772769] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (2208348 kB) free!
[1563772828] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1563772828] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2043588 kB) free!
[1563772854] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.9% (12400840 kB) free!
[1563772860] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 411.91, 4931.57, 4068.90
[1563772889] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (615732 kB) free!
[1563772974] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (10711900 kB) free!
[1563773382] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2815660 kB) free!
[1563773502] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 25.7% (4182156 kB) free.
[1563773574] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (1430168 kB) free!
[1563773848] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563773852] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1192.36, 542.42, 210.26
[1563773968] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563773972] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1262.25, 784.12, 339.10
[1563774088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563774089] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1563774089] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 39.4% (9662776 kB) free.
[1563774092] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1272.33, 946.33, 452.51
[1563774260] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (344020 kB) free!
[1563774380] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (160956 kB) free!
[1563774481] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9790952 kB) free!
[1563774500] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (186420 kB) free!
[1563774543] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8909.79, 2911.96, 1029.04
[1563774579] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173020 kB) free!
[1563774602] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8560688 kB) free!
[1563774663] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 11368.30, 5836.83, 2329.17
[1563774689] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3617152 kB) free!
[1563774699] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193068 kB) free!
[1563774702] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2533400 kB) free!
[1563774722] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7292344 kB) free!
[1563774774] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (11020212 kB) free!
[1563774783] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11699.68, 7793.82, 3471.61
[1563774809] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (2312156 kB) free!
[1563774819] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167652 kB) free!
[1563774822] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.6% (2377200 kB) free!
[1563774929] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (2910068 kB) free!
[1563774942] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.5% (2197400 kB) free!
[1563775374] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1563775374] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (2202916 kB) free!
[1563775542] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1323452 kB) free!
[1563776488] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563776522] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11124012 kB) free.
[1563776574] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (13054276 kB) free!
[1563776599] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (1372388 kB) free!
[1563776719] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.2% (1115988 kB) free!
[1563776750] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (266520 kB) free!
[1563776839] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.9% (878648 kB) free!
[1563776870] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216176 kB) free!
[1563776990] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237020 kB) free!
[1563777088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563777122] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9665880 kB) free!
[1563777174] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (3224560 kB) free!
[1563777208] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563777242] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8259308 kB) free!
[1563777278] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 555 processes
[1563777328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563777329] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 22.7% (5573692 kB) free.
[1563777342] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 64.7% (10521936 kB) free.
[1563777362] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6770264 kB) free!
[1563777398] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 550 processes
[1563777440] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6832440 kB) free.
[1563777692] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.06, 3.86, 15.29
[1563777774] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (9018440 kB) free!
[1563777928] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1563777928] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563777962] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1563777962] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10871888 kB) free.
[1563777999] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 556 processes
[1563778119] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 556 processes
[1563778240] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 559 processes
[1563778260] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.06, 2.23, 14.91
[1563778374] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (14240212 kB) free.
[1563778528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563778529] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (896616 kB) free!
[1563778562] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9327016 kB) free!
[1563778649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563778650] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (1105852 kB) free!
[1563778682] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (8001900 kB) free!
[1563778769] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563778770] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (523524 kB) free!
[1563778802] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6244792 kB) free!
[1563778840] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 482 processes
[1563778974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (4842776 kB) free!
[1563779094] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (10211108 kB) free!
[1563779214] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (10389136 kB) free!
[1563779492] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.81, 2.90, 4.73
[1563779732] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (13153428 kB) free!
[1563779814] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (576036 kB) free!
[1563779852] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 20.0% (13157952 kB) free!
[1563779970] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (3839760 kB) free!
[1563779972] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.8% (16994836 kB) free.
[1563780002] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11096048 kB) free.
[1563780060] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.46, 1.64, 3.62
[1563780414] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (9446472 kB) free!
[1563780569] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563780570] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 54.7% (13406620 kB) free.
[1563780626] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.53, 6.12, 5.42
[1563780746] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.44, 6.21, 5.55
[1563780783] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 0.56, 11.16
[1563780866] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.60, 6.24, 5.65
[1563781014] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (866408 kB) free!
[1563781169] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563781170] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.0% (2705980 kB) free!
[1563781172] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.5% (11515404 kB) free!
[1563781289] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563781290] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 54.2% (13286540 kB) free.
[1563781292] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.6% (11561088 kB) free!
[1563781412] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (15362028 kB) free.
[1563781466] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.11, 4.07, 4.89
[1563781614] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (7462620 kB) free!
[1563781890] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.5% (3067572 kB) free!
[1563781983] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.90, 0.62, 3.46
[1563782010] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (2396364 kB) free!
[1563782130] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1786132 kB) free!
[1563782215] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (13264140 kB) free.
[1563782341] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2413280 kB) free!
[1563782402] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6495552 kB) free!
[1563782522] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11102268 kB) free.
[1563782730] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3888100 kB) free!
[1563782815] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (5372640 kB) free!
[1563782935] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (578300 kB) free!
[1563783055] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (13019076 kB) free!
[1563783089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563783209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563783329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563783330] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 23.1% (5653252 kB) free.
[1563783628] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563783655] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1691928 kB) free!
[1563784040] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (278592 kB) free!
[1563784118] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563784160] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (114936 kB) free!
[1563784255] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (9292088 kB) free!
[1563784280] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (126088 kB) free!
[1563784855] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (14857568 kB) free.
[1563785012] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (12786932 kB) free!
[1563785129] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563785131] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4734708 kB) free!
[1563785132] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.2% (12623740 kB) free!
[1563785251] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4462424 kB) free!
[1563785252] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.1% (16539140 kB) free.
[1563785371] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 60.4% (14802464 kB) free.
[1563785455] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (7737020 kB) free!
[1563785522] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7718580 kB) free!
[1563785575] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (2926416 kB) free!
[1563785620] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 60.8% (15009376 kB) free.
[1563785643] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5977360 kB) free!
[1563785695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (13632960 kB) free.
[1563785729] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563785742] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2874492 kB) free!
[1563785763] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10529428 kB) free.
[1563785850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563785862] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.6% (2692112 kB) free!
[1563785970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563785971] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (212888 kB) free!
[1563785982] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.7% (2544548 kB) free!
[1563786091] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169120 kB) free!
[1563786211] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (269924 kB) free!
[1563786295] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (4610640 kB) free!
[1563786415] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (13976188 kB) free.
[1563786452] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (12764404 kB) free!
[1563786572] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (13295548 kB) free.
[1563786895] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.8% (21529436 kB) free.
[1563786915] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1563786941] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.6% (21238292 kB) free.
[1563786988] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1563786988] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.8% (6826388 kB) free.
[1563787015] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (7582928 kB) free!
[1563787035] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1563787052] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.9% (38901604 kB) free.
[1563787135] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (2650700 kB) free!
[1563787172] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (13155664 kB) free!
[1563787183] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1563787183] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1193424 kB) free!
[1563787250] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 47.5% (31225232 kB) free.
[1563787292] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (13958000 kB) free.
[1563787296] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.3% (13868020 kB) free.
[1563787394] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 62.4% (15293036 kB) free.
[1563787436] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 67.2% (10923280 kB) free.
[1563787443] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 315 processes
[1563787486] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.8% (14595276 kB) free.
[1563787563] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1563787568] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.13, 4.13, 4.64
[1563787653] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243164 kB) free!
[1563787770] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563787773] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228620 kB) free!
[1563787893] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226172 kB) free!
[1563788092] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172756 kB) free!
[1563788163] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7627700 kB) free!
[1563788212] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (151676 kB) free!
[1563788283] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11092008 kB) free.
[1563788333] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164484 kB) free!
[1563788601] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181732 kB) free!
[1563788721] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (221376 kB) free!
[1563788841] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (924960 kB) free!
[1563788883] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7759480 kB) free!
[1563788970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563789003] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6123572 kB) free!
[1563789090] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563789123] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10603304 kB) free.
[1563789210] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563789441] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2557200 kB) free!
[1563789692] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (12186800 kB) free!
[1563789812] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.8% (13001240 kB) free!
[1563789932] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (13280512 kB) free.
[1563790041] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170376 kB) free!
[1563790262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13057924 kB) free!
[1563790382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (8756640 kB) free!
[1563790410] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563790420] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4725696 kB) free!
[1563790502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (18915960 kB) free.
[1563790540] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4676060 kB) free!
[1563790548] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171444 kB) free!
[1563790568] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.58, 7.91, 5.66
[1563790641] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (3808716 kB) free!
[1563790660] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4627772 kB) free!
[1563790668] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166744 kB) free!
[1563790688] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.25, 9.93, 6.68
[1563790788] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166088 kB) free!
[1563790808] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.60, 10.96, 7.44
[1563790924] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6481608 kB) free!
[1563790990] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.92, 5.71, 5.04
[1563791044] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11086340 kB) free.
[1563791102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (10449324 kB) free!
[1563791110] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.64, 5.68, 5.11
[1563791222] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (21256328 kB) free.
[1563791230] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.43, 5.62, 5.16
[1563791241] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 62.6% (15349972 kB) free.
[1563791610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563791644] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6613420 kB) free!
[1563791730] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563791764] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563791764] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11077000 kB) free.
[1563791795] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1563791822] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (13134164 kB) free!
[1563791841] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (1181780 kB) free!
[1563791915] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1563791942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (8481484 kB) free!
[1563791961] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (238764 kB) free!
[1563792035] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1563792062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (19182020 kB) free.
[1563792081] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (290856 kB) free!
[1563792330] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563792332] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (12759968 kB) free!
[1563792364] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6376812 kB) free!
[1563792450] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563792452] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.6% (12867916 kB) free!
[1563792484] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11082344 kB) free.
[1563792572] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (13599120 kB) free.
[1563792635] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1563792662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (9391216 kB) free!
[1563792755] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1563792782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (20260684 kB) free.
[1563792875] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1563793050] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563793084] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6350748 kB) free!
[1563793171] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563793204] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11071764 kB) free.
[1563793281] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4399660 kB) free!
[1563793301] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 11.7% (1903044 kB) free!
[1563793382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (10989844 kB) free!
[1563793421] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (1489884 kB) free!
[1563793502] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (6101048 kB) free!
[1563793541] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1047160 kB) free!
[1563793622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (19357672 kB) free.
[1563793660] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2311660 kB) free!
[1563793770] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563793804] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6231280 kB) free!
[1563793882] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.7% (5061644 kB) free.
[1563793890] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563793924] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11080896 kB) free.
[1563794036] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3030292 kB) free!
[1563794156] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2791500 kB) free!
[1563794222] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (9676932 kB) free!
[1563794230] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.67, 4.69, 4.90
[1563794276] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.0% (2603824 kB) free!
[1563794342] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (4898728 kB) free!
[1563794462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (16601348 kB) free.
[1563794482] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (1174388 kB) free!
[1563794490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563794524] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6112172 kB) free!
[1563794602] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 57.9% (14193344 kB) free.
[1563794610] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563794644] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11082576 kB) free.
[1563794802] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.7% (788116 kB) free!
[1563794876] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 21.1% (3427628 kB) free.
[1563794923] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (501172 kB) free!
[1563794972] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.8% (9097552 kB) free!
[1563795043] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (213368 kB) free!
[1563795062] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (7581184 kB) free!
[1563795092] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.3% (10070344 kB) free!
[1563795182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (16971080 kB) free.
[1563795210] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563795212] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.7% (11663736 kB) free!
[1563795245] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5947448 kB) free!
[1563795299] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4514964 kB) free!
[1563795330] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563795365] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11073832 kB) free.
[1563795419] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4483160 kB) free!
[1563795539] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4434080 kB) free!
[1563795782] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (2020708 kB) free!
[1563795802] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (3767604 kB) free!
[1563795812] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (13809136 kB) free.
[1563795866] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.95, 5.41, 5.01
[1563795902] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (14363336 kB) free.
[1563795922] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (4023544 kB) free!
[1563795930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563795986] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.56, 5.99, 5.26
[1563796030] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.30, 5.41, 5.14
[1563796042] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (3832108 kB) free!
[1563796051] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563796076] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.4% (2671984 kB) free!
[1563796107] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.97, 5.60, 5.21
[1563796150] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.82, 5.54, 5.22
[1563796197] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2419408 kB) free!
[1563796270] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.11, 5.35, 5.19
[1563796317] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.1% (2130916 kB) free!
[1563796412] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.4% (11415268 kB) free!
[1563796503] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (4724560 kB) free!
[1563796532] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.4% (12120980 kB) free!
[1563796565] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7526388 kB) free!
[1563796623] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (386692 kB) free!
[1563796642] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 63.3% (15518836 kB) free.
[1563796650] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563796652] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.0% (11846504 kB) free!
[1563796685] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11078660 kB) free.
[1563796707] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.35, 4.04, 4.60
[1563796743] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (11514440 kB) free!
[1563796770] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563796871] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1563796871] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.11, 3.79, 4.54
[1563796917] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (880672 kB) free!
[1563797285] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7325860 kB) free!
[1563797343] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (3766736 kB) free!
[1563797370] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563797405] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11066452 kB) free.
[1563797490] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563797842] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196640 kB) free!
[1563797943] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (11626440 kB) free!
[1563797962] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178472 kB) free!
[1563798005] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7190460 kB) free!
[1563798071] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.51, 5.33, 5.02
[1563798082] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (819348 kB) free!
[1563798125] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11074332 kB) free.
[1563798191] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.86, 5.51, 5.12
[1563798311] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.63, 5.49, 5.16
[1563798543] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (3451744 kB) free!
[1563798644] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.5% (6799612 kB) free.
[1563798690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563798717] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 77.0% (12517460 kB) free.
[1563798725] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7094672 kB) free!
[1563798810] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563798845] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11075892 kB) free.
[1563798875] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1563798911] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.65, 4.33, 4.78
[1563799143] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (10189244 kB) free!
[1563799411] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563799445] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6957128 kB) free!
[1563799531] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563799566] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11066860 kB) free.
[1563799743] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2804024 kB) free!
[1563800252] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1563800252] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.0% (14438164 kB) free.
[1563800343] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (9147676 kB) free!
[1563800852] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (11378844 kB) free!
[1563800943] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (915572 kB) free!
[1563800972] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.9% (11801848 kB) free!
[1563801092] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.2% (12657664 kB) free!
[1563801294] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 528 processes
[1563801354] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 524 processes
[1563801415] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 523 processes
[1563801475] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 0 warn out of 518 processes
[1563801482] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4818228 kB) free!
[1563801539] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2419504 kB) free!
[1563801542] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4841460 kB) free!
[1563801543] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (7876380 kB) free!
[1563801602] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.8% (4851172 kB) free!
[1563801609] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.93, 4.01, 4.51
[1563801662] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.8% (4853620 kB) free!
[1563801682] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3328520 kB) free!
[1563801775] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 1 warn out of 496 processes
[1563801931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563801966] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6698072 kB) free!
[1563802051] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563802086] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11061304 kB) free.
[1563802292] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.7% (14899160 kB) free.
[1563802651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563802675] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 484 processes
[1563802686] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6500560 kB) free!
[1563802743] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (6258692 kB) free!
[1563802772] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563802772] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563802806] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11052608 kB) free.
[1563802892] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (11407676 kB) free!
[1563803012] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.0% (11839464 kB) free!
[1563803132] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (12409724 kB) free!
[1563803343] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (11883608 kB) free!
[1563803371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563803406] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6926512 kB) free!
[1563803491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563803526] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11061024 kB) free.
[1563803733] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (14187788 kB) free.
[1563803944] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (4586832 kB) free!
[1563804091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563804126] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7471188 kB) free!
[1563804211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563804247] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5901320 kB) free!
[1563804331] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563804367] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11053804 kB) free.
[1563804544] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (10062652 kB) free!
[1563804931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563804933] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (11333680 kB) free!
[1563805008] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 360 processes
[1563805051] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563805053] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.2% (11314348 kB) free!
[1563805128] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1563805144] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (2127696 kB) free!
[1563805173] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.1% (11246140 kB) free!
[1563805567] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7378540 kB) free!
[1563805651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563805687] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11059656 kB) free.
[1563805744] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (8362832 kB) free!
[1563805772] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563805773] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (15820076 kB) free.
[1563806287] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7268168 kB) free!
[1563806371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563806407] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11047172 kB) free.
[1563806491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563806944] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (6034956 kB) free!
[1563806973] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (12481052 kB) free!
[1563807007] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7113384 kB) free!
[1563807044] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1563807060] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 493.36, 169.44, 60.89
[1563807093] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.2% (12640372 kB) free!
[1563807127] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11056600 kB) free.
[1563807180] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 603.00, 318.55, 128.69
[1563807213] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (12413004 kB) free!
[1563807300] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 97.41, 227.56, 119.09
[1563807544] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (11381408 kB) free!
[1563807691] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563807811] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563808144] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (3375948 kB) free!
[1563808327] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8368084 kB) free!
[1563808411] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563808413] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (15863024 kB) free.
[1563808447] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7306068 kB) free!
[1563808531] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563808568] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11043772 kB) free.
[1563808652] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563808744] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (8679208 kB) free!
[1563808882] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2356736 kB) free!
[1563808917] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563809100] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.35, 1.02, 17.68
[1563809167] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8181228 kB) free!
[1563809251] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563809287] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (6984652 kB) free!
[1563809344] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1593496 kB) free!
[1563809371] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563809407] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11042512 kB) free.
[1563809971] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563810007] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7834292 kB) free!
[1563810091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563810127] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6722488 kB) free!
[1563810211] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563810247] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11047952 kB) free.
[1563810311] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563810544] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (10871940 kB) free!
[1563810588] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 31.7% (7765012 kB) free.
[1563810682] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 65.9% (16158008 kB) free.
[1563810811] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563810847] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7685896 kB) free!
[1563810901] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.45, 2.94
[1563810931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563810967] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6228800 kB) free!
[1563811052] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563811087] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11048412 kB) free.
[1563811144] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (4079252 kB) free!
[1563811188] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199832 kB) free!
[1563811308] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (812552 kB) free!
[1563811428] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (531932 kB) free!
[1563811651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563811687] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7286432 kB) free!
[1563811728] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 363 processes
[1563811744] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (9239288 kB) free!
[1563811771] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563811807] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5869372 kB) free!
[1563811848] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 362 processes
[1563811891] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563811928] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11049628 kB) free.
[1563811968] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 359 processes
[1563812344] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2980528 kB) free!
[1563812409] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.7% (21629624 kB) free.
[1563812442] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 83.1% (20370612 kB) free.
[1563812478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 44.6% (29348396 kB) free.
[1563812491] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563812526] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 62.3% (30705344 kB) free.
[1563812527] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6998248 kB) free!
[1563812557] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.3% (13860580 kB) free.
[1563812568] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1563812584] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 80.8% (19797872 kB) free.
[1563812611] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563812624] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.0% (14791480 kB) free.
[1563812647] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11052616 kB) free.
[1563813211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563813247] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7105916 kB) free!
[1563813331] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563813367] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11043288 kB) free.
[1563813650] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175980 kB) free!
[1563813770] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176796 kB) free!
[1563813890] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1563813890] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199284 kB) free!
[1563813931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563813967] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7037184 kB) free!
[1563814052] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563814087] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11050092 kB) free.
[1563814651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563814687] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6868264 kB) free!
[1563814771] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563814807] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11041560 kB) free.
[1563815371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1563815407] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6740992 kB) free!
[1563815491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563815527] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11050412 kB) free.
[1563816091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563816128] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6558504 kB) free!
[1563816211] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563816248] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11053548 kB) free.
[1563816344] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1563816374] HOST ALERT: rq-2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1563816385] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'rq-2' was out of bounds)
[1563816398] SERVICE ALERT: rq-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'rq-2' was out of bounds)
[1563816398] SERVICE ALERT: rq-1;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'rq-1' was out of bounds)
[1563816398] SERVICE ALERT: rq-1;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'rq-1' was out of bounds)
[1563816401] HOST ALERT: rq-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.203)
[1563816456] SERVICE ALERT: rq-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'rq-2' was out of bounds)
[1563816467] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'rq-2' was out of bounds)
[1563816521] SERVICE ALERT: rq-1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'rq-1' was out of bounds)
[1563816521] SERVICE ALERT: rq-1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'rq-1' was out of bounds)
[1563816522] SERVICE ALERT: rq-2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'rq-2' was out of bounds)
[1563816564] HOST ALERT: rq-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1563816576] SERVICE ALERT: rq-2;CPU Load;OK;HARD;1;OK - load average: 0.62, 0.32, 0.12
[1563816591] HOST ALERT: rq-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1563816642] SERVICE ALERT: rq-2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1563816811] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563816848] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6536576 kB) free!
[1563816932] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563816968] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11041496 kB) free.
[1563817064] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 98.6% (129964800 kB) free.
[1563817105] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1563817118] SERVICE ALERT: rq-2;Total Processes;OK;HARD;1;PROCS OK: 336 processes
[1563817118] SERVICE ALERT: rq-1;Memory;OK;HARD;1;OK - 98.6% (129940696 kB) free.
[1563817118] SERVICE ALERT: rq-1;Total Processes;OK;HARD;1;PROCS OK: 341 processes
[1563817531] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563817568] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6315536 kB) free!
[1563817651] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563817667] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'rq-2' was out of bounds)
[1563817670] HOST ALERT: rq-2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.204)
[1563817688] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11051644 kB) free.
[1563817708] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'rq-2' was out of bounds)
[1563817719] SERVICE ALERT: rq-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'rq-2' was out of bounds)
[1563817721] SERVICE ALERT: rq-1;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'rq-1' was out of bounds)
[1563817721] SERVICE ALERT: rq-1;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'rq-1' was out of bounds)
[1563817724] HOST ALERT: rq-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.203)
[1563817780] HOST ALERT: rq-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1563817787] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 98.7% (130129332 kB) free.
[1563817828] SERVICE FLAPPING ALERT: rq-2;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1563817828] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1563817839] SERVICE FLAPPING ALERT: rq-2;Total Processes;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1563817839] SERVICE ALERT: rq-2;Total Processes;OK;HARD;1;PROCS OK: 354 processes
[1563817844] SERVICE ALERT: rq-1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'rq-1' was out of bounds)
[1563817844] SERVICE ALERT: rq-1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'rq-1' was out of bounds)
[1563817914] HOST ALERT: rq-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1563817968] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 363 processes
[1563818088] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1563818251] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563818288] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6600156 kB) free!
[1563818371] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563818386] SERVICE ALERT: rq-2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1563818389] SERVICE ALERT: rq-2;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'rq-2' was out of bounds)
[1563818389] SERVICE ALERT: rq-2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'rq-2' was out of bounds)
[1563818406] HOST ALERT: rq-2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1563818408] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11040036 kB) free.
[1563818432] HOST ALERT: rq-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1563818506] SERVICE ALERT: rq-2;CPU Load;OK;HARD;1;OK - load average: 0.10, 0.10, 0.05
[1563818510] SERVICE ALERT: rq-2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1563818510] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 98.8% (130186260 kB) free.
[1563818971] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563819008] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6722216 kB) free!
[1563819041] SERVICE FLAPPING ALERT: rq-1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1563819041] SERVICE ALERT: rq-1;Memory;OK;HARD;1;OK - 98.8% (130179060 kB) free.
[1563819041] SERVICE FLAPPING ALERT: rq-1;Total Processes;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1563819041] SERVICE ALERT: rq-1;Total Processes;OK;HARD;1;PROCS OK: 341 processes
[1563819091] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563819128] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11052004 kB) free.
[1563819213] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563819691] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563819728] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6489752 kB) free!
[1563819812] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563819848] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11050684 kB) free.