[1575052318] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.1% (17283752 kB) free!
[1575052500] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15320544 kB) free!
[1575052620] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37253336 kB) free.
[1575052641] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/11/30 00:03:42 [error] 24130#0: *10970239 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575052704] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195560 kB) free!
[1575052761] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575052824] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185804 kB) free!
[1575052944] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181512 kB) free!
[1575053103] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184212 kB) free!
[1575053220] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16324740 kB) free!
[1575053223] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170352 kB) free!
[1575053340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12675016 kB) free!
[1575053343] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198508 kB) free!
[1575053361] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2019/11/30 00:19:05 [error] 24130#0: *10974363 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575053460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35307972 kB) free.
[1575053481] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(58) < 2019/11/30 00:20:28 [error] 24130#0: *10974741 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:38/0, bytes from/to upstream:0/38
[1575053601] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575054060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15826700 kB) free!
[1575054180] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37154272 kB) free.
[1575054201] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2019/11/30 00:28:16 [error] 24130#0: *10976787 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575054322] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575054780] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16218520 kB) free!
[1575054900] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12694072 kB) free!
[1575054921] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2019/11/30 00:45:20 [error] 24132#0: *10981453 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575055020] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34830104 kB) free.
[1575055041] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(36) < 2019/11/30 00:46:08 [error] 24130#0: *10981608 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575055161] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575055620] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15047608 kB) free!
[1575055677] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.9% (10350408 kB) free.
[1575055740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37036252 kB) free.
[1575055761] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2019/11/30 00:54:10 [error] 24130#0: *10983797 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575055881] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575056277] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191240 kB) free!
[1575056340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15187416 kB) free!
[1575056397] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177644 kB) free!
[1575056416] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2878384 kB) free!
[1575056460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36947316 kB) free.
[1575056481] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(57) < 2019/11/30 01:11:08 [error] 24128#0: *10988331 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575056517] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204444 kB) free!
[1575056602] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(24) < 2019/11/30 01:11:54 [error] 24134#0: *10988497 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575056722] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575057060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18248396 kB) free!
[1575057181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15317048 kB) free!
[1575057301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36569404 kB) free.
[1575057322] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2019/11/30 01:18:57 [error] 24130#0: *10990485 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575057442] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575057901] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15978224 kB) free!
[1575058021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36931740 kB) free.
[1575058042] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(60) < 2019/11/30 01:37:03 [error] 24130#0: *10995298 recv() failed (104: Connection reset by peer) while proxying connection, client: 210.212.235.210, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575058162] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(22) < 2019/11/30 01:37:42 [error] 24130#0: *10995473 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575058282] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575058621] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19015980 kB) free!
[1575058741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16089332 kB) free!
[1575058861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36125492 kB) free.
[1575058883] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2019/11/30 01:43:37 [error] 24130#0: *10997079 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575059003] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575059461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16133416 kB) free!
[1575059581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36910608 kB) free.
[1575059603] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(66) < 2019/11/30 02:03:21 [error] 24130#0: *11002495 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575059723] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2019/11/30 02:03:27 [error] 24130#0: *11002323 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575059843] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575060181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18661568 kB) free!
[1575060301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15562648 kB) free!
[1575060421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34480880 kB) free.
[1575060443] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/11/30 02:09:29 [error] 24130#0: *11004046 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575060563] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575060657] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232604 kB) free!
[1575060934] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214416 kB) free!
[1575061021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13838712 kB) free!
[1575061054] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227772 kB) free!
[1575061141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34939528 kB) free.
[1575061164] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(82) < 2019/11/30 02:29:15 [error] 24130#0: *11009339 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575061174] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242416 kB) free!
[1575061244] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1575061284] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575061364] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1575061741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16032356 kB) free!
[1575061774] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 59.4% (29274764 kB) free.
[1575061857] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 53.9% (26559588 kB) free.
[1575061862] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35617424 kB) free.
[1575061884] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2019/11/30 02:34:22 [error] 24130#0: *11010736 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575062004] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575062461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14257732 kB) free!
[1575062581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33899572 kB) free.
[1575062663] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1575062783] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1575063204] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(105) < 2019/11/30 03:00:29 [error] 24130#0: *11017769 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575063324] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575063781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17560052 kB) free!
[1575063901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36750592 kB) free.
[1575064501] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17061040 kB) free!
[1575064524] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(100) < 2019/11/30 03:25:11 [error] 24130#0: *11024505 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575064621] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36765016 kB) free.
[1575064644] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575065221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18104396 kB) free!
[1575065341] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36745832 kB) free.
[1575065844] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(83) < 2019/11/30 03:47:20 [error] 24130#0: *11030493 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575065941] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16836520 kB) free!
[1575065965] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575066061] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36363588 kB) free.
[1575066564] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2019/11/30 03:50:47 [error] 24132#0: *11031361 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575066661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17050656 kB) free!
[1575066684] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575066781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33691656 kB) free.
[1575067284] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(57) < 2019/11/30 04:09:23 [error] 24129#0: *11036391 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575067404] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(28) < 2019/11/30 04:12:42 [error] 24130#0: *11037299 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575067524] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575067981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18837972 kB) free!
[1575068102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35543972 kB) free.
[1575068124] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2019/11/30 04:15:35 [error] 24131#0: *11038085 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575068244] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575068701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26079528 kB) free!
[1575068821] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33688860 kB) free.
[1575068845] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(33) < 2019/11/30 04:33:56 [error] 24130#0: *11043107 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575068965] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(21) < 2019/11/30 04:38:53 [error] 24130#0: *11044340 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575069085] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(25) < 2019/11/30 04:41:22 [error] 24130#0: *11045078 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575069764] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1575069884] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1575070021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18091884 kB) free!
[1575070141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34073964 kB) free.
[1575071341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19026604 kB) free!
[1575071461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34033924 kB) free.
[1575071485] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575072086] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(71) < 2019/11/30 05:29:52 [error] 24130#0: *11058104 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575072206] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(23) < 2019/11/30 05:31:32 [error] 24130#0: *11058555 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575072326] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575073526] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(82) < 2019/11/30 05:55:20 [error] 24129#0: *11064967 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575073646] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(28) < 2019/11/30 05:56:02 [error] 24134#0: *11065154 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575073766] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575073861] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14390348 kB) free!
[1575073981] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35716972 kB) free.
[1575074518] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.8% (35263800 kB) free.
[1575074943] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.0% (20837784 kB) free.
[1575074967] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(57) < 2019/11/30 06:17:26 [error] 24130#0: *11070914 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575075087] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575075109] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 62.1% (15215220 kB) free.
[1575075144] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.2% (21629176 kB) free.
[1575075182] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19426080 kB) free!
[1575075302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35556612 kB) free.
[1575075687] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(53) < 2019/11/30 06:23:14 [error] 24130#0: *11072517 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575075807] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575076407] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(50) < 2019/11/30 06:42:51 [error] 24130#0: *11077669 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575076448] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.3% (21539352 kB) free.
[1575076478] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.6% (37767840 kB) free.
[1575076527] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/11/30 06:45:21 [error] 24135#0: *11078501 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575076647] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575077084] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1575077204] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1575077247] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(46) < 2019/11/30 06:49:14 [error] 24130#0: *11079532 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575077368] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575077967] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(55) < 2019/11/30 07:07:49 [error] 24131#0: *11084437 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575078087] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575078687] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(51) < 2019/11/30 07:13:43 [error] 24130#0: *11086020 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575078807] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575079407] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/11/30 07:29:00 [error] 24130#0: *11090160 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575079527] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(29) < 2019/11/30 07:34:11 [error] 24130#0: *11091567 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575079648] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(28) < 2019/11/30 07:37:21 [error] 24130#0: *11092435 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575080102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20541916 kB) free!
[1575080222] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34675764 kB) free.
[1575082048] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1575082048] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575082648] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(52) < 2019/11/30 08:24:44 [error] 24128#0: *11105212 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575082768] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(53) < 2019/11/30 08:29:07 [error] 24133#0: *11106325 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575082888] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575084012] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 372 processes
[1575084088] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(57) < 2019/11/30 08:49:54 [error] 24128#0: *11111963 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575084132] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 376 processes
[1575084208] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(23) < 2019/11/30 08:53:27 [error] 24130#0: *11112992 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575084252] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1575084255] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189000 kB) free!
[1575084284] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237424 kB) free!
[1575084328] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(24) < 2019/11/30 08:54:29 [error] 24130#0: *11113235 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575084375] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185412 kB) free!
[1575084404] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (246300 kB) free!
[1575084495] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189764 kB) free!
[1575084524] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235028 kB) free!
[1575084928] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575085104] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(167) < 2019/11/30 09:07:38 [error] 26850#0: *12889485 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575085224] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575085528] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(61) < 2019/11/30 09:14:55 [error] 24130#0: *11118677 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575085648] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575086249] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(59) < 2019/11/30 09:19:50 [error] 24130#0: *11120006 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575086369] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575086969] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(31) < 2019/11/30 09:35:00 [error] 24130#0: *11124096 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575087089] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(34) < 2019/11/30 09:39:49 [error] 24130#0: *11125367 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575087209] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575087660] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(37) < {writer,send_failed,{error,timeout}}
[1575087781] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(15) < {writer,send_failed,{error,timeout}}
[1575087809] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(57) < 2019/11/30 09:44:16 [error] 24132#0: *11126603 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575087901] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(3) < {writer,send_failed,{error,timeout}}
[1575087929] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575088004] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1575088125] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1575088530] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(31) < 2019/11/30 10:00:16 [error] 24133#0: *11130823 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575088622] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8624264 kB) free!
[1575088650] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(24) < 2019/11/30 10:05:55 [error] 24130#0: *11132367 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575088742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27273688 kB) free.
[1575088770] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(36) < 2019/11/30 10:08:58 [error] 24129#0: *11133236 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575089370] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575089473] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(2) < {writer,send_failed,{error,timeout}}
[1575089593] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1575089876] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.6% (19880676 kB) free.
[1575089913] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 49.6% (24449808 kB) free.
[1575089942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24527656 kB) free!
[1575089970] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/11/30 10:26:21 [error] 24131#0: *11137870 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575090062] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18134896 kB) free!
[1575090090] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(29) < 2019/11/30 10:31:04 [error] 24129#0: *11139150 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575090182] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13255104 kB) free!
[1575090194] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(16) < {writer,send_failed,{error,timeout}}
[1575090210] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575090314] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(2) < {writer,send_failed,{error,timeout}}
[1575090434] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;3;(13) < {writer,send_failed,{error,timeout}}
[1575090439] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(32) < {writer,send_failed,{error,timeout}}
[1575090561] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(16) < {writer,send_failed,{error,timeout}}
[1575090683] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;HARD;3;(24) < {writer,send_failed,{error,timeout}}
[1575090783] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35107368 kB) free.
[1575090811] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(52) < 2019/11/30 10:34:26 [error] 24134#0: *11140156 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575090931] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575091531] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2019/11/30 10:51:45 [error] 24132#0: *11144711 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575091651] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(36) < 2019/11/30 10:56:33 [error] 24131#0: *11145980 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575091743] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4809100 kB) free!
[1575091771] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(25) < 2019/11/30 10:59:25 [error] 24130#0: *11146801 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575091863] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 26.6% (6509836 kB) free.
[1575092457] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8964468 kB) free!
[1575092578] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8604660 kB) free!
[1575092583] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13323436 kB) free!
[1575092698] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.7% (8231396 kB) free!
[1575092703] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38510116 kB) free.
[1575092884] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.8% (3415604 kB) free!
[1575093004] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.3% (3281100 kB) free!
[1575093030] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.62, 6.38, 5.29
[1575093063] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4715300 kB) free!
[1575093124] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.0% (3207244 kB) free!
[1575093150] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.44, 6.02, 5.29
[1575093183] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4556456 kB) free!
[1575093270] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.95, 5.70, 5.26
[1575093303] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10864512 kB) free!
[1575093303] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4307392 kB) free!
[1575093328] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2993432 kB) free!
[1575093423] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38726472 kB) free.
[1575093448] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.3% (2528428 kB) free!
[1575093568] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2056544 kB) free!
[1575093718] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.9% (23532008 kB) free!
[1575093724] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575093724] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197820 kB) free!
[1575093839] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.1% (36999484 kB) free.
[1575093870] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.95, 3.97, 4.53
[1575093898] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4423888 kB) free!
[1575093920] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.3% (3775088 kB) free!
[1575094023] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9456780 kB) free!
[1575094040] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.2% (3249960 kB) free!
[1575094143] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37981692 kB) free.
[1575094160] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 11.1% (2735116 kB) free!
[1575094172] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575094498] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 41.9% (20647708 kB) free.
[1575094743] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7260912 kB) free!
[1575094760] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189960 kB) free!
[1575094772] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(88) < 2019/11/30 11:49:25 [error] 24133#0: *11160270 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575094863] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3609668 kB) free!
[1575094892] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575094909] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4399336 kB) free!
[1575094983] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30513056 kB) free.
[1575095029] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (4273052 kB) free!
[1575095038] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.7% (20693984 kB) free!
[1575095103] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2066208 kB) free!
[1575095149] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (4175612 kB) free!
[1575095158] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.7% (32618888 kB) free.
[1575095374] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9261676 kB) free!
[1575095395] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.0% (5963992 kB) free!
[1575095395] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 13.6% (4724940 kB) free!
[1575095492] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(294) < 2019/11/30 12:01:12 [error] 24125#0: *11163918 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:74/17, bytes from/to upstream:17/74
[1575095494] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (9086588 kB) free!
[1575095515] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.8% (31798124 kB) free.
[1575095515] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.8% (32321964 kB) free.
[1575095583] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3955964 kB) free!
[1575095612] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/11/30 12:03:04 [error] 24133#0: *11164938 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:5063/17, bytes from/to upstream:17/5063
[1575095614] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (8799472 kB) free!
[1575095703] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28748696 kB) free.
[1575095732] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575095849] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 36.96, 33.65, 30.48
[1575095924] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1575095969] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 83.4% (18734700 kB) free.
[1575095969] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 34.60, 33.34, 30.72
[1575096045] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1575096089] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 35.08, 33.64, 31.12
[1575096303] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (608544 kB) free!
[1575096332] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(86) < 2019/11/30 12:14:26 [error] 24130#0: *11167929 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575096358] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.9% (22217084 kB) free!
[1575096424] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42597868 kB) free.
[1575096453] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(23) < 2019/11/30 12:17:05 [error] 24132#0: *11168647 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575096478] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.6% (33751076 kB) free.
[1575096560] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 94.2% (23245104 kB) free.
[1575096573] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575096949] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1911968 kB) free!
[1575097023] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9235448 kB) free!
[1575097143] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4386260 kB) free!
[1575097224] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.2% (22992884 kB) free.
[1575097257] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.4% (20696936 kB) free.
[1575097263] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35647628 kB) free.
[1575097302] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 75.2% (18433612 kB) free.
[1575097414] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4432760 kB) free!
[1575097470] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.64, 5.61, 5.03
[1575097590] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.01, 6.44, 5.40
[1575097710] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.25, 7.17, 5.80
[1575097773] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(92) < 2019/11/30 12:39:07 [error] 24130#0: *11174538 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575097863] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5731460 kB) free!
[1575097889] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 26.85, 27.56, 29.41
[1575097893] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575097983] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2052048 kB) free!
[1575098103] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31816952 kB) free.
[1575098278] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (26019916 kB) free!
[1575098398] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.6% (17970260 kB) free!
[1575098490] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 41.05, 36.40, 33.27
[1575098493] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/11/30 12:42:35 [error] 24130#0: *11175521 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575098518] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.6% (31062288 kB) free.
[1575098610] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 44.46, 39.56, 34.82
[1575098613] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575098699] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9795156 kB) free!
[1575098703] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5206512 kB) free!
[1575098730] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 42.45, 40.69, 35.81
[1575098819] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9438984 kB) free!
[1575098823] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36235760 kB) free.
[1575098910] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1575098910] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.07, 3.48, 4.49
[1575098939] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1575098939] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9082568 kB) free!
[1575099024] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187744 kB) free!
[1575099144] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191568 kB) free!
[1575099214] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(65) < 2019/11/30 12:59:24 [error] 24132#0: *11179993 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575099264] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199724 kB) free!
[1575099334] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(32) < 2019/11/30 13:04:58 [error] 24131#0: *11181507 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575099423] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4256756 kB) free!
[1575099454] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575099543] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395584 kB) free!
[1575099663] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36570736 kB) free.
[1575099718] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (22869076 kB) free!
[1575099838] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12901120 kB) free!
[1575099958] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4082864 kB) free!
[1575100054] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2019/11/30 13:08:11 [error] 24131#0: *11182368 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575100174] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575100263] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5225348 kB) free!
[1575100383] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1295616 kB) free!
[1575100503] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35810488 kB) free.
[1575100558] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.9% (34181148 kB) free.
[1575100739] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3825688 kB) free!
[1575100774] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(74) < 2019/11/30 13:25:16 [error] 24133#0: *11186907 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575100894] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(31) < 2019/11/30 13:30:12 [error] 24133#0: *11188318 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575101015] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(23) < 2019/11/30 13:33:35 [error] 24134#0: *11189286 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575101041] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1575101059] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1575101104] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4915100 kB) free!
[1575101130] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1575101130] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1267.88, 733.09, 320.42
[1575101161] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1575101179] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2850 processes
[1575101224] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (523752 kB) free!
[1575101234] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1575101317] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8571052 kB) free!
[1575101344] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33821232 kB) free.
[1575101437] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8570324 kB) free!
[1575101557] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (8569640 kB) free!
[1575101758] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (24775940 kB) free!
[1575101834] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(6) < {writer,send_failed,{error,timeout}}
[1575101878] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.4% (16298344 kB) free!
[1575101944] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1725380 kB) free!
[1575101954] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(34) < {writer,send_failed,{error,timeout}}
[1575101998] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.4% (22918040 kB) free!
[1575102064] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401216 kB) free!
[1575102074] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1575102184] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35851308 kB) free.
[1575102598] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1575102598] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28060040 kB) free.
[1575102674] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(125) < {writer,send_failed,{error,timeout}}
[1575102784] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3342456 kB) free!
[1575102795] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(10) < {writer,send_failed,{error,timeout}}
[1575102832] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.7% (22871236 kB) free.
[1575102877] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 73.4% (36192652 kB) free.
[1575102904] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (388640 kB) free!
[1575102915] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;3;(13) < {writer,send_failed,{error,timeout}}
[1575102932] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 29.5% (14514220 kB) free.
[1575102967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.6% (52137860 kB) free.
[1575102971] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1575102989] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1575103092] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1575103119] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1575103199] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25099876 kB) free!
[1575103249] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1575103319] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.8% (33950860 kB) free.
[1575103769] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.9% (3113048 kB) free!
[1575103850] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2652 processes
[1575103889] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.5% (2813912 kB) free!
[1575104009] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.1% (2499032 kB) free!
[1575104519] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (21191096 kB) free!
[1575104557] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4701164 kB) free!
[1575104609] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1575104609] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (897812 kB) free!
[1575104639] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.1% (13337376 kB) free!
[1575104759] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.1% (30488680 kB) free.
[1575104960] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4699636 kB) free!
[1575105080] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (4377628 kB) free!
[1575105201] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.4% (4049268 kB) free!
[1575105251] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 73.8% (36378692 kB) free.
[1575105291] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 83.5% (18759680 kB) free.
[1575105359] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.6% (23152444 kB) free!
[1575105479] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.3% (34731480 kB) free.
[1575105975] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13136192 kB) free!
[1575106095] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9105620 kB) free!
[1575106216] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37126032 kB) free.
[1575106401] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1575106401] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1170000 kB) free!
[1575106679] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (25994568 kB) free!
[1575106799] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.3% (18897892 kB) free!
[1575106815] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10542796 kB) free!
[1575106919] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (26767952 kB) free.
[1575106935] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6815332 kB) free!
[1575107001] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 94.6% (23360388 kB) free.
[1575107032] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189248 kB) free!
[1575107055] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35997196 kB) free.
[1575107152] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182176 kB) free!
[1575107272] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (844640 kB) free!
[1575107615] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1575107655] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10425208 kB) free!
[1575107775] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6625244 kB) free!
[1575107895] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34401720 kB) free.
[1575108315] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1575108495] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10224316 kB) free!
[1575108615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6762020 kB) free!
[1575108706] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4864440 kB) free!
[1575108735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32532712 kB) free.
[1575108744] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4682952 kB) free!
[1575108826] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4777164 kB) free!
[1575108864] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4625264 kB) free!
[1575108901] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1575108930] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 18.57, 20.95, 34.59
[1575108947] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (4703600 kB) free!
[1575108984] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (4564272 kB) free!
[1575109285] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1575109335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11332616 kB) free!
[1575109455] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33839560 kB) free.
[1575109531] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 13.65, 15.41, 25.02
[1575109710] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575110055] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13966508 kB) free!
[1575110175] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9714040 kB) free!
[1575110296] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33462388 kB) free.
[1575110615] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575110895] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9817116 kB) free!
[1575111015] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6037240 kB) free!
[1575111135] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31624708 kB) free.
[1575111215] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(66) < 2019/11/30 16:18:55 [error] 24130#0: *11234391 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1575111286] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575111335] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(30) < 2019/11/30 16:24:22 [error] 24134#0: *11235942 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575111455] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575111735] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7909828 kB) free!
[1575111855] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20581764 kB) free!
[1575111947] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2380360 kB) free!
[1575111975] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29180764 kB) free.
[1575112055] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/11/30 16:27:48 [error] 24129#0: *11236927 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575112176] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575112575] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5552760 kB) free!
[1575112695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31559012 kB) free.
[1575112775] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(54) < 2019/11/30 16:42:10 [error] 24130#0: *11240793 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575112895] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(23) < 2019/11/30 16:49:53 [error] 24132#0: *11242964 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575113015] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(33) < 2019/11/30 16:53:24 [error] 24130#0: *11243861 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575113295] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7628780 kB) free!
[1575113415] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3766452 kB) free!
[1575113535] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30769320 kB) free.
[1575113615] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575114066] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4606420 kB) free!
[1575114135] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5204852 kB) free!
[1575114186] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4486328 kB) free!
[1575114215] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(67) < 2019/11/30 17:07:04 [error] 24132#0: *11247443 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575114255] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1320052 kB) free!
[1575114306] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4238188 kB) free!
[1575114335] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(27) < 2019/11/30 17:15:20 [error] 24130#0: *11249773 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575114375] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29002740 kB) free.
[1575114456] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575114719] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575114976] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3560544 kB) free!
[1575114984] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2426828 kB) free!
[1575115055] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2019/11/30 17:18:29 [error] 24130#0: *11250702 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575115072] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1575115096] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397480 kB) free!
[1575115175] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(33) < 2019/11/30 17:29:24 [error] 24132#0: *11253609 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575115216] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28351932 kB) free.
[1575115295] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575115454] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8945748 kB) free!
[1575115494] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575115574] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8937672 kB) free!
[1575115695] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (8938500 kB) free!
[1575115816] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4633304 kB) free!
[1575115895] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(58) < 2019/11/30 17:39:29 [error] 24132#0: *11256458 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1575115936] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1123092 kB) free!
[1575116015] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1575116056] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27337440 kB) free.
[1575116615] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(52) < 2019/11/30 17:53:35 [error] 24130#0: *11260293 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:38/0, bytes from/to upstream:0/38
[1575116656] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2730932 kB) free!
[1575116706] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2337296 kB) free!
[1575116736] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(4) < 2019/11/30 17:53:47 [error] 24131#0: *11260244 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1575116776] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (16712828 kB) free!
[1575116856] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1575116856] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(38) < 2019/11/30 17:56:14 [error] 24132#0: *11260917 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1575116896] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27221976 kB) free.
[1575117496] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2348340 kB) free!
[1575117616] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30647152 kB) free.
[1575118216] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6373788 kB) free!
[1575118336] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2348540 kB) free!
[1575118401] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575118456] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30993948 kB) free.
[1575119056] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6335456 kB) free!
[1575119177] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2358552 kB) free!
[1575119297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31342916 kB) free.
[1575119337] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (3112724 kB) free!
[1575119457] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (3111576 kB) free!
[1575119577] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (3117288 kB) free!
[1575119897] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6339776 kB) free!
[1575120017] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2149188 kB) free!
[1575120137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31944644 kB) free.
[1575120198] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3114980 kB) free!
[1575120318] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 36.8% (5988628 kB) free.
[1575120737] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6630288 kB) free!
[1575120857] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2762188 kB) free!
[1575120881] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22098256 kB) free.
[1575120920] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.4% (22641852 kB) free.
[1575120931] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575120963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 48.1% (63422524 kB) free.
[1575120997] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 75.7% (18555160 kB) free.
[1575121039] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.9% (22774780 kB) free.
[1575121486] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (703252 kB) free!
[1575121524] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197968 kB) free!
[1575121605] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193656 kB) free!
[1575121606] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (675952 kB) free!
[1575121644] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185820 kB) free!
[1575121647] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.1% (2483572 kB) free!
[1575121695] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4715444 kB) free!
[1575121725] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190800 kB) free!
[1575121726] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (642472 kB) free!
[1575121764] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575121764] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (350872 kB) free!
[1575121767] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (2451580 kB) free!
[1575121845] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195484 kB) free!
[1575121887] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2428852 kB) free!
[1575126087] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (2602020 kB) free!
[1575126456] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575127805] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.1% (23231212 kB) free.
[1575127850] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.3% (23124276 kB) free.
[1575127887] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575127887] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2376296 kB) free!
[1575127889] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39636808 kB) free.
[1575127924] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 80.3% (19673028 kB) free.
[1575127956] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 95.3% (23359592 kB) free.
[1575128166] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14842064 kB) free!
[1575128286] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40369800 kB) free.
[1575128418] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176088 kB) free!
[1575128457] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186684 kB) free!
[1575128527] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171692 kB) free!
[1575128538] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183792 kB) free!
[1575128577] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195348 kB) free!
[1575128647] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (324224 kB) free!
[1575128658] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1575128658] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196180 kB) free!
[1575128697] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185144 kB) free!
[1575128767] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1575128767] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (271236 kB) free!
[1575128886] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17497748 kB) free!
[1575129006] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13975056 kB) free!
[1575129126] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39670880 kB) free.
[1575129167] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.2% (2506884 kB) free!
[1575129287] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.1% (2479472 kB) free!
[1575129407] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (2468136 kB) free!
[1575129727] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16334496 kB) free!
[1575129847] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40083312 kB) free.
[1575130447] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18118472 kB) free!
[1575130567] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (14529656 kB) free!
[1575130607] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2382300 kB) free!
[1575130687] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38579880 kB) free.
[1575131287] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16173008 kB) free!
[1575131407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39746820 kB) free.
[1575131652] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 379 processes
[1575131772] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 379 processes
[1575132007] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18878924 kB) free!
[1575132127] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15509016 kB) free!
[1575132247] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38968412 kB) free.
[1575132847] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18135428 kB) free!
[1575132967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39598092 kB) free.
[1575133567] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19580756 kB) free!
[1575133687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16325832 kB) free!
[1575133807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37874564 kB) free.
[1575134408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17716428 kB) free!
[1575134528] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39523344 kB) free.
[1575135128] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19413788 kB) free!
[1575135248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33196836 kB) free.
[1575135617] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.8% (23153704 kB) free.
[1575135676] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.4% (23127416 kB) free.
[1575135718] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 80.8% (19811744 kB) free.
[1575135766] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 95.3% (23361228 kB) free.
[1575135848] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20711028 kB) free!
[1575135968] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17476124 kB) free!
[1575136088] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38282272 kB) free.
[1575136688] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18461204 kB) free!
[1575136808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39463416 kB) free.
[1575137183] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1575137303] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1575137408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19430636 kB) free!
[1575137444] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1575137528] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39419516 kB) free.
[1575137564] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1575138128] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20521908 kB) free!
[1575138248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39425312 kB) free.