[1553625147] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.5% (5653936 kB) free!
[1553625267] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9943700 kB) free.
[1553625293] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4699928 kB) free!
[1553625413] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2981396 kB) free!
[1553625533] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11361748 kB) free.
[1553625719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553625867] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9041456 kB) free!
[1553625987] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.6% (6708988 kB) free!
[1553626107] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4162464 kB) free!
[1553626133] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5642360 kB) free!
[1553626253] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4314476 kB) free!
[1553626320] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553626373] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2239408 kB) free!
[1553626440] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553626556] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 322 processes
[1553626707] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9936248 kB) free.
[1553626973] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6598244 kB) free!
[1553627040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553627160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553627280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553627573] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11359768 kB) free.
[1553628173] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3615348 kB) free!
[1553628294] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11352352 kB) free.
[1553628480] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1553628893] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6112324 kB) free!
[1553629013] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4834384 kB) free!
[1553629080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553629133] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3226616 kB) free!
[1553629201] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553629733] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7184160 kB) free!
[1553629800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553629920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553630040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553630333] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11349916 kB) free.
[1553630933] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4347188 kB) free!
[1553631053] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (2403668 kB) free!
[1553631173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11357812 kB) free.
[1553631240] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553631773] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5405796 kB) free!
[1553631840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553631893] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (4111228 kB) free!
[1553631960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553632013] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11350288 kB) free.
[1553632081] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553632614] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6407504 kB) free!
[1553632680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553632734] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.3% (5099200 kB) free!
[1553632800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553632854] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3688688 kB) free!
[1553632920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553633454] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7736816 kB) free!
[1553634054] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11344244 kB) free.
[1553634120] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553634507] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.2% (7470596 kB) free!
[1553634627] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4906628 kB) free!
[1553634654] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4555484 kB) free!
[1553634720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553634748] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9948844 kB) free.
[1553634774] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (2675252 kB) free!
[1553634840] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553634894] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11356024 kB) free.
[1553635239] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 335 processes
[1553635441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553635494] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5534900 kB) free!
[1553635561] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1553635614] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (4215012 kB) free!
[1553635681] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553635734] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11348524 kB) free.
[1553636335] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6522268 kB) free!
[1553636455] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5226748 kB) free!
[1553636575] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3907680 kB) free!
[1553636881] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553637175] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.2% (7983824 kB) free!
[1553637481] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553637591] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43182 MB (19% inode=99%):
[1553637601] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1553637711] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55475 MB (24% inode=99%):
[1553637721] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553637775] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11338848 kB) free.
[1553638322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553638347] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1553638375] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4810580 kB) free!
[1553638442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553638495] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (3197528 kB) free!
[1553638562] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1553638615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11346824 kB) free.
[1553639162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553639215] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5864576 kB) free!
[1553639282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553639335] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4570216 kB) free!
[1553639402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553639455] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2743348 kB) free!
[1553640055] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6919436 kB) free!
[1553640602] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553640656] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11350888 kB) free.
[1553641203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553641256] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4090276 kB) free!
[1553641323] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553641376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11347004 kB) free.
[1553641443] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553641976] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6443588 kB) free!
[1553642043] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553642096] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5134624 kB) free!
[1553642163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553642216] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3755796 kB) free!
[1553642283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553642816] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.8% (7771596 kB) free!
[1553643111] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45289 MB (20% inode=99%):
[1553643231] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45281 MB (20% inode=99%):
[1553643351] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55481 MB (24% inode=99%):
[1553643416] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11343884 kB) free.
[1553643483] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553644016] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4737072 kB) free!
[1553644083] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553644136] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (3058000 kB) free!
[1553644204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553644256] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11347088 kB) free.
[1553644347] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9372764 kB) free!
[1553644467] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.7% (6744920 kB) free!
[1553644587] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.9% (9798020 kB) free!
[1553644803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553644857] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.7% (5764828 kB) free!
[1553644923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553644977] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4473564 kB) free!
[1553645043] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553645097] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (2634624 kB) free!
[1553645187] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9971792 kB) free.
[1553645697] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.1% (6942512 kB) free!
[1553646243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553646297] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11350472 kB) free.
[1553646387] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9704340 kB) free!
[1553646507] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.6% (7178484 kB) free!
[1553646627] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4404860 kB) free!
[1553646843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553646897] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4313676 kB) free!
[1553646963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553647017] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (2335692 kB) free!
[1553647084] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553647137] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11342872 kB) free.
[1553647227] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1553647227] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.6% (5720756 kB) free!
[1553647683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553647737] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5459020 kB) free!
[1553647803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553647827] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4120108 kB) free!
[1553647857] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4174224 kB) free!
[1553647923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553647977] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11350772 kB) free.
[1553648427] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.8% (7309424 kB) free!
[1553648577] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6528692 kB) free!
[1553648698] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5225004 kB) free!
[1553648818] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3910576 kB) free!
[1553649028] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9979196 kB) free.
[1553649123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553649418] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.3% (8023092 kB) free!
[1553649628] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8171872 kB) free!
[1553649723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553649748] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.3% (5567416 kB) free!
[1553649843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553649868] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9879556 kB) free!
[1553649964] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553650018] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11348100 kB) free.
[1553650563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553650618] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4899848 kB) free!
[1553650683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553650738] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (3349716 kB) free!
[1553650803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553650858] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11334196 kB) free.
[1553651458] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5966760 kB) free!
[1553651578] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4696084 kB) free!
[1553651668] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4362792 kB) free!
[1553651698] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2971884 kB) free!
[1553652003] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553652268] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.6% (8659820 kB) free!
[1553652298] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (7019680 kB) free!
[1553652603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553652723] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553652868] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4686208 kB) free!
[1553652899] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11339332 kB) free.
[1553653324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553653444] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553653468] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.0% (9386388 kB) free!
[1553653499] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4594612 kB) free!
[1553653564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553653619] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2862004 kB) free!
[1553653739] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11344564 kB) free.
[1553654339] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6548464 kB) free!
[1553654459] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5254932 kB) free!
[1553654579] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3948612 kB) free!
[1553654764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553655179] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.9% (8326980 kB) free!
[1553655364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553655484] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553655604] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553655779] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11341488 kB) free.
[1553655868] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4403472 kB) free!
[1553656379] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4924376 kB) free!
[1553656468] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.6% (6220592 kB) free!
[1553656499] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3393824 kB) free!
[1553656552] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43174 MB (19% inode=99%):
[1553656620] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332380 kB) free.
[1553656672] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55467 MB (24% inode=99%):
[1553656805] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553657219] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5981176 kB) free!
[1553657339] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4691256 kB) free!
[1553657405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553657459] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (3014088 kB) free!
[1553657525] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553658125] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553658245] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1553658659] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.2% (5512664 kB) free!
[1553658845] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553658965] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1553659468] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10187556 kB) free.
[1553659565] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553659686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553659806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553660272] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46738 MB (20% inode=99%):
[1553660392] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55185 MB (24% inode=99%):
[1553660459] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10122760 kB) free.
[1553660568] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.1% (6992592 kB) free.
[1553660615] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.3% (7007196 kB) free.
[1553660668] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.7% (6765180 kB) free!
[1553660788] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.2% (10916748 kB) free.
[1553660992] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44184 MB (19% inode=99%):
[1553661112] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55448 MB (24% inode=99%):
[1553661388] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5777304 kB) free!
[1553661508] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10668932 kB) free.
[1553661523] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1553661643] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1553661659] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4173448 kB) free!
[1553661712] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46424 MB (20% inode=99%):
[1553661763] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1553661779] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (2679624 kB) free!
[1553661832] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55449 MB (24% inode=99%):
[1553661900] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11197636 kB) free.
[1553662109] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.5% (5180640 kB) free!
[1553662229] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.5% (10097780 kB) free.
[1553662253] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47182 MB (20% inode=99%):
[1553662373] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47162 MB (20% inode=99%):
[1553662493] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47162 MB (20% inode=99%):
[1553662499] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3091632 kB) free!
[1553662619] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11286312 kB) free.
[1553662829] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4719652 kB) free!
[1553662949] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9699432 kB) free!
[1553663032] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46455 MB (20% inode=99%):
[1553663069] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.3% (7045516 kB) free!
[1553663152] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46455 MB (20% inode=99%):
[1553663219] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2795676 kB) free!
[1553663272] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55416 MB (24% inode=99%):
[1553663339] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11167888 kB) free.
[1553663524] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 19.6% (1445028 kB) free!
[1553663592] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 53.2% (26221548 kB) free.
[1553663644] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 19.3% (1418856 kB) free!
[1553663705] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 19.8% (1457992 kB) free!
[1553663764] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 19.0% (1395040 kB) free!
[1553663825] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.5% (1436556 kB) free!
[1553663939] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (2108684 kB) free!
[1553663945] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 19.2% (1413756 kB) free!
[1553664007] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1553664059] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11006268 kB) free.
[1553664607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1553664659] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (1724208 kB) free!
[1553664727] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1553664779] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10785632 kB) free.
[1553665072] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47152 MB (20% inode=99%):
[1553665192] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47152 MB (20% inode=99%):
[1553665293] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47154 MB (20% inode=99%):
[1553665312] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47152 MB (20% inode=99%):
[1553665327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1553665379] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (1471856 kB) free!
[1553665413] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47146 MB (20% inode=99%):
[1553665447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1553665499] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10404916 kB) free.
[1553665533] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47146 MB (20% inode=99%):
[1553666003] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7525012 kB) free!
[1553666047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553666123] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.1% (12886896 kB) free.
[1553666167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1553666287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1553666512] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49132 MB (21% inode=99%):
[1553666699] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3469104 kB) free!
[1553666723] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7751616 kB) free!
[1553666819] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11135968 kB) free.
[1553666843] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.3% (12956316 kB) free.
[1553666846] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;1;WARNING - 19.9% (1466920 kB) free!
[1553666966] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;2;WARNING - 19.9% (1464076 kB) free!
[1553667086] SERVICE ALERT: sig5-vm1;Memory;WARNING;HARD;3;WARNING - 19.9% (1462224 kB) free!
[1553667419] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3383384 kB) free!
[1553667443] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.4% (7580920 kB) free!
[1553667539] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11139304 kB) free.
[1553667564] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.2% (12893956 kB) free.
[1553668139] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4025128 kB) free!
[1553668163] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.9% (7855316 kB) free!
[1553668259] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (2102900 kB) free!
[1553668283] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.5% (13074768 kB) free.
[1553668379] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10622832 kB) free.
[1553668398] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1553668518] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1553668883] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8066728 kB) free!
[1553668912] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47019 MB (20% inode=99%):
[1553668979] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3619524 kB) free!
[1553669003] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.9% (13252596 kB) free.
[1553669032] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43718 MB (19% inode=99%):
[1553669099] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11123808 kB) free.
[1553669152] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55211 MB (24% inode=99%):
[1553669603] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8327032 kB) free!
[1553669699] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3476712 kB) free!
[1553669723] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.1% (13366924 kB) free.
[1553669752] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47019 MB (20% inode=99%):
[1553669765] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 95.0% (6988416 kB) free.
[1553669819] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11111192 kB) free.
[1553669826] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 95.0% (6987392 kB) free.
[1553669872] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54698 MB (24% inode=99%):
[1553669875] SERVICE ALERT: sig5-vm1;Memory;OK;HARD;1;OK - 94.7% (6961088 kB) free.
[1553670323] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8568372 kB) free!
[1553670420] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4095652 kB) free!
[1553670443] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.5% (13542920 kB) free.
[1553670472] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47017 MB (20% inode=99%):
[1553670540] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (2464980 kB) free!
[1553670592] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55210 MB (24% inode=99%):
[1553670660] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10895272 kB) free.
[1553671043] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8719000 kB) free!
[1553671163] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.4% (13507808 kB) free.
[1553671192] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47018 MB (20% inode=99%):
[1553671260] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (2320328 kB) free!
[1553671312] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55209 MB (24% inode=99%):
[1553671380] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10679152 kB) free.
[1553671763] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8658968 kB) free!
[1553671883] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.3% (13478052 kB) free.
[1553671912] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47016 MB (20% inode=99%):
[1553671980] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2062508 kB) free!
[1553672032] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55209 MB (24% inode=99%):
[1553672100] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10414604 kB) free.
[1553672483] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8718848 kB) free!
[1553672604] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.2% (13408288 kB) free.
[1553672632] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47009 MB (20% inode=99%):
[1553672700] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1553672752] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55201 MB (24% inode=99%):
[1553673203] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8920256 kB) free!
[1553673300] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3701988 kB) free!
[1553673323] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.0% (13330512 kB) free.
[1553673352] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47008 MB (20% inode=99%):
[1553673420] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11091248 kB) free.
[1553673473] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55168 MB (24% inode=99%):
[1553673923] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9008456 kB) free!
[1553674020] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3869064 kB) free!
[1553674043] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.9% (13280584 kB) free.
[1553674072] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47016 MB (20% inode=99%):
[1553674140] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10922308 kB) free.
[1553674192] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55207 MB (24% inode=99%):
[1553674643] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9357088 kB) free!
[1553674741] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3917036 kB) free!
[1553674763] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.8% (13226076 kB) free.
[1553674792] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47006 MB (20% inode=99%):
[1553674861] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (1870644 kB) free!
[1553674912] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55199 MB (24% inode=99%):
[1553674981] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10432284 kB) free.
[1553675363] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9569396 kB) free!
[1553675483] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.9% (6863520 kB) free!
[1553675512] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47015 MB (20% inode=99%):
[1553675581] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2084472 kB) free!
[1553675603] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1553675603] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.7% (12188464 kB) free.
[1553675632] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55207 MB (24% inode=99%):
[1553675701] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10593144 kB) free.
[1553676203] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.1% (7453672 kB) free!
[1553676232] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47021 MB (20% inode=99%):
[1553676301] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2594228 kB) free!
[1553676323] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.6% (12626320 kB) free.
[1553676352] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55214 MB (24% inode=99%):
[1553676421] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10763004 kB) free.
[1553676487] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1553676924] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7760772 kB) free!
[1553676952] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47030 MB (20% inode=99%):
[1553677021] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2342572 kB) free!
[1553677044] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.1% (12856844 kB) free.
[1553677072] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55222 MB (24% inode=99%):
[1553677088] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1553677141] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10533908 kB) free.
[1553677644] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8114292 kB) free!
[1553677672] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47017 MB (20% inode=99%):
[1553677688] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1553677764] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.6% (13104916 kB) free.
[1553677792] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55218 MB (24% inode=99%):
[1553677808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1553677928] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1553678341] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3595728 kB) free!
[1553678364] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8269980 kB) free!
[1553678392] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47026 MB (20% inode=99%):
[1553678462] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11051100 kB) free.
[1553678484] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.5% (13075256 kB) free.
[1553678512] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54962 MB (24% inode=99%):
[1553679061] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3111964 kB) free!
[1553679084] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8437428 kB) free!
[1553679112] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45010 MB (20% inode=99%):
[1553679128] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1553679181] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11059128 kB) free.
[1553679204] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.5% (13040524 kB) free.
[1553679232] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54194 MB (24% inode=99%):
[1553679728] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1553679781] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3659976 kB) free!
[1553679804] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8740904 kB) free!
[1553679848] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1553679901] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11044096 kB) free.
[1553679924] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.4% (13010560 kB) free.
[1553679969] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1553679969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1553680432] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47064 MB (20% inode=99%):
[1553680501] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3769768 kB) free!
[1553680524] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9106784 kB) free!
[1553680552] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47064 MB (20% inode=99%):
[1553680621] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11048412 kB) free.
[1553680644] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.9% (12758008 kB) free.
[1553680672] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55257 MB (24% inode=99%):
[1553681221] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3536744 kB) free!
[1553681245] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9398312 kB) free!
[1553681273] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47057 MB (20% inode=99%):
[1553681341] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11045608 kB) free.
[1553681365] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6793496 kB) free!
[1553681393] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55000 MB (24% inode=99%):
[1553681485] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.6% (12118488 kB) free.
[1553681941] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3487832 kB) free!
[1553681993] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43727 MB (19% inode=99%):
[1553682061] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11039336 kB) free.
[1553682085] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.1% (6941628 kB) free!
[1553682113] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54997 MB (24% inode=99%):
[1553682205] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.1% (12371540 kB) free.
[1553682321] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4876768 kB) free!
[1553682441] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4873524 kB) free!
[1553682562] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4867580 kB) free!
[1553682661] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3039368 kB) free!
[1553682713] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45237 MB (20% inode=99%):
[1553682782] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11024032 kB) free.
[1553682805] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7730344 kB) free!
[1553682833] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54995 MB (24% inode=99%):
[1553682925] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.9% (12751900 kB) free.
[1553683381] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3229008 kB) free!
[1553683433] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47020 MB (20% inode=99%):
[1553683501] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11009724 kB) free.
[1553683525] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7653776 kB) free!
[1553683553] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54707 MB (24% inode=99%):
[1553683645] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.6% (12605368 kB) free.
[1553684101] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (2291884 kB) free!
[1553684153] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46545 MB (20% inode=99%):
[1553684221] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (1570884 kB) free!
[1553684245] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7904444 kB) free!
[1553684273] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 59437 MB (26% inode=99%):
[1553684341] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240164 kB) free!
[1553684365] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.2% (12408924 kB) free.
[1553684941] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5281064 kB) free!
[1553684965] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8406100 kB) free!
[1553685085] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.3% (5573652 kB) free!
[1553685206] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11343076 kB) free.
[1553685391] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3126804 kB) free!
[1553685511] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.0% (3091412 kB) free!
[1553685631] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (3043684 kB) free!
[1553685806] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.3% (5580656 kB) free!
[1553685926] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.6% (11649660 kB) free.
[1553686141] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223668 kB) free!
[1553686363] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1553686526] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.2% (6035032 kB) free!
[1553686646] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11655256 kB) free.
[1553687246] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.2% (6521532 kB) free!
[1553687366] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.4% (11540888 kB) free.
[1553687941] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1553687941] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.0% (14792952 kB) free.
[1553687962] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2368404 kB) free!
[1553687966] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.7% (6733976 kB) free!
[1553688086] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.9% (11291388 kB) free.
[1553688370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1553688541] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2145116 kB) free!
[1553688626] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.1% (22087236 kB) free.
[1553688662] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239904 kB) free!
[1553688686] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.5% (7127876 kB) free!
[1553688782] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217976 kB) free!
[1553688806] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4270068 kB) free!
[1553688926] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10140564 kB) free.
[1553688970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553689023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11101928 kB) free.
[1553689090] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553689109] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (14794532 kB) free.
[1553689273] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2943420 kB) free!
[1553689393] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2866804 kB) free!
[1553689513] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.2% (2793752 kB) free!
[1553689526] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4640748 kB) free!
[1553689646] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10537632 kB) free.
[1553689690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553689810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553689930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553690231] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5734920 kB) free!
[1553690247] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.8% (5323832 kB) free!
[1553690351] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4303340 kB) free!
[1553690367] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10469800 kB) free.
[1553690471] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2557516 kB) free!
[1553690967] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5787392 kB) free!
[1553691087] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (2931920 kB) free!
[1553691207] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.3% (9018708 kB) free!
[1553691313] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1591624 kB) free!
[1553691671] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5577124 kB) free!
[1553691807] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3401592 kB) free!
[1553692407] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.7% (7267464 kB) free!
[1553693007] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1553693007] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2134976 kB) free!
[1553693471] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11535496 kB) free.
[1553693607] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.2% (6016956 kB) free!
[1553694072] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3332936 kB) free!
[1553694192] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11531224 kB) free.
[1553694792] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3755120 kB) free!
[1553694912] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11506164 kB) free.
[1553695512] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3663896 kB) free!
[1553695632] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4483516 kB) free!
[1553695752] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10548020 kB) free.
[1553696007] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4371216 kB) free!
[1553696608] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.1% (8444832 kB) free!
[1553696952] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3740052 kB) free!
[1553697072] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11483308 kB) free.
[1553697208] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4018000 kB) free!
[1553697672] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3883384 kB) free!
[1553697793] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11474612 kB) free.
[1553697808] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.5% (8142460 kB) free!
[1553698392] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3864088 kB) free!
[1553698408] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4025588 kB) free!
[1553698512] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11466084 kB) free.
[1553698953] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;1;WARNING - 20.0% (1471096 kB) free!
[1553699008] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.5% (8124832 kB) free!
[1553699073] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;2;WARNING - 20.0% (1467988 kB) free!
[1553699112] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4269752 kB) free!
[1553699193] SERVICE ALERT: sig6-vm1;Memory;WARNING;HARD;3;WARNING - 19.9% (1464904 kB) free!
[1553699232] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (2401852 kB) free!
[1553699352] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11153836 kB) free.
[1553699608] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3916032 kB) free!
[1553699952] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3501016 kB) free!
[1553700072] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11443216 kB) free.
[1553700208] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.2% (7971668 kB) free!
[1553700672] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3565572 kB) free!
[1553700730] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1553700792] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11433564 kB) free.
[1553700808] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3900596 kB) free!
[1553701392] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3502324 kB) free!
[1553701408] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.0% (7895072 kB) free!
[1553701513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11431140 kB) free.
[1553702008] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4017068 kB) free!
[1553702112] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3464380 kB) free!
[1553702232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11426516 kB) free.
[1553702608] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.9% (7859724 kB) free!
[1553702832] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (3309116 kB) free!
[1553702952] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11431512 kB) free.
[1553703208] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3675252 kB) free!
[1553703552] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3062500 kB) free!
[1553703672] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11298704 kB) free.
[1553703808] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.6% (7682456 kB) free!
[1553704272] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2804448 kB) free!
[1553704330] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553704392] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11426140 kB) free.
[1553704409] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3546332 kB) free!
[1553704930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553704992] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5057556 kB) free!
[1553705009] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.5% (7662932 kB) free!
[1553705050] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553705112] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (3622788 kB) free!
[1553705171] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553705232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11419156 kB) free.
[1553705609] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3159024 kB) free!
[1553705833] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6238736 kB) free!
[1553705953] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4926684 kB) free!
[1553706073] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3390704 kB) free!
[1553706209] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.6% (7219380 kB) free!
[1553706371] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553706673] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7095532 kB) free!
[1553706809] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2551992 kB) free!
[1553706971] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553707091] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553707409] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.6% (6695616 kB) free!
[1553707691] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553707811] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553707873] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11415848 kB) free.
[1553708009] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (1929536 kB) free!
[1553708411] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553708473] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5606196 kB) free!
[1553708532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553708593] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4311832 kB) free!
[1553708609] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.8% (6305732 kB) free!
[1553708652] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1553708652] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553708713] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2375956 kB) free!
[1553709209] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2643388 kB) free!
[1553709313] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6769672 kB) free!
[1553709809] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.9% (7847616 kB) free!
[1553709852] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553709913] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11409352 kB) free.
[1553710452] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553710514] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3986320 kB) free!
[1553710572] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553710634] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11405256 kB) free.
[1553710692] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553711234] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6500592 kB) free!
[1553711292] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553711354] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5216784 kB) free!