[1591814241] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1591816200] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 62.07, 53.42, 44.83
[1591817400] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 18.99, 34.01, 39.84
[1591817774] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.2% (3237704 kB) free!
[1591817835] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (2035888 kB) free!
[1591817895] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (191140 kB) free!
[1591817960] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (190940 kB) free!
[1591818849] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7718108 kB) free!
[1591818969] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (896304 kB) free!
[1591819089] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397456 kB) free!
[1591819160] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 79.0% (19362584 kB) free.
[1591819419] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236400 kB) free!
[1591819539] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241568 kB) free!
[1591819659] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235880 kB) free!
[1591821459] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.2% (29684920 kB) free.
[1591822852] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (5685516 kB) free!
[1591823173] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.81, 5.39, 5.16
[1591823293] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.94, 4.88, 5.00
[1591823452] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1166592 kB) free!
[1591824052] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.1% (44401188 kB) free.
[1591824847] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1339) < 172.105.59.173 [11/Jun/2020:03:02:12 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1591824849] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (45240044 kB) free.
[1591824967] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591824990] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(175) < 103.16.101.72 [11/Jun/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1591825110] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591825126] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(230) < 2020/06/11 03:01:22 [error] 12477#0: *94154187 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.62.98, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1591825160] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1591825246] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591827493] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.32, 6.24, 5.28
[1591827613] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.61, 6.08, 5.35
[1591827733] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.83, 5.32, 5.16
[1591828333] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.03, 3.53, 4.32
[1591830089] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (201596 kB) free!
[1591831200] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.64, 34.54, 35.67
[1591831329] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.1% (10071596 kB) free.
[1591831800] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 32.35, 31.25, 33.04
[1591831929] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174972 kB) free!
[1591832049] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176188 kB) free!
[1591832169] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (210116 kB) free!
[1591833560] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 67.5% (16537844 kB) free.
[1591833619] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.1% (10646700 kB) free.
[1591836019] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163956 kB) free!
[1591836139] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180048 kB) free!
[1591836259] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188336 kB) free!
[1591836489] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.8% (15538008 kB) free!
[1591836600] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 57.44, 49.39, 45.27
[1591836669] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 73.7% (18061328 kB) free.
[1591836861] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (1242976 kB) free!
[1591836921] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (1240288 kB) free!
[1591836981] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 5.0% (1232768 kB) free!
[1591837041] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 5.0% (1236736 kB) free!
[1591837200] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1591837200] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 53.31, 44.45, 44.23
[1591837800] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 73.24, 73.18, 59.01
[1591841960] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199984 kB) free!
[1591841960] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1591842046] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/06/11 07:43:33 [error] 12489#0: *94160278 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1591842080] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (203848 kB) free!
[1591842080] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1591842102] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1591842166] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591842200] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165684 kB) free!
[1591842201] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1591842223] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1591842343] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1591842801] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1591842882] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.87, 6.51, 5.30
[1591842943] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1591843002] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.73, 5.79, 5.18
[1591843122] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.31, 5.59, 5.20
[1591843362] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 75.2% (18429392 kB) free.
[1591849200] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591851069] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (4082312 kB) free!
[1591851162] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.5% (4053320 kB) free!
[1591851189] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3988800 kB) free!
[1591851282] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.8% (3883628 kB) free!
[1591851309] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (3737156 kB) free!
[1591851367] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(49) < 203.212.70.170 [11/Jun/2020:10:25:23 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1591851402] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.6% (3570000 kB) free!
[1591851488] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(2) < 203.212.70.170 [11/Jun/2020:10:26:14 +0530] TCP 502 0 0 187.015 "10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "7.015, 60.000, 60.000, 60.000"
[1591851608] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591851801] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1591851909] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2356576 kB) free!
[1591851921] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1591852002] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2112632 kB) free!
[1591852041] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1591852808] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [11/Jun/2020:10:45:59 +0530] TCP 502 0 0 180.001 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 60.000, 0.000"
[1591852928] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591853709] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (2917796 kB) free!
[1591853840] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.3% (2541612 kB) free!
[1591853960] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (2131040 kB) free!
[1591854080] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1772296 kB) free!
[1591854166] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/06/11 11:07:19 [error] 12484#0: *94168000 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1591854287] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591854309] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1591854309] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1847300 kB) free!
[1591855289] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1591855349] SERVICE ALERT: TempHP1;CPU-STATISTICS;UNKNOWN;SOFT;1;NRPE: Command 'check_cpu_stats.sh' not defined
[1591855469] SERVICE ALERT: TempHP1;CPU-STATISTICS;UNKNOWN;SOFT;2;NRPE: Command 'check_cpu_stats.sh' not defined
[1591855513] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1591855965] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1591856372] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1591856743] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1591856863] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1591856983] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1591857080] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.5% (16061048 kB) free.
[1591857249] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9788184 kB) free!
[1591857369] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9411952 kB) free!
[1591857402] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 332 processes
[1591857483] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1302 processes
[1591857489] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9047616 kB) free!
[1591857522] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1591857584] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1591857603] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1302 processes
[1591857723] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 324 processes
[1591857729] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 203.212.70.170 [11/Jun/2020:12:06:18 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1591857733] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.39, 6.71, 5.11
[1591857849] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591857853] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.47, 6.26, 5.15
[1591857973] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.53, 6.36, 5.31
[1591858002] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 57.4% (14057220 kB) free.
[1591858689] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4935884 kB) free!
[1591859371] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(5) < Channel error on connection <0.16186.4097> (10.147.212.158:49932 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1591859492] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1591860159] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.9% (6519900 kB) free!
[1591860279] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.4% (6356060 kB) free!
[1591860373] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.50, 3.39, 4.28
[1591860399] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.9% (6187412 kB) free!
[1591860999] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 25.1% (8210088 kB) free.
[1591861602] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.8% (4121144 kB) free!
[1591861722] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.6% (3832568 kB) free!
[1591861842] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.4% (3526868 kB) free!
[1591862322] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 337 processes
[1591862442] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1591862442] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1591862442] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2009508 kB) free!
[1591862538] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 770 processes
[1591862658] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 771 processes
[1591862773] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.34, 7.59, 5.85
[1591862778] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 343 processes
[1591862799] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 18.3% (5981212 kB) free!
[1591862893] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.14, 7.54, 6.05
[1591862919] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 17.8% (5817404 kB) free!
[1591863013] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.62, 6.86, 5.98
[1591863039] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 17.3% (5659360 kB) free!
[1591863281] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1591863613] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1591863613] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.44, 4.08, 4.91
[1591863635] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1591863635] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1591864035] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1591864035] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1591864280] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4188652 kB) free!
[1591864400] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3666228 kB) free!
[1591864520] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (3194640 kB) free!
[1591864741] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1591865121] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1591865121] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (873908 kB) free!
[1591865254] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1591865254] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1591865571] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9699156 kB) free!
[1591865691] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9513596 kB) free!
[1591865786] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1591865786] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1591865812] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9340816 kB) free!
[1591866516] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1591866526] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1591866647] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1591866647] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (3032204 kB) free!
[1591867419] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1591867847] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 11.1% (3646492 kB) free!
[1591867991] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (26.9% change >= 20.0% threshold)
[1591868002] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (4048192 kB) free!
[1591868049] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 185.255.8.58 [11/Jun/2020:15:00:04 +0530] TCP 502 0 0 0.002 "10.147.212.154:13001, 10.147.212.143:13001, 10.147.212.155:13001, 10.147.212.142:13001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 0.001, 0.000"
[1591868110] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3885512 kB) free!
[1591868122] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1591868122] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.0% (16166172 kB) free.
[1591868169] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591868170] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.8% (3374128 kB) free!
[1591868274] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (26.9% change >= 20.0% threshold)
[1591868274] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1591868447] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2956764 kB) free!
[1591868602] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180036 kB) free!
[1591868704] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (32.5% change >= 20.0% threshold)
[1591868704] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1591869370] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191412 kB) free!
[1591869412] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4641960 kB) free!
[1591869802] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.1% (3692536 kB) free!
[1591869914] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2267032 kB) free!
[1591869990] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (30.9% change >= 20.0% threshold)
[1591870402] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1591870402] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2259608 kB) free!
[1591870514] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 41.4% (10141428 kB) free.
[1591870610] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1591870610] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1591870788] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1591870788] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1591870798] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1591870798] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1591871114] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4819684 kB) free!
[1591871170] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4143700 kB) free!
[1591871234] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4657036 kB) free!
[1591871246] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1591871276] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (29.3% change >= 20.0% threshold)
[1591871276] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1591871354] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1591871354] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4475352 kB) free!
[1591871461] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1591871461] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1591871461] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1591871602] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4431348 kB) free!
[1591871647] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1591871647] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1591871647] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (29.1% change >= 20.0% threshold)
[1591871659] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (3393008 kB) free!
[1591871769] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 172.105.59.173 [11/Jun/2020:16:02:23 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591871822] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1591871822] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1591871822] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (29.1% change >= 20.0% threshold)
[1591871890] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591871954] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 27.3% (6680392 kB) free.
[1591872201] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (27.8% change >= 20.0% threshold)
[1591872201] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (29.7% change >= 20.0% threshold)
[1591872201] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (29.1% change >= 20.0% threshold)
[1591872260] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198452 kB) free!
[1591872348] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (27.8% change >= 20.0% threshold)
[1591872348] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (29.7% change >= 20.0% threshold)
[1591872348] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (28.5% change >= 20.0% threshold)
[1591872370] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1591872370] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2178624 kB) free!
[1591872446] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1591872489] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 172.105.57.57 [11/Jun/2020:16:10:40 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1591872566] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1591872608] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (27.8% change >= 20.0% threshold)
[1591872608] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (29.0% change >= 20.0% threshold)
[1591872608] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1591872608] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (28.5% change >= 20.0% threshold)
[1591872618] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591872686] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1591872892] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/06/11 16:21:59 [error] 12489#0: *94191209 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1591873012] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591873412] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1623592 kB) free!
[1591873754] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4762780 kB) free!
[1591873874] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4642428 kB) free!
[1591873994] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (4559244 kB) free!
[1591874250] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (33.9% change >= 20.0% threshold)
[1591874250] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1591874250] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (32.6% change >= 20.0% threshold)
[1591876394] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1875120 kB) free!
[1591876522] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2985632 kB) free!
[1591876642] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.6% (2596652 kB) free!
[1591876762] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2214688 kB) free!
[1591877660] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (3778556 kB) free!
[1591878260] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1591878260] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183764 kB) free!
[1591879163] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16200404 kB) free.
[1591883170] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1591883613] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1591885785] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.25, 4.20, 4.96
[1591885994] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1591886385] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.79, 4.98, 5.03
[1591886505] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.96, 5.42, 5.19
[1591886625] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.35, 5.17, 5.11
[1591888089] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3313736 kB) free!
[1591888218] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 172.105.57.57 [11/Jun/2020:20:36:00 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1591888218] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.69, 5.92, 5.58
[1591888248] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1591888286] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1591888338] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.14, 5.31, 5.39
[1591888338] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591888368] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1591888425] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.57, 4.58, 4.95
[1591888458] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.04, 5.26, 5.37
[1591888764] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4119212 kB) free!
[1591888884] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.4% (3897660 kB) free!
[1591888887] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1591889004] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.6% (3722700 kB) free!
[1591889007] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 340 processes
[1591889058] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.83, 3.74, 4.60
[1591889127] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1591889447] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 32.3% (10563740 kB) free.
[1591889604] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16117844 kB) free.
[1591889626] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.94, 5.49, 5.18
[1591889660] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591889746] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.56, 6.11, 5.43
[1591889866] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1591889866] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.39, 5.42, 5.27
[1591890466] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.07, 4.22, 4.83
[1591893692] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.22889.4129> (10.147.212.155:42700 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1591893812] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1591894527] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1591896667] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 38.33, 36.30, 40.89