[1580581878] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.2% (29669976 kB) free.
[1580582309] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184264 kB) free!
[1580582430] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183208 kB) free!
[1580582550] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176940 kB) free!
[1580582713] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (1457336 kB) free!
[1580582833] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (1453828 kB) free!
[1580582953] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1453064 kB) free!
[1580584950] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.0% (11595872 kB) free.
[1580586150] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (1293332 kB) free!
[1580586270] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167996 kB) free!
[1580586390] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200988 kB) free!
[1580586576] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2020/02/02 01:15:57 [error] 10214#0: *47269102 connect() failed (111: Connection refused) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580586697] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580587447] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4587824 kB) free!
[1580587567] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4587112 kB) free!
[1580587687] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4586668 kB) free!
[1580589678] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.7% (10225352 kB) free.
[1580590400] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.3% (14878760 kB) free!
[1580590940] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 69.8% (34407052 kB) free.
[1580592696] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(252) < 2020/02/02 03:01:32 [error] 10221#0: *47297130 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580592816] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580592852] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(203) < 103.16.101.72 [02/Feb/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1580592973] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580593181] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3025) < 203.212.70.170 [02/Feb/2020:03:01:41 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1580593301] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580599878] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179388 kB) free!
[1580599998] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (160360 kB) free!
[1580600118] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175792 kB) free!
[1580600790] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.6% (11757752 kB) free.
[1580600953] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.3% (20665476 kB) free.
[1580604200] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.7% (37849652 kB) free.
[1580604448] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39945836 kB) free.
[1580604504] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22003432 kB) free.
[1580605591] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.6% (3355372 kB) free!
[1580605711] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.6% (3348488 kB) free!
[1580605831] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.5% (3341772 kB) free!
[1580608831] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2408188 kB) free!
[1580610816] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(72) < 2020/02/02 08:01:18 [error] 10221#0: *47381327 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580610936] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580610973] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1580610973] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.0% (21717520 kB) free.
[1580612228] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1580612349] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580612501] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [02/Feb/2020:08:22:39 +0530] TCP 502 0 0 63.008 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 3.006, 60.001, 0.000"
[1580612621] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580614650] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4912900 kB) free!
[1580614770] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.6% (4814848 kB) free!
[1580614890] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.2% (4714892 kB) free!
[1580616336] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(45) < 2020/02/02 09:31:07 [error] 10210#0: *47411092 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1580616456] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580617890] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2016208 kB) free!
[1580618377] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21653028 kB) free.
[1580618622] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(10) < 203.212.70.170 [02/Feb/2020:10:11:37 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1580618742] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580618856] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/02/02 10:10:32 [error] 10216#0: *47424917 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580618976] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580619028] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4903408 kB) free!
[1580619148] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.8% (4862596 kB) free!
[1580619268] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.7% (4825436 kB) free!
[1580619977] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1580622377] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4765808 kB) free!
[1580622498] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4581736 kB) free!
[1580622618] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (4487824 kB) free!
[1580623818] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1194384 kB) free!
[1580624060] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (3795952 kB) free!
[1580624180] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3571096 kB) free!
[1580624300] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3340532 kB) free!
[1580624900] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2156808 kB) free!
[1580625342] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(19) < 203.212.70.170 [02/Feb/2020:12:03:54 +0530] TCP 502 0 0 180.001 "10.147.212.154:3001, 10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 60.000, 60.000, 60.000"
[1580625463] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580625500] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.6% (21616088 kB) free.
[1580625500] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.5% (16053156 kB) free.
[1580625576] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(214) < 2020/02/02 12:02:46 [error] 10219#0: *47462900 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580625697] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580627326] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1580627447] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1580629110] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4625192 kB) free!
[1580629230] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4432504 kB) free!
[1580629350] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (4281820 kB) free!
[1580630859] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9410924 kB) free!
[1580630979] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9411080 kB) free!
[1580631099] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (9412068 kB) free!
[1580631150] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1577448 kB) free!
[1580632109] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (4504808 kB) free!
[1580632229] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (4401284 kB) free!
[1580632349] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (4038524 kB) free!
[1580632468] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.8% (2412988 kB) free!
[1580632949] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1580632949] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (1082392 kB) free!
[1580633379] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.0% (4667380 kB) free!
[1580633499] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.7% (4581140 kB) free!
[1580633619] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.3% (4487116 kB) free!
[1580633669] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9054680 kB) free!
[1580633789] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8844100 kB) free!
[1580633909] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8602780 kB) free!
[1580635100] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.6% (3943384 kB) free!
[1580635220] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.0% (3813088 kB) free!
[1580635340] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.7% (3745808 kB) free!
[1580635645] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.5% (22193336 kB) free.
[1580635680] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21228140 kB) free.
[1580635722] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 59.7% (14628840 kB) free.
[1580635941] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1580635941] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.5% (16056668 kB) free.
[1580636354] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4686384 kB) free!
[1580636474] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4645668 kB) free!
[1580636499] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3669920 kB) free!
[1580636594] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (4620920 kB) free!
[1580637819] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2302344 kB) free!
[1580638109] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4702720 kB) free!
[1580639540] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9689560 kB) free!
[1580639660] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9624404 kB) free!
[1580639780] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9550720 kB) free!
[1580642891] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1580644249] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21656460 kB) free.
[1580644288] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 75.6% (37269828 kB) free.
[1580644325] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 41.3% (20364780 kB) free.
[1580644994] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2291164 kB) free!
[1580645290] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166832 kB) free!
[1580645410] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199416 kB) free!
[1580645530] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188984 kB) free!
[1580647062] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(170) < 92.42.104.166 [02/Feb/2020:18:04:32 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1580647183] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580647341] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1580648180] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4678732 kB) free!
[1580649531] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4661336 kB) free!
[1580649651] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (4583864 kB) free!
[1580649771] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4514328 kB) free!
[1580651043] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.1% (21733824 kB) free.
[1580651075] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 39.7% (19547448 kB) free.
[1580651103] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.9% (20079224 kB) free.
[1580652771] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1948860 kB) free!
[1580653896] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(20) < 2020/02/02 20:01:28 [error] 10221#0: *47656217 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580654016] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580654383] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 172.105.59.173 [02/Feb/2020:20:02:04 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1580654503] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580655688] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13413836 kB) free!
[1580655808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44244248 kB) free.
[1580656408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13945768 kB) free!
[1580656528] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43725548 kB) free.
[1580657128] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12314708 kB) free!
[1580657248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43040640 kB) free.
[1580657848] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13676180 kB) free!
[1580657968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43141196 kB) free.
[1580658568] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14828660 kB) free!
[1580658688] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43404992 kB) free.
[1580659888] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21167824 kB) free!
[1580660008] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43165360 kB) free.
[1580660311] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 58.3% (14285304 kB) free.
[1580660348] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580660468] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580661208] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20626776 kB) free!
[1580661328] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42831564 kB) free.
[1580662528] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20348976 kB) free!
[1580662648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40613304 kB) free.
[1580662868] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1580662989] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580667079] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (248060 kB) free!
[1580667199] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240472 kB) free!
[1580667200] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.4% (17625404 kB) free!
[1580667319] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250512 kB) free!
[1580667320] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10058212 kB) free!
[1580667440] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2286392 kB) free!
[1580667448] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23448384 kB) free!
[1580667568] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40826872 kB) free.