[1586284325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12883792 kB) free!
[1586284445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8733552 kB) free!
[1586284565] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4937208 kB) free!
[1586284601] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (2667244 kB) free!
[1586285165] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19144636 kB) free!
[1586285204] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151528 kB) free!
[1586285208] SERVICE ALERT: sigqscluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1586285328] SERVICE ALERT: sigqscluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1586285531] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1586285765] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36090360 kB) free.
[1586285804] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (1984936 kB) free!
[1586286311] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1586286317] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (737152 kB) free!
[1586286365] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11784132 kB) free!
[1586286377] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (549776 kB) free!
[1586286437] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 2.0% (491116 kB) free!
[1586286486] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7832256 kB) free!
[1586286497] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (180388 kB) free!
[1586286606] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4027788 kB) free!
[1586287004] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1068364 kB) free!
[1586287206] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15683752 kB) free!
[1586287397] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 82.6% (20241032 kB) free.
[1586287604] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (1895868 kB) free!
[1586287806] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30467032 kB) free.
[1586288406] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8132756 kB) free!
[1586288526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4673900 kB) free!
[1586288646] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36713380 kB) free.
[1586288987] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (4020164 kB) free!
[1586289030] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 361 processes
[1586289151] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1586289246] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13568656 kB) free!
[1586289348] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4152656 kB) free!
[1586289366] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9764044 kB) free!
[1586289469] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241548 kB) free!
[1586289486] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6238564 kB) free!
[1586289509] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8620564 kB) free!
[1586289587] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1310556 kB) free!
[1586289589] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (245720 kB) free!
[1586289629] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9763332 kB) free!
[1586289749] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4083544 kB) free!
[1586290086] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16666584 kB) free!
[1586290686] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29008312 kB) free.
[1586291286] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8044776 kB) free!
[1586291406] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4608744 kB) free!
[1586291526] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32957928 kB) free.
[1586291549] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.2% (14735516 kB) free!
[1586291989] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 22.6% (11153680 kB) free.
[1586292127] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11415236 kB) free!
[1586292190] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(163) < 92.42.104.166 [08/Apr/2020:02:05:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1586292247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7811072 kB) free!
[1586292310] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586292367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37820180 kB) free.
[1586292749] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12767612 kB) free!
[1586292967] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14846020 kB) free!
[1586293087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11078136 kB) free!
[1586293207] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7642944 kB) free!
[1586293349] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1586293349] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.5% (15217744 kB) free!
[1586293807] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17235800 kB) free!
[1586293911] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 73.8% (36384568 kB) free.
[1586294407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27732356 kB) free.
[1586294566] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.8% (45734108 kB) free.
[1586295007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8183672 kB) free!
[1586295048] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(241) < 103.16.101.72 [08/Apr/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1586295127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35469972 kB) free.
[1586295169] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586295310] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2077) < 203.212.70.170 [08/Apr/2020:03:02:23 +0530] TCP 502 0 0 120.001 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1586295428] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(301) < 2020/04/08 03:01:20 [error] 22128#0: *82149450 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.144.246, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:229/17, bytes from/to upstream:17/229
[1586295431] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586295548] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586295587] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (3147332 kB) free!
[1586295727] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13275544 kB) free!
[1586295847] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9845432 kB) free!
[1586295967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37181504 kB) free.
[1586296230] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 22.4% (11042376 kB) free.
[1586296511] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.92, 4.41, 4.70
[1586296567] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14572332 kB) free!
[1586296687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10997208 kB) free!
[1586296807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37680316 kB) free.
[1586296830] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9658264 kB) free!
[1586296950] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (4755168 kB) free!
[1586297070] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3442668 kB) free!
[1586297408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15476740 kB) free!
[1586297528] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11904168 kB) free!
[1586297648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37683976 kB) free.
[1586298248] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17050724 kB) free!
[1586298368] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13386132 kB) free!
[1586298488] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35852468 kB) free.
[1586299004] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586299088] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17722360 kB) free!
[1586299187] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1982056 kB) free!
[1586299208] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13967016 kB) free!
[1586299328] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37683220 kB) free.
[1586299928] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18019988 kB) free!
[1586300048] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14352016 kB) free!
[1586300168] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37666064 kB) free.
[1586300768] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18164476 kB) free!
[1586300804] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3867000 kB) free.
[1586300888] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (14492936 kB) free!
[1586301008] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37663576 kB) free.
[1586301609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19209336 kB) free!
[1586301729] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15662024 kB) free!
[1586301849] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12364960 kB) free!
[1586302449] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16713372 kB) free!
[1586303049] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (13218024 kB) free!
[1586303531] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.6% (18771128 kB) free.
[1586303649] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31085252 kB) free.
[1586303976] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.5% (10981580 kB) free.
[1586303987] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 58.4% (14323952 kB) free.
[1586304249] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8088956 kB) free!
[1586304369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33650932 kB) free.
[1586304749] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586304969] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20453916 kB) free!
[1586305089] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16886600 kB) free!
[1586305209] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39810136 kB) free.
[1586305809] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21369620 kB) free!
[1586305929] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39916404 kB) free.
[1586306529] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24887072 kB) free!
[1586306649] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21436816 kB) free!
[1586306770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40361756 kB) free.
[1586307149] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.7% (47038160 kB) free.
[1586307369] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21915560 kB) free!
[1586307489] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41289256 kB) free.
[1586308089] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22959060 kB) free!
[1586308209] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41872340 kB) free.
[1586308809] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24001076 kB) free!
[1586308929] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41870756 kB) free.
[1586309529] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23480184 kB) free!
[1586309649] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19373744 kB) free!
[1586309769] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40097448 kB) free.
[1586310369] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20354668 kB) free!
[1586310489] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40681436 kB) free.
[1586311089] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20425268 kB) free!
[1586311209] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40952092 kB) free.
[1586311809] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19633704 kB) free!
[1586311930] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40906700 kB) free.
[1586312529] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22274976 kB) free!
[1586312649] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40691236 kB) free.
[1586313249] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20984108 kB) free!
[1586313369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40337668 kB) free.
[1586313497] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4696428 kB) free!
[1586313557] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4696148 kB) free!
[1586313617] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.1% (4692332 kB) free!
[1586313678] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.1% (4693216 kB) free!
[1586313911] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.16, 8.39, 6.09
[1586313969] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16460768 kB) free!
[1586314031] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.94, 7.91, 6.20
[1586314089] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (12053992 kB) free!
[1586314151] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.01, 7.23, 6.15
[1586314209] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37721088 kB) free.
[1586314776] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4876940 kB) free!
[1586314809] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13479912 kB) free!
[1586314897] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4245500 kB) free!
[1586314929] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9100252 kB) free!
[1586315017] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (4076172 kB) free!
[1586315049] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34310584 kB) free.
[1586315649] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9673660 kB) free!
[1586315769] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5318012 kB) free!
[1586315889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32641256 kB) free.
[1586315987] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4692956 kB) free!
[1586316108] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (4347240 kB) free!
[1586316111] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.7% (2843364 kB) free!
[1586316217] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2030572 kB) free!
[1586316228] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (4022692 kB) free!
[1586316231] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.6% (2384992 kB) free!
[1586316351] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1935360 kB) free!
[1586316489] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8029412 kB) free!
[1586316609] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3735648 kB) free!
[1586316730] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31335712 kB) free.
[1586316828] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2048440 kB) free!
[1586316951] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.4% (16027196 kB) free.
[1586317329] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5229576 kB) free!
[1586317449] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (943400 kB) free!
[1586317569] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30920104 kB) free.
[1586318169] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7909032 kB) free!
[1586318289] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3575156 kB) free!
[1586318409] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384044 kB) free!
[1586319609] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15413524 kB) free!
[1586321409] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8490184 kB) free!
[1586322009] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16925108 kB) free!
[1586322314] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1586322434] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1586322609] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29612800 kB) free.
[1586323004] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3186376 kB) free!
[1586323034] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1586323124] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3862236 kB) free.
[1586323154] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1586323209] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4285400 kB) free!
[1586323274] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1586323330] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399224 kB) free!
[1586323450] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396560 kB) free!
[1586323932] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4365536 kB) free!
[1586323949] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26360568 kB) free!
[1586324052] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (4199372 kB) free!
[1586324069] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.3% (29332864 kB) free.
[1586324151] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4219824 kB) free!
[1586324172] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (4018640 kB) free!
[1586324271] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (3731188 kB) free!
[1586324324] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2096172 kB) free!
[1586324392] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.5% (3265364 kB) free!
[1586324444] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3249844 kB) free!
[1586324564] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2527888 kB) free!
[1586324650] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24687140 kB) free!
[1586324669] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.5% (23099680 kB) free!
[1586324789] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.9% (27549956 kB) free.
[1586324992] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1586324992] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (968608 kB) free!
[1586325250] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394296 kB) free!
[1586325372] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2329388 kB) free!
[1586325389] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.0% (18406816 kB) free!
[1586325509] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.3% (24178848 kB) free!
[1586325629] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.6% (17892220 kB) free!
[1586325850] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22461388 kB) free!
[1586326450] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (11442804 kB) free!
[1586326911] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9569724 kB) free!
[1586327031] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9367124 kB) free!
[1586327151] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (9171812 kB) free!
[1586327429] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28075480 kB) free.
[1586327650] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (25044596 kB) free!
[1586327992] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.5% (16273556 kB) free.
[1586328029] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.0% (21049456 kB) free!
[1586328149] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26478428 kB) free.
[1586328250] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2119504 kB) free!
[1586328674] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 335 processes
[1586328749] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.0% (18396076 kB) free!
[1586328794] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1586328850] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13256656 kB) free!
[1586328869] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.5% (23039444 kB) free!
[1586328989] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.3% (25414888 kB) free!
[1586329148] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/08 12:27:19 [error] 22135#0: *82267121 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1586329268] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586329364] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163356 kB) free!
[1586329589] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1586329589] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.0% (31608128 kB) free.
[1586329994] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1586330050] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2351676 kB) free!
[1586330114] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1586330566] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9441704 kB) free!
[1586330650] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14641572 kB) free!
[1586330686] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9348012 kB) free!
[1586330789] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25033900 kB) free!
[1586330806] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9240728 kB) free!
[1586330909] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (28857856 kB) free.
[1586331251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26526516 kB) free.
[1586331851] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3468352 kB) free!
[1586331971] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396912 kB) free!
[1586332030] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(876) < 103.255.146.164 [08/Apr/2020:13:15:16 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1586332091] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30102316 kB) free.
[1586332150] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586332364] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 35.2% (5718204 kB) free.
[1586332691] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5007476 kB) free!
[1586332709] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.6% (23233216 kB) free!
[1586332811] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (919196 kB) free!
[1586332829] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.7% (24667908 kB) free!
[1586332931] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31798288 kB) free.
[1586332949] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.8% (31335524 kB) free.
[1586333151] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4645988 kB) free!
[1586333531] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8107320 kB) free!
[1586333651] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3677712 kB) free!
[1586333771] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391236 kB) free!
[1586334149] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.1% (23812168 kB) free!
[1586334164] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2164972 kB) free!
[1586334269] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (28265876 kB) free.
[1586334285] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.0% (1951996 kB) free!
[1586334405] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1586334405] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (1710652 kB) free!
[1586334971] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16019804 kB) free!
[1586335006] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4436996 kB) free!
[1586335469] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.1% (19904056 kB) free!
[1586335571] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (413900 kB) free!
[1586335589] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (26755032 kB) free.
[1586335818] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1586335938] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1586336058] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1586336068] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1586336188] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1586336308] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1586336771] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26821264 kB) free.
[1586336789] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.7% (21952472 kB) free!
[1586336805] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1184832 kB) free!
[1586336908] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1586336909] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (28595072 kB) free.
[1586336992] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1586337224] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 13 users currently logged in
[1586337258] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1586337304] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 73.7% (18073084 kB) free.
[1586337344] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 13 users currently logged in
[1586337372] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1686664 kB) free!
[1586337464] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1586337492] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (420084 kB) free!
[1586337509] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.3% (20222452 kB) free!
[1586337592] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3417496 kB) free!
[1586337612] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400084 kB) free!
[1586337629] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (29676636 kB) free.
[1586337668] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/04/08 14:50:55 [error] 22130#0: *82300115 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1586337713] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3210872 kB) free!
[1586337788] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586337833] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.4% (3000368 kB) free!
[1586338005] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (1783576 kB) free!
[1586338229] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.7% (23288216 kB) free!
[1586338349] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.4% (29511580 kB) free.
[1586338433] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.2% (16222836 kB) free.
[1586338750] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(128) < 92.42.104.166 [08/Apr/2020:15:00:15 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1586338812] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25801576 kB) free!
[1586338870] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586338949] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (24850244 kB) free!
[1586338988] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/04/08 15:12:10 [error] 22131#0: *82304572 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1586339069] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.9% (18295056 kB) free!
[1586339108] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586339190] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.5% (25680452 kB) free!
[1586339412] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (851856 kB) free!
[1586339865] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 13 users currently logged in
[1586339985] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 13 users currently logged in
[1586340012] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13696328 kB) free!
[1586340105] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 13 users currently logged in
[1586340612] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30764280 kB) free.
[1586340705] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 11 users currently logged in
[1586341212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4682844 kB) free!
[1586341332] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374892 kB) free!
[1586341452] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400088 kB) free!
[1586341590] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.2% (29252928 kB) free.
[1586342108] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/08 16:03:21 [error] 22128#0: *82315695 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1586342190] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.8% (22114156 kB) free!
[1586342228] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586342310] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.7% (15412940 kB) free!
[1586342430] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.6% (23254756 kB) free!
[1586342652] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29189908 kB) free.
[1586342805] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (949860 kB) free!
[1586343252] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4275480 kB) free!
[1586343337] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4892928 kB) free!
[1586343372] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397028 kB) free!
[1586343405] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2846464 kB) free!
[1586343428] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/04/08 16:18:55 [error] 22135#0: *82319330 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1586343457] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.9% (4877996 kB) free!
[1586343493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34515724 kB) free.
[1586343548] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586343577] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.8% (4855164 kB) free!
[1586344006] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1619668 kB) free!
[1586344092] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9346504 kB) free!
[1586344212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4682348 kB) free!
[1586344230] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (31812496 kB) free.
[1586344332] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405228 kB) free!
[1586344606] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.8% (2401784 kB) free!
[1586344830] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.7% (24594772 kB) free!
[1586344932] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15808364 kB) free!
[1586344950] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (30253168 kB) free.
[1586345206] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (269360 kB) free!
[1586345532] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35325960 kB) free.
[1586345550] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.7% (23338652 kB) free!
[1586345670] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (28670376 kB) free.
[1586345806] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (2056776 kB) free!
[1586346132] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8330040 kB) free!
[1586346252] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3694368 kB) free!
[1586346270] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.8% (22180384 kB) free!
[1586346372] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400948 kB) free!
[1586346390] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27167072 kB) free.
[1586346972] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15896280 kB) free!
[1586347006] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1155436 kB) free!
[1586347572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36433208 kB) free.
[1586347590] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (25173924 kB) free!
[1586347606] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (2507040 kB) free!
[1586347710] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (30815068 kB) free.
[1586348033] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3658504 kB) free!
[1586348153] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.0% (3144344 kB) free!
[1586348172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8354176 kB) free!
[1586348206] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1513472 kB) free!
[1586348273] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.7% (2633064 kB) free!
[1586348292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3446796 kB) free!
[1586348310] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.8% (22080416 kB) free!
[1586348412] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382692 kB) free!
[1586348430] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (28228112 kB) free.
[1586348806] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (2927676 kB) free!
[1586348873] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1586348873] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186332 kB) free!
[1586348951] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 37.74, 35.88, 22.79
[1586349012] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17836188 kB) free!
[1586349031] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (22633188 kB) free!
[1586349151] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.7% (24606292 kB) free!
[1586349271] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.0% (17083216 kB) free!
[1586349304] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.0% (3430836 kB) free!
[1586349406] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (811592 kB) free!
[1586349425] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 12.2% (2995932 kB) free!
[1586349474] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16125312 kB) free.
[1586349545] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 10.4% (2559576 kB) free!
[1586349612] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387148 kB) free!
[1586350006] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (2976392 kB) free!
[1586350145] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (448212 kB) free!
[1586350212] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15685920 kB) free!
[1586350606] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3290276 kB) free.
[1586350812] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407328 kB) free!
[1586351071] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12720216 kB) free!
[1586351206] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (1002492 kB) free!
[1586351327] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (1566656 kB) free!
[1586351447] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (2087856 kB) free!
[1586352047] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (660924 kB) free!
[1586352647] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (3097932 kB) free!
[1586352871] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.7% (14097284 kB) free!
[1586353847] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (868264 kB) free!
[1586354447] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3271336 kB) free.
[1586355047] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3037292 kB) free!
[1586355167] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3002104 kB) free!
[1586355287] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (653156 kB) free!
[1586355551] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.02, 13.55, 17.13
[1586355612] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35804928 kB) free.
[1586355887] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3029432 kB) free!
[1586356212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6210072 kB) free!
[1586356332] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1200364 kB) free!
[1586356452] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401068 kB) free!
[1586356674] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4150912 kB) free!
[1586356794] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.6% (3942644 kB) free!
[1586356914] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.7% (3745228 kB) free!
[1586357052] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22200100 kB) free!
[1586357087] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1617752 kB) free!
[1586357653] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (463636 kB) free!
[1586357687] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.0% (4227148 kB) free.
[1586358114] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (197396 kB) free!
[1586358253] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20310884 kB) free!
[1586358853] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412628 kB) free!
[1586359453] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (19989156 kB) free!
[1586359515] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1586359635] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1586359914] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.7% (16313812 kB) free.
[1586360053] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412168 kB) free!
[1586360653] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20533488 kB) free!
[1586360687] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3182576 kB) free!
[1586360807] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (3173232 kB) free!
[1586360927] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1442204 kB) free!
[1586361253] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390504 kB) free!
[1586361527] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (3045504 kB) free!
[1586361853] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21131000 kB) free!
[1586362453] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411568 kB) free!
[1586363053] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19391628 kB) free!
[1586363071] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1586363120] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1600 processes
[1586363653] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37689648 kB) free.
[1586363671] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (28849608 kB) free.
[1586363720] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1795 processes
[1586363951] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.71, 3.24, 4.25
[1586364253] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15948276 kB) free!
[1586364373] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11814972 kB) free!
[1586364493] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (7867656 kB) free!
[1586364871] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25027388 kB) free!
[1586364991] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.1% (19845980 kB) free!
[1586365093] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24915308 kB) free!
[1586365111] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.1% (25212412 kB) free!
[1586365694] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3978676 kB) free!
[1586365727] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (939512 kB) free!
[1586366294] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20291508 kB) free!
[1586366311] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (27937536 kB) free.
[1586366327] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2767880 kB) free!
[1586366352] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 16.34, 13.29, 8.78
[1586366472] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 16.70, 14.19, 9.64
[1586366592] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 21.97, 16.30, 10.92
[1586366894] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397616 kB) free!
[1586366911] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (24742816 kB) free!
[1586366928] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3265556 kB) free.
[1586367031] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.6% (19299888 kB) free!
[1586367151] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1586367151] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.0% (25032356 kB) free!
[1586367494] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14999536 kB) free!
[1586367528] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (2265476 kB) free!
[1586367648] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2911000 kB) free!
[1586367768] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3521760 kB) free.
[1586368094] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37118440 kB) free.
[1586368351] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.2% (31948776 kB) free.
[1586368368] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.4% (2185248 kB) free!
[1586368488] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2838144 kB) free!
[1586368608] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3538508 kB) free.
[1586368694] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14584612 kB) free!
[1586368814] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10524524 kB) free!
[1586368934] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6620728 kB) free!
[1586369208] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2085468 kB) free!
[1586369328] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (2685924 kB) free!
[1586369448] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (3254944 kB) free!
[1586369534] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23942936 kB) free!
[1586370048] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.1% (4396884 kB) free.
[1586370114] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1586370134] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4753500 kB) free!