[1569523175] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187596 kB) free!
[1569523214] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199332 kB) free!
[1569523248] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190936 kB) free!
[1569523297] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184832 kB) free!
[1569523334] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183508 kB) free!
[1569523369] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182468 kB) free!
[1569523417] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183000 kB) free!
[1569523454] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173756 kB) free!
[1569523489] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187400 kB) free!
[1569525500] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569526101] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569527864] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (7664056 kB) free!
[1569527984] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.5% (7131588 kB) free!
[1569528104] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.0% (5918716 kB) free!
[1569529102] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569529702] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1569529702] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569531497] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244176 kB) free!
[1569531617] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244332 kB) free!
[1569531738] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220472 kB) free!
[1569532338] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.5% (45589952 kB) free.
[1569532703] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569533303] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569534254] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 64.3% (15763768 kB) free.
[1569534289] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 53.6% (13136200 kB) free.
[1569537503] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569538103] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569541103] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569541703] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569543503] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569544103] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569544703] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569545303] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569545367] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 505 processes
[1569545427] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 502 processes
[1569545487] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 502 processes
[1569545547] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 498 processes
[1569546504] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569547104] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569547647] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 506 processes
[1569547707] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 509 processes
[1569547767] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 511 processes
[1569547827] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 502 processes
[1569548018] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.9% (10837848 kB) free.
[1569548618] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188744 kB) free!
[1569548664] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569548695] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187792 kB) free!
[1569548738] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176732 kB) free!
[1569548785] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172608 kB) free!
[1569548815] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188440 kB) free!
[1569548859] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178112 kB) free!
[1569548905] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173312 kB) free!
[1569548935] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195620 kB) free!
[1569549104] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1759816 kB) free!
[1569552504] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569553104] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569557304] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569557904] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569559704] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569560305] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569562671] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.71, 7.20, 5.73
[1569562791] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.75, 6.29, 5.58
[1569562911] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.78, 7.75, 6.18
[1569564505] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1569565527] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 558 processes
[1569565587] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 558 processes
[1569565648] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 556 processes
[1569565708] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;HARD;4;CPU CRITICAL: 1 crit, 0 warn out of 560 processes
[1569566008] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 568 processes
[1569566687] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.71, 13.83, 6.73
[1569566760] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.95, 13.83, 8.43
[1569566807] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.38, 9.93, 6.16
[1569566876] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.56, 11.48, 8.26
[1569566880] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.87, 10.09, 7.72
[1569566927] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.44, 7.39, 5.69
[1569566996] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.50, 8.09, 7.41
[1569567000] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.58, 7.60, 7.09
[1569567116] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.21, 5.85, 6.67
[1569567505] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569567527] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.78, 2.55, 3.84
[1569567600] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.29, 3.09, 4.88
[1569567716] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.20, 2.96, 4.58
[1569567738] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9459648 kB) free!
[1569567858] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8852236 kB) free!
[1569567978] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8264976 kB) free!
[1569568316] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 122.47, 34.43, 14.33
[1569568436] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 17.93, 23.55, 12.77
[1569568556] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.68, 16.28, 11.42
[1569568727] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.53, 6.81, 7.33
[1569568788] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 601 processes
[1569568800] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.84, 17.62, 12.07
[1569568848] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.20, 4.89, 6.56
[1569568909] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 598 processes
[1569568920] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.84, 12.06, 10.71
[1569568968] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 65.29, 17.41, 10.38
[1569569029] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1569569040] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 66.43, 27.30, 16.01
[1569569156] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1569569156] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 27.49, 27.59, 18.88
[1569569178] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2839824 kB) free!
[1569569336] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4332896 kB) free!
[1569569511] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 21.09, 22.65, 20.47
[1569569568] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1569569568] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.83, 12.00, 12.13
[1569569640] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1569569640] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.70, 11.58, 13.72
[1569570164] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.7% (22388348 kB) free.
[1569570168] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.79, 27.29, 20.29
[1569570219] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.3% (38096604 kB) free.
[1569570229] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 633 processes
[1569570240] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 18.29, 26.93, 21.06
[1569570251] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 74.5% (18257064 kB) free.
[1569570307] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 25.6% (12638980 kB) free.
[1569570711] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.90, 13.44, 16.94
[1569570774] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180628 kB) free!
[1569570894] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198244 kB) free!
[1569571014] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177540 kB) free!
[1569571069] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1569571136] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2111020 kB) free!
[1569571189] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 331 processes
[1569571309] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1569571913] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 752 processes
[1569571968] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.49, 17.14, 19.50
[1569572033] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 340 processes
[1569572640] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.08, 8.70, 15.73
[1569572757] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.89, 5.91, 15.31
[1569573859] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4048568 kB) free!
[1569573979] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.2% (3576632 kB) free!
[1569574099] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (3060172 kB) free!
[1569574368] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.79, 0.90, 3.23
[1569574440] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.65, 1.00, 3.56
[1569574557] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.87, 1.77, 3.81
[1569574699] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (838308 kB) free!
[1569574906] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1569575026] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1569575146] SERVICE ALERT: primary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1569575151] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1569576505] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1569576951] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1569577099] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.6% (23014136 kB) free.
[1569577188] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1569577308] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569577308] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1569577311] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.96, 3.91, 4.99
[1569577368] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1569577428] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1569577428] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1569577440] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.05, 7.77, 5.73
[1569577488] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1569577560] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.61, 7.69, 5.95
[1569577680] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.56, 7.62, 6.14
[1569578511] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 14.64, 9.92, 7.10
[1569578631] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.43, 9.72, 7.37
[1569578751] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1569578751] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.94, 7.90, 6.99
[1569578880] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.86, 2.50, 3.87
[1569578988] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1569579226] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4157284 kB) free!
[1569579228] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1569579346] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (4158092 kB) free!
[1569579351] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.37, 3.23, 4.89
[1569579466] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4157624 kB) free!
[1569579588] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569579648] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1569579708] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1569579768] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1569580309] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1569580429] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1569580551] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.44, 8.03, 6.32
[1569580595] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4883552 kB) free!
[1569580671] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.51, 8.02, 6.52
[1569580715] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.5% (4292876 kB) free!
[1569580791] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.93, 8.73, 6.96
[1569580835] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.5% (4286392 kB) free!
[1569581814] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.4% (21805556 kB) free.
[1569581833] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.5% (23407256 kB) free.
[1569581866] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 68.6% (33810468 kB) free.
[1569581888] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 93.5% (22911888 kB) free.
[1569582256] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4885636 kB) free!
[1569582346] SERVICE ALERT: primary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1569582376] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4745140 kB) free!
[1569582423] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193060 kB) free!
[1569582467] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (222632 kB) free!
[1569582496] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (4600344 kB) free!
[1569582544] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164740 kB) free!
[1569582587] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237276 kB) free!
[1569582664] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193600 kB) free!
[1569582707] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232304 kB) free!
[1569583191] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 6.60, 4.88, 4.95
[1569583968] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1569584391] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.61, 4.66, 5.01
[1569584511] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.22, 4.50, 4.91
[1569584757] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1569584896] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2248008 kB) free!
[1569585314] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9788224 kB) free!
[1569585434] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9674564 kB) free!
[1569585554] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.2% (9473576 kB) free!
[1569586968] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 51 users currently logged in
[1569588238] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 331 processes
[1569589438] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1569589680] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569589754] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4312124 kB) free!
[1569590511] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.32, 10.71, 6.94
[1569590631] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 15.39, 12.10, 7.91
[1569590752] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 20.67, 14.81, 9.37
[1569592430] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1569592550] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 335 processes
[1569592670] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1569595070] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1569595190] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1569596752] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1569597352] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.35, 3.47, 4.94
[1569601633] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (4576900 kB) free!
[1569601753] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4491104 kB) free!
[1569601873] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4385664 kB) free!
[1569606073] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2273520 kB) free!