[1558118071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11563188 kB) free.
[1558118632] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558119832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558119871] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9648108 kB) free!
[1558119953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558119991] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8100004 kB) free!
[1558120073] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558120111] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6425200 kB) free!
[1558123111] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10367064 kB) free.
[1558123673] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558124873] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558124911] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6193368 kB) free!
[1558124993] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558125031] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11541616 kB) free.
[1558125593] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558125631] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6417428 kB) free!
[1558125713] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558125751] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11554056 kB) free.
[1558126313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558126351] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6779432 kB) free!
[1558126433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558126471] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11556636 kB) free.
[1558126553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558127071] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6925864 kB) free!
[1558127154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558127192] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547068 kB) free.
[1558127274] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558127719] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 209.99, 80.80, 33.78
[1558127743] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558127791] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7251436 kB) free!
[1558127839] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 266.55, 145.17, 63.07
[1558127874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558127911] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11562248 kB) free.
[1558127934] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558127960] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 274.11, 188.21, 88.80
[1558127994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1558128511] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7339068 kB) free!
[1558128536] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.6% (3615104 kB) free!
[1558128594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558128631] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11551844 kB) free.
[1558128714] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558129231] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7233796 kB) free!
[1558129314] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558129351] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10827440 kB) free.
[1558129435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558129760] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.25, 1.00, 16.65
[1558129951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7366600 kB) free!
[1558130034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558130071] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5256284 kB) free!
[1558130154] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558130191] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11531408 kB) free.
[1558130754] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558130791] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5409300 kB) free!
[1558130874] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558130911] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11540084 kB) free.
[1558130960] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.86, 0.63, 4.97
[1558131474] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558131512] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5602824 kB) free!
[1558131594] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558131632] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11528932 kB) free.
[1558132194] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558132232] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5696404 kB) free!
[1558132315] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558132352] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11541040 kB) free.
[1558132914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558132952] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5855364 kB) free!
[1558133034] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558133072] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11537828 kB) free.
[1558133634] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558133672] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6067816 kB) free!
[1558133754] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1558133754] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558133792] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1558133792] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11527992 kB) free.
[1558134354] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558134392] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6661744 kB) free!
[1558134474] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558134512] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11540312 kB) free.
[1558134594] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558134977] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 45.23, 11.04, 4.79
[1558135097] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 253.33, 101.95, 38.86
[1558135112] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6962364 kB) free!
[1558135194] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558135217] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 281.03, 162.68, 68.76
[1558135233] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11534320 kB) free.
[1558135315] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558135767] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2779568 kB) free!
[1558135832] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7103216 kB) free!
[1558135914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558135952] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11545056 kB) free.
[1558136034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558136154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558137152] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5314372 kB) free!
[1558137272] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11538528 kB) free.
[1558137354] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558137617] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.50, 1.75, 14.85
[1558137872] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5607188 kB) free!
[1558137992] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11527176 kB) free.
[1558138555] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558138675] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558138767] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2138132 kB) free!
[1558138795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558139192] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8497992 kB) free!
[1558139312] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6866132 kB) free!
[1558139417] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.68, 1.84, 3.73
[1558139433] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11212212 kB) free.
[1558139995] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1558139995] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558140006] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2671592 kB) free!
[1558140595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558140632] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9272940 kB) free!
[1558140715] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558140752] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7687664 kB) free!
[1558140835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558140872] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6012120 kB) free!
[1558141472] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11440852 kB) free.
[1558142672] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8855152 kB) free!
[1558142792] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7307784 kB) free!
[1558142912] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5297780 kB) free!
[1558143512] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1558143512] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10912544 kB) free.
[1558144112] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9395500 kB) free!
[1558144206] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2424424 kB) free!
[1558144233] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7934212 kB) free!
[1558144353] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6247688 kB) free!
[1558144435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558144953] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11270292 kB) free.
[1558145035] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558145037] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558145155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558145275] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558145553] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9704024 kB) free!
[1558145673] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8201744 kB) free!
[1558145793] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.5% (6656060 kB) free!
[1558145991] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 96.19, 22.78, 7.69
[1558146111] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 257.29, 108.71, 41.01
[1558146231] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 279.01, 166.15, 70.27
[1558146393] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11252124 kB) free.
[1558146606] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.4% (3036436 kB) free!
[1558147136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (668556 kB) free!
[1558147206] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1558147206] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (572104 kB) free!
[1558147593] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8705548 kB) free!
[1558147713] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7155852 kB) free!
[1558147833] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4987564 kB) free!
[1558148031] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 1.06, 17.18
[1558148433] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10754328 kB) free.
[1558148875] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1558149034] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9227640 kB) free!
[1558149154] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7703960 kB) free!
[1558149231] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.27, 4.88
[1558149274] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6003456 kB) free!
[1558149475] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558149595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558149716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558149874] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11215788 kB) free.
[1558150916] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1558151074] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9266080 kB) free!
[1558151194] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7578652 kB) free!
[1558151314] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5912964 kB) free!
[1558151516] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558151636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558151756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558151914] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4267004 kB) free!
[1558152514] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10544900 kB) free.
[1558152674] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1558152744] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22943904 kB) free.
[1558152780] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.7% (22728108 kB) free.
[1558152825] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.6% (14719320 kB) free.
[1558152864] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.7% (22481400 kB) free.
[1558152905] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.9% (14282964 kB) free.
[1558152956] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558153114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8967016 kB) free!
[1558153234] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7353620 kB) free!
[1558153355] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.0% (5445272 kB) free!
[1558153557] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558153677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558153797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558153955] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11074612 kB) free.
[1558154555] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9477508 kB) free!
[1558154675] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7689688 kB) free!
[1558154795] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6007376 kB) free!
[1558155395] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3949720 kB) free!
[1558155995] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10742016 kB) free.
[1558156595] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9261088 kB) free!
[1558156715] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7575604 kB) free!
[1558156760] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 41.72, 9.52, 3.51
[1558156835] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5931596 kB) free!
[1558156880] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 236.79, 94.78, 35.48
[1558157000] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 263.30, 151.92, 63.61
[1558157109] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.8% (2239768 kB) free!
[1558157229] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (163080 kB) free!
[1558157349] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166860 kB) free!
[1558157631] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.32, 6.09, 5.68
[1558157751] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.48, 4.29, 5.07
[1558157871] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 2.98, 4.50
[1558157998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1558158035] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10228140 kB) free.
[1558158269] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1558158343] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558158463] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558158598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1558158636] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8620776 kB) free!
[1558158718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558158756] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6884336 kB) free!
[1558158800] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 1.62, 18.53
[1558158838] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558158876] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5185592 kB) free!
[1558159063] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558159126] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.4% (14532744 kB) free.
[1558159184] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558159476] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4395540 kB) free!
[1558159671] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.88, 6.36, 5.32
[1558159783] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558159791] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 4.29, 4.69
[1558159903] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1558160038] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1558160076] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10026328 kB) free.
[1558160330] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166548 kB) free!
[1558160450] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165604 kB) free!
[1558160570] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1558160570] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165560 kB) free!
[1558160600] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.50, 0.79, 3.43
[1558160638] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558160676] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8829584 kB) free!
[1558160758] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558160796] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7056480 kB) free!
[1558160878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558160916] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5294636 kB) free!
[1558161516] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3166596 kB) free!
[1558161829] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.8% (2400800 kB) free!
[1558161938] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.20, 5.90, 5.34
[1558161949] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 12.9% (2103068 kB) free!
[1558161961] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.3% (14521768 kB) free.
[1558162016] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 43.4% (21369140 kB) free.
[1558162058] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.28, 5.83, 5.39
[1558162069] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.3% (1832160 kB) free!
[1558162078] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558162178] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.97, 5.25, 5.24
[1558162567] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178736 kB) free!
[1558162669] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (660604 kB) free!
[1558162678] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558162687] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176992 kB) free!
[1558162779] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.22, 4.38, 4.79
[1558162798] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558162807] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176352 kB) free!
[1558162918] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558163379] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.07, 5.47, 5.16
[1558163499] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.35, 5.77, 5.30
[1558163620] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.93, 5.59, 5.30
[1558164118] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558164719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558164820] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1558164820] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.61, 4.90, 4.95
[1558164839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558164959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558165613] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.7% (14749520 kB) free.
[1558165656] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.5% (14550752 kB) free.
[1558166159] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558166257] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178048 kB) free!
[1558166377] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178708 kB) free!
[1558166497] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (177660 kB) free!
[1558166620] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.08, 5.81, 5.30
[1558166740] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.84, 5.19, 5.15
[1558166759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558166860] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.34, 4.51, 4.90
[1558166879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558166999] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558168199] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1558168799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558168919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558169039] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558169417] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 317.67, 199.74, 87.37
[1558169537] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 319.77, 239.56, 115.58
[1558169657] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 319.73, 266.16, 140.37
[1558170223] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (1360584 kB) free!
[1558170239] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1558170343] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.1% (1259508 kB) free!
[1558170463] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 16.0% (1178668 kB) free!
[1558170839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558170960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558171063] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (737088 kB) free!
[1558171080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558172057] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.67, 1.99, 13.91
[1558172263] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 89.1% (6552688 kB) free.
[1558172678] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.3% (14520784 kB) free.
[1558172880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558173103] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 354 processes
[1558173223] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558173286] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176412 kB) free!
[1558173406] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173432 kB) free!
[1558173416] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558173480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558173526] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174268 kB) free!
[1558173600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558173720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558173857] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.81, 1.88, 3.65
[1558174018] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3244156 kB) free!
[1558174138] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (3123948 kB) free!
[1558174258] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (2986060 kB) free!
[1558175860] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558176058] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1074236 kB) free!
[1558176200] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 294.30, 186.33, 81.18
[1558176320] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 195.37, 199.06, 99.44
[1558176440] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 27.30, 133.57, 87.52
[1558176534] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2756024 kB) free!
[1558176654] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2613100 kB) free!
[1558176774] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.2% (2470924 kB) free!
[1558177493] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1558177614] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1558177766] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46665 MB (20% inode=99%):
[1558177886] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46593 MB (20% inode=99%):
[1558177920] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1558177974] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1344228 kB) free!
[1558178006] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46417 MB (20% inode=99%):
[1558178240] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.64, 1.00, 13.34
[1558178520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558178641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558178761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558179417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9558820 kB) free!
[1558179440] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.42, 3.97
[1558179537] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.4% (15975008 kB) free.
[1558179946] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.7% (14573200 kB) free.
[1558179961] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1558180005] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 86.5% (14060512 kB) free.
[1558180057] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.3% (14846476 kB) free.
[1558180561] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558180681] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558180801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558185601] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558185946] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1558186201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1558186321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558186442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558187391] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 322.33, 207.82, 91.18
[1558187511] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 324.02, 246.35, 119.44
[1558187631] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.60, 168.55, 106.52
[1558189137] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9562892 kB) free!
[1558189257] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.9% (15710192 kB) free.
[1558189431] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.72, 2.46, 16.59
[1558189988] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4701832 kB) free!
[1558190109] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4635336 kB) free!
[1558190229] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4581540 kB) free!
[1558190631] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 0.38, 4.81
[1558191842] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1558192066] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (2911912 kB) free!
[1558192186] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (198036 kB) free!
[1558192257] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9220824 kB) free!
[1558192307] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168992 kB) free!
[1558192378] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.6% (15584504 kB) free.
[1558192423] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 346 processes
[1558192442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558192544] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 346 processes
[1558192546] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171464 kB) free!
[1558192562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558192664] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 345 processes
[1558192666] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (147916 kB) free!
[1558192682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558192786] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169136 kB) free!
[1558192907] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 53.4% (8675068 kB) free.
[1558193006] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 91344 MB (40% inode=99%):
[1558193264] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558193882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1558194483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558194603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558194723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558195923] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1558196523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558196643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558196763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558197963] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558198563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558198683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558198803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558200003] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558200603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558200723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558200844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558201977] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9338152 kB) free!
[1558202044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1558202097] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.6% (15089848 kB) free.
[1558202644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558202764] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558202884] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1558204085] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes