[1584124335] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23769692 kB) free!
[1584124767] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (203196 kB) free!
[1584124788] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584124885] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199928 kB) free!
[1584124887] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192084 kB) free!
[1584124935] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3365748 kB) free!
[1584125005] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206980 kB) free!
[1584125007] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200768 kB) free!
[1584125125] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193028 kB) free!
[1584125535] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27389976 kB) free.
[1584126146] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5775020 kB) free!
[1584126275] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1793188 kB) free!
[1584126396] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395144 kB) free!
[1584126448] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(247) < 103.255.146.164 [14/Mar/2020:00:31:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584126490] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195736 kB) free!
[1584126550] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197620 kB) free!
[1584126568] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584126610] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (187032 kB) free!
[1584126670] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (196068 kB) free!
[1584126970] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 77.8% (19077200 kB) free.
[1584126996] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18569980 kB) free!
[1584127596] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404376 kB) free!
[1584127804] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (2472676 kB) free!
[1584127924] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.5% (2201828 kB) free!
[1584127934] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437204 kB) free!
[1584127964] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584128044] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 13.5% (2190148 kB) free!
[1584128196] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20037880 kB) free!
[1584128347] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584128796] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1378208 kB) free!
[1584129223] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (22711840 kB) free!
[1584129343] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.7% (14145828 kB) free!
[1584129396] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21491328 kB) free!
[1584129463] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (4015076 kB) free!
[1584129673] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233828 kB) free!
[1584129793] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238680 kB) free!
[1584129913] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240812 kB) free!
[1584129996] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (2987076 kB) free!
[1584130597] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21227260 kB) free!
[1584131197] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2630384 kB) free!
[1584131714] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.3% (30206560 kB) free.
[1584131797] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20389160 kB) free!
[1584131968] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(205) < 172.105.59.173 [14/Mar/2020:02:04:57 +0530] TCP 502 0 0 240.002 "10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 60.001, 60.000"
[1584132088] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584132397] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (2035500 kB) free!
[1584132439] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1584132559] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1584132679] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1584132912] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (247716 kB) free!
[1584132997] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19515436 kB) free!
[1584133032] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213052 kB) free!
[1584133152] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219744 kB) free!
[1584133436] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232936 kB) free!
[1584133556] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214088 kB) free!
[1584133597] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1384944 kB) free!
[1584133676] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224116 kB) free!
[1584134197] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18373436 kB) free!
[1584134276] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.3% (41524768 kB) free.
[1584134352] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 78.2% (38539312 kB) free.
[1584134797] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (883412 kB) free!
[1584135079] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584135088] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1832) < 182.18.144.246 [14/Mar/2020:03:01:26 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584135209] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(108) < 172.105.59.173 [14/Mar/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1584135329] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584135392] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(336) < 2020/03/14 03:01:32 [error] 12253#0: *70258451 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584135397] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15572000 kB) free!
[1584135464] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(272) < 103.16.101.72 [14/Mar/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1584135512] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584135585] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584135997] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34623312 kB) free.
[1584136407] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 48.1% (11796844 kB) free.
[1584136597] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12641592 kB) free!
[1584136717] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9247452 kB) free!
[1584136837] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5896004 kB) free!
[1584137437] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14376744 kB) free!
[1584138038] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417032 kB) free!
[1584138638] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17156788 kB) free!
[1584138679] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1584138799] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1584140438] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32383676 kB) free.
[1584141038] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12977636 kB) free!
[1584141158] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9644288 kB) free!
[1584141278] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6385100 kB) free!
[1584141878] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22892220 kB) free!
[1584142407] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179796 kB) free!
[1584142478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35311332 kB) free.
[1584142527] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (140596 kB) free!
[1584142647] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168988 kB) free!
[1584143078] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16276752 kB) free!
[1584143198] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12978708 kB) free!
[1584143318] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (9734936 kB) free!
[1584143919] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23273348 kB) free!
[1584143947] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 51.4% (12676948 kB) free.
[1584144325] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.3% (18706172 kB) free.
[1584144519] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37098836 kB) free.
[1584145119] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20560304 kB) free!
[1584145239] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17347760 kB) free!
[1584145359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40229420 kB) free.
[1584145959] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24261108 kB) free!
[1584146079] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (21004144 kB) free!
[1584146199] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17696756 kB) free!
[1584146799] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27593188 kB) free.
[1584146863] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.4% (33426704 kB) free.
[1584147999] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20597880 kB) free!
[1584148063] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1289572 kB) free!
[1584148119] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (17040528 kB) free!
[1584148183] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (642376 kB) free!
[1584148239] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13529564 kB) free!
[1584148303] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389696 kB) free!
[1584148393] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 88.9% (117163240 kB) free.
[1584148448] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.3% (21406932 kB) free.
[1584148505] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39288924 kB) free.
[1584148566] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.2% (21133224 kB) free.
[1584149439] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35582772 kB) free.
[1584150039] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17295176 kB) free!
[1584150160] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13668992 kB) free!
[1584150280] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38339480 kB) free.
[1584150547] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190536 kB) free!
[1584150667] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182172 kB) free!
[1584150787] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194008 kB) free!
[1584150880] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20769012 kB) free!
[1584151000] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17487276 kB) free!
[1584151120] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14185952 kB) free!
[1584151270] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4797308 kB) free!
[1584151330] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 13.2% (3244436 kB) free!
[1584151390] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 7.0% (1720868 kB) free!
[1584151450] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (206520 kB) free!
[1584151587] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.9% (20962632 kB) free.
[1584151750] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 11.9% (2912520 kB) free!
[1584152320] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33846292 kB) free.
[1584152920] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14888648 kB) free!
[1584153040] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33048716 kB) free.
[1584153640] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21716532 kB) free!
[1584153760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18051644 kB) free!
[1584153880] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14430944 kB) free!
[1584155080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36345628 kB) free.
[1584155680] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16610232 kB) free!
[1584155800] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12840728 kB) free!
[1584155921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37841284 kB) free.
[1584156329] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(479) < 103.255.146.164 [14/Mar/2020:08:53:11 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584156449] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 203.212.70.170 [14/Mar/2020:08:56:33 +0530] TCP 502 0 0 75.029 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "15.028, 60.001, 0.000"
[1584156520] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18815816 kB) free!
[1584156569] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584156640] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14782308 kB) free!
[1584156760] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (10411644 kB) free!
[1584157360] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19489000 kB) free!
[1584157960] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32998172 kB) free.
[1584158560] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11680136 kB) free!
[1584158680] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7694508 kB) free!
[1584158715] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9109780 kB) free!
[1584158800] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3952204 kB) free!
[1584158836] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9108912 kB) free!
[1584158956] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (9106624 kB) free!
[1584159400] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13967516 kB) free!
[1584160000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28801256 kB) free.
[1584160400] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 727 processes
[1584160521] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 726 processes
[1584160525] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 33.21, 35.04, 26.45
[1584160600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4383504 kB) free!
[1584160641] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 728 processes
[1584160645] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 28.79, 32.60, 26.58
[1584160720] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404424 kB) free!
[1584160840] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (450972 kB) free!
[1584161441] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13474664 kB) free!
[1584161632] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 368 processes
[1584161650] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.2% (2245848 kB) free!
[1584161725] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4716152 kB) free!
[1584161752] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1584161798] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (4027680 kB) free!
[1584161845] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4385416 kB) free!
[1584161918] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.2% (2519504 kB) free!
[1584161965] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (4044256 kB) free!
[1584162038] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1880540 kB) free!
[1584162041] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379392 kB) free!
[1584162244] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (201752 kB) free!
[1584162253] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.7% (4578364 kB) free!
[1584162297] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.5% (21110972 kB) free.
[1584162358] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.8% (55137072 kB) free.
[1584162373] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.7% (4103156 kB) free!
[1584162448] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.7% (14090824 kB) free.
[1584162493] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.8% (3633796 kB) free!
[1584162518] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 90.7% (22219992 kB) free.
[1584162556] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3284976 kB) free!
[1584162565] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2153728 kB) free!
[1584162736] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.0% (38423520 kB) free.
[1584162907] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4400752 kB) free!
[1584163027] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.3% (3778884 kB) free!
[1584163093] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1541068 kB) free!
[1584163096] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.5% (20210804 kB) free.
[1584163147] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (3616264 kB) free!
[1584163167] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.7% (21499712 kB) free.
[1584163562] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22430256 kB) free!
[1584163683] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16501764 kB) free!
[1584163803] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (11013320 kB) free!
[1584163809] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3326264 kB) free!
[1584163929] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.0% (2925392 kB) free!
[1584164001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37604352 kB) free.
[1584164049] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.2% (2504428 kB) free!
[1584164348] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1584164348] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180944 kB) free!
[1584164649] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (619016 kB) free!
[1584164819] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1584164939] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1584165211] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13047484 kB) free!
[1584165331] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8739424 kB) free!
[1584165376] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (4076364 kB) free!
[1584165451] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4663996 kB) free!
[1584165496] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.7% (3853636 kB) free!
[1584165586] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 13.3% (3276844 kB) free!
[1584165616] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.6% (3573204 kB) free!
[1584165706] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 10.9% (2702024 kB) free!
[1584165826] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2166768 kB) free!
[1584166216] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2133872 kB) free!
[1584166344] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7910100 kB) free!
[1584166395] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41520 MB (18% inode=99%):
[1584166423] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (5887140 kB) free!
[1584166464] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (7898896 kB) free!
[1584166514] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7181496 kB) free!
[1584166515] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41520 MB (18% inode=99%):
[1584166543] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (5431704 kB) free!
[1584166584] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1584166584] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (7896388 kB) free!
[1584166634] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6707912 kB) free!
[1584166635] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41520 MB (18% inode=99%):
[1584166663] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 14.3% (4975736 kB) free!
[1584166704] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (3743608 kB) free!
[1584166754] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.5% (6147224 kB) free!
[1584166824] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.0% (3437720 kB) free!
[1584166944] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1584166944] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (3177588 kB) free!
[1584167245] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 31.57, 31.73, 30.73
[1584167263] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (2420996 kB) free!
[1584167354] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3437032 kB) free!
[1584167365] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 27.16, 31.16, 30.74
[1584167485] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 23.37, 28.10, 29.64
[1584167544] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1223784 kB) free!
[1584167649] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.8% (15663020 kB) free.
[1584167764] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.1% (20764128 kB) free.
[1584167820] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42003492 kB) free.
[1584167864] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 91.9% (32011188 kB) free.
[1584167876] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8715916 kB) free!
[1584167952] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.2% (30171592 kB) free.
[1584167997] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8348396 kB) free!
[1584168005] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.2% (20646304 kB) free.
[1584168064] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.9% (20079512 kB) free.
[1584168117] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.1% (7928072 kB) free!
[1584168138] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22858872 kB) free.
[1584168384] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3733288 kB) free!
[1584168577] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188052 kB) free!
[1584168664] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (1025664 kB) free!
[1584168685] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 35.31, 36.07, 32.63
[1584168689] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584168697] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180716 kB) free!
[1584168707] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584168784] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167744 kB) free!
[1584168805] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 31.14, 34.14, 32.33
[1584168809] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1584168817] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584168817] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188680 kB) free!
[1584168827] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584168904] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192808 kB) free!
[1584168925] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 30.37, 32.57, 31.96
[1584168929] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584168947] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584169035] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16759116 kB) free!
[1584169152] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7363316 kB) free!
[1584169155] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12948160 kB) free!
[1584169272] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (4857916 kB) free!
[1584169275] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9187880 kB) free!
[1584169317] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234684 kB) free!
[1584169392] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (2393544 kB) free!
[1584169526] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 24.96, 26.88, 29.42
[1584169529] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1584169547] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1584169564] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188620 kB) free!
[1584169635] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 104837 MB (46% inode=99%):
[1584169684] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188752 kB) free!
[1584169804] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202048 kB) free!
[1584169875] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16116880 kB) free!
[1584169989] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21678496 kB) free.
[1584170048] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 68.9% (17010796 kB) free.
[1584170106] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.6% (44632636 kB) free.
[1584170157] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.5% (39164828 kB) free.
[1584170237] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.4% (45541444 kB) free.
[1584170283] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.9% (20063812 kB) free.
[1584170475] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387404 kB) free!
[1584170653] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188816 kB) free!
[1584170707] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.7% (20663592 kB) free.
[1584170760] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39201476 kB) free.
[1584171965] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13730668 kB) free!
[1584172085] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10016204 kB) free!
[1584172093] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 26458 MB (20% inode=99%):
[1584172205] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29840556 kB) free.
[1584172213] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 26458 MB (20% inode=99%):
[1584172333] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 26458 MB (20% inode=99%):
[1584172805] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4349088 kB) free!
[1584172925] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391188 kB) free!
[1584173045] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395892 kB) free!
[1584173119] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191724 kB) free!
[1584173167] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4717384 kB) free!
[1584173239] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179280 kB) free!
[1584173287] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (4716584 kB) free!
[1584173359] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196088 kB) free!
[1584173407] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4714308 kB) free!
[1584173412] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4660684 kB) free!
[1584173532] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (4515060 kB) free!
[1584173645] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27807480 kB) free.
[1584173652] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1584173652] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4241524 kB) free!
[1584174245] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2563864 kB) free!
[1584174365] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397544 kB) free!
[1584174485] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383096 kB) free!
[1584174852] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (706568 kB) free!
[1584174882] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.3% (20548900 kB) free.
[1584174934] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 42.9% (56603548 kB) free.
[1584174947] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1584174982] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39232992 kB) free.
[1584175036] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.2% (20396804 kB) free.
[1584175068] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1584175087] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.1% (3698608 kB) free!
[1584175187] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3091680 kB) free!
[1584175188] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1584175207] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (3519980 kB) free!
[1584175307] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3562144 kB) free.
[1584175327] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.7% (3349608 kB) free!
[1584176049] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3647676 kB) free!
[1584176105] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189000 kB) free!
[1584176134] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10279868 kB) free!
[1584176169] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.5% (3484928 kB) free!
[1584176225] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182888 kB) free!
[1584176254] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29385584 kB) free.
[1584176289] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.7% (3298152 kB) free!
[1584176345] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185256 kB) free!
[1584176527] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1893484 kB) free!
[1584176788] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4922660 kB) free!
[1584176854] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2540276 kB) free!
[1584176908] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4827372 kB) free!
[1584176974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (419136 kB) free!
[1584177028] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4714772 kB) free!
[1584177094] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406748 kB) free!
[1584177329] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1584177449] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1584177489] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1549892 kB) free!
[1584177569] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1584177588] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1584177588] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1584177694] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26066752 kB) free!
[1584177733] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.2% (20524648 kB) free.
[1584177783] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1584177783] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.6% (15638944 kB) free.
[1584177795] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191612 kB) free!
[1584177854] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39220984 kB) free.
[1584177899] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.9% (19822184 kB) free.
[1584177915] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182072 kB) free!
[1584177970] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21659208 kB) free.
[1584178169] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1584178169] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1584178188] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1584178294] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (909816 kB) free!
[1584178769] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(161) < 92.42.104.166 [14/Mar/2020:15:00:22 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584178788] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 346 processes
[1584178889] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584178894] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29840612 kB) free.
[1584178907] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3114012 kB) free!
[1584178937] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172980 kB) free!
[1584178945] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4788004 kB) free!
[1584179027] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3818672 kB) free.
[1584179057] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182080 kB) free!
[1584179065] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4519668 kB) free!
[1584179177] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172332 kB) free!
[1584179185] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22880256 kB) free.
[1584179495] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3246404 kB) free!
[1584179615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405972 kB) free!
[1584179704] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179820 kB) free!
[1584179735] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381772 kB) free!
[1584179777] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584179824] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192840 kB) free!
[1584179944] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178212 kB) free!
[1584179972] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.1% (20759088 kB) free.
[1584180000] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14511812 kB) free!
[1584180005] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32417792 kB) free.
[1584180026] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.7% (20268976 kB) free.
[1584180227] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2580700 kB) free!
[1584180254] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.5% (6668676 kB) free!
[1584180347] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3425772 kB) free.
[1584180374] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.2% (6514708 kB) free!
[1584180494] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.9% (6368792 kB) free!
[1584180947] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2898224 kB) free!
[1584181067] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.9% (3877236 kB) free.
[1584181205] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9221068 kB) free!
[1584181235] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163636 kB) free!
[1584181325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5180684 kB) free!
[1584181333] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 91349 MB (71% inode=99%):
[1584181355] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179752 kB) free!
[1584181445] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1175260 kB) free!
[1584181475] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190760 kB) free!
[1584182867] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.4% (2344260 kB) free!
[1584182895] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3688484 kB) free!
[1584182987] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.0% (5356100 kB) free.
[1584183575] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (3882512 kB) free!
[1584183587] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.4% (2182056 kB) free!
[1584183695] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2378404 kB) free!
[1584183707] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.9% (5511196 kB) free.
[1584183815] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1733188 kB) free!
[1584184044] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (4030584 kB) free!
[1584184164] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.3% (3741580 kB) free!
[1584184284] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (3529056 kB) free!
[1584184307] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.2% (1663584 kB) free!
[1584184427] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3273128 kB) free.
[1584184445] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35070208 kB) free.
[1584184884] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2169900 kB) free!
[1584185046] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6511200 kB) free!
[1584185166] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1650860 kB) free!
[1584185286] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (451860 kB) free!
[1584185627] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.5% (2526644 kB) free!
[1584185748] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.3% (4757412 kB) free.
[1584185777] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.8% (3881688 kB) free!
[1584185886] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30370240 kB) free.
[1584185897] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.9% (3657456 kB) free!
[1584186017] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.2% (3481772 kB) free!
[1584186486] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1808028 kB) free!
[1584186606] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (449788 kB) free!
[1584186726] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (457916 kB) free!
[1584186947] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.1% (1634080 kB) free!
[1584187067] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3417692 kB) free.
[1584187217] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1497292 kB) free!
[1584187326] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31127228 kB) free.
[1584187389] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3405696 kB) free!
[1584187509] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3182604 kB) free!
[1584187629] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.1% (2936816 kB) free!
[1584187641] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584187926] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2394716 kB) free!
[1584188046] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408120 kB) free!
[1584188166] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424468 kB) free!
[1584188229] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1768708 kB) free!
[1584188241] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584188361] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1584188481] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1584188766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32147716 kB) free.
[1584188867] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (962532 kB) free!
[1584188987] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.5% (2351124 kB) free!
[1584189107] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3868640 kB) free.
[1584189343] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21716376 kB) free.
[1584189366] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4162104 kB) free!
[1584189372] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.6% (20380468 kB) free.
[1584189405] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31422780 kB) free.
[1584189430] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.6% (15628856 kB) free.
[1584189446] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9231192 kB) free!
[1584189487] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.7% (20759016 kB) free.
[1584189518] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.1% (19873980 kB) free.
[1584189550] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39227492 kB) free.
[1584189566] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9109988 kB) free!
[1584189569] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584189681] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584189686] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (8994760 kB) free!
[1584189976] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193424 kB) free!
[1584190017] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20268060 kB) free!
[1584190096] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178184 kB) free!
[1584190137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40983628 kB) free.
[1584190154] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.9% (7343844 kB) free!
[1584190188] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584190216] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163992 kB) free!
[1584190274] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (8043520 kB) free!
[1584190282] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584190307] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (1589168 kB) free!
[1584190394] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.3% (8045044 kB) free!
[1584190402] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1584190427] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2899540 kB) free!
[1584190522] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1584190522] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1584190547] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.8% (4512912 kB) free.
[1584190724] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184388 kB) free!
[1584190738] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13609944 kB) free!
[1584190844] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171164 kB) free!
[1584190858] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8514968 kB) free!
[1584190964] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190740 kB) free!
[1584190978] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3794668 kB) free!
[1584190994] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3517072 kB) free!
[1584191122] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1584191339] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1584191459] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1584191578] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26953280 kB) free.
[1584191747] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2975312 kB) free!
[1584191867] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3983012 kB) free.
[1584192178] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (396096 kB) free!
[1584192298] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400148 kB) free!
[1584192418] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389284 kB) free!
[1584192467] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2855760 kB) free!
[1584192588] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3751944 kB) free.
[1584193018] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31083684 kB) free.
[1584193187] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3053844 kB) free!
[1584193286] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4538528 kB) free!
[1584193307] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3781672 kB) free.
[1584193618] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (442744 kB) free!
[1584193738] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (388844 kB) free!
[1584193858] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406716 kB) free!
[1584193907] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2111616 kB) free!
[1584194027] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3030716 kB) free!
[1584194147] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3602232 kB) free.
[1584194388] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1584194459] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1584194580] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1584194727] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.0% (19750736 kB) free.
[1584194760] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32938896 kB) free.
[1584194785] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 59.4% (29275380 kB) free.
[1584194815] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 42.5% (20931212 kB) free.
[1584194844] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.8% (19805492 kB) free.
[1584195336] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181284 kB) free!
[1584195347] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (2683000 kB) free!
[1584195364] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (491916 kB) free!
[1584195455] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186948 kB) free!
[1584195467] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3238156 kB) free!
[1584195484] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386656 kB) free!
[1584195494] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1584195575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191976 kB) free!
[1584195587] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3988600 kB) free.
[1584195604] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376268 kB) free!
[1584195996] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8561552 kB) free!
[1584196116] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.4% (8076576 kB) free!
[1584196187] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.1% (1807452 kB) free!
[1584196237] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (8075972 kB) free!
[1584196307] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (2334316 kB) free!
[1584196427] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (3157188 kB) free!
[1584197027] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (908900 kB) free!
[1584197116] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9141392 kB) free!
[1584197236] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (8979252 kB) free!
[1584197356] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (8810200 kB) free!
[1584197627] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3757900 kB) free.
[1584198037] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (2614448 kB) free!
[1584199095] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4769264 kB) free!
[1584199215] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4648468 kB) free!
[1584199335] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4511560 kB) free!
[1584199428] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2712824 kB) free!
[1584199548] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3210088 kB) free!
[1584199660] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584199668] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.2% (3939440 kB) free.
[1584200145] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.9% (4378976 kB) free!
[1584200265] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.1% (4184492 kB) free!
[1584200268] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.3% (1839836 kB) free!
[1584200385] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.6% (4066764 kB) free!
[1584200388] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (2674484 kB) free!
[1584200508] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1584200508] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (3093328 kB) free!
[1584200956] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4914028 kB) free!
[1584201108] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.1% (4250368 kB) free.
[1584201604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28635336 kB) free.
[1584201735] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2061796 kB) free!
[1584202186] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2385480 kB) free!
[1584202204] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (390508 kB) free!
[1584202325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (387784 kB) free!
[1584202445] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (458236 kB) free!
[1584202522] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584202644] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1584202867] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.4% (20097972 kB) free.
[1584202898] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21723688 kB) free.
[1584202908] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3133524 kB) free!
[1584202935] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 35.3% (17377224 kB) free.
[1584202964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43804124 kB) free.
[1584202985] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 56.4% (27773648 kB) free.
[1584203021] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 82.4% (20196772 kB) free.
[1584203028] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1197112 kB) free!
[1584203148] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 34.9% (5678156 kB) free.
[1584203475] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172872 kB) free!
[1584203566] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21845532 kB) free!
[1584203595] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186652 kB) free!
[1584203686] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17267344 kB) free!
[1584203715] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168692 kB) free!
[1584203748] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2971968 kB) free!
[1584203806] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12891228 kB) free!
[1584203868] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3226544 kB) free!
[1584203988] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3656384 kB) free.
[1584204195] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.9% (6839392 kB) free!
[1584204315] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.4% (5134772 kB) free!
[1584204436] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (5367080 kB) free!
[1584205036] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4010532 kB) free!
[1584206206] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15823744 kB) free!
[1584206807] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399840 kB) free!
[1584208607] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20052224 kB) free!
[1584209207] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380792 kB) free!