[1543950033] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1543950089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1543950153] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1543950209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1543950329] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1543950754] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1543950874] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1543951493] SERVICE ALERT: localhost;Current Users;OK;HARD;1;USERS OK - 39 users currently logged in
[1543976093] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1543976153] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1543976213] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1543976273] SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 41 users currently logged in
[1543977929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1543978049] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1543978180] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.71, 5.45, 7.94
[1543980698] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 20 users currently logged in
[1543980818] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 20 users currently logged in
[1543980938] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 14 users currently logged in
[1543982074] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1543982194] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1543982249] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1543982314] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1543982369] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1543982489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1543982914] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1543983089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1543983209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1543983329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1543984114] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1543984235] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1543984355] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1543984513] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 309 processes
[1543984633] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 307 processes
[1543984753] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 307 processes
[1543984955] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1543984955] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1543985729] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1543986155] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1543986275] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1543986338] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1543986395] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1543986458] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1543986553] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1543986578] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1543986866] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4760660 kB) free!
[1543986986] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4719076 kB) free!
[1543987106] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4680216 kB) free!
[1543987529] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1543987649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1543987769] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1543988196] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1543988378] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1543988573] SERVICE ALERT: localhost;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 51 users currently logged in
[1543988969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1543988978] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1543989089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1543989098] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1543989209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1543989218] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1543989809] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1543989809] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1543989996] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1543990116] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1543990236] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1543990409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1543990529] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1543990836] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1543991074] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1543991129] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1543991194] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1543991249] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1543991314] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1543991370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1543991507] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 360 processes
[1543991618] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1543991618] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1543992107] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 360 processes
[1543993769] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1543993836] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1543993889] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1543993956] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1543994076] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1543994353] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1543994507] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1543995218] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1543995338] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1543995458] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1543996476] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1543996597] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1543997196] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1543997316] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1543997507] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1543997507] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 358 processes
[1544000258] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1544001689] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1544002058] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1544002117] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1544002179] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1544002237] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1544002299] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1544002837] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1544002907] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 358 processes
[1544004689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544004699] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1544004810] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1544005114] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1544005308] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 355 processes
[1544005837] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1544005957] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1544006273] SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 50 users currently logged in
[1544006573] SERVICE ALERT: localhost;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 51 users currently logged in
[1544006983] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1544007757] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1544007878] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1544008558] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 299 processes
[1544008678] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 298 processes
[1544008798] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 299 processes
[1544008812] SERVICE ALERT: pservercluster1;Sensors Status;UNKNOWN;SOFT;1;SENSORS UNKNOWN - command not found (did you install lmsensors?)
[1544008932] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1544009009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544009077] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1544009129] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544009197] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1544009249] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544009797] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1544009849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544009917] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1544009957] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 305 processes
[1544009970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544010038] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1544010077] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 305 processes
[1544010090] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544010169] SERVICE ALERT: secondary-server;Sensors Status;UNKNOWN;SOFT;1;SENSORS UNKNOWN - command not found (did you install lmsensors?)
[1544010197] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 306 processes
[1544010216] SERVICE ALERT: secondary-server;Sensors Status;OK;HARD;1;SENSORS OK
[1544010598] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1544010638] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1544010711] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1544011010] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;NRPE: Command 'check_sensors' not defined
[1544011095] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;1;NRPE: Command 'check_sensors' not defined
[1544011130] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;NRPE: Command 'check_sensors' not defined
[1544011174] SERVICE ALERT: localhost;Current Users;OK;HARD;1;USERS OK - 37 users currently logged in
[1544011215] SERVICE ALERT: sservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1544011250] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;NRPE: Command 'check_sensors' not defined
[1544011310] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1544011715] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1544012057] SERVICE ALERT: smscapp1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1544012268] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1544012324] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;SOFT;1;NRPE: Command 'check_sensors' not defined
[1544012328] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;2;USERS WARNING - 42 users currently logged in
[1544012388] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;3;USERS WARNING - 42 users currently logged in
[1544012438] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544012444] SERVICE ALERT: webservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1544012448] SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 42 users currently logged in
[1544012512] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1544012512] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 3 warn out of 339 processes
[1544012559] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544012679] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1544012804] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1544012842] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1544012907] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;NRPE: Command 'check_sensors' not defined
[1544012987] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1544013027] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1544013048] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1544013115] SERVICE ALERT: localhost;Current Users;OK;HARD;1;USERS OK - 39 users currently logged in
[1544013279] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1544013399] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1544013542] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1544013609] SERVICE ALERT: sig4;Sensors Status;CRITICAL;SOFT;1;NRPE: Command 'check_sensors' not defined
[1544013680] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1544013709] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1544013719] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1544013729] SERVICE ALERT: sig4;Sensors Status;OK;HARD;1;SENSORS OK
[1544013779] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;2;USERS WARNING - 42 users currently logged in
[1544013839] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;3;USERS WARNING - 42 users currently logged in
[1544013899] SERVICE FLAPPING ALERT: localhost;Current Users;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1544013899] SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 42 users currently logged in
[1544013950] SERVICE FLAPPING ALERT: localhost;Current Users;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1544013960] SERVICE FLAPPING ALERT: localhost;Current Users;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1544013978] SERVICE FLAPPING ALERT: localhost;Current Users;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1544014043] SERVICE ALERT: db2;Sensors Status;CRITICAL;SOFT;1;NRPE: Command 'check_sensors' not defined
[1544014120] SERVICE FLAPPING ALERT: localhost;Current Users;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1544014163] SERVICE ALERT: db2;Sensors Status;OK;HARD;1;SENSORS OK
[1544014499] SERVICE ALERT: localhost;Current Users;OK;HARD;1;USERS OK - 40 users currently logged in
[1544015199] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544015319] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544015439] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544015906] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2314788 kB) free!
[1544016039] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544016159] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544016280] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544016599] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1544016659] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1544016719] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1544016779] SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 41 users currently logged in
[1544017079] SERVICE ALERT: localhost;Current Users;OK;HARD;1;USERS OK - 39 users currently logged in
[1544017115] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1544017235] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1544017356] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1544017480] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1544017480] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544017891] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544018011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1544018131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544018307] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.3% (10600664 kB) free.
[1544018665] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.20, 10.40, 10.85
[1544018680] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544018785] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.14, 7.70, 9.81
[1544018800] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544018905] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.04, 6.17, 8.99
[1544018920] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544018980] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.99, 12.93, 11.35
[1544019100] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.51, 12.74, 11.47
[1544019220] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.38, 11.84, 11.32
[1544019331] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544019520] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544019640] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1544020105] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.92, 3.32
[1544020420] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.35, 7.55, 9.24
[1544021440] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1544021560] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544021879] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1544021939] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1544021999] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1544022059] SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 41 users currently logged in
[1544022331] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544022359] SERVICE ALERT: localhost;Current Users;OK;HARD;1;USERS OK - 39 users currently logged in
[1544022451] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544022571] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1544022760] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1544022880] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1544023171] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544023291] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1544024081] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1544024201] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544024321] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544024397] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 338 processes
[1544024921] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544024997] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 337 processes
[1544025041] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544025092] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544025212] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1544025332] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1544025641] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544025761] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544026961] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544027081] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544028059] SERVICE FLAPPING ALERT: localhost;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1544028881] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1544028933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1544029001] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544029053] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1544029121] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544029173] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1544029721] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544029773] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1544029773] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544029841] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1544029841] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544029859] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1544029919] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1544029979] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1544030039] SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 41 users currently logged in
[1544030639] SERVICE ALERT: localhost;Current Users;OK;HARD;1;USERS OK - 36 users currently logged in
[1544031641] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544031761] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1544033561] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544033681] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544033973] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544034093] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544034213] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1544034281] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544034401] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes