[1573410831] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 92.9% (45773164 kB) free.
[1573411062] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190976 kB) free!
[1573411182] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166536 kB) free!
[1573411302] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194380 kB) free!
[1573411430] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178128 kB) free!
[1573411550] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176048 kB) free!
[1573411670] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165040 kB) free!
[1573412535] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.5% (4281768 kB) free!
[1573412539] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4668852 kB) free!
[1573412595] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 67.5% (16531988 kB) free.
[1573412659] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.2% (3226176 kB) free!
[1573412747] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.7% (38278960 kB) free.
[1573412780] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.1% (22338868 kB) free.
[1573412870] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 52.9% (12960020 kB) free.
[1573413103] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.6% (22688480 kB) free.
[1573414350] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.6% (5580256 kB) free.
[1573414950] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195920 kB) free!
[1573415070] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179568 kB) free!
[1573415148] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7550216 kB) free!
[1573415190] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192916 kB) free!
[1573415268] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5953700 kB) free!
[1573415388] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4671284 kB) free!
[1573416588] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39949136 kB) free.
[1573416991] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (3945004 kB) free!
[1573417591] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193208 kB) free!
[1573419734] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219172 kB) free!
[1573419854] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214864 kB) free!
[1573419974] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.2% (44944404 kB) free.
[1573420591] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4180176 kB) free!
[1573421075] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1573421195] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573422096] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/11 03:11:22 [error] 23644#0: *2268167 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573422216] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573422392] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.4% (23285620 kB) free.
[1573425231] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 299 processes
[1573425351] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1573427388] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1573432592] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167368 kB) free!
[1573432670] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178116 kB) free!
[1573432712] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173336 kB) free!
[1573432790] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177664 kB) free!
[1573432833] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180976 kB) free!
[1573432910] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2644144 kB) free!
[1573433080] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22844372 kB) free.
[1573433510] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.2% (12309672 kB) free.
[1573443289] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573446396] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1573446516] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1573448890] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 327 processes
[1573448917] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/11 10:30:19 [error] 24129#0: *5539705 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573449010] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1573449037] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573449289] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193288 kB) free!
[1573449409] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174436 kB) free!
[1573449529] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173308 kB) free!
[1573449689] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21228760 kB) free.
[1573450116] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1573450236] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573450295] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187076 kB) free!
[1573450416] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182996 kB) free!
[1573450536] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189136 kB) free!
[1573450897] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1573450897] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.2% (22255564 kB) free.
[1573450919] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.6% (3275852 kB) free!
[1573451039] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.2% (2965140 kB) free!
[1573451160] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 57.2% (12844556 kB) free.
[1573451511] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4628380 kB) free!
[1573451631] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (4469136 kB) free!
[1573451751] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (4313816 kB) free!
[1573452660] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25789472 kB) free!
[1573452780] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.0% (21038948 kB) free!
[1573452900] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16358200 kB) free!
[1573453210] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1573453308] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195252 kB) free!
[1573453330] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1573453428] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166212 kB) free!
[1573453548] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178688 kB) free!
[1573453836] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1573453957] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1573454151] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1904492 kB) free!
[1573454338] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22344804 kB) free.
[1573454406] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 52.6% (12895508 kB) free.
[1573454701] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43188652 kB) free.
[1573454988] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9330864 kB) free!
[1573455108] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9331708 kB) free!
[1573455228] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (9331360 kB) free!
[1573456147] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193036 kB) free!
[1573456267] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173136 kB) free!
[1573456387] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190596 kB) free!
[1573457180] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4907652 kB) free!
[1573457300] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4756836 kB) free!
[1573457420] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (4607724 kB) free!
[1573457986] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.3% (18587200 kB) free.
[1573458106] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 224 processes
[1573458226] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 230 processes
[1573458347] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 232 processes
[1573458621] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.5% (21929148 kB) free.
[1573459187] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176352 kB) free!
[1573459307] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171244 kB) free!
[1573459427] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174128 kB) free!
[1573459429] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4345804 kB) free!
[1573460174] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9705820 kB) free!
[1573460294] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9421768 kB) free!
[1573460414] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.5% (9129908 kB) free!
[1573461359] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4051848 kB) free!
[1573461479] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.1% (3842156 kB) free!
[1573461599] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.0% (3602368 kB) free!
[1573461963] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21954780 kB) free.
[1573461979] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 56.5% (12688984 kB) free.
[1573462000] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39981924 kB) free.
[1573462226] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4901004 kB) free!
[1573462347] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.9% (4878820 kB) free!
[1573462467] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.8% (4860724 kB) free!
[1573462814] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4240244 kB) free!
[1573463169] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (2431428 kB) free!
[1573463289] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (2340320 kB) free!
[1573463409] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2251060 kB) free!
[1573464006] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4544376 kB) free!
[1573464037] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/11 14:45:30 [error] 24130#0: *5459011 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:22315/22900, bytes from/to upstream:22900/22315
[1573464126] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4416584 kB) free!
[1573464157] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573464246] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (4311636 kB) free!
[1573464757] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2019/11/11 14:56:58 [error] 24131#0: *5589334 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573464877] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573465214] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.9% (47247012 kB) free.
[1573465477] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2019/11/11 15:13:38 [error] 24132#0: *5589885 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573465598] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(16) < 2019/11/11 15:16:22 [error] 24130#0: *5590035 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573465718] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573466250] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 231 processes
[1573466318] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(20) < 2019/11/11 15:24:08 [error] 24130#0: *5590367 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573466438] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573466616] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(194) < 2019/11/11 15:32:43 [error] 23643#0: *2345144 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573466646] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2087860 kB) free!
[1573466701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26137928 kB) free!
[1573466737] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573466821] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21373836 kB) free!
[1573466850] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 238 processes
[1573466941] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16806600 kB) free!
[1573466970] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 235 processes
[1573467090] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 231 processes
[1573467541] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27427296 kB) free.
[1573468606] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224888 kB) free!
[1573468726] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239708 kB) free!
[1573468821] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (535856 kB) free!
[1573468846] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237656 kB) free!
[1573468941] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (382340 kB) free!
[1573469061] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (247164 kB) free!
[1573469341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26396036 kB) free!
[1573469461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.8% (22151224 kB) free!
[1573469581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40710308 kB) free.
[1573471078] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.5% (20120436 kB) free.
[1573471116] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.3% (22862364 kB) free.
[1573471170] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39525196 kB) free.
[1573471210] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 55.8% (13688108 kB) free.
[1573471250] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 92.5% (22666428 kB) free.
[1573471688] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202788 kB) free!
[1573471809] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192376 kB) free!
[1573471929] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 34.1% (8422180 kB) free.
[1573472191] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (3886088 kB) free!
[1573472311] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.2% (3639164 kB) free!
[1573472378] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (247020 kB) free!
[1573472431] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.0% (3363364 kB) free!
[1573472498] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213240 kB) free!
[1573472529] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194188 kB) free!
[1573472618] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241052 kB) free!
[1573472649] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180324 kB) free!
[1573472769] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189512 kB) free!
[1573473031] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1762620 kB) free!
[1573473130] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1573473181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16129136 kB) free!
[1573473251] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1573473301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35965592 kB) free.
[1573474231] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 52.7% (11842832 kB) free.
[1573475701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16667600 kB) free!
[1573475821] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34263904 kB) free.
[1573476090] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 236 processes
[1573476617] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4705220 kB) free!
[1573476690] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 234 processes
[1573476737] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (4590852 kB) free!
[1573476810] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 233 processes
[1573476857] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4423184 kB) free!
[1573476930] SERVICE FLAPPING ALERT: sig4-vm2;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1573476930] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 237 processes
[1573478821] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15525972 kB) free!
[1573478941] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32455512 kB) free.
[1573479258] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1860900 kB) free!
[1573480012] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1573480132] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1573480252] SERVICE ALERT: primary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1573480530] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 229 processes
[1573480741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14225772 kB) free!
[1573480861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31306040 kB) free.
[1573481769] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573482930] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 223 processes
[1573483050] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 223 processes
[1573483170] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 227 processes
[1573483771] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 224 processes
[1573484556] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1573484676] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573485661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25677492 kB) free!
[1573485781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41822628 kB) free.
[1573487581] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25505652 kB) free!
[1573487701] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41661696 kB) free.
[1573488519] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (1368648 kB) free!
[1573488639] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180040 kB) free!
[1573488759] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187216 kB) free!
[1573490676] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573490796] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1573491761] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1325 processes
[1573491881] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1325 processes
[1573492002] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 363 processes
[1573492596] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573492716] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1573494231] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8259172 kB) free!
[1573494304] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.6% (16603508 kB) free!
[1573494351] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8240948 kB) free!
[1573494424] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.5% (17854720 kB) free!
[1573494471] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.7% (8214348 kB) free!
[1573494545] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.5% (17848536 kB) free!
[1573495171] SERVICE FLAPPING ALERT: sig4-vm2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)