[1573238252] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182932 kB) free!
[1573238356] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235472 kB) free!
[1573238375] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186836 kB) free!
[1573238412] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186372 kB) free!
[1573238476] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239092 kB) free!
[1573238496] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189276 kB) free!
[1573238532] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191296 kB) free!
[1573238596] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218368 kB) free!
[1573238619] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188272 kB) free!
[1573238630] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191012 kB) free!
[1573238652] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171980 kB) free!
[1573238740] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189504 kB) free!
[1573238750] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184116 kB) free!
[1573238758] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23594016 kB) free!
[1573238860] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183924 kB) free!
[1573238870] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1573238870] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191848 kB) free!
[1573238878] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.7% (20752956 kB) free!
[1573238998] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17926704 kB) free!
[1573239598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29814736 kB) free.
[1573239840] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21925140 kB) free.
[1573240070] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.4% (12097388 kB) free.
[1573240296] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.2% (22844948 kB) free.
[1573240396] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 36.9% (18194088 kB) free.
[1573240455] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195004 kB) free!
[1573240575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191948 kB) free!
[1573240695] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171188 kB) free!
[1573241180] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/09 00:55:21 [error] 24130#0: *4314575 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573241296] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 21.0% (5184240 kB) free.
[1573241300] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/11/09 00:57:07 [error] 24130#0: *4314911 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573241420] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2019/11/09 00:58:53 [error] 24130#0: *4315249 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573241896] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174160 kB) free!
[1573241998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23253292 kB) free!
[1573242016] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196220 kB) free!
[1573242118] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20293292 kB) free!
[1573242136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179780 kB) free!
[1573242238] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17603568 kB) free!
[1573242838] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29080932 kB) free.
[1573243396] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.1% (6942124 kB) free!
[1573243516] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (6076212 kB) free!
[1573243636] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4928852 kB) free!
[1573244038] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24598816 kB) free!
[1573244158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.0% (21064848 kB) free!
[1573244278] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41695572 kB) free.
[1573244536] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4461760 kB) free!
[1573244878] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24803176 kB) free!
[1573244951] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573244998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (21775820 kB) free!
[1573245118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39492604 kB) free.
[1573245136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191972 kB) free!
[1573245436] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.9% (35435100 kB) free.
[1573245460] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.8% (22735296 kB) free.
[1573245564] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2019/11/09 02:09:04 [error] 24130#0: *4329932 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573245619] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/11/09 02:09:36 [error] 24131#0: *4330065 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573245622] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573245733] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21671608 kB) free.
[1573246231] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2019/11/09 02:20:18 [error] 24130#0: *4332298 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573246348] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185228 kB) free!
[1573246351] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(5) < 2019/11/09 02:22:05 [error] 24130#0: *4332684 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573246468] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187524 kB) free!
[1573246471] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(5) < 2019/11/09 02:23:49 [error] 24131#0: *4333054 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573246570] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (221668 kB) free!
[1573246588] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174524 kB) free!
[1573246690] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233912 kB) free!
[1573246810] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227772 kB) free!
[1573246954] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573247075] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573248010] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 61.9% (30486492 kB) free.
[1573249318] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24885004 kB) free!
[1573249438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21529136 kB) free!
[1573249467] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1023) < 103.255.146.164 [09/Nov/2019:03:10:12 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1573249558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41698584 kB) free.
[1573249587] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573249744] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1304) < 203.212.70.170 [09/Nov/2019:03:11:21 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1573249774] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1089) < 2019/11/09 03:11:22 [error] 23640#0: *2008025 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573249864] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573249894] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573250071] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573250158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26273048 kB) free!
[1573250188] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.7% (20893056 kB) free.
[1573250278] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.4% (23001284 kB) free!
[1573250398] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19600392 kB) free!
[1573250870] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1573250998] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26942580 kB) free.
[1573252798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23823532 kB) free!
[1573252918] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.3% (20147448 kB) free!
[1573253038] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40857504 kB) free.
[1573253639] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23033196 kB) free!
[1573253759] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19606748 kB) free!
[1573253879] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40293352 kB) free.
[1573254479] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22873796 kB) free!
[1573254599] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.2% (23966088 kB) free!
[1573254719] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40151672 kB) free.
[1573255319] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23748200 kB) free!
[1573255439] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20503636 kB) free!
[1573255559] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41658040 kB) free.
[1573255636] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1573256159] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24927696 kB) free!
[1573256279] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21359680 kB) free!
[1573256399] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41647564 kB) free.
[1573256999] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24759292 kB) free!
[1573257119] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21310128 kB) free!
[1573257239] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41643656 kB) free.
[1573257839] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24457512 kB) free!
[1573257960] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.1% (21170992 kB) free!
[1573258080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41628120 kB) free.
[1573258588] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1573258680] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25807616 kB) free!
[1573258800] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22575356 kB) free!
[1573258920] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19442248 kB) free!
[1573259520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28624680 kB) free.
[1573259788] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173460 kB) free!
[1573259870] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182908 kB) free!
[1573259908] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172812 kB) free!
[1573259990] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177316 kB) free!
[1573260029] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.5% (17882996 kB) free.
[1573260072] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.9% (12474492 kB) free.
[1573260096] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191796 kB) free!
[1573260163] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.2% (22363692 kB) free.
[1573260274] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 407 processes
[1573260394] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1573260640] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196648 kB) free!
[1573260720] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25531392 kB) free!
[1573260760] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201064 kB) free!
[1573260840] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.8% (22210044 kB) free!
[1573260880] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200312 kB) free!
[1573260960] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18905224 kB) free!
[1573261560] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28472460 kB) free.
[1573262760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21591044 kB) free!
[1573262880] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18502560 kB) free!
[1573263000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40676808 kB) free.
[1573263600] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20812060 kB) free!
[1573263616] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22506444 kB) free.
[1573263721] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17162408 kB) free!
[1573263841] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13536428 kB) free!
[1573265041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32463212 kB) free.
[1573265641] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14446720 kB) free!
[1573265761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38708700 kB) free.
[1573266361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20979368 kB) free!
[1573266481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17621128 kB) free!
[1573266601] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14393684 kB) free!
[1573267801] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32716512 kB) free.
[1573268401] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15788492 kB) free!
[1573268521] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38658868 kB) free.
[1573268636] SERVICE ALERT: qservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268648] SERVICE ALERT: qservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268661] SERVICE ALERT: qservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268692] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268750] SERVICE ALERT: qservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268766] SERVICE ALERT: qservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268772] SERVICE ALERT: qservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268778] SERVICE ALERT: qservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268791] SERVICE ALERT: qservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268819] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1573268877] SERVICE ALERT: qservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1573268887] SERVICE ALERT: qservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1573268887] SERVICE ALERT: qservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1573268902] SERVICE ALERT: qservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268908] SERVICE ALERT: qservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573268915] SERVICE ALERT: qservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1573268944] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1573269003] SERVICE ALERT: qservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1573269011] SERVICE ALERT: qservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1573269025] SERVICE ALERT: qservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1573269076] SERVICE ALERT: qservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573269121] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21525124 kB) free!
[1573269141] SERVICE ALERT: qservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573269206] SERVICE ALERT: qservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573269241] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18198496 kB) free!
[1573269336] SERVICE ALERT: qservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573269362] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14794172 kB) free!
[1573269394] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1573269514] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1573270210] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9728696 kB) free!
[1573270330] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9515156 kB) free!
[1573270332] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 401055 MB (37% inode=99%):
[1573270450] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9292264 kB) free!
[1573270562] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33733180 kB) free.
[1573271162] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14408940 kB) free!
[1573271282] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37984252 kB) free.
[1573271650] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4509740 kB) free!
[1573271882] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19951980 kB) free!
[1573272002] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15990240 kB) free!
[1573272122] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (12082284 kB) free!
[1573272247] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 97.8% (48183380 kB) free.
[1573272250] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 77.5% (38199988 kB) free.
[1573272336] SERVICE ALERT: qservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1573272488] SERVICE ALERT: qservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1573272508] SERVICE ALERT: qservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 27081 MB (52% inode=98%):
[1573272515] SERVICE ALERT: qservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1573272544] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 328 processes
[1573272603] SERVICE ALERT: qservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 171810 MB (76% inode=99%):
[1573272625] SERVICE ALERT: qservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1573272722] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20934596 kB) free!
[1573272741] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.60, 0.58, 0.36
[1573273146] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.11680.427> (10.147.212.158:48854 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573273266] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573273322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28105072 kB) free.
[1573273922] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6592092 kB) free!
[1573274042] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.6% (23211692 kB) free!
[1573274162] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18529756 kB) free!
[1573274762] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419140 kB) free!
[1573275017] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.8% (4863196 kB) free!
[1573275137] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.7% (4822576 kB) free!
[1573275257] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.5% (4785408 kB) free!
[1573275514] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1573275634] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1573275676] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (4163072 kB) free!
[1573275796] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (3966680 kB) free!
[1573275885] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27437124 kB) free.
[1573275916] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.4% (3778296 kB) free!
[1573276485] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (2978332 kB) free!
[1573276605] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (380088 kB) free!
[1573276725] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33632460 kB) free.
[1573277116] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2343088 kB) free!
[1573277325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10841080 kB) free!
[1573277445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6121776 kB) free!
[1573277565] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1481872 kB) free!
[1573278616] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194928 kB) free!
[1573278737] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178108 kB) free!
[1573278766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29288456 kB) free.
[1573278857] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (220060 kB) free!
[1573279036] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9432304 kB) free!
[1573279099] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4389700 kB) free!
[1573279156] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9432304 kB) free!
[1573279219] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4087924 kB) free!
[1573279276] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (8759824 kB) free!
[1573279339] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (3654424 kB) free!
[1573279366] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5508048 kB) free!
[1573279486] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1181512 kB) free!
[1573279606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33280540 kB) free.
[1573279939] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1615640 kB) free!
[1573280206] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8335776 kB) free!
[1573280326] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4262640 kB) free!
[1573280446] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399064 kB) free!
[1573280658] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.5% (2332552 kB) free!
[1573280671] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/11/09 11:54:29 [error] 24132#0: *4572730 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:35/0, bytes from/to upstream:0/35
[1573280740] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.7% (21150684 kB) free.
[1573280791] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(167) < 2019/11/09 11:56:31 [error] 24132#0: *4573794 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573280801] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 92.4% (22651952 kB) free.
[1573280842] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 51.6% (12659012 kB) free.
[1573280874] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.3% (16898956 kB) free.
[1573280911] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(51) < 2019/11/09 11:58:24 [error] 24130#0: *4574403 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573280914] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 44.0% (58046012 kB) free.
[1573282112] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573282367] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (4073760 kB) free!
[1573282487] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (4082748 kB) free!
[1573282607] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3892244 kB) free!
[1573282712] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(368) < 2019/11/09 12:24:56 [error] 24135#0: *4587590 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:86/0, bytes from/to upstream:0/86
[1573282832] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573282865] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(215) < 182.18.182.41 [09/Nov/2019:12:24:21 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1573282876] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4935432 kB) free!
[1573282985] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573283320] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10961112 kB) free!
[1573283440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41881808 kB) free.
[1573283650] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1573283770] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1573283807] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2036788 kB) free!
[1573283846] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4679856 kB) free!
[1573283890] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1573283966] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (4518940 kB) free!
[1573284040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17651212 kB) free!
[1573284086] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1573284086] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4339048 kB) free!
[1573284160] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13259480 kB) free!
[1573284280] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9184976 kB) free!
[1573284386] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.3% (3743484 kB) free!
[1573284506] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.2% (3717552 kB) free!
[1573284626] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.1% (3698764 kB) free!
[1573284880] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21305004 kB) free!
[1573285090] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1573285226] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 65.8% (16126184 kB) free.
[1573285480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32671820 kB) free.
[1573285690] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1573285811] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1573285854] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.2% (21613256 kB) free.
[1573285886] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2126300 kB) free!
[1573285893] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.1% (37979720 kB) free.
[1573285927] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.9% (12486620 kB) free.
[1573286080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12142728 kB) free!
[1573286200] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40471068 kB) free.
[1573286650] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9435132 kB) free!
[1573286747] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181880 kB) free!
[1573286770] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9306412 kB) free!
[1573286800] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19149608 kB) free!
[1573286867] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174436 kB) free!
[1573286890] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.5% (9100056 kB) free!
[1573286920] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13947604 kB) free!
[1573286987] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170952 kB) free!
[1573287041] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9218532 kB) free!
[1573287641] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21200156 kB) free!
[1573288241] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33816344 kB) free.
[1573288298] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (618804 kB) free!
[1573288418] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237924 kB) free!
[1573288538] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1573288538] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215532 kB) free!
[1573288688] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1573288691] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.0% (46833268 kB) free.
[1573288808] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1573288841] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11549716 kB) free!
[1573288937] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (3553480 kB) free!
[1573288961] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7516292 kB) free!
[1573289058] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.0% (3419828 kB) free!
[1573289081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38107144 kB) free.
[1573289178] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.5% (3305236 kB) free!
[1573289252] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.8% (21429008 kB) free.
[1573289295] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39099212 kB) free.
[1573289681] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15602912 kB) free!
[1573289778] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.3% (4979724 kB) free.
[1573289801] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (11410888 kB) free!
[1573289858] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168644 kB) free!
[1573289901] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (247936 kB) free!
[1573289921] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7374144 kB) free!
[1573289978] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196600 kB) free!
[1573290021] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (220820 kB) free!
[1573290098] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1573290098] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183436 kB) free!
[1573290141] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (263060 kB) free!
[1573290378] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (4106848 kB) free!
[1573290460] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4474496 kB) free!
[1573290498] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.1% (3941912 kB) free!
[1573290521] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14763264 kB) free!
[1573290580] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4355840 kB) free!
[1573290618] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3728560 kB) free!
[1573290700] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (4260660 kB) free!
[1573291121] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31998696 kB) free.
[1573291294] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(210) < 2019/11/09 14:47:29 [error] 23645#0: *2071081 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573291414] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573291721] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9752216 kB) free!
[1573291841] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5802580 kB) free!
[1573291961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36138356 kB) free.
[1573292418] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1774688 kB) free!
[1573292561] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11767916 kB) free!
[1573292682] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7693112 kB) free!
[1573292802] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3808560 kB) free!
[1573292895] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.6% (20880724 kB) free.
[1573292963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 47.4% (62539776 kB) free.
[1573293012] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39102068 kB) free.
[1573293032] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2019/11/09 15:20:31 [error] 24130#0: *4677191 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573293048] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.9% (12235344 kB) free.
[1573293152] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/11/09 15:22:04 [error] 24131#0: *4677791 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573293273] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(10) < 2019/11/09 15:24:24 [error] 24130#0: *4678715 recv() failed (104: Connection reset by peer) while proxying connection, client: 121.241.242.109, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:44/0, bytes from/to upstream:0/44
[1573293368] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1061.74, 872.86, 453.46
[1573293488] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 201.36, 633.75, 418.97
[1573293496] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170372 kB) free!
[1573293608] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 28.55, 424.59, 368.37
[1573293609] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1573293616] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184972 kB) free!
[1573293617] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216568 kB) free!
[1573293729] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1573293736] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192728 kB) free!
[1573293737] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227224 kB) free!
[1573293849] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 404 processes
[1573293857] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227172 kB) free!
[1573294208] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.90, 58.56, 193.94
[1573294449] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 383 processes
[1573295073] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573295368] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22084004 kB) free!
[1573295488] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.5% (17765200 kB) free!
[1573295500] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2392864 kB) free!
[1573295513] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1573295608] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13548756 kB) free!
[1573295633] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1573295753] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 651 processes
[1573295862] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.6% (21608884 kB) free.
[1573295919] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.4% (39140564 kB) free.
[1573295973] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.1% (22821136 kB) free.
[1573296467] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185936 kB) free!
[1573296530] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219032 kB) free!
[1573296588] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187264 kB) free!
[1573296650] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219152 kB) free!
[1573296708] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187832 kB) free!
[1573296770] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234712 kB) free!
[1573296808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38586840 kB) free.
[1573297408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14846316 kB) free!
[1573297528] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10280564 kB) free!
[1573297648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42440608 kB) free.
[1573298248] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15949336 kB) free!
[1573298368] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11691384 kB) free!
[1573298488] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (8096424 kB) free!
[1573299088] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19757760 kB) free!
[1573299350] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21956500 kB) free.
[1573299400] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39320188 kB) free.
[1573299960] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184300 kB) free!
[1573300009] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238472 kB) free!
[1573300080] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187792 kB) free!
[1573300129] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217872 kB) free!
[1573300200] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191024 kB) free!
[1573300250] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227668 kB) free!
[1573300288] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27018756 kB) free.
[1573300859] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4626784 kB) free!
[1573300888] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6244732 kB) free!
[1573300979] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (4516300 kB) free!
[1573301008] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2325360 kB) free!
[1573301099] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4405304 kB) free!
[1573301128] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381012 kB) free!
[1573301728] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14179676 kB) free!
[1573302328] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27539736 kB) free.
[1573302928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5010396 kB) free!
[1573303048] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (934400 kB) free!
[1573303168] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31935268 kB) free.
[1573303344] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.1% (19769196 kB) free.
[1573303399] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39607040 kB) free.
[1573303768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8005212 kB) free!
[1573303888] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3934928 kB) free!
[1573303949] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193476 kB) free!
[1573304008] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392700 kB) free!
[1573304069] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163008 kB) free!
[1573304099] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2064676 kB) free!
[1573304190] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199356 kB) free!
[1573304609] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (263568 kB) free!
[1573304729] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236896 kB) free!
[1573304849] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222660 kB) free!
[1573305208] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23962592 kB) free!
[1573305809] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3447996 kB) free!
[1573306409] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13258172 kB) free!
[1573306735] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.1% (20996184 kB) free.
[1573306777] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39246732 kB) free.
[1573306820] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.2% (12304452 kB) free.
[1573306860] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (968488 kB) free!
[1573306981] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (3540744 kB) free!
[1573307101] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (3129604 kB) free!
[1573307341] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196856 kB) free!
[1573307389] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (253304 kB) free!
[1573307461] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198888 kB) free!
[1573307509] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215908 kB) free!
[1573307581] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187652 kB) free!
[1573307609] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3458180 kB) free!
[1573307629] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232340 kB) free!
[1573308809] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23783004 kB) free!
[1573308901] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2301280 kB) free!
[1573309409] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2170960 kB) free!
[1573310271] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.8% (21177416 kB) free.
[1573310407] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.1% (21842900 kB) free.
[1573310463] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.8% (38840500 kB) free.
[1573310500] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 46.6% (61444832 kB) free.
[1573310882] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194456 kB) free!
[1573311002] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197584 kB) free!
[1573311068] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215004 kB) free!
[1573311123] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192496 kB) free!
[1573311188] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (224068 kB) free!
[1573311308] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (251120 kB) free!
[1573311629] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.6% (3573336 kB) free!
[1573311749] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.3% (3495612 kB) free!
[1573311869] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (3410700 kB) free!
[1573312340] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.7% (22137136 kB) free.
[1573312395] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.5% (39647448 kB) free.
[1573312473] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(292) < 2019/11/09 20:44:29 [error] 24132#0: *4804292 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:83/0, bytes from/to upstream:0/83
[1573312593] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(11) < 2019/11/09 20:46:33 [error] 24132#0: *4804760 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:87/0, bytes from/to upstream:0/87
[1573312713] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(7) < 2019/11/09 20:48:26 [error] 24132#0: *4805188 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:87/0, bytes from/to upstream:0/87
[1573312905] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23469816 kB) free!
[1573312950] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192308 kB) free!
[1573313026] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44750380 kB) free.
[1573313070] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195984 kB) free!
[1573313190] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182488 kB) free!
[1573313596] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (250988 kB) free!
[1573313625] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23958876 kB) free!
[1573313716] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242812 kB) free!
[1573313745] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20421776 kB) free!
[1573313836] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248480 kB) free!
[1573313865] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (46038328 kB) free.
[1573313879] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39504200 kB) free.
[1573313939] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.3% (22287112 kB) free.
[1573314465] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24933048 kB) free!
[1573314489] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220144 kB) free!
[1573314550] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176456 kB) free!
[1573314585] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.1% (21180340 kB) free!
[1573314609] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247624 kB) free!
[1573314670] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180744 kB) free!
[1573314705] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.0% (46184396 kB) free.
[1573314729] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (244876 kB) free!
[1573314790] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192672 kB) free!
[1573315305] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26189108 kB) free!
[1573315425] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.0% (22415568 kB) free!
[1573315545] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46785968 kB) free.
[1573316069] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2423724 kB) free!
[1573317343] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21922416 kB) free.
[1573317390] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39278424 kB) free.
[1573317428] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.5% (12137428 kB) free.
[1573317945] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24284296 kB) free!
[1573317949] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182104 kB) free!
[1573317999] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233020 kB) free!
[1573318065] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46630056 kB) free.
[1573318069] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187932 kB) free!
[1573318119] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (250852 kB) free!
[1573318189] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189696 kB) free!
[1573318239] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248956 kB) free!
[1573320785] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(69) < 49.50.64.122 [09/Nov/2019:22:58:24 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1573320905] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573321190] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.8% (20441172 kB) free.
[1573321240] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39408752 kB) free.
[1573321665] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1573321714] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573321790] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186816 kB) free!
[1573321866] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232960 kB) free!
[1573321910] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193408 kB) free!
[1573321986] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (249376 kB) free!
[1573322030] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190524 kB) free!
[1573322106] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230916 kB) free!