[1569436683] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (204136 kB) free!
[1569436803] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194884 kB) free!
[1569436925] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187596 kB) free!
[1569437084] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180168 kB) free!
[1569437204] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166840 kB) free!
[1569437324] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169156 kB) free!
[1569438973] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569440150] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 279 processes
[1569440173] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 53.5% (13204204 kB) free.
[1569440498] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (236016 kB) free!
[1569440618] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (245760 kB) free!
[1569440738] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (234304 kB) free!
[1569440773] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4552264 kB) free!
[1569440893] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.5% (3823756 kB) free!
[1569441000] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2459272 kB) free!
[1569441013] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.2% (3252220 kB) free!
[1569441120] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (485356 kB) free!
[1569441240] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382388 kB) free!
[1569441350] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 274 processes
[1569441613] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190520 kB) free!
[1569445348] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (228332 kB) free!
[1569445469] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235840 kB) free!
[1569445589] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230552 kB) free!
[1569446189] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 94.1% (46351736 kB) free.
[1569447727] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 52.0% (12739344 kB) free.
[1569448124] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 45.6% (11182976 kB) free.
[1569456116] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 136.31, 83.31, 36.01
[1569456236] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 137.29, 101.25, 48.32
[1569456356] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 137.48, 113.24, 59.13
[1569458138] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 94.3% (33034828 kB) free.
[1569458156] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.94, 1.12, 11.32
[1569458640] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.8% (20858860 kB) free!
[1569459356] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.59, 3.61
[1569460556] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 65.38, 18.63, 7.60
[1569460676] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.01, 12.82, 6.80
[1569460796] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1569460796] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.95, 8.79, 6.05
[1569460813] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 46.4% (11446128 kB) free.
[1569461396] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.97, 1.91, 3.55
[1569462016] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 501 processes
[1569462076] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 494 processes
[1569462226] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4890800 kB) free!
[1569462346] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.6% (4804512 kB) free!
[1569462466] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 20.0% (4904632 kB) free!
[1569463124] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185600 kB) free!
[1569463213] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179540 kB) free!
[1569463245] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186316 kB) free!
[1569463334] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173524 kB) free!
[1569463337] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569463365] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183148 kB) free!
[1569463454] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184224 kB) free!
[1569463460] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194156 kB) free!
[1569463580] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188284 kB) free!
[1569463876] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 527 processes
[1569463936] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 526 processes
[1569463997] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 530 processes
[1569464057] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 528 processes
[1569464266] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 20.1% (4917472 kB) free.
[1569464633] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1569464866] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4905280 kB) free!
[1569464986] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 20.0% (4904920 kB) free!
[1569465107] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.9% (4887776 kB) free!
[1569465233] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 8679 MB (16% inode=98%):
[1569465486] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569465606] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1569466307] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1569466307] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 20.7% (5066124 kB) free.
[1569467400] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569467520] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1569468120] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569468240] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1569468840] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1569468960] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569469907] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4892668 kB) free!
[1569470027] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.9% (4885000 kB) free!
[1569470147] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.9% (4878496 kB) free!
[1569472438] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.27, 14.68, 8.04
[1569472558] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.28, 10.04, 7.14
[1569472678] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 6.77, 6.30
[1569472796] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569473191] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 93.6% (22947868 kB) free.
[1569473235] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21854460 kB) free.
[1569473279] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.58, 2.99, 4.35
[1569473314] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569473347] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 67.9% (16651416 kB) free.
[1569473405] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.2% (40040332 kB) free.
[1569473438] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 380 processes
[1569473474] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.9% (23499560 kB) free.
[1569473547] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569473557] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 88.2% (116298800 kB) free.
[1569473670] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1569473846] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196388 kB) free!
[1569473959] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188980 kB) free!
[1569473966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188864 kB) free!
[1569474079] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164688 kB) free!
[1569474086] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194608 kB) free!
[1569474199] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199112 kB) free!
[1569475551] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.96, 6.78, 5.70
[1569475671] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.43, 6.27, 5.64
[1569475791] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.05, 6.25, 5.71
[1569477479] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 139.65, 118.74, 64.68
[1569477599] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 52.17, 100.45, 65.09
[1569477719] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.06, 67.62, 57.35
[1569478548] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1569478668] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1569478699] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569478788] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1569478919] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1569478919] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.92, 2.68, 16.82
[1569479148] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569479208] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1569479268] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1569479299] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569479328] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1569479988] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 14 users currently logged in
[1569480228] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1569480719] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.54, 1.78, 4.00
[1569480995] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1569482789] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8668724 kB) free!
[1569482909] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.8% (8272412 kB) free!
[1569483029] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.1% (7947620 kB) free!
[1569484229] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2977512 kB) free!
[1569485028] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569485088] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1569485148] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1569485208] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1569485760] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1569485880] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1569485899] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569486499] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1569486499] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569486799] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2542788 kB) free!
[1569487284] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4273624 kB) free!
[1569487399] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1828648 kB) free!
[1569487404] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4095776 kB) free!
[1569487524] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.1% (3950220 kB) free!
[1569487819] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (3916796 kB) free!
[1569487939] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.2% (3590336 kB) free!
[1569488059] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3186436 kB) free!
[1569488659] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (762956 kB) free!
[1569488724] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2001616 kB) free!
[1569490459] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.6% (23007000 kB) free.
[1569490699] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569490796] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 114.73, 109.10, 56.61
[1569490916] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 20.19, 74.74, 50.37
[1569490919] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1569491036] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.73, 50.71, 44.51
[1569491207] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 95.50, 37.18, 14.23
[1569491208] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 57 users currently logged in
[1569491327] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.76, 25.26, 12.65
[1569491435] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 627 processes
[1569491447] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.94, 17.28, 11.25
[1569491555] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 367 processes
[1569491899] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569492236] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.51, 2.38, 13.27
[1569492499] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569492616] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9809372 kB) free!
[1569492647] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.97, 1.20, 3.76
[1569492736] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9810712 kB) free!
[1569492856] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9810692 kB) free!
[1569493099] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569493436] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.38, 1.56, 4.78
[1569497899] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569498500] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569499456] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4173488 kB) free!
[1569499791] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.00, 3.08, 4.16
[1569501119] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 60.64, 38.37, 15.92
[1569501236] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 24.50, 39.60, 19.00
[1569501239] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 9.61, 26.28, 14.22
[1569501356] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.31, 26.84, 16.81
[1569501359] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.33, 18.33, 12.77
[1569501476] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.98, 18.46, 14.96
[1569501647] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.83, 8.81, 9.06
[1569501767] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.63, 6.15, 8.06
[1569501887] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.66, 4.35, 7.17
[1569502191] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.26, 6.78, 5.33
[1569502311] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.11, 7.11, 5.65
[1569502431] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.04, 8.34, 6.28
[1569502487] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1569502487] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.19, 1.39, 4.18
[1569502560] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.74, 1.92, 4.56
[1569503276] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1569503276] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.09, 1.19, 3.19
[1569503300] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569503900] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569504231] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.79, 3.79, 4.43
[1569504686] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (3577004 kB) free!
[1569504799] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (3027220 kB) free!
[1569504924] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (3531604 kB) free!
[1569505286] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2194168 kB) free!
[1569505399] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1530724 kB) free!
[1569505578] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 69.5% (17157808 kB) free.
[1569505614] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 77.6% (19014220 kB) free.
[1569505654] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.2% (40019340 kB) free.
[1569505693] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 40.5% (19953752 kB) free.
[1569506031] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.71, 6.02, 5.29
[1569506151] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.46, 5.43, 5.16
[1569506271] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1569506271] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.26, 5.62, 5.25
[1569506724] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2031460 kB) free!
[1569506871] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.13, 3.95, 4.51
[1569508680] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569508800] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569509185] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (2346540 kB) free!
[1569509305] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (2250452 kB) free!
[1569509425] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1569509425] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2186220 kB) free!
[1569510432] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190988 kB) free!
[1569510552] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190220 kB) free!
[1569510672] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1569510672] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191168 kB) free!
[1569513476] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1569513887] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569514170] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.8% (23382592 kB) free.
[1569514205] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.2% (22361024 kB) free.
[1569514238] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.0% (22796112 kB) free.
[1569515300] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569516600] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569516720] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569518271] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569520773] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1569521412] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)