[1546716447] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1546720647] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2.27, 31.78, 21.72
[1546720767] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.38, 21.31, 19.11
[1546720887] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 14.36, 16.83
[1546722087] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.48, 4.78
[1546725679] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 780 processes
[1546725687] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8.13, 38.21, 21.26
[1546725799] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 350 processes
[1546725807] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.18, 25.61, 18.70
[1546725927] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 17.27, 16.48
[1546727127] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1546727127] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.08, 0.45, 4.63
[1546730727] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 138.02, 74.80, 28.79
[1546730847] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 18.88, 50.12, 25.32
[1546730967] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.83, 33.66, 22.29
[1546731567] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.35, 4.72, 11.79
[1546732767] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.35, 3.42
[1546736967] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.57, 25.92, 22.62
[1546737087] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.29, 17.42, 19.90
[1546737207] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 11.72, 17.51
[1546738407] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.11, 0.33, 4.92
[1546742007] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15.10, 53.77, 28.54
[1546742127] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.17, 36.02, 25.10
[1546742247] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.35, 24.15, 22.07
[1546742847] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.12, 3.37, 11.65
[1546744047] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.27, 0.30, 3.36
[1546747047] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.22, 41.92, 21.18
[1546747167] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.47, 28.16, 18.65
[1546747287] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 18.91, 16.42
[1546748487] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.27, 0.57, 4.68
[1546752087] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.72, 30.73, 19.20
[1546752207] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.73, 20.67, 16.91
[1546752327] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.89, 14.03, 14.93
[1546753527] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.47, 4.27
[1546755758] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546755878] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546757127] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.97, 31.29, 21.74
[1546757247] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.41, 21.02, 19.14
[1546757367] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.26, 14.14, 16.85
[1546758567] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.17, 0.52, 4.82
[1546760309] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 351 processes
[1546761278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546761398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546761519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1546762109] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 397 processes
[1546762167] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.83, 26.79, 18.32
[1546762287] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.49, 18.05, 16.15
[1546762407] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 12.20, 14.23
[1546763318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546763438] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1546763607] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.45, 4.08
[1546763885] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 446 processes
[1546763945] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 452 processes
[1546767207] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.69, 19.08, 16.49
[1546767327] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.23, 12.82, 14.51
[1546767447] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.12, 8.62, 12.77
[1546768647] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.12, 0.44, 3.72
[1546768838] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546768958] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546769078] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546769679] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1546769909] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 358 processes
[1546772079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546772199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546772247] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 17.89, 49.47, 25.43
[1546772309] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 374 processes
[1546772319] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1546772367] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.90, 33.29, 22.42
[1546772487] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 22.41, 19.75
[1546772909] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1546772909] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 354 processes
[1546774287] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.51, 0.57, 3.30
[1546777109] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 351 processes
[1546777887] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.09, 26.99, 20.81
[1546778007] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.76, 18.28, 18.37
[1546778128] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.27, 12.29, 16.16
[1546779328] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.53, 4.65
[1546782928] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.37, 27.89, 20.97
[1546783048] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.32, 18.72, 18.44
[1546783168] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.19, 12.58, 16.23
[1546784368] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.50, 4.64
[1546784909] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1546794519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546794639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546794759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1546795168] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1546796909] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 353 processes