[1595961579] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192792 kB) free!
[1595961591] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.7% (3348436 kB) free!
[1595961699] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194056 kB) free!
[1595961711] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3332864 kB) free!
[1595961819] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197836 kB) free!
[1595961831] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (3313800 kB) free!
[1595961888] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175012 kB) free!
[1595962008] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22007140 kB) free.
[1595962331] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 42.75, 33.43, 36.89
[1595962431] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.7% (21977544 kB) free.
[1595963019] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21999104 kB) free.
[1595964131] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1595964131] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 37.41, 48.29, 46.70
[1595964731] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 43.68, 40.64, 42.94
[1595964784] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (1501520 kB) free!
[1595964822] SERVICE ALERT: webserver2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595964844] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172508 kB) free!
[1595964887] SERVICE ALERT: webserver2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=0.10% system=3.28% iowait=3.34% idle=93.28%
[1595964906] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (183152 kB) free!
[1595964966] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (171504 kB) free!
[1595965080] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1503 processes
[1595965201] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1500 processes
[1595965266] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 75.1% (18406864 kB) free.
[1595965331] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 52.72, 48.56, 45.43
[1595965419] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4421756 kB) free!
[1595965539] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.2% (20774200 kB) free.
[1595965931] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 40.66, 41.32, 43.41
[1595966052] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.0% (19782720 kB) free!
[1595966172] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.1% (13317424 kB) free!
[1595966209] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (432884 kB) free!
[1595966292] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6327228 kB) free!
[1595966329] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185928 kB) free!
[1595966450] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.5% (20457560 kB) free.
[1595966876] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237760 kB) free!
[1595966996] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236640 kB) free!
[1595967049] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (3062836 kB) free!
[1595967116] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (244312 kB) free!
[1595967169] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193976 kB) free!
[1595967289] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187004 kB) free!
[1595967889] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (3756484 kB) free!
[1595968316] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 42.7% (21051476 kB) free.
[1595968489] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176508 kB) free!
[1595970009] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 75.3% (37109780 kB) free.
[1595970363] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595970592] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 81.6% (40212524 kB) free.
[1595970889] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1595970889] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 35.8% (8769272 kB) free.
[1595971331] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 56.63, 51.08, 46.65
[1595972689] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (740616 kB) free!
[1595972809] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 77.2% (18927316 kB) free.
[1595973131] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 20.79, 34.83, 43.12
[1595974931] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 56.80, 48.51, 44.62
[1595975763] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16224.36, 13334.64, 6910.01
[1595975883] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16227.17, 14291.83, 8039.75
[1595976003] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16228.00, 14932.44, 9032.56
[1595976731] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 36.66, 42.00, 44.23
[1595978531] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 60.08, 51.79, 46.50
[1595979003] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.41, 2.25, 478.07
[1595980009] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.3% (3515364 kB) free!
[1595980130] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169132 kB) free!
[1595980250] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175892 kB) free!
[1595980800] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1502 processes
[1595980850] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.9% (21796652 kB) free.
[1595980921] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1501 processes
[1595981041] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1501 processes
[1595981641] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1499 processes
[1595982731] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 45.68, 42.97, 44.75
[1595982842] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1507 processes
[1595982962] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1503 processes
[1595983082] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1502 processes
[1595983331] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 48.47, 50.73, 47.78
[1595984882] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1595984882] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1500 processes
[1595986092] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 51.2% (67551108 kB) free.
[1595992250] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595994121] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994166] SERVICE ALERT: sservercluster2;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595994194] SERVICE ALERT: pservercluster2;Messages Log Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595994205] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994214] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994242] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994251] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994293] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994296] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994322] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994324] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994325] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994326] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994335] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994344] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994371] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994372] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994381] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994397] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994409] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994420] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994423] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994426] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994441] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994452] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994454] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994455] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994456] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994458] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994466] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994474] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994478] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994501] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994502] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994516] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994520] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994527] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994536] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994539] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994550] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994553] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994571] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994582] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994585] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994586] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994588] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994588] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994602] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994608] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994631] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994646] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994647] SERVICE ALERT: pservercluster2;zombie_procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595994650] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994659] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;Could not open pipe: /usr/bin/ps
[1595994666] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994670] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1595994697] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994701] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994708] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595994718] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994719] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595994721] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994722] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;SOFT;2;Error opening /usr/bin/uptime
[1595994724] SERVICE ALERT: pservercluster2;Sensors Status;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1595994728] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38287 MB (74% inode=98%):
[1595994731] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595994766] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 321 processes
[1595994771] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595994786] SERVICE ALERT: sservercluster2;Open-Files;OK;HARD;1;OK: 6592 open files (0% of max 2424092)
[1595994827] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994849] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994851] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994852] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994854] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994957] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994979] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994981] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1595994981] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595994984] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995026] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 205489 MB (93% inode=99%):
[1595995066] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.93, 1.13, 1.19
[1595995075] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 33086 MB (64% inode=98%):
[1595995082] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1595995102] SERVICE ALERT: pservercluster2;Memory;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595995158] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.57% system=5.10% iowait=0.00% idle=93.32%
[1595995182] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1595995186] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.8% (17595424 kB) free.
[1595995186] SERVICE ALERT: sservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1595995257] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995269] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995280] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995338] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995410] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995468] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995540] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995598] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995712] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995768] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995792] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995796] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995796] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995898] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995918] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995922] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995926] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995926] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995976] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995980] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595995996] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996028] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996048] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996052] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996056] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996056] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996106] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996110] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996126] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996178] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996191] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996229] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16144.67, 10577.48, 4670.77
[1595996236] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996236] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996240] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996256] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996276] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996285] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996321] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996349] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16217.56, 12447.39, 6072.18
[1595996366] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996406] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996415] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996451] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996469] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16227.50, 13698.65, 7303.70
[1595996496] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996536] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595996545] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997050] SERVICE ALERT: sservercluster1;Memory;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595997180] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997277] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997278] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997305] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595997305] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595997306] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595997326] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997361] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595997361] SERVICE ALERT: sservercluster1;zombie_procs;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595997361] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595997363] SERVICE ALERT: sservercluster1;Open-Files;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595997392] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997397] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595997398] SERVICE ALERT: sservercluster1;HDD_Root status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595997430] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595997436] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997446] SERVICE ALERT: sservercluster1;CPU-STATISTICS;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595997481] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595997483] SERVICE ALERT: sservercluster1;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595997486] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595997491] SERVICE ALERT: sservercluster1;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997522] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595997527] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997528] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997546] SERVICE ALERT: sservercluster1;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997560] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997566] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997576] SERVICE ALERT: sservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997606] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595997611] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595997613] SERVICE ALERT: sservercluster1;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997616] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997621] SERVICE ALERT: sservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997652] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997676] SERVICE ALERT: sservercluster1;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997741] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997806] SERVICE ALERT: sservercluster1;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595997871] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595998336] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1595998336] SERVICE ALERT: sservercluster2;CPU Load;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595998345] SERVICE FLAPPING ALERT: sservercluster2;HDD_Root status;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1595998345] SERVICE ALERT: sservercluster2;HDD_Root status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595998946] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595998955] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595999328] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595999328] SERVICE ALERT: sservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 40505 MB (79% inode=99%):
[1595999360] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1595999366] SERVICE ALERT: sservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1595999382] SERVICE ALERT: sservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.66% system=5.72% iowait=0.00% idle=92.62%
[1595999407] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 315 processes
[1595999414] SERVICE ALERT: sservercluster1;Open-Files;OK;HARD;1;OK: 6848 open files (0% of max 2424092)
[1595999416] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 174889 MB (79% inode=99%):
[1595999421] SERVICE ALERT: sservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1595999451] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1595999451] SERVICE ALERT: sservercluster2;Current Users;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595999452] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1595999457] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1595999469] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595999501] SERVICE FLAPPING ALERT: sservercluster2;HDD_Home status;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1595999501] SERVICE ALERT: sservercluster2;HDD_Home status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595999501] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 183099 MB (81% inode=99%):
[1595999518] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 330 processes
[1595999546] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.00, 1.95, 1.40
[1595999555] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 33085 MB (64% inode=98%):
[1595999611] SERVICE ALERT: sservercluster1;Messages Log Status;UNKNOWN;HARD;3;NRPE: Unable to read output
[1595999652] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.89, 1.17, 1.27
[1595999671] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.40, 11.64, 10.17
[1595999706] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.0% (20355196 kB) free.
[1595999741] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1595999758] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1595999781] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 10592 open files (0% of max 2424156)
[1595999784] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.95% system=6.45% iowait=0.12% idle=91.47%
[1595999784] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595999798] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38287 MB (74% inode=98%):
[1595999846] SERVICE ALERT: sservercluster2;Total Processes;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595999856] SERVICE ALERT: sservercluster2;Open-Files;OK;HARD;1;OK: 8928 open files (0% of max 2424092)
[1595999912] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1595999912] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 81.7% (20017700 kB) free.
[1595999928] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595999928] SERVICE ALERT: sservercluster1;HDD_Root status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595999976] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000048] SERVICE ALERT: sservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 40504 MB (79% inode=99%):
[1596000048] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596000051] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596000069] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.05, 22.06, 285.52
[1596000096] SERVICE ALERT: sservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1596000101] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 203552 MB (92% inode=99%):
[1596000212] SERVICE ALERT: sservercluster1;Messages Log Status;CRITICAL;HARD;3;(235) < Jul 29 10:49:26 sservercluster1 xinetd[1585]: nrpe: fork failed: Cannot allocate memory (errno = 12)
[1596000233] SERVICE FLAPPING ALERT: sservercluster2;CPU-STATISTICS;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1596000233] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=0.59% system=2.56% iowait=0.00% idle=96.85%
[1596000252] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1596000252] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1596000256] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1596000256] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 78.9% (19342452 kB) free.
[1596000256] SERVICE FLAPPING ALERT: sservercluster2;zombie_procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1596000256] SERVICE ALERT: sservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596000378] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1596000440] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596000446] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 322 processes
[1596000454] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596000512] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1596000632] SERVICE ALERT: sservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596000642] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000667] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000679] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000711] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000728] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000746] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15678.57, 8078.47, 3263.73
[1596000772] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000787] SERVICE ALERT: pservercluster2;zombie_procs;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596000809] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000812] SERVICE FLAPPING ALERT: sservercluster1;Messages Log Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1596000812] SERVICE ALERT: sservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596000831] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 182981 MB (81% inode=99%):
[1596000858] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000862] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000866] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16135.74, 10768.01, 4833.10
[1596000917] SERVICE FLAPPING ALERT: pservercluster2;zombie_procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1596000917] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000939] SERVICE FLAPPING ALERT: pservercluster2;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1596000939] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000951] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000968] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000986] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16197.21, 12567.59, 6212.16
[1596000988] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000991] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000992] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000994] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596000994] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001008] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001064] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001081] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001098] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001121] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001122] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001124] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001124] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001138] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001194] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001211] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001228] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001251] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001254] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001254] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001268] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596001314] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;(192) < Jul 29 11:11:16 pservercluster2 xinetd[1483]: nrpe: fork failed: Cannot allocate memory (errno = 12)
[1596001372] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 77.6% (19025572 kB) free.
[1596001471] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15998.36, 9425.67, 3986.58
[1596001517] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596001539] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596001591] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16182.10, 11670.69, 5468.76
[1596001711] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16205.83, 13172.43, 6771.05
[1596001811] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1596001811] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1596001851] SERVICE FLAPPING ALERT: pservercluster2;Open-Files;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1596001851] SERVICE ALERT: pservercluster2;Open-Files;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596001854] SERVICE FLAPPING ALERT: pservercluster2;CPU-STATISTICS;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1596001854] SERVICE ALERT: pservercluster2;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596001854] SERVICE FLAPPING ALERT: pservercluster2;Sensors Status;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1596001854] SERVICE ALERT: pservercluster2;Sensors Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596001982] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596002041] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596002107] SERVICE ALERT: pservercluster2;Memory;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596002127] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596002149] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596002161] SERVICE ALERT: pservercluster2;HDD_Home status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596002188] SERVICE FLAPPING ALERT: pservercluster2;Total Processes;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1596002188] SERVICE ALERT: pservercluster2;Total Processes;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596002237] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596002247] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596002269] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596002281] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 174147 MB (77% inode=99%):
[1596002428] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1596002428] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1596002452] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 9216 open files (0% of max 2424156)
[1596002454] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1596002459] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.85% system=5.24% iowait=0.12% idle=92.79%
[1596002468] SERVICE FLAPPING ALERT: pservercluster2;HDD_Root status;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1596002468] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38287 MB (74% inode=98%):
[1596002789] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 333 processes
[1596002837] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.9% (21310592 kB) free.
[1596003114] SERVICE FLAPPING ALERT: pservercluster2;Messages Log Status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1596003114] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596003986] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.82, 2.76, 510.41
[1596004722] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1596004722] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.24, 4.93, 675.51
[1596005312] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1596005312] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.87, 2.17, 349.09
[1596005322] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.79, 1.42, 354.40
[1596008038] SERVICE FLAPPING ALERT: sservercluster2;CPU-STATISTICS;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1596008052] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1596008056] SERVICE FLAPPING ALERT: sservercluster2;zombie_procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1596008057] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1596010492] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.0% (6228324 kB) free!
[1596010612] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 18.8% (6160184 kB) free!
[1596010732] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.6% (6081132 kB) free!
[1596010955] SERVICE FLAPPING ALERT: sservercluster2;HDD_Root status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596011451] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596011501] SERVICE FLAPPING ALERT: sservercluster2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596012212] SERVICE FLAPPING ALERT: sservercluster1;Messages Log Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596012940] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (1742436 kB) free!
[1596012986] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1596013047] SERVICE FLAPPING ALERT: pservercluster2;zombie_procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596013060] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (1404540 kB) free!
[1596013069] SERVICE FLAPPING ALERT: pservercluster2;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596013180] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (1332152 kB) free!
[1596013211] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596013228] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1596013268] SERVICE FLAPPING ALERT: pservercluster2;HDD_Root status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1596013314] SERVICE FLAPPING ALERT: pservercluster2;Messages Log Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1596013712] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1596013780] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 70.8% (17482668 kB) free.
[1596013852] SERVICE FLAPPING ALERT: pservercluster2;Open-Files;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596013854] SERVICE FLAPPING ALERT: pservercluster2;Sensors Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596013864] SERVICE FLAPPING ALERT: pservercluster2;CPU-STATISTICS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596014190] SERVICE FLAPPING ALERT: pservercluster2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596014237] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596015482] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1508 processes
[1596015592] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9798660 kB) free!
[1596015602] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1504 processes
[1596015712] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9725264 kB) free!
[1596015723] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1499 processes
[1596015832] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (9652136 kB) free!
[1596016722] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596016922] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1502 processes
[1596017043] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1505 processes
[1596017163] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1504 processes
[1596017763] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1500 processes
[1596018363] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1502 processes
[1596018484] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1506 processes
[1596018604] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1502 processes
[1596018998] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 321 processes
[1596019118] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1596020632] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4935952 kB) free!
[1596020932] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (3263540 kB) free!
[1596022718] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 319 processes
[1596022732] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 10.8% (3520604 kB) free!
[1596022838] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1596023332] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (3234036 kB) free!
[1596024604] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1596024604] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1498 processes
[1596027009] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9571768 kB) free!
[1596027129] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9489672 kB) free!
[1596027249] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (9408376 kB) free!
[1596028269] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14856.21, 6389.11, 2415.56
[1596028389] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15999.21, 9627.85, 4084.19
[1596028509] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16153.12, 11794.79, 5550.43
[1596028666] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596028712] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15851.27, 13074.80, 6800.25
[1596028786] SERVICE ALERT: sservercluster2;Open-Files;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596028832] SERVICE ALERT: sservercluster2;Messages Log Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596028832] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15854.21, 13994.58, 7898.09
[1596028861] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596028901] SERVICE ALERT: sservercluster2;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596028916] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596028952] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15854.39, 14610.01, 8862.81
[1596028962] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596028965] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596028991] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029031] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029043] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029062] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029066] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029066] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029092] SERVICE ALERT: sservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029095] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029121] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029161] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029173] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029188] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029192] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029196] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029196] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029196] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029225] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029250] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029256] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029303] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029318] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029322] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029326] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029326] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029326] SERVICE ALERT: sservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029380] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029386] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029448] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029456] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029510] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029516] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596029721] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596029761] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 191098 MB (86% inode=99%):
[1596029825] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 33083 MB (64% inode=98%):
[1596029908] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=0.82% system=3.57% iowait=0.00% idle=95.61%
[1596029922] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1596029926] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.8% (20787748 kB) free.
[1596029926] SERVICE ALERT: sservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596030048] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1596030056] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.16, 4.77, 4.00
[1596030110] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596030116] SERVICE ALERT: sservercluster2;Open-Files;OK;HARD;1;OK: 6944 open files (0% of max 2424092)
[1596030116] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 320 processes
[1596030292] SERVICE ALERT: sservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596030399] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030447] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030452] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596030479] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030491] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030529] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030532] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030577] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030582] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030604] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1596030609] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030621] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030621] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030638] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030656] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16121.47, 10739.52, 4785.03
[1596030659] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596030659] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596030662] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030664] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596030669] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030678] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030707] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030712] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030724] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030739] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030751] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030751] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030768] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030776] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16183.34, 12543.87, 6168.21
[1596030789] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030792] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030794] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030799] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030808] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030854] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030881] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030896] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16191.55, 13751.23, 7383.69
[1596030898] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030922] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030924] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030929] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030938] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596030984] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596031952] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.37, 3.18, 547.24
[1596032081] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1596032098] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1596032109] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.93, 1.74, 363.74
[1596032122] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 10080 open files (0% of max 2424156)
[1596032124] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1596032134] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.69% system=5.65% iowait=0.00% idle=92.66%
[1596032138] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38285 MB (74% inode=98%):
[1596032459] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 334 processes
[1596032507] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.3% (21154876 kB) free.
[1596032512] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596032539] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596032551] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 171900 MB (76% inode=99%):
[1596032552] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.62, 2.25, 287.79
[1596032589] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.21, 1.31, 1.28
[1596032709] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.27, 1.36, 191.18
[1596032784] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596033296] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1596033296] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.79, 10.11, 843.78
[1596033897] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.54, 1.81, 442.39
[1596037133] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 44.6% (14598008 kB) free.
[1596039905] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596040035] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596040161] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=10.42% system=32.84% iowait=6.20% idle=50.54%
[1596045131] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 32.82, 39.53, 43.80
[1596045297] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)