[1582914702] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.1% (29147240 kB) free.
[1582914973] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7557712 kB) free!
[1582915093] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36534856 kB) free.
[1582915693] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12892236 kB) free!
[1582915813] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8895748 kB) free!
[1582915933] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5351388 kB) free!
[1582916533] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15869268 kB) free!
[1582916581] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (1343816 kB) free!
[1582916655] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582916715] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (185760 kB) free!
[1582916776] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (182636 kB) free!
[1582916884] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1705.92, 666.96, 246.88
[1582916992] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582917004] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 401.81, 584.23, 273.24
[1582917076] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 66.7% (16342528 kB) free.
[1582917134] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26862912 kB) free.
[1582917734] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6163196 kB) free!
[1582917854] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33052976 kB) free.
[1582918192] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1582918313] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1582918433] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1582918454] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11025616 kB) free!
[1582918574] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7585616 kB) free!
[1582918694] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35606436 kB) free.
[1582919033] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582919244] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.1% (3715620 kB) free!
[1582919294] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12203368 kB) free!
[1582919414] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8509968 kB) free!
[1582919534] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35864836 kB) free.
[1582919633] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582919753] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1582919844] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (744380 kB) free!
[1582919874] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1582920134] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12846864 kB) free!
[1582920254] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9186964 kB) free!
[1582920374] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35802772 kB) free.
[1582920474] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582920974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13143868 kB) free!
[1582921074] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582921094] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9575488 kB) free!
[1582921194] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1582921214] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35719312 kB) free.
[1582921314] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1582921752] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46726 MB (20% inode=99%):
[1582921814] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12699748 kB) free!
[1582921872] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46726 MB (20% inode=99%):
[1582921935] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8970948 kB) free!
[1582921992] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46726 MB (20% inode=99%):
[1582922055] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35645812 kB) free.
[1582922655] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12348824 kB) free!
[1582922775] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8926064 kB) free!
[1582922895] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34449636 kB) free.
[1582923495] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12426332 kB) free!
[1582923615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9119436 kB) free!
[1582923735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32510328 kB) free.
[1582924320] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 84.4% (41599492 kB) free.
[1582924335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6052152 kB) free!
[1582924455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (2039152 kB) free!
[1582924560] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.3% (10112872 kB) free.
[1582924575] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35948588 kB) free.
[1582924856] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.4% (45053200 kB) free.
[1582925175] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9124092 kB) free!
[1582925295] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5404636 kB) free!
[1582925415] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35908220 kB) free.
[1582925464] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(200) < 103.16.101.72 [29/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582925586] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582925767] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(260) < 2020/02/29 03:00:22 [error] 12249#0: *61994202 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582925788] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4166) < 203.212.70.170 [29/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582925887] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582925908] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582926015] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11550656 kB) free!
[1582926135] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7659560 kB) free!
[1582926256] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35845796 kB) free.
[1582926855] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12744080 kB) free!
[1582926975] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8742836 kB) free!
[1582927095] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35838532 kB) free.
[1582927695] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9926076 kB) free!
[1582927815] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5937284 kB) free!
[1582927935] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33943880 kB) free.
[1582928514] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582928535] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12536368 kB) free!
[1582928655] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37805540 kB) free.
[1582929114] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1582929235] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1582929255] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19355812 kB) free!
[1582929355] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582929375] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15829516 kB) free!
[1582929487] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2020/02/29 04:00:52 [error] 12256#0: *62015086 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582929495] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37812604 kB) free.
[1582929508] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 203.212.70.170 [29/Feb/2020:04:03:00 +0530] TCP 502 0 0 135.027 "10.147.212.155:3001, 10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "15.026, 60.000, 60.001, 0.000"
[1582929607] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582929628] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582929955] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582930075] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582930095] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18456960 kB) free!
[1582930215] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15186948 kB) free!
[1582930335] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36081240 kB) free.
[1582930675] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582930795] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1582930915] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1582930936] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17558056 kB) free!
[1582931056] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39891784 kB) free.
[1582931656] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22106464 kB) free!
[1582931776] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39903240 kB) free.
[1582932115] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582932169] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.6% (13930748 kB) free!
[1582932376] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23108340 kB) free!
[1582932496] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19816396 kB) free!
[1582932616] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582932616] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38332364 kB) free.
[1582932715] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1582932835] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1582932955] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1582932960] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202508 kB) free!
[1582933080] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (207492 kB) free!
[1582933201] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201280 kB) free!
[1582933216] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19631960 kB) free!
[1582933336] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38800156 kB) free.
[1582933936] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21304188 kB) free!
[1582934056] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39797552 kB) free.
[1582934061] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.0% (20842664 kB) free.
[1582934244] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 54.3% (13399168 kB) free.
[1582934656] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22923992 kB) free!
[1582934776] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19616784 kB) free!
[1582934896] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38310464 kB) free.
[1582935496] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19708856 kB) free!
[1582935616] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38240184 kB) free.
[1582936217] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20243328 kB) free!
[1582936337] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38689464 kB) free.
[1582936369] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (13191680 kB) free!
[1582936937] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19156508 kB) free!
[1582937057] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38872368 kB) free.
[1582937569] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 29.2% (38508148 kB) free.
[1582937657] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20251976 kB) free!
[1582937777] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39462372 kB) free.
[1582938377] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20850996 kB) free!
[1582938497] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39323080 kB) free.
[1582938693] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1582938813] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1582939097] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21478896 kB) free!
[1582939217] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39217964 kB) free.
[1582939817] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22326924 kB) free!
[1582939937] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18676472 kB) free!
[1582940057] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37750288 kB) free.
[1582940155] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582940657] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19384096 kB) free!
[1582940755] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582940777] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38474784 kB) free.
[1582940876] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1582940996] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582941377] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19370972 kB) free!
[1582941498] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38347876 kB) free.
[1582941596] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1582941716] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1582941836] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582942097] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19212652 kB) free!
[1582942217] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38566800 kB) free.
[1582942267] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.0% (14136828 kB) free.
[1582942295] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 63.5% (31298296 kB) free.
[1582942349] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 57.1% (13984644 kB) free.
[1582942378] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.3% (21402332 kB) free.
[1582942436] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582942557] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1582942677] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582942766] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.0% (22053376 kB) free.
[1582942817] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19087600 kB) free!
[1582942937] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38447256 kB) free.
[1582943277] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1582943397] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1582943466] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1582943517] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1582943537] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18012776 kB) free!
[1582943582] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199396 kB) free!
[1582943586] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582943657] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37712288 kB) free.
[1582943661] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196372 kB) free!
[1582943702] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194740 kB) free!
[1582943706] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582943781] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197612 kB) free!
[1582943822] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164616 kB) free!
[1582943829] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582943901] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193904 kB) free!
[1582943949] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582944028] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [29/Feb/2020:08:02:35 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.000, 0.000"
[1582944069] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1582944148] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582944192] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 105281 MB (46% inode=99%):
[1582944194] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.7% (21506548 kB) free.
[1582944251] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.5% (21944580 kB) free.
[1582944257] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17762592 kB) free!
[1582944307] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1582944377] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37450028 kB) free.
[1582944444] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163188 kB) free!
[1582944564] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (204756 kB) free!
[1582944628] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21932264 kB) free.
[1582944718] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1582944977] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17635456 kB) free!
[1582945097] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37045344 kB) free.
[1582945697] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17187068 kB) free!
[1582945817] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36849560 kB) free.
[1582946407] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(72) < 2020/02/29 08:42:07 [error] 12257#0: *62100585 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1582946417] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16535344 kB) free!
[1582946527] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582946537] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35973184 kB) free.
[1582947137] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15440144 kB) free!
[1582947258] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35369664 kB) free.
[1582947857] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14017880 kB) free!
[1582947977] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34963244 kB) free.
[1582948060] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 182.26, 258.50, 124.13
[1582948180] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 675.93, 369.23, 177.22
[1582948301] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 948.47, 569.17, 273.10
[1582948577] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13345876 kB) free!
[1582948697] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34990600 kB) free.
[1582949297] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12228056 kB) free!
[1582949417] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8323252 kB) free!
[1582949537] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31541492 kB) free.
[1582950137] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9649008 kB) free!
[1582950257] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29722332 kB) free.
[1582950857] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6249676 kB) free!
[1582950977] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40816148 kB) free.
[1582951577] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18005568 kB) free!
[1582951697] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36739084 kB) free.
[1582951869] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1582951901] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 21.53, 18.81, 34.17
[1582951989] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1582952109] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1582952298] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18606188 kB) free!
[1582952418] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14430456 kB) free!
[1582952501] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 22.21, 20.68, 27.80
[1582952538] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38745136 kB) free.
[1582953138] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16073476 kB) free!
[1582953258] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (12014632 kB) free!
[1582953378] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36778424 kB) free.
[1582953909] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582953978] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14050576 kB) free!
[1582954098] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (9993436 kB) free!
[1582954218] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34721616 kB) free.
[1582954818] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11835284 kB) free!
[1582954901] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9270560 kB) free!
[1582954938] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7917256 kB) free!
[1582954972] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2783652 kB) free!
[1582955021] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8398760 kB) free!
[1582955058] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4166404 kB) free!
[1582955092] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.9% (2451920 kB) free!
[1582955141] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (9201420 kB) free!
[1582955212] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2088536 kB) free!
[1582955518] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582956032] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4694644 kB) free!
[1582956148] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(11) < 103.255.146.164 [29/Feb/2020:11:31:38 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1582956152] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.2% (2766684 kB) free!
[1582956258] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17108276 kB) free!
[1582956267] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195328 kB) free!
[1582956268] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582956272] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2625432 kB) free!
[1582956341] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4694100 kB) free!
[1582956387] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172496 kB) free!
[1582956479] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.0% (15717632 kB) free.
[1582956504] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.6% (38756860 kB) free.
[1582956508] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187972 kB) free!
[1582956528] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.8% (21526112 kB) free.
[1582956858] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26825492 kB) free.
[1582956872] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186500 kB) free!
[1582957130] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187516 kB) free!
[1582957251] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188120 kB) free!
[1582957302] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7469908 kB) free!
[1582957371] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190968 kB) free!
[1582957403] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.1% (4181080 kB) free!
[1582957422] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.2% (6997316 kB) free!
[1582957459] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3406032 kB) free!
[1582957523] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.8% (3869804 kB) free!
[1582957525] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3845280 kB) free!
[1582957542] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.4% (6594452 kB) free!
[1582957579] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394596 kB) free!
[1582957643] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.6% (3576236 kB) free!
[1582957645] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.1% (3482904 kB) free!
[1582957699] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27997484 kB) free.
[1582957765] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (3208668 kB) free!
[1582957831] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582957912] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582957927] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(58) < 2020/02/29 12:00:56 [error] 12255#0: *62181735 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582957951] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582958042] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582958047] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582958071] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582958109] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1582958152] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582958153] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582958163] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 24.22, 21.10, 19.44
[1582958283] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582958299] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (391136 kB) free!
[1582958365] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1937612 kB) free!
[1582958408] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2181 processes
[1582958419] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400372 kB) free!
[1582958539] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37520436 kB) free.
[1582958753] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.1% (30124168 kB) free.
[1582958843] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1337540 kB) free!
[1582958966] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22890728 kB) free.
[1582958992] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582959122] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582959139] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7012780 kB) free!
[1582959152] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163152 kB) free!
[1582959247] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582959259] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2944360 kB) free!
[1582959271] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582959272] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180012 kB) free!
[1582959363] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582959379] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1582959379] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393892 kB) free!
[1582959392] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190416 kB) free!
[1582959493] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582959580] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582959585] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582959618] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.8% (29980344 kB) free.
[1582959672] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (2538592 kB) free!
[1582959709] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1582959710] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582959793] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.6% (2538408 kB) free!
[1582959839] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2182 processes
[1582959847] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/29 12:27:24 [error] 12254#0: *62196147 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/uuid HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/DlrLoad/uuid"
[1582959852] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582959913] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (2536344 kB) free!
[1582959968] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582959979] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14910020 kB) free!
[1582959982] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960106] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582960223] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960315] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960353] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960435] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1582960483] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1582960483] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960509] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582960529] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960572] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960579] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31086552 kB) free.
[1582960629] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582960655] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1582960696] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 26.35, 26.62, 24.62
[1582960712] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960749] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1582960749] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582960813] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960842] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960943] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582960971] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582961040] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582961073] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582961170] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582961179] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7235516 kB) free!
[1582961259] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582961299] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3351360 kB) free!
[1582961300] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582961302] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582961389] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582961419] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30706176 kB) free.
[1582961422] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 255.25, 87.54, 47.04
[1582961456] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8797192 kB) free!
[1582961514] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1582961542] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 57.01, 66.96, 44.39
[1582961572] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582961577] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8438228 kB) free!
[1582961674] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2200 processes
[1582961691] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.1% (28622920 kB) free.
[1582961692] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582961697] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.4% (8092768 kB) free!
[1582961900] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2194 processes
[1582961949] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1582962019] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5128136 kB) free!
[1582962140] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (997348 kB) free!
[1582962142] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1582962142] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 20.72, 27.96, 33.84
[1582962260] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402172 kB) free!
[1582962313] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (944600 kB) free!
[1582962549] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582962669] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582962742] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 18.41, 20.34, 26.77
[1582962790] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1582962860] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13793792 kB) free!
[1582962868] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(812) < 103.255.146.164 [29/Feb/2020:13:18:04 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1582962897] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4716352 kB) free!
[1582962907] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582962988] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582963027] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582963147] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582963709] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.9% (10841316 kB) free.
[1582963710] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8397988 kB) free!
[1582963727] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 28.1% (13847444 kB) free.
[1582963757] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.9% (21291152 kB) free.
[1582963782] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1582963782] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.9% (20311560 kB) free.
[1582963810] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.9% (14123660 kB) free.
[1582963830] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8394168 kB) free!
[1582963832] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.5% (21700160 kB) free.
[1582963950] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (8394224 kB) free!
[1582964060] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416584 kB) free!
[1582964347] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582964362] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (2401076 kB) free!
[1582964392] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173924 kB) free!
[1582964467] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582964482] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (2196796 kB) free!
[1582964513] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190044 kB) free!
[1582964587] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582964602] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2011552 kB) free!
[1582964633] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175360 kB) free!
[1582964660] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14812796 kB) free!
[1582965037] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.8% (4855080 kB) free!
[1582965157] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.7% (4584076 kB) free!
[1582965188] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1582965260] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30164852 kB) free.
[1582965277] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1582965277] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.7% (4339060 kB) free!
[1582965648] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582965648] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.9% (21292856 kB) free.
[1582965682] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.9% (21310520 kB) free.
[1582965789] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1582965860] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5225204 kB) free!
[1582965871] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582965877] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 55.7% (13651472 kB) free.
[1582965909] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1582965980] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1082992 kB) free!
[1582965988] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(681) < 103.255.146.164 [29/Feb/2020:14:14:38 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1582965991] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582966081] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3296140 kB) free!
[1582966100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35677932 kB) free.
[1582966108] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582966111] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (854884 kB) free!
[1582966201] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.0% (2917920 kB) free!
[1582966231] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (710608 kB) free!
[1582966282] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175112 kB) free!
[1582966321] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9589780 kB) free!
[1582966321] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.7% (2618720 kB) free!
[1582966351] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (566592 kB) free!
[1582966402] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183420 kB) free!
[1582966441] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9403936 kB) free!
[1582966510] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582966522] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195452 kB) free!
[1582966561] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (9222400 kB) free!
[1582966630] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1582966700] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10461452 kB) free!
[1582966820] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6051884 kB) free!
[1582966918] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1582966921] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1198048 kB) free!
[1582966940] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1834376 kB) free!
[1582966950] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2702936 kB) free!
[1582967038] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1582967158] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1582967308] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(593) < 103.255.146.164 [29/Feb/2020:14:29:52 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1582967428] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582967540] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15990368 kB) free!
[1582967759] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 694 processes
[1582968141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33574304 kB) free.
[1582968291] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.3% (8995404 kB) free!
[1582968411] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.3% (8548020 kB) free!
[1582968430] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582968531] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.6% (8159612 kB) free!
[1582968550] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1582968628] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(149) < 92.42.104.166 [29/Feb/2020:15:00:12 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582968722] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.9% (15917792 kB) free.
[1582968741] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6950668 kB) free!
[1582968748] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582968845] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582968861] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2619536 kB) free!
[1582968970] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1582968981] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396404 kB) free!
[1582969132] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.0% (29577568 kB) free.
[1582969150] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582969166] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4494796 kB) free!
[1582969270] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582969286] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (4125224 kB) free!
[1582969391] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582969406] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 15.4% (3807000 kB) free!
[1582969484] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582969611] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2203 processes
[1582969742] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582969872] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582969988] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582970002] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582970006] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1582970006] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22883696 kB) free.
[1582970077] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184304 kB) free!
[1582970108] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582970181] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30523084 kB) free.
[1582970197] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192464 kB) free!
[1582970228] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582970317] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185860 kB) free!
[1582970591] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1582970761] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4579828 kB) free!
[1582970781] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5334892 kB) free!
[1582970828] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582970828] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1582970901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1224692 kB) free!
[1582971021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36304500 kB) free.
[1582971191] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582971202] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.9% (29540080 kB) free.
[1582971311] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582971391] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582971431] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1582971511] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582971517] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 58.3% (14295080 kB) free.
[1582971621] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9185356 kB) free!
[1582971631] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582971741] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4780440 kB) free!
[1582971861] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (699668 kB) free!
[1582972461] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15075708 kB) free!
[1582972567] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(137) < 2020/02/29 16:02:10 [error] 12250#0: *62318848 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582972687] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582973061] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36682564 kB) free.
[1582973661] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8882984 kB) free!
[1582973782] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4421652 kB) free!
[1582973887] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.1% (21743080 kB) free.
[1582973902] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386608 kB) free!
[1582973922] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.3% (45006420 kB) free.
[1582973957] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32779948 kB) free.
[1582973994] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39253876 kB) free.
[1582974035] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.1% (19382116 kB) free.
[1582974142] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582974490] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176896 kB) free!
[1582974610] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189136 kB) free!
[1582974641] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187912 kB) free!
[1582974650] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181964 kB) free!
[1582974730] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182624 kB) free!
[1582974761] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206892 kB) free!
[1582974770] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188504 kB) free!
[1582974881] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185544 kB) free!
[1582974890] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184780 kB) free!
[1582974942] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2974988 kB) free!
[1582975062] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3246576 kB) free!
[1582975182] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3520900 kB) free.
[1582975632] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1582975760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21397356 kB) free!
[1582975800] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9538824 kB) free!
[1582975880] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (16990868 kB) free!
[1582975921] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9539440 kB) free!
[1582976000] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12800860 kB) free!
[1582976041] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (9534620 kB) free!
[1582976232] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 356 processes
[1582976352] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 356 processes
[1582976472] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 362 processes
[1582976522] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (3689332 kB) free!
[1582976600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26551624 kB) free.
[1582976642] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3375708 kB) free!
[1582976762] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.6% (3052612 kB) free!
[1582976982] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2558808 kB) free!
[1582977102] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2993008 kB) free!
[1582977200] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4170412 kB) free!
[1582977222] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3430520 kB) free.
[1582977320] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386580 kB) free!
[1582977362] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1582977362] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1274296 kB) free!
[1582977440] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21427140 kB) free!
[1582977517] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.5% (3797796 kB) free!
[1582977637] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.4% (3517928 kB) free!
[1582977672] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 362 processes
[1582977757] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.1% (3215888 kB) free!
[1582977841] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (601980 kB) free!
[1582978040] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396408 kB) free!
[1582978357] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1969208 kB) free!
[1582978402] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7861500 kB) free!
[1582978522] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7394280 kB) free!
[1582978559] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1582978640] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23577460 kB) free!
[1582978642] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.2% (6985232 kB) free!
[1582978679] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 706 processes
[1582978799] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 706 processes
[1582979006] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (4117588 kB) free!
[1582979023] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (2466288 kB) free!
[1582979126] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.4% (3806832 kB) free!
[1582979143] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3142632 kB) free!
[1582979240] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382152 kB) free!
[1582979242] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4839752 kB) free!
[1582979246] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (3497288 kB) free!
[1582979263] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3683140 kB) free.
[1582979302] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (2476232 kB) free!
[1582979422] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (1553712 kB) free!
[1582979519] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979542] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.5% (5116448 kB) free.
[1582979649] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979702] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979762] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.2% (15992872 kB) free.
[1582979775] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979779] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979815] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979827] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979832] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29455680 kB) free.
[1582979846] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1715472 kB) free!
[1582979905] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979910] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582979932] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8915580 kB) free!
[1582979936] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2208 processes
[1582979947] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 93523 MB (41% inode=99%):
[1582979962] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582980035] SERVICE ALERT: smscapp2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582980040] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582980052] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8650544 kB) free!
[1582980071] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1582980071] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39075652 kB) free.
[1582980108] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21724120 kB) free.
[1582980143] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21234292 kB) free.
[1582980166] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2214 processes
[1582980172] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.0% (8357312 kB) free!
[1582980187] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.9% (21780484 kB) free.
[1582980249] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.1% (21601768 kB) free.
[1582980441] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5506960 kB) free!
[1582980446] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22879896 kB) free.
[1582980447] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.6% (28874196 kB) free.
[1582980561] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1503584 kB) free!
[1582980599] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 693 processes
[1582980681] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390608 kB) free!
[1582980752] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582980754] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167548 kB) free!
[1582980770] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582980851] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187796 kB) free!
[1582980874] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181068 kB) free!
[1582980882] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582980893] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2224 processes
[1582980972] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178900 kB) free!
[1582980994] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1582980994] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166124 kB) free!
[1582981012] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981057] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981092] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166628 kB) free!
[1582981144] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981145] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981187] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981274] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981275] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981281] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16980828 kB) free!
[1582981312] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.8% (28954088 kB) free.
[1582981342] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3077956 kB) free!
[1582981404] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981405] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981462] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.9% (4048708 kB) free.
[1582981616] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 26.75, 28.66, 27.68
[1582981757] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981763] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582981881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40332892 kB) free.
[1582981881] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8111472 kB) free!
[1582981882] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 89773 MB (39% inode=99%):
[1582981917] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582981973] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4368724 kB) free!
[1582982001] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (8043140 kB) free!
[1582982047] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582982062] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3014328 kB) free!
[1582982121] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (7999592 kB) free!
[1582982177] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582982182] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3738376 kB) free.
[1582982187] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1582982228] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582982373] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582982481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13287716 kB) free!
[1582982501] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582982548] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(44) < 103.255.146.164 [29/Feb/2020:18:52:08 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1582982601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8835608 kB) free!
[1582982668] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582982721] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4583676 kB) free!
[1582982778] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.5% (28828660 kB) free.
[1582982782] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (2751076 kB) free!
[1582982796] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 622 processes
[1582982902] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3500488 kB) free.
[1582982917] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 363 processes
[1582983043] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1582983204] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38756 MB (75% inode=98%):
[1582983207] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2226 processes
[1582983321] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23419344 kB) free!
[1582983502] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (295044 kB) free!
[1582983622] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.8% (5168288 kB) free.
[1582983652] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21783068 kB) free.
[1582983706] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.3% (45010376 kB) free.
[1582983739] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.7% (21749132 kB) free.
[1582983921] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385428 kB) free!
[1582983921] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4913036 kB) free!
[1582984342] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169452 kB) free!
[1582984462] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181720 kB) free!
[1582984521] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20006052 kB) free!
[1582984521] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (5250280 kB) free!
[1582984582] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168672 kB) free!
[1582984822] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3020132 kB) free!
[1582984942] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3775012 kB) free.
[1582985121] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380332 kB) free!
[1582985121] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (1628468 kB) free!
[1582985542] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.1% (2295900 kB) free!
[1582985662] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2912840 kB) free!
[1582985721] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17590348 kB) free!
[1582985782] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3608360 kB) free.
[1582986190] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1582986321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38998908 kB) free.
[1582986383] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.2% (1977940 kB) free!
[1582986503] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2582040 kB) free!
[1582986623] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3199840 kB) free!
[1582986672] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 383 processes
[1582986868] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(55) < 203.212.70.170 [29/Feb/2020:20:03:04 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1582986922] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14331068 kB) free!
[1582986943] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.4% (19940380 kB) free.
[1582986984] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39250176 kB) free.
[1582986988] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582987042] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9797948 kB) free!
[1582987088] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/02/29 20:01:56 [error] 12253#0: *62416339 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:32440/454934, bytes from/to upstream:454934/32440
[1582987162] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5662536 kB) free!
[1582987208] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582987272] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1582987550] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167396 kB) free!
[1582987670] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174504 kB) free!
[1582987762] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26143156 kB) free!
[1582987790] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185388 kB) free!
[1582987823] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4029028 kB) free.
[1582987864] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180708 kB) free!
[1582987984] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188260 kB) free!
[1582988104] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170316 kB) free!
[1582988362] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (4016552 kB) free!
[1582988423] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.8% (1913612 kB) free!
[1582988472] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1582988543] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (2337816 kB) free!
[1582988663] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (2719116 kB) free!
[1582988791] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9532848 kB) free!
[1582988911] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9457428 kB) free!
[1582988962] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20313324 kB) free!
[1582989031] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (9349804 kB) free!
[1582989263] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1582989263] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3821184 kB) free.
[1582989362] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582989562] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376852 kB) free!
[1582989672] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1582989711] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582989792] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1582989863] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (2258464 kB) free!
[1582989983] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.8% (2575664 kB) free!
[1582990103] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2851780 kB) free!
[1582990162] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18630004 kB) free!
[1582990703] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3640564 kB) free.
[1582990762] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38949512 kB) free.
[1582991362] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13842204 kB) free!
[1582991482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9502636 kB) free!
[1582991593] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582991602] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (5515708 kB) free!
[1582991713] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1582991846] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582992202] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23218880 kB) free!
[1582992802] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1170240 kB) free!
[1582993402] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19320832 kB) free!
[1582993513] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1582993633] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1582994003] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400180 kB) free!
[1582994057] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.7% (21886528 kB) free.
[1582994111] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 53.6% (70651888 kB) free.
[1582994149] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.7% (21983784 kB) free.
[1582994178] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582994752] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181960 kB) free!
[1582994873] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193836 kB) free!
[1582994993] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183016 kB) free!
[1582995190] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4901536 kB) free!
[1582995310] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4824296 kB) free!
[1582995430] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4749388 kB) free!
[1582995433] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1582995553] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1582995922] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16881668 kB) free!
[1582996024] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 725 processes
[1582996042] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12942848 kB) free!
[1582996144] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 725 processes
[1582996153] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582996162] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9128884 kB) free!
[1582996208] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2020/02/29 22:36:44 [error] 12253#0: *62528864 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582996264] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 334 processes
[1582996273] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1582996328] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582996393] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1582996630] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172396 kB) free!
[1582996762] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.7% (26010084 kB) free!
[1582996993] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1582997061] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181916 kB) free!
[1582997181] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174176 kB) free!
[1582997301] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200008 kB) free!
[1582997362] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.9% (47389308 kB) free.
[1582997431] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582997963] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22453804 kB) free!
[1582998031] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1582998083] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.9% (18276516 kB) free!
[1582998194] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582998203] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14428592 kB) free!
[1582998314] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582998434] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582998803] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31050768 kB) free.
[1582999231] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4263284 kB) free!
[1582999403] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10747744 kB) free!
[1582999523] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.2% (46365376 kB) free.
[1582999578] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3536088 kB) free!
[1582999698] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (246320 kB) free!
[1582999818] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (244356 kB) free!
[1582999969] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.3% (16219252 kB) free!
[1583000089] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.7% (14152328 kB) free!
[1583000123] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21825760 kB) free!
[1583000209] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9260480 kB) free!
[1583000243] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17688456 kB) free!
[1583000363] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13970304 kB) free!
[1583000834] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1583000955] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1583000963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29013664 kB) free.