[1591641641] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45831 MB (20% inode=99%):
[1591641762] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45829 MB (20% inode=99%):
[1591641838] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 6.95, 27.50, 42.39
[1591641884] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45826 MB (20% inode=99%):
[1591643638] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 64.18, 52.18, 46.40
[1591644238] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 29.42, 35.39, 40.26
[1591645075] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183528 kB) free!
[1591645137] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167724 kB) free!
[1591645197] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (192184 kB) free!
[1591645257] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (197552 kB) free!
[1591645857] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 83.9% (20560760 kB) free.
[1591646091] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1591646091] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49829 MB (22% inode=99%):
[1591646409] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (498112 kB) free!
[1591646726] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216740 kB) free!
[1591646846] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242408 kB) free!
[1591646966] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243896 kB) free!
[1591647235] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 220097 MB (20% inode=99%):
[1591647355] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 214025 MB (20% inode=99%):
[1591647475] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 208116 MB (19% inode=99%):
[1591648767] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.0% (9354048 kB) free!
[1591650766] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 76.2% (37576468 kB) free.
[1591651569] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.9% (41339560 kB) free.
[1591651827] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(187) < 103.16.101.72 [09/Jun/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1591651846] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(228) < 2020/06/09 03:00:21 [error] 12482#0: *93863909 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1591651925] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1167) < 203.212.70.170 [09/Jun/2020:03:01:23 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591651948] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 49.50.86.152 [09/Jun/2020:03:01:22 +0530] TCP 502 0 0 60.001 "10.147.212.61:4000, backend3" "0, 0" "0, 0" "60.001, 0.000"
[1591651966] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(23) < 2020/06/09 03:01:22 [error] 12478#0: *93864023 no live upstreams while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1591652045] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(3) < 172.105.59.173 [09/Jun/2020:03:02:22 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1591652068] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591652086] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591652165] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591652369] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 85.2% (41981112 kB) free.
[1591653893] SERVICE FLAPPING ALERT: smscapp2;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1591654565] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.2% (10336608 kB) free.
[1591655638] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591655765] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (1174536 kB) free!
[1591655885] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183664 kB) free!
[1591656006] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176260 kB) free!
[1591656238] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 22.19, 26.68, 29.36
[1591656838] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 39.76, 31.91, 30.17
[1591656958] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 32.94, 30.45, 29.78
[1591657039] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (488816 kB) free!
[1591657159] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (177784 kB) free!
[1591657206] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.1% (10312184 kB) free.
[1591657279] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176868 kB) free!
[1591657905] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (5166684 kB) free!
[1591658158] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 38.28, 33.93, 32.84
[1591658278] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 36.43, 33.17, 32.61
[1591658398] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 22.49, 30.67, 31.89
[1591659606] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (203328 kB) free!
[1591659726] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175160 kB) free!
[1591659846] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1591659846] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178692 kB) free!
[1591660149] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.6% (18533960 kB) free.
[1591660305] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4854964 kB) free!
[1591660697] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 56.0% (13811084 kB) free.
[1591661999] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 57.57, 40.18, 37.14
[1591662599] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1591662599] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 37.28, 39.41, 37.97
[1591663617] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 71.1% (17416372 kB) free.
[1591665009] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 41.1% (54240940 kB) free.
[1591666075] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 240556 MB (22% inode=99%):
[1591666199] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.11, 48.73, 45.87
[1591667897] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.2% (3262624 kB) free!
[1591668017] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.7% (2898592 kB) free!
[1591668137] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (2488240 kB) free!
[1591668737] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2308960 kB) free!
[1591669446] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1591673079] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.61, 5.46, 5.03
[1591673199] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.06, 5.89, 5.27
[1591673319] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.03, 5.12, 5.06
[1591673784] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1591673905] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 12.63, 12.87, 11.81
[1591674599] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1591677417] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4503312 kB) free!
[1591677537] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (4280052 kB) free!
[1591677658] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (4055320 kB) free!
[1591678858] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2102616 kB) free!
[1591679078] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1591679198] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1591679609] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(5) < Channel error on connection <0.9842.3839> (10.147.212.158:34862 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1591679730] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1591679925] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.55, 5.92, 5.07
[1591680045] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.03, 5.65, 5.10
[1591680166] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.03, 5.70, 5.19
[1591680315] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.7% (2635740 kB) free!
[1591680435] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (2025000 kB) free!
[1591680555] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1443948 kB) free!
[1591680658] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (3450728 kB) free!
[1591681149] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (4051308 kB) free!
[1591681258] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1591681258] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1277288 kB) free!
[1591681270] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (956020 kB) free!
[1591681338] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1591681366] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.98, 4.76, 4.85
[1591681390] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178380 kB) free!
[1591681420] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1591681458] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1591681540] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1591681578] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1591681660] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1591681966] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.00, 6.16, 5.27
[1591682086] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.78, 6.21, 5.39
[1591682178] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1591682206] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.48, 6.58, 5.62
[1591682260] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1591683873] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (6563748 kB) free!
[1591683969] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8231724 kB) free!
[1591683994] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (6224108 kB) free!
[1591684089] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.0% (7889692 kB) free!
[1591684114] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 16.7% (5858688 kB) free!
[1591684209] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.3% (7542788 kB) free!
[1591684478] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 52.6% (12881612 kB) free.
[1591684606] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1591684606] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.75, 3.60, 4.45
[1591684714] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.6% (31732916 kB) free.
[1591684755] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.4% (16037976 kB) free.
[1591685116] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4723116 kB) free!
[1591685236] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4405084 kB) free!
[1591685356] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.6% (4093436 kB) free!
[1591685409] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4032060 kB) free!
[1591685765] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 172.105.59.173 [09/Jun/2020:12:26:01 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1591685885] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(9) < 203.212.70.170 [09/Jun/2020:12:26:07 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1591685956] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.3% (22775036 kB) free.
[1591686005] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591686378] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1591686406] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.97, 6.10, 5.14
[1591686498] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1591686526] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.84, 5.97, 5.23
[1591686619] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1591686646] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.27, 5.35, 5.09
[1591687548] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (26.9% change >= 20.0% threshold)
[1591687598] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1591687607] SERVICE ALERT: TempHP2;CPU-STATISTICS;UNKNOWN;SOFT;1;NRPE: Command 'check_cpu_stats.sh' not defined
[1591687718] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1591687732] SERVICE ALERT: TempHP2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=68.96% system=27.84% iowait=0.00% idle=3.20%
[1591687787] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (26.9% change >= 20.0% threshold)
[1591688078] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.1% (4428948 kB) free!
[1591688199] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.9% (4142044 kB) free!
[1591688319] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.7% (3844176 kB) free!
[1591688468] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1591688478] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1591688651] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.55, 3.94, 4.51
[1591688920] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2300688 kB) free!
[1591690755] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.4% (3446672 kB) free!
[1591690875] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.3% (2994876 kB) free!
[1591690966] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9254884 kB) free!
[1591690995] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.5% (2575860 kB) free!
[1591691086] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (9085748 kB) free!
[1591691206] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8926480 kB) free!
[1591691315] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 18.2% (5968000 kB) free!
[1591691351] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1591691351] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1591691405] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 203.212.70.170 [09/Jun/2020:13:57:20 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1591691435] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 25.1% (8226728 kB) free.
[1591691525] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591691595] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1591691595] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (848348 kB) free!
[1591691918] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1591691951] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1591692038] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1591692071] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 346 processes
[1591692191] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 346 processes
[1591692635] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 20.0% (6561760 kB) free!
[1591692755] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.6% (6416172 kB) free!
[1591692852] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.03, 6.57, 5.47
[1591692875] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.2% (6287936 kB) free!
[1591692972] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.63, 6.32, 5.51
[1591693060] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1591693092] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.98, 6.00, 5.50
[1591693180] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1591693300] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1591693838] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1591693958] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1591694206] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4349928 kB) free!
[1591694292] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.02, 3.20, 4.04
[1591694525] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [09/Jun/2020:14:45:25 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1591694646] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591694953] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1591695073] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 359 processes
[1591695193] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 356 processes
[1591695196] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.9% (16138776 kB) free.
[1591695245] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 185.255.8.58 [09/Jun/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1591695365] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591696359] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1591696479] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 335 processes
[1591696599] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1591697075] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (3244548 kB) free!
[1591697799] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1591699900] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1591699993] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1591700165] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 172.105.59.173 [09/Jun/2020:16:18:57 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1591700285] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591701192] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1591701796] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.5% (3488208 kB) free!
[1591701891] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46807 MB (20% inode=99%):
[1591701916] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3184808 kB) free!
[1591702011] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46805 MB (20% inode=99%):
[1591702036] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (2851852 kB) free!
[1591702131] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46801 MB (20% inode=99%):
[1591702636] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (834604 kB) free!
[1591702692] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.65, 6.71, 5.72
[1591702812] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.72, 6.33, 5.70
[1591702932] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.47, 5.70, 5.55
[1591704531] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 193703 MB (86% inode=99%):
[1591704700] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1591704792] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 356 processes
[1591704820] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1591704912] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 356 processes
[1591704940] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1591705032] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 356 processes
[1591705197] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 26449 MB (20% inode=99%):
[1591705317] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 26449 MB (20% inode=99%):
[1591705437] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 26447 MB (20% inode=99%):
[1591705932] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.75, 2.50, 3.98
[1591706129] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(5) < Channel error on connection <0.8719.3941> (10.147.212.158:37580 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1591706237] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16256548 kB) free.
[1591706249] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1591707238] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 93753 MB (72% inode=99%):
[1591707790] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 36.2% (8883624 kB) free.
[1591708032] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1591708420] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1591708540] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1591709140] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1591709260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1591709380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1591709980] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1591710432] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 4 crit, 0 warn out of 364 processes
[1591710552] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 7 crit, 0 warn out of 371 processes
[1591710580] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1591710672] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 7 crit, 0 warn out of 373 processes
[1591710701] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1591710790] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187588 kB) free!
[1591710910] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166344 kB) free!
[1591711030] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194008 kB) free!
[1591714637] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1591714685] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 182.18.144.246 [09/Jun/2020:20:19:16 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1591714806] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591714932] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1591715837] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4046748 kB) free!
[1591715919] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 2.72, 4.25, 4.97
[1591715957] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (3866496 kB) free!
[1591716009] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.9% (17068176 kB) free!
[1591716077] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (3665984 kB) free!
[1591716129] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.0% (17070952 kB) free!
[1591716249] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.0% (17073704 kB) free!
[1591716677] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.0% (16172372 kB) free.
[1591716875] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 65.6% (21488796 kB) free.
[1591717030] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 34.2% (8386704 kB) free.
[1591717272] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1591717720] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.39, 5.76, 5.45
[1591717840] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.00, 6.60, 5.83
[1591717960] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.34, 5.84, 5.65
[1591722160] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.06, 3.92, 4.55
[1591722340] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1591722430] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4760888 kB) free!
[1591722550] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4693784 kB) free!
[1591722670] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (4640568 kB) free!
[1591723272] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)