[1558724181] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558724799] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558725456] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558730199] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2800696 kB) free!
[1558730294] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 31.54, 143.17, 93.51
[1558730415] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.59, 95.93, 82.22
[1558730535] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.82, 64.25, 72.27
[1558731185] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558731380] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 268.79, 110.25, 42.23
[1558731500] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 295.65, 173.00, 73.48
[1558731620] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 298.89, 214.85, 100.88
[1558731731] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6494152 kB) free!
[1558731785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558731851] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13202088 kB) free.
[1558731905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558732335] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 0.52, 10.66
[1558732451] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8446060 kB) free!
[1558732505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558732571] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6383140 kB) free!
[1558732625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558732691] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13201072 kB) free.
[1558732746] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558733291] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8355480 kB) free!
[1558733345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558733411] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6335028 kB) free!
[1558733465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558733531] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13191436 kB) free.
[1558733535] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.25, 3.12
[1558733585] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558734020] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.26, 2.13, 15.41
[1558734131] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5725344 kB) free!
[1558734252] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12900572 kB) free.
[1558735385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558735451] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6150364 kB) free!
[1558735506] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558735571] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13171288 kB) free.
[1558735599] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2373580 kB) free!
[1558735626] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558735820] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.88, 1.94, 3.91
[1558736171] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8151164 kB) free!
[1558736226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558736291] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5981152 kB) free!
[1558736346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558736411] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13163460 kB) free.
[1558736466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558737011] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8061672 kB) free!
[1558737066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558737131] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5878484 kB) free!
[1558737186] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558737251] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13176880 kB) free.
[1558737787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558737851] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7925876 kB) free!
[1558737907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558737971] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.4% (5601504 kB) free!
[1558738027] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558738092] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13173988 kB) free.
[1558738627] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558738691] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7742424 kB) free!
[1558738747] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558738811] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.0% (5398088 kB) free!
[1558738867] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558738931] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13173700 kB) free.
[1558739467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558739531] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7530044 kB) free!
[1558739587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558739651] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.1% (5483984 kB) free!
[1558739707] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558739771] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13165404 kB) free.
[1558740308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558740371] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7567012 kB) free!
[1558740428] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558740491] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.8% (5307624 kB) free!
[1558740548] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558740611] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13161940 kB) free.
[1558741148] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558741211] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7386492 kB) free!
[1558741268] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558741332] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13171672 kB) free.
[1558741388] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558742531] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5602912 kB) free!
[1558742589] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558742651] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13136720 kB) free.
[1558742709] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558742829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558746441] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3122532 kB) free!
[1558746561] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (3104884 kB) free!
[1558746681] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.0% (3095676 kB) free!
[1558749775] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1558749895] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1558750335] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 220.68, 72.31, 25.73
[1558750455] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 281.62, 144.71, 57.91
[1558750575] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 290.06, 193.21, 86.20
[1558751199] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2762300 kB) free!
[1558752375] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.78, 1.15, 16.96
[1558752399] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190260 kB) free!
[1558752754] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 583 processes
[1558752790] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1558752814] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 578 processes
[1558752851] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5101664 kB) free!
[1558752910] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1558752971] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (3600864 kB) free!
[1558753030] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1558753056] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 274.22, 184.30, 82.67
[1558753091] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1570376 kB) free!
[1558753167] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2854292 kB) free!
[1558753176] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 166.85, 188.14, 97.28
[1558753296] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.71, 125.95, 85.51
[1558753575] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.17, 0.20, 4.80
[1558754830] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 6 users currently logged in
[1558755096] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.62, 0.70, 12.60
[1558755681] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1600540 kB) free!
[1558756296] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.43, 0.39, 3.74
[1558756767] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2463016 kB) free!
[1558757240] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9690532 kB) free!
[1558757360] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 32.7% (16121660 kB) free.
[1558758491] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.2% (9484588 kB) free!
[1558758895] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1558759015] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1558759091] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10281904 kB) free.
[1558760291] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (1074216 kB) free!
[1558760411] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9324152 kB) free!
[1558760532] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1558760532] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7230556 kB) free!
[1558761560] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9091612 kB) free!
[1558761680] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7097400 kB) free!
[1558761800] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (5256192 kB) free!
[1558763000] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4172772 kB) free!
[1558765438] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (156260 kB) free!
[1558765558] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155468 kB) free!
[1558765678] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (153736 kB) free!
[1558766505] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.54, 5.64, 5.18
[1558766625] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.41, 5.73, 5.29
[1558766745] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.43, 5.87, 5.38
[1558769831] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1558769951] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1558770071] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1558770868] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.5% (14865676 kB) free.
[1558770919] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.4% (23057936 kB) free.
[1558770970] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.1% (23065176 kB) free.
[1558771037] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 45.0% (22181912 kB) free.
[1558771091] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (14924976 kB) free.
[1558771131] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.1% (22824960 kB) free.
[1558771271] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1558771332] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558771560] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558771680] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558771689] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 355 processes
[1558771736] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (3474936 kB) free!
[1558771800] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558771809] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558771856] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173820 kB) free!
[1558771871] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1558771879] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558771977] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185124 kB) free!
[1558771991] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1558771999] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558772078] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (170544 kB) free!
[1558772111] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1558772119] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558772175] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (1813472 kB) free!
[1558772198] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (195380 kB) free!
[1558772247] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9240124 kB) free!
[1558772297] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194688 kB) free!
[1558772318] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (195192 kB) free!
[1558772367] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.0% (11812380 kB) free.
[1558772419] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191204 kB) free!
[1558772672] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.0% (1178828 kB) free!
[1558772746] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.58, 4.53, 4.82
[1558772792] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.2% (972880 kB) free!
[1558772912] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.7% (785608 kB) free!
[1558773009] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 359 processes
[1558773130] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 359 processes
[1558773250] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 358 processes
[1558773512] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (127300 kB) free!
[1558773850] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558773946] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.39, 5.87, 5.35
[1558774066] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.10, 6.26, 5.56
[1558774186] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.15, 6.22, 5.63
[1558774977] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.5% (3787812 kB) free!
[1558775367] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1016796 kB) free!
[1558775488] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.5% (11100320 kB) free.
[1558775711] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1558775711] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1558776497] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.7% (2720152 kB) free!
[1558776617] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.2% (2470636 kB) free!
[1558776687] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (314424 kB) free!
[1558776732] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10016756 kB) free.
[1558776737] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.6% (2203012 kB) free!
[1558776777] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1558776777] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2018628 kB) free!
[1558776807] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.9% (10275636 kB) free.
[1558776930] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4655052 kB) free!
[1558777050] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4564852 kB) free!
[1558777112] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.8% (6822400 kB) free.
[1558777170] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4679760 kB) free!
[1558777337] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (743132 kB) free!
[1558778607] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (926544 kB) free!
[1558778650] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1558778711] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1558778727] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.3% (10966544 kB) free.
[1558778770] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 356 processes
[1558778832] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1558778890] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 357 processes
[1558778952] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1558778970] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (1060472 kB) free!
[1558778986] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.69, 3.92, 4.55
[1558779019] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 35.5% (8692608 kB) free.
[1558779132] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.7% (5288140 kB) free!
[1558779252] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 18.5% (9131972 kB) free!
[1558779372] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6710952 kB) free!
[1558779552] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1558780090] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1558780090] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 365 processes
[1558780152] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1558780200] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1558780272] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1558780392] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1558780519] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558780527] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (478948 kB) free!
[1558780647] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.8% (10252140 kB) free.
[1558780691] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 357 processes
[1558781368] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.3% (22615852 kB) free.
[1558781386] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.91, 5.34, 5.30
[1558781411] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.2% (23499724 kB) free.
[1558781454] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1558781454] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.2% (14985648 kB) free.
[1558781506] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.30, 4.91, 5.14
[1558781543] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 80.6% (13108892 kB) free.
[1558781626] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1558781626] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.89, 4.68, 5.04
[1558781848] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5314088 kB) free!
[1558781891] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1558781968] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 21.6% (10643348 kB) free.
[1558782147] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175216 kB) free!
[1558782226] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.44, 3.30, 4.19
[1558782268] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (179644 kB) free!
[1558782388] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165296 kB) free!
[1558782491] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 357 processes
[1558782611] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 356 processes
[1558782731] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558783539] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46244 MB (20% inode=99%):
[1558783660] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46227 MB (20% inode=99%):
[1558783780] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46243 MB (20% inode=99%):
[1558784188] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (2220176 kB) free!
[1558784312] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (1219260 kB) free!
[1558784368] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7614076 kB) free!
[1558784380] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48004 MB (21% inode=99%):
[1558784432] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.8% (1091144 kB) free!
[1558784488] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.4% (10032088 kB) free.
[1558784552] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.3% (980412 kB) free!
[1558784980] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47153 MB (20% inode=99%):
[1558785100] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46684 MB (20% inode=99%):
[1558785152] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1558785152] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (375284 kB) free!
[1558785220] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46484 MB (20% inode=99%):
[1558785375] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 25.71, 132.86, 88.84
[1558785495] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.11, 89.18, 78.17
[1558785615] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.09, 59.85, 68.76
[1558786218] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183296 kB) free!
[1558786219] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (4121788 kB) free!
[1558786220] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 26.36, 21.29, 11.74
[1558786251] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5396 MB (10% inode=97%):
[1558786338] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184480 kB) free!
[1558786339] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3680680 kB) free!
[1558786340] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 26.32, 22.93, 13.50
[1558786458] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180744 kB) free!
[1558786459] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4295672 kB) free!
[1558786460] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.94, 16.71, 12.40
[1558786815] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.18, 1.43, 19.16
[1558786851] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7011 MB (13% inode=98%):
[1558787377] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1558787480] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182904 kB) free!
[1558787552] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.5% (6873388 kB) free.
[1558787600] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167200 kB) free!
[1558787660] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.94, 2.12, 4.70
[1558787720] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175544 kB) free!
[1558787772] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10362504 kB) free.
[1558788088] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8935996 kB) free!
[1558788209] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6555296 kB) free!
[1558788329] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4454744 kB) free!
[1558788615] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.17, 0.46, 4.38
[1558788660] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1558788929] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (7855784 kB) free!
[1558789459] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2435700 kB) free!
[1558789460] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 278.68, 202.92, 96.52
[1558789580] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 279.20, 228.15, 118.66
[1558789700] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 236.70, 235.91, 135.15
[1558790059] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1558790059] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (4156392 kB) free!
[1558791792] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558791972] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (670096 kB) free!
[1558792030] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558792092] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 18.8% (9289160 kB) free!
[1558792100] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1558792100] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.59, 1.97, 12.01
[1558792212] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7190976 kB) free!
[1558792931] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558793300] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.75, 1.83, 4.59
[1558793459] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (3090652 kB) free!
[1558793579] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.8% (3049064 kB) free!
[1558793626] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558793699] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (2998272 kB) free!
[1558795520] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (2568844 kB) free!
[1558796230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558796351] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558796471] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558797152] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1558797299] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1407256 kB) free!
[1558797329] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.7% (13153512 kB) free.
[1558797415] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1558797536] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1558801459] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558801529] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9449200 kB) free!
[1558801649] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7611988 kB) free!
[1558801769] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 11.8% (5831020 kB) free!
[1558801812] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9948880 kB) free.
[1558802369] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.5% (13050824 kB) free.
[1558804100] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558805369] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9838620 kB) free!
[1558805489] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8046924 kB) free!
[1558805609] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1558805609] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 12.5% (6163908 kB) free!
[1558806209] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.5% (13046464 kB) free.