[1585161223] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6467416 kB) free!
[1585161343] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2436988 kB) free!
[1585161463] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32914208 kB) free.
[1585162063] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9826504 kB) free!
[1585162183] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5791768 kB) free!
[1585162303] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2083280 kB) free!
[1585162903] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13316060 kB) free!
[1585163156] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 28.1% (6884760 kB) free.
[1585164103] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6274776 kB) free!
[1585164703] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16401704 kB) free!
[1585165303] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27798212 kB) free.
[1585165903] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8122908 kB) free!
[1585166023] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4761928 kB) free!
[1585166119] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (4490308 kB) free!
[1585166143] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4468068 kB) free!
[1585166143] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30829608 kB) free.
[1585166263] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236992 kB) free!
[1585166383] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229076 kB) free!
[1585166567] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.75, 6.16, 5.16
[1585166687] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.46, 6.52, 5.41
[1585166720] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (227812 kB) free!
[1585166743] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10612240 kB) free!
[1585166807] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.70, 6.70, 5.62
[1585166863] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7052796 kB) free!
[1585166984] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32637700 kB) free.
[1585167583] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13144684 kB) free!
[1585167703] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9476964 kB) free!
[1585167823] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6030364 kB) free!
[1585168423] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14730048 kB) free!
[1585168783] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 70.8% (34869944 kB) free.
[1585169623] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32686604 kB) free.
[1585170103] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 65.5% (32268264 kB) free.
[1585170223] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12646024 kB) free!
[1585170343] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9068864 kB) free!
[1585170463] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23284820 kB) free!
[1585171228] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.1% (43419728 kB) free.
[1585172263] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31160940 kB) free.
[1585172807] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.88, 4.45, 4.84
[1585172863] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (11001096 kB) free!
[1585172949] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.3% (9883560 kB) free.
[1585172983] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7640308 kB) free!
[1585173103] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32059228 kB) free.
[1585173704] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12027140 kB) free!
[1585173824] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8518064 kB) free!
[1585173944] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32592568 kB) free.
[1585174544] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5699276 kB) free!
[1585174664] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1928048 kB) free!
[1585174784] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32667888 kB) free.
[1585175384] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8379680 kB) free!
[1585175504] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4626580 kB) free!
[1585175624] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32775428 kB) free.
[1585176224] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11433748 kB) free!
[1585176344] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7730032 kB) free!
[1585176464] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32860424 kB) free.
[1585177064] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14382440 kB) free!
[1585177184] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10938040 kB) free!
[1585177304] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34042572 kB) free.
[1585177905] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18466280 kB) free!
[1585178025] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15104120 kB) free!
[1585178145] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35164144 kB) free.
[1585178745] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17838792 kB) free!
[1585178865] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14418856 kB) free!
[1585178985] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34736696 kB) free.
[1585179549] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188660 kB) free!
[1585179585] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16854996 kB) free!
[1585179669] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186236 kB) free!
[1585179705] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13516176 kB) free!
[1585179790] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189760 kB) free!
[1585179825] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33565464 kB) free.
[1585180424] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.5% (19005904 kB) free.
[1585180425] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15437020 kB) free!
[1585180520] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.5% (12456524 kB) free.
[1585180545] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35412780 kB) free.
[1585180780] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 72.0% (17643508 kB) free.
[1585181145] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16651364 kB) free!
[1585181265] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35421772 kB) free.
[1585181865] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18982592 kB) free!
[1585181985] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15590368 kB) free!
[1585182106] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35506584 kB) free.
[1585182705] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17646008 kB) free!
[1585182825] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13918152 kB) free!
[1585182945] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34370644 kB) free.
[1585182977] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (24936196 kB) free!
[1585183097] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.0% (14471916 kB) free!
[1585183217] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (4003724 kB) free!
[1585183545] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14111924 kB) free!
[1585183665] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34861868 kB) free.
[1585184265] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15954068 kB) free!
[1585184385] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31290864 kB) free.
[1585184642] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) <                  {amqp_error,frame_error,
[1585184764] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1585184985] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16706380 kB) free!
[1585185017] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 83.3% (109825540 kB) free.
[1585185105] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12707232 kB) free!
[1585185225] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35065656 kB) free.
[1585185825] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13416668 kB) free!
[1585185945] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1585185945] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34719708 kB) free.
[1585186545] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14587208 kB) free!
[1585186608] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.54, 5.56, 5.25
[1585186665] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10544544 kB) free!
[1585186728] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.62, 5.34, 5.21
[1585186786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33719284 kB) free.
[1585186848] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.30, 4.96, 5.09
[1585187385] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13436884 kB) free!
[1585187448] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.06, 3.60, 4.32
[1585187505] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (13196396 kB) free!
[1585187625] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31942912 kB) free.
[1585188225] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11715484 kB) free!
[1585188345] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33611904 kB) free.
[1585188945] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13749524 kB) free!
[1585189065] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9847956 kB) free!
[1585189185] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32189516 kB) free.
[1585189785] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11694616 kB) free!
[1585189848] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.34, 5.30, 5.10
[1585189905] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32771432 kB) free.
[1585189968] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.53, 5.80, 5.31
[1585190088] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.79, 6.46, 5.61
[1585190505] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12046480 kB) free!
[1585190625] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7821064 kB) free!
[1585190745] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30016512 kB) free.
[1585191320] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4426888 kB) free!
[1585191345] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8564468 kB) free!
[1585191440] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.5% (3836260 kB) free!
[1585191465] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16915668 kB) free!
[1585191560] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3739888 kB) free!
[1585191585] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (18953804 kB) free!
[1585191768] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(158) < 92.42.104.166 [26/Mar/2020:08:30:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585191888] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585192186] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376556 kB) free!
[1585193088] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1585193088] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.20, 3.85, 4.66
[1585193360] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1748916 kB) free!
[1585193386] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15267044 kB) free!
[1585193980] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.1% (2464800 kB) free!
[1585194100] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196336 kB) free!
[1585194220] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183680 kB) free!
[1585194288] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.52, 6.25, 5.46
[1585194408] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.84, 5.92, 5.43
[1585194528] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.79, 5.50, 5.34
[1585194586] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1585194586] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1689244 kB) free!
[1585194820] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 29.8% (7300984 kB) free.
[1585195682] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 271.47, 144.77, 60.83
[1585195786] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16678472 kB) free!
[1585195802] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 42.08, 98.90, 54.19
[1585195924] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 13.92, 68.50, 48.62
[1585196986] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2916944 kB) free!
[1585197124] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 6.88, 11.76, 25.60
[1585197220] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3874784 kB) free!
[1585197288] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(323) < 103.255.146.164 [26/Mar/2020:10:01:35 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585197340] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (3706980 kB) free!
[1585197408] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585197460] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (3467056 kB) free!
[1585198008] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(697) < 103.255.146.164 [26/Mar/2020:10:14:46 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585198128] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585198186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28691968 kB) free.
[1585198660] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1585198660] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1892488 kB) free!
[1585198786] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3590516 kB) free!
[1585198906] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30044656 kB) free.
[1585199506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8983804 kB) free!
[1585199626] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4801776 kB) free!
[1585199747] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15401344 kB) free!
[1585200347] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (8541208 kB) free!
[1585200947] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16366340 kB) free!
[1585201547] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26866200 kB) free.
[1585201681] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585201682] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1585201728] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(672) < 103.255.146.164 [26/Mar/2020:11:16:43 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585201801] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1585201802] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1585201849] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585201922] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1585201922] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1585202024] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166400 kB) free!
[1585202144] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189736 kB) free!
[1585202147] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3092380 kB) free!
[1585202264] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176008 kB) free!
[1585202267] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393280 kB) free!
[1585202387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27773648 kB) free.
[1585202522] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1585202522] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1585202987] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3045660 kB) free!
[1585203107] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394524 kB) free!
[1585203227] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28935520 kB) free.
[1585203771] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 57.3% (14040744 kB) free.
[1585203827] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6023508 kB) free!
[1585203947] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1549332 kB) free!
[1585204067] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389356 kB) free!
[1585204323] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1585204443] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1585205267] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21824900 kB) free!
[1585205328] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585205868] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404976 kB) free!
[1585206906] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1585207026] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1585207068] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23483168 kB) free!
[1585207325] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 215.30, 264.01, 137.50
[1585207445] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 37.25, 178.14, 121.83
[1585207565] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.10, 122.14, 108.15
[1585207668] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414240 kB) free!
[1585207828] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8368.87, 4983.94, 2105.74
[1585207948] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7155.42, 5858.87, 2786.79
[1585208068] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 967.80, 3921.22, 2449.22
[1585208571] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.4% (4275160 kB) free!
[1585208691] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.8% (4106228 kB) free!
[1585208765] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.23, 8.18, 34.66
[1585208811] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.3% (3991468 kB) free!
[1585208868] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14195612 kB) free!
[1585209268] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.90, 72.59, 674.11
[1585209365] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.39, 5.12, 20.33
[1585209460] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585209468] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401128 kB) free!
[1585209868] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.07, 11.60, 354.27
[1585210011] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2207988 kB) free!
[1585210026] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 333 processes
[1585210146] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1585210668] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22773148 kB) free!
[1585211268] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416876 kB) free!
[1585211503] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9689824 kB) free!
[1585211623] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9515748 kB) free!
[1585211643] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1585211743] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (9273572 kB) free!
[1585211763] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1585211868] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13724384 kB) free!
[1585211883] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1585211928] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 40.02, 26.70, 19.01
[1585211946] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1585212028] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8325348 kB) free!
[1585212066] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1585212148] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.2% (7989152 kB) free!
[1585212268] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.5% (7647820 kB) free!
[1585212434] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (3515920 kB) free!
[1585212468] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448324 kB) free!
[1585212528] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.31, 18.82, 18.65
[1585212554] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3119068 kB) free!
[1585212657] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.2866.4437> (10.147.212.181:57633 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1585212674] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.3% (2761600 kB) free!
[1585212777] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1585213274] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1145900 kB) free!
[1585213468] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4669020 kB) free!
[1585213668] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27350592 kB) free.
[1585214268] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (769436 kB) free!
[1585214388] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384744 kB) free!
[1585214509] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32943804 kB) free.
[1585214743] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4797948 kB) free!
[1585215075] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.5% (15600712 kB) free.
[1585215108] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5894244 kB) free!
[1585215228] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1135052 kB) free!
[1585215348] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403748 kB) free!
[1585216058] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 13 users currently logged in
[1585216178] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 14 users currently logged in
[1585216298] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 14 users currently logged in
[1585216548] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31880624 kB) free.
[1585217148] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4668952 kB) free!
[1585217268] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399360 kB) free!
[1585217388] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384020 kB) free!
[1585218588] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33242220 kB) free.
[1585219064] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (4163468 kB) free!
[1585219188] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3435492 kB) free!
[1585219308] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (436308 kB) free!
[1585219428] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402472 kB) free!
[1585219664] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195356 kB) free!
[1585219728] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.84, 28.01, 21.83
[1585219898] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1585220293] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 25.6% (4164428 kB) free.
[1585220328] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1585220328] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 20.47, 17.86, 19.19
[1585220406] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2166844 kB) free!
[1585220498] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 13 users currently logged in
[1585220526] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (2374728 kB) free!
[1585220618] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 14 users currently logged in
[1585220646] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (2683908 kB) free!
[1585220738] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 15 users currently logged in
[1585220928] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.09, 21.03, 20.28
[1585221246] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3920428 kB) free.
[1585221528] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 15.32, 16.52, 18.62
[1585221938] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1585221938] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1585222275] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4487276 kB) free!
[1585222395] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4265704 kB) free!
[1585222446] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2639696 kB) free!
[1585222515] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.3% (4110860 kB) free!
[1585222567] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3130988 kB) free!
[1585222687] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3524448 kB) free.
[1585224315] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1585224315] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1314304 kB) free!
[1585224487] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.8% (1917068 kB) free!
[1585224607] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2611828 kB) free!
[1585224727] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3075564 kB) free!
[1585225327] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1585225327] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1489432 kB) free!
[1585225428] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28298580 kB) free.
[1585225927] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3597320 kB) free.
[1585226028] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (409292 kB) free!
[1585226116] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.6% (15637204 kB) free.
[1585226148] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (426552 kB) free!
[1585226268] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392192 kB) free!
[1585226456] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.3% (4004360 kB) free!
[1585226516] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 15.4% (3770648 kB) free!
[1585226576] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 15.3% (3761144 kB) free!
[1585226636] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 15.3% (3755672 kB) free!
[1585228327] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2579900 kB) free!
[1585228447] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3259968 kB) free.
[1585228669] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27768324 kB) free.
[1585229047] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.4% (2341848 kB) free!
[1585229167] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2975864 kB) free!
[1585229269] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (386424 kB) free!
[1585229287] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3542816 kB) free.
[1585229389] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401376 kB) free!
[1585229509] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403408 kB) free!
[1585229887] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2120724 kB) free!
[1585230008] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2762872 kB) free!
[1585230128] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3422412 kB) free.
[1585230709] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31968600 kB) free.
[1585230728] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (2006540 kB) free!
[1585230848] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2531076 kB) free!
[1585230968] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3368476 kB) free.
[1585231309] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (382316 kB) free!
[1585231429] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (429440 kB) free!
[1585231549] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397056 kB) free!
[1585231568] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (1727284 kB) free!
[1585231688] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.0% (2434320 kB) free!
[1585231808] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (3087764 kB) free!
[1585232408] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1579128 kB) free!
[1585232738] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585232749] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31533116 kB) free.
[1585232928] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585233008] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3685980 kB) free.
[1585233349] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3459660 kB) free!
[1585233469] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399880 kB) free!
[1585233528] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.62, 3.79, 4.77
[1585233589] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400536 kB) free!
[1585233608] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2625284 kB) free!
[1585233728] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3204788 kB) free!
[1585233848] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3579564 kB) free.
[1585234448] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (1993064 kB) free!
[1585234568] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.5% (2519496 kB) free!
[1585234688] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3059648 kB) free!
[1585234790] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31843936 kB) free.
[1585235288] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (803816 kB) free!
[1585235328] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.79, 6.54, 5.61
[1585235390] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3285932 kB) free!
[1585235448] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.87, 6.45, 5.68
[1585235510] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (392736 kB) free!
[1585235568] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.28, 6.31, 5.72
[1585235630] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408980 kB) free!
[1585235889] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3313304 kB) free.
[1585236048] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(143) < 92.42.104.166 [26/Mar/2020:20:50:15 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585236168] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585236316] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1585236830] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31990272 kB) free.
[1585237430] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4727668 kB) free!
[1585237550] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (411460 kB) free!
[1585237670] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419620 kB) free!
[1585237968] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.31, 3.30, 4.11
[1585238284] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1585238870] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26892496 kB) free.
[1585239470] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2408980 kB) free!
[1585239590] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389320 kB) free!
[1585239711] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32932888 kB) free.
[1585240084] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 388 processes
[1585240197] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3259.98, 3127.30, 1413.66
[1585240204] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 423 processes
[1585240310] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6084368 kB) free!
[1585240317] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 441.38, 2093.18, 1242.49
[1585240430] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1266348 kB) free!
[1585240437] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 61.60, 1401.36, 1092.11
[1585240550] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408916 kB) free!
[1585240665] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 30.2% (7404388 kB) free.
[1585241037] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.33, 189.23, 572.97
[1585241289] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3038804 kB) free!
[1585241409] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3158748 kB) free!
[1585241529] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3304280 kB) free.
[1585241638] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.52, 26.63, 300.91
[1585241750] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26789560 kB) free.
[1585242350] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1275568 kB) free!
[1585242470] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391044 kB) free!
[1585242590] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33435892 kB) free.
[1585243190] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7817952 kB) free!
[1585243310] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3890412 kB) free!
[1585243430] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391104 kB) free!
[1585243930] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2996996 kB) free!
[1585244050] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.5% (4800028 kB) free.
[1585244630] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.7% (25947740 kB) free!
[1585245230] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4096516 kB) free!
[1585245830] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14254472 kB) free!
[1585246431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26826020 kB) free.
[1585247031] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4812588 kB) free!
[1585247151] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33573276 kB) free.