[1585420458] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4814684 kB) free!
[1585420578] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28464700 kB) free.
[1585421179] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5988812 kB) free!
[1585421299] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28410204 kB) free.
[1585421899] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7423092 kB) free!
[1585421931] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1585422019] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4027128 kB) free!
[1585422125] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4402864 kB) free!
[1585422139] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26516432 kB) free.
[1585422185] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1812248 kB) free!
[1585422246] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (182144 kB) free!
[1585422306] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (179840 kB) free!
[1585422606] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 70.5% (17271264 kB) free.
[1585422739] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5269660 kB) free!
[1585422859] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28384528 kB) free.
[1585423459] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6477880 kB) free!
[1585423579] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28376900 kB) free.
[1585424179] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7206608 kB) free!
[1585424299] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28365888 kB) free.
[1585424899] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7393164 kB) free!
[1585425019] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28325876 kB) free.
[1585425619] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8108076 kB) free!
[1585425739] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4861192 kB) free!
[1585425860] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26427172 kB) free!
[1585426131] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.85, 6.52, 5.70
[1585426252] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.69, 6.58, 5.82
[1585426372] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.57, 6.58, 5.91
[1585426460] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6023696 kB) free!
[1585427572] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.31, 4.10, 4.99
[1585427660] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13389044 kB) free!
[1585429285] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 36.8% (9028236 kB) free.
[1585429460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27715492 kB) free.
[1585429665] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 63.9% (31479176 kB) free.
[1585429956] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.8% (44236820 kB) free.
[1585430060] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1419016 kB) free!
[1585430180] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28376128 kB) free.
[1585430780] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2879292 kB) free!
[1585430900] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28709168 kB) free.
[1585431085] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2759616 kB) free!
[1585431231] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2764000 kB) free!
[1585431341] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.7% (3351060 kB) free!
[1585431404] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(332) < 2020/03/29 03:01:33 [error] 22121#0: *78323335 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585431414] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2981) < 203.212.70.170 [29/Mar/2020:03:06:03 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 0.000, 60.000, 0.000"
[1585431500] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4248284 kB) free!
[1585431507] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(259) < 103.16.101.72 [29/Mar/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1585431525] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585431534] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585431620] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28715860 kB) free.
[1585431633] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585431685] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1733056 kB) free!
[1585432093] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1585432213] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1585432220] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5820400 kB) free!
[1585432340] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28736020 kB) free.
[1585432940] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6854724 kB) free!
[1585433061] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28667156 kB) free.
[1585433631] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452832 kB) free!
[1585433660] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7579572 kB) free!
[1585433780] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28671432 kB) free.
[1585434380] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6589140 kB) free!
[1585434500] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26850624 kB) free.
[1585434941] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1996660 kB) free!
[1585435700] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13236572 kB) free!
[1585435820] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30645732 kB) free.
[1585436420] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25764284 kB) free!
[1585436540] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30835124 kB) free.
[1585438285] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187332 kB) free!
[1585438340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16822348 kB) free!
[1585438405] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169080 kB) free!
[1585438460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32048932 kB) free.
[1585438525] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204588 kB) free!
[1585439031] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.5% (20946344 kB) free.
[1585439141] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 55.6% (13629120 kB) free.
[1585439485] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.0% (12342804 kB) free.
[1585439660] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16374356 kB) free!
[1585439781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32185404 kB) free.
[1585440980] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16234932 kB) free!
[1585441100] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1585441100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31379420 kB) free.
[1585444100] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14550808 kB) free!
[1585444220] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31641540 kB) free.
[1585448068] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15226.77, 6786.13, 2591.64
[1585448188] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16204.38, 9952.84, 4260.67
[1585448308] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16336.91, 12071.96, 5727.39
[1585448334] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(40) < 103.13.97.241 [29/Mar/2020:07:48:54 +0530] TCP 502 0 0 60.001 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 0.000, 0.000"
[1585448420] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16448680 kB) free!
[1585448455] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(67) < 172.105.59.173 [29/Mar/2020:07:50:38 +0530] TCP 502 0 0 120.002 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.154:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 0.000, 0.001, 60.000"
[1585448540] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30856072 kB) free.
[1585448575] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(186) < 203.212.70.170 [29/Mar/2020:07:52:55 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1585448817] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16374.33, 13987.71, 7562.01
[1585448937] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16375.91, 14777.94, 8630.70
[1585449057] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6182.90, 12396.82, 8570.67
[1585450375] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585450972] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.12, 5.54, 5.01
[1585451092] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.99, 5.33, 5.00
[1585451308] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.87, 6.39, 709.67
[1585451457] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.45, 6.53, 648.27
[1585451908] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.55, 3.22, 375.16
[1585452057] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.16, 2.80, 340.74
[1585452140] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12326360 kB) free!
[1585452260] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28229012 kB) free.
[1585452942] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4537532 kB) free!
[1585453062] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4405432 kB) free!
[1585453182] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4297984 kB) free!
[1585454982] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2310856 kB) free!
[1585455136] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 20.5% (10091348 kB) free.
[1585455293] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.14, 5.45, 5.06
[1585455413] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.36, 4.99, 4.93
[1585455736] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9808996 kB) free!
[1585455856] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9742856 kB) free!
[1585455860] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9161656 kB) free!
[1585455976] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (9673892 kB) free!
[1585455981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.9% (26287204 kB) free!
[1585456013] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.71, 6.46, 5.83
[1585456101] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21850140 kB) free!
[1585456133] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.43, 6.00, 5.73
[1585456253] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.62, 6.84, 6.06
[1585456576] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 21.3% (10487368 kB) free.
[1585456926] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1585457046] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1585457166] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1585457408] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1585457528] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1585457631] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181624 kB) free!
[1585457751] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199372 kB) free!
[1585457766] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1585457776] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9861888 kB) free!
[1585457871] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185788 kB) free!
[1585457896] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9804996 kB) free!
[1585458016] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1585458016] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (9748072 kB) free!
[1585458501] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4807804 kB) free!
[1585458653] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.66, 4.31, 4.65
[1585458966] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1585459086] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1585459206] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1585459373] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 61.2% (15011728 kB) free.
[1585459816] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 21.1% (10383996 kB) free.
[1585459853] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.04, 6.13, 5.25
[1585459973] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.28, 6.69, 5.59
[1585460093] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.55, 7.01, 5.82
[1585461016] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9706760 kB) free!
[1585461136] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9660452 kB) free!
[1585461256] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (9597624 kB) free!
[1585462101] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17982360 kB) free!
[1585462206] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1585462206] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1585463301] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1585463301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27254048 kB) free.
[1585464502] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5977840 kB) free!
[1585464622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27973708 kB) free.
[1585465222] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22573736 kB) free!
[1585465342] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.4% (25518172 kB) free!
[1585465462] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20067448 kB) free!
[1585465493] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1585465493] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.51, 4.50, 4.80
[1585465806] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1585465926] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1585466046] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1585466093] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.45, 5.35, 5.17
[1585466213] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.56, 5.84, 5.37
[1585466308] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16309.50, 12517.86, 6107.06
[1585466333] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.46, 5.36, 5.26
[1585466428] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16319.46, 13776.22, 7345.49
[1585466457] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15754.57, 7974.10, 3167.41
[1585466548] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16320.48, 14618.13, 8433.73
[1585466575] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1226) < 203.212.70.170 [29/Mar/2020:12:52:54 +0530] TCP 502 0 0 120.002 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 0.000, 0.001, 60.000"
[1585466577] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16255.45, 10740.03, 4763.80
[1585466662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26864996 kB) free.
[1585466695] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(244) < 182.18.144.246 [29/Mar/2020:12:54:55 +0530] TCP 502 0 0 0.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "0.000, 0.000"
[1585466697] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16322.97, 12590.78, 6166.62
[1585466816] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(1023) < 172.105.59.173 [29/Mar/2020:12:56:53 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1585466933] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.49, 4.15, 4.68
[1585467325] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2828752 kB) free!
[1585467862] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3872948 kB) free!
[1585467925] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2245836 kB) free!
[1585467982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.8% (26104060 kB) free!
[1585467987] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 361 processes
[1585468016] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585468102] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20629872 kB) free!
[1585468107] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 361 processes
[1585468227] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 362 processes
[1585469046] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1585469302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26454916 kB) free.
[1585469548] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.01, 4.83, 498.72
[1585469902] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2886844 kB) free!
[1585470022] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.2% (25358648 kB) free!
[1585470142] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20061564 kB) free!
[1585470297] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.52, 2.20, 359.67
[1585471071] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 69.0% (16919288 kB) free.
[1585471133] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.58, 9.55, 7.05
[1585471253] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.82, 8.65, 7.03
[1585471373] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.75, 7.34, 6.76
[1585471942] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3415532 kB) free!
[1585471973] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.6% (4557048 kB) free!
[1585472056] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585472093] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.9% (4396620 kB) free!
[1585472213] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.6% (4324004 kB) free!
[1585473173] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.63, 3.02, 4.30
[1585473413] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 20.8% (5094972 kB) free.
[1585473742] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16301572 kB) free!
[1585474014] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4798848 kB) free!
[1585474134] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4751764 kB) free!
[1585474254] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.2% (4695444 kB) free!
[1585474617] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(108) < 185.255.8.58 [29/Mar/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585474737] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585476142] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2685888 kB) free!
[1585478046] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1585478454] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1585478454] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2431396 kB) free!
[1585479142] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16936928 kB) free!
[1585480973] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.98, 6.60, 5.50
[1585481093] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.40, 6.62, 5.65
[1585481213] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.45, 6.69, 5.78
[1585481542] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (810800 kB) free!
[1585483065] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9858260 kB) free!
[1585483185] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (9758320 kB) free!
[1585483305] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (9660724 kB) free!
[1585483342] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16384852 kB) free!
[1585486254] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1585486342] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28592196 kB) free.
[1585486942] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6840172 kB) free!
[1585487062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28264520 kB) free.
[1585487213] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.20, 3.67, 4.45
[1585487662] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6458696 kB) free!
[1585487782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26826656 kB) free.
[1585488027] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1585488309] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1585488429] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1585488983] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6512788 kB) free!
[1585489103] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27866104 kB) free.
[1585489137] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(369) < 203.212.70.170 [29/Mar/2020:19:06:32 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 0.000, 0.001, 60.000"
[1585489258] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585489703] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5058820 kB) free!
[1585489823] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26677724 kB) free.
[1585490097] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16392.11, 14231.03, 7850.48
[1585490217] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16394.01, 14946.78, 8886.40
[1585490337] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6390.54, 12601.18, 8829.07
[1585490423] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3376784 kB) free!
[1585490457] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(373) < 203.212.70.170 [29/Mar/2020:19:30:29 +0530] TCP 502 0 0 120.003 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.001, 0.000, 0.001"
[1585490543] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.1% (25154056 kB) free!
[1585490577] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(17) < 182.18.144.246 [29/Mar/2020:19:31:21 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 0.001, 0.000"
[1585490663] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19921960 kB) free!
[1585490698] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585492285] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4926792 kB) free!
[1585492405] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4925956 kB) free!
[1585492463] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27036296 kB) free.
[1585492525] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (4932700 kB) free!
[1585493063] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5547404 kB) free!
[1585493183] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.5% (25666736 kB) free!
[1585493213] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1585493303] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20551692 kB) free!
[1585493337] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.84, 2.32, 350.46
[1585495703] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26545404 kB) free.
[1585496303] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5874624 kB) free!
[1585496423] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.8% (24722056 kB) free!
[1585496543] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19932508 kB) free!
[1585498956] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9869632 kB) free!
[1585499076] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9847352 kB) free!
[1585499196] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.9% (9826904 kB) free!
[1585502543] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6206040 kB) free!
[1585506105] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4949656 kB) free!