[1561141939] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561142335] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561142391] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8849448 kB) free!
[1561142455] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561142511] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13194228 kB) free.
[1561142582] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (215888 kB) free!
[1561142666] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196240 kB) free!
[1561142702] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206276 kB) free!
[1561142753] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187652 kB) free!
[1561142786] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184520 kB) free!
[1561142822] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (206844 kB) free!
[1561142874] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188012 kB) free!
[1561142906] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182860 kB) free!
[1561142994] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180284 kB) free!
[1561143055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561143111] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8561048 kB) free!
[1561143175] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561143231] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13196848 kB) free.
[1561143665] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 121.01, 162.37, 94.52
[1561143775] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561143785] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.71, 109.16, 83.24
[1561143895] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561143905] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.92, 73.59, 73.35
[1561144495] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561144615] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561145216] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561145336] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561145475] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (1482288 kB) free!
[1561145535] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (2292292 kB) free!
[1561145595] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 9.4% (2294600 kB) free!
[1561145655] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.4% (2295512 kB) free!
[1561145705] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.60, 1.84, 11.97
[1561145822] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.9% (2943664 kB) free!
[1561145936] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561145976] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.96, 8.49, 6.62
[1561146056] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561146096] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.70, 5.96, 5.93
[1561146216] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.51, 4.16, 5.27
[1561146656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561146776] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561147106] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (2602660 kB) free!
[1561147376] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561147416] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.65, 1.61, 3.40
[1561147432] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9379536 kB) free!
[1561147496] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561147552] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13185152 kB) free.
[1561148097] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561148105] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.15, 2.15, 4.33
[1561148152] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9473428 kB) free!
[1561148217] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561148272] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13192568 kB) free.
[1561148817] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561148872] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9611036 kB) free!
[1561148937] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561148992] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13197248 kB) free.
[1561149057] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561149506] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2445460 kB) free!
[1561149592] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9829892 kB) free!
[1561149657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561149667] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (154984 kB) free!
[1561149712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13192636 kB) free.
[1561149714] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (182744 kB) free!
[1561149777] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561149787] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 41.3% (6713128 kB) free.
[1561149834] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (294320 kB) free!
[1561149954] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179660 kB) free!
[1561150377] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561150498] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561150672] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1561150792] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 372 processes
[1561150912] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1561151097] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561151217] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561151817] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561151937] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561152537] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561152657] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561152739] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 28.13, 60.16, 34.34
[1561152816] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 199.21, 110.92, 45.78
[1561152859] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.97, 40.32, 30.20
[1561152936] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 202.95, 141.53, 64.90
[1561152979] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.79, 27.05, 26.57
[1561153056] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 58.24, 116.87, 65.92
[1561153194] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2805272 kB) free!
[1561153258] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561153378] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561153580] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 3.85, 14.05
[1561153912] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9788696 kB) free!
[1561153978] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561154032] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13195840 kB) free.
[1561154098] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561154256] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1561154256] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.51, 2.46, 18.35
[1561154632] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9559932 kB) free!
[1561154698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561154752] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13174672 kB) free.
[1561154780] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 0.32, 4.04
[1561154814] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1561154818] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561154934] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561155352] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9192264 kB) free!
[1561155472] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13183156 kB) free.
[1561156018] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561156056] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.10, 0.68, 3.11
[1561156072] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8872712 kB) free!
[1561156139] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561156192] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13183912 kB) free.
[1561156738] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561156858] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561157392] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9808804 kB) free!
[1561157458] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561157513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13179692 kB) free.
[1561157578] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561158112] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9612984 kB) free!
[1561158178] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561158232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13183660 kB) free.
[1561158298] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561158832] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9376900 kB) free!
[1561158952] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13171692 kB) free.
[1561159194] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2451336 kB) free!
[1561159499] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561159552] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9032268 kB) free!
[1561159619] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561159672] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13180752 kB) free.
[1561160219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561160272] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8706688 kB) free!
[1561160339] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561160392] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13180620 kB) free.
[1561160939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561161059] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561161592] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9803544 kB) free!
[1561161659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561161712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13171148 kB) free.
[1561161779] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561161905] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 60.52, 19.68, 8.05
[1561162025] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 83.18, 41.86, 17.59
[1561162145] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 86.20, 56.68, 25.96
[1561162313] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9546568 kB) free!
[1561162433] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13174308 kB) free.
[1561162979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561163033] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9251228 kB) free!
[1561163099] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561163153] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13163308 kB) free.
[1561163345] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.17, 5.39, 16.86
[1561163699] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561163753] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8816612 kB) free!
[1561163819] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561163823] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2380332 kB) free!
[1561163873] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13170996 kB) free.
[1561164354] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 39.0% (6341716 kB) free.
[1561164419] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561164539] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561164954] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171636 kB) free!
[1561165073] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9833180 kB) free!
[1561165074] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (177932 kB) free!
[1561165139] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561165145] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.62, 1.73, 3.90
[1561165193] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13173352 kB) free.
[1561165194] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151184 kB) free!
[1561165260] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561165793] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9384468 kB) free!
[1561165859] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561165913] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13020876 kB) free.
[1561165979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561166099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561166256] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561166355] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 27.1% (6643196 kB) free.
[1561166394] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 61.9% (10064912 kB) free.
[1561167299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561167594] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (149932 kB) free!
[1561167715] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176148 kB) free!
[1561167765] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (148104 kB) free!
[1561167835] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1561167835] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (177096 kB) free!
[1561167885] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171572 kB) free!
[1561167899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561168005] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171340 kB) free!
[1561168019] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561168140] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561168605] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.5% (13739204 kB) free.
[1561169256] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 64.52, 29.45, 11.60
[1561169376] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 67.70, 42.30, 18.47
[1561169496] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.03, 49.87, 24.18
[1561169513] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9576652 kB) free!
[1561169633] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12962420 kB) free.
[1561169823] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2851752 kB) free!
[1561169945] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 218.57, 145.26, 65.71
[1561170065] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 219.50, 169.89, 84.38
[1561170096] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 7.02, 12.84
[1561170185] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1561170185] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 219.47, 186.31, 100.77
[1561170506] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2769736 kB) free!
[1561171140] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1561171140] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561171296] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 0.33, 3.71
[1561171706] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2409828 kB) free!
[1561171740] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561171860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561171981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561171985] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.78, 2.38, 17.29
[1561172033] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9070648 kB) free!
[1561172153] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12734696 kB) free.
[1561172224] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2364052 kB) free!
[1561172581] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561173181] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561173301] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561173421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561173534] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561173654] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1561173785] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.48, 1.75, 4.02
[1561177021] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561178153] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9241592 kB) free!
[1561178221] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561178273] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12939708 kB) free.
[1561178341] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561178461] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561178635] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561179035] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (153184 kB) free!
[1561179155] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155416 kB) free!
[1561179222] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.48, 6.07, 5.58
[1561179275] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155468 kB) free!
[1561179342] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.04, 6.12, 5.66
[1561179463] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.47, 6.05, 5.70
[1561179662] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561180073] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9118692 kB) free!
[1561180193] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12934832 kB) free.
[1561180581] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 199.23, 96.51, 38.02
[1561180606] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.8% (2568152 kB) free!
[1561180701] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 207.09, 133.49, 58.67
[1561180726] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 13.9% (2263120 kB) free!
[1561180793] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8706228 kB) free!
[1561180821] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 208.06, 158.24, 76.81
[1561180846] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.3% (1995176 kB) free!
[1561180854] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561180913] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12933368 kB) free.
[1561180974] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561181394] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2670924 kB) free!
[1561181446] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (418752 kB) free!
[1561181462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561181513] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9111168 kB) free!
[1561181574] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1561181582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561181633] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12934472 kB) free.
[1561181694] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1561181994] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2056720 kB) free!
[1561182112] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 371 processes
[1561182182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561182232] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1561182233] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9136092 kB) free!
[1561182302] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561182353] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12924740 kB) free.
[1561182621] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 1.02, 14.89
[1561182674] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1561182894] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1561182897] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 375 processes
[1561182902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561182953] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8943828 kB) free!
[1561183014] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561183018] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 377 processes
[1561183023] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561183073] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12931428 kB) free.
[1561183138] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 378 processes
[1561183244] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (1310728 kB) free!
[1561183364] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.4% (915168 kB) free!
[1561183484] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (517872 kB) free!
[1561183614] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1561183734] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561183738] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1561183821] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.64, 0.59, 4.50
[1561183985] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561184222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561184342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561184462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561184874] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9703048 kB) free!
[1561184881] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.36, 5.97, 5.38
[1561184994] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12906504 kB) free.
[1561185001] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.68, 5.94, 5.45
[1561185121] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.53, 6.54, 5.73
[1561185185] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 70.68, 27.15, 11.75
[1561185285] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6830076 kB) free.
[1561185305] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 83.00, 46.29, 20.63
[1561185425] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 84.84, 59.15, 28.45
[1561185538] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 377 processes
[1561185658] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 379 processes
[1561185778] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1561186321] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.43, 3.79, 4.55
[1561186379] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 375 processes
[1561186499] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1561186625] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.13, 5.38, 17.35
[1561187394] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9025220 kB) free!
[1561187462] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561187514] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12750128 kB) free.
[1561187632] SERVICE ALERT: redis-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1561187752] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1561188062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561188182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561188299] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 376 processes
[1561188302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561188419] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 378 processes
[1561188425] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.77, 2.73, 4.74
[1561188534] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1561188539] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1561188654] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1561189502] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561189886] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.1% (21734108 kB) free.
[1561189896] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 168.38, 64.86, 24.30
[1561189957] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.0% (22802504 kB) free.
[1561190016] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 190.70, 107.67, 44.91
[1561190030] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.4% (14690040 kB) free.
[1561190096] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.6% (21972252 kB) free.
[1561190102] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561190136] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 194.74, 136.61, 63.12
[1561190162] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.0% (14946612 kB) free.
[1561190212] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 70.8% (11511064 kB) free.
[1561190223] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561190340] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 378 processes
[1561190454] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1561190460] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1561190514] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1561190569] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (2591652 kB) free!
[1561190574] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1561190634] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1561190685] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (1091264 kB) free!
[1561190689] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (2455848 kB) free!
[1561190769] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (155740 kB) free!
[1561190805] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 11.0% (807064 kB) free!
[1561190809] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2337508 kB) free!
[1561190822] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561190889] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161272 kB) free!
[1561190925] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (578664 kB) free!
[1561190942] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561191009] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151736 kB) free!
[1561191534] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1561191542] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561191660] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 380 processes
[1561191662] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561191780] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 374 processes
[1561191834] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1561191894] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1561191936] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.89, 1.38, 12.41
[1561191954] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1561192254] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1561192294] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4464184 kB) free!
[1561192374] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1561192414] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (4343124 kB) free!
[1561192494] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1561192534] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4249460 kB) free!
[1561192862] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561192914] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9276584 kB) free!
[1561192982] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561192987] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173688 kB) free!
[1561193034] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12881668 kB) free.
[1561193103] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561193107] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (177384 kB) free!
[1561193136] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.74, 3.99
[1561193227] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172360 kB) free!
[1561193634] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9061480 kB) free!
[1561193754] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12867036 kB) free.
[1561194302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561194354] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9099068 kB) free!
[1561194422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561194474] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12874808 kB) free.
[1561195022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561195063] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.62, 4.51, 4.94
[1561195074] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8786236 kB) free!
[1561195142] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561195195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12868544 kB) free.
[1561195221] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 189.16, 124.11, 55.03
[1561195341] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 125.39, 130.78, 66.39
[1561195366] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 585 processes
[1561195461] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 17.16, 87.58, 58.37
[1561195486] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 587 processes
[1561195488] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.6% (22598220 kB) free.
[1561195606] SERVICE ALERT: sservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 587 processes
[1561195641] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.9% (22535056 kB) free.
[1561195742] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561195766] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4469404 kB) free!
[1561195794] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8474820 kB) free!
[1561195862] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561195886] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (4395940 kB) free!
[1561195914] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12866052 kB) free.
[1561195980] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 379 processes
[1561196006] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4052616 kB) free!
[1561196100] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1561196207] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 328 processes
[1561196325] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.4% (6798504 kB) free.
[1561196463] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561196514] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8275920 kB) free!
[1561196583] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561196634] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12867192 kB) free.
[1561196661] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.09, 1.79, 16.22
[1561196736] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 122.58, 64.23, 26.19
[1561196856] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 125.98, 84.87, 38.37
[1561196976] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1561196976] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 76.94, 86.64, 45.06
[1561197183] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561197234] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8249528 kB) free!
[1561197303] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561197354] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 36.0% (17761116 kB) free.
[1561197417] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.9% (39359116 kB) free.
[1561197473] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.4% (13710368 kB) free.
[1561197646] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.9% (14938788 kB) free.
[1561197861] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.47, 4.80
[1561197903] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561198023] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561198176] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.77, 2.28, 12.91
[1561198443] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.2% (2640352 kB) free!
[1561198562] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9652448 kB) free!
[1561198563] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.5% (2519176 kB) free!
[1561198623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561198682] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13221688 kB) free.
[1561198684] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.7% (2388420 kB) free!
[1561198743] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561198794] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1561198914] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1561199034] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1561199282] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9837616 kB) free!
[1561199344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561199376] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 0.66, 3.99
[1561199402] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13211152 kB) free.
[1561199464] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561199849] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181108 kB) free!
[1561199884] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1023164 kB) free!
[1561199969] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193620 kB) free!
[1561200003] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9661288 kB) free!
[1561200064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561200089] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186360 kB) free!
[1561200123] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13223672 kB) free.
[1561200184] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561200294] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1561200414] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1561200425] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 7.34, 44.27, 30.88
[1561200534] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1561200545] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.29, 30.49, 27.45
[1561200601] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1561200665] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.43, 21.58, 24.56
[1561200721] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1561200723] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9514024 kB) free!
[1561200841] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1561200843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13223144 kB) free.
[1561201134] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1561201265] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.96, 4.62, 13.85
[1561201289] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1561201289] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 79.8% (19548948 kB) free.
[1561201384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561201443] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9618188 kB) free!
[1561201504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561201563] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13207300 kB) free.
[1561201624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561201890] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168952 kB) free!
[1561202010] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166572 kB) free!
[1561202093] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173344 kB) free!
[1561202130] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197620 kB) free!
[1561202163] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9645488 kB) free!
[1561202214] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167716 kB) free!
[1561202225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561202283] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13217212 kB) free.
[1561202334] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183160 kB) free!
[1561202345] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561202345] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561202474] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1561202883] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9127276 kB) free!
[1561203003] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13203964 kB) free.
[1561203261] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 286.33, 193.94, 87.36
[1561203381] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 114.88, 176.59, 95.10
[1561203501] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1561203501] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.77, 118.28, 83.61
[1561203545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561203665] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561203665] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.17, 2.64, 4.00
[1561204203] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9844528 kB) free!
[1561204265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561204324] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13217228 kB) free.
[1561204385] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561204865] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.46, 6.16, 5.39
[1561204923] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9710392 kB) free!
[1561204985] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561204986] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.40, 5.71, 5.32
[1561205010] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 35.4% (17455096 kB) free.
[1561205093] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.6% (21118780 kB) free.
[1561205106] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561205106] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1561205106] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.92, 5.10, 5.15
[1561205108] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167724 kB) free!
[1561205148] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 81.3% (19923808 kB) free.
[1561205202] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.4% (14696224 kB) free.
[1561205228] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191040 kB) free!
[1561205258] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21930596 kB) free.
[1561205301] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.23, 1.34, 12.66
[1561205699] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181820 kB) free!
[1561205705] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561205706] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.99, 3.37, 4.17
[1561205755] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (416208 kB) free!
[1561205819] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167736 kB) free!
[1561205825] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561205866] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190960 kB) free!
[1561205875] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (223016 kB) free!
[1561205925] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (1114468 kB) free!
[1561205939] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1561205939] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184560 kB) free!
[1561205986] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170056 kB) free!
[1561205995] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163372 kB) free!
[1561206045] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (1006332 kB) free!
[1561206106] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190736 kB) free!
[1561206165] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.5% (922780 kB) free!
[1561206234] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1561206306] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.93, 6.71, 5.58
[1561206425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561206436] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1561206501] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.31, 1.40, 4.50
[1561206545] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561206556] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.06, 6.35, 5.72
[1561206766] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.8% (6827240 kB) free.
[1561207145] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561207265] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561207756] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.64, 2.92, 4.28
[1561207795] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.7% (20762224 kB) free.
[1561207866] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561207925] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165068 kB) free!
[1561207986] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561208045] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185708 kB) free!
[1561208165] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192948 kB) free!
[1561208586] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561208617] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8761256 kB) free!
[1561208706] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561208737] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13210916 kB) free.
[1561208934] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1561209054] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1561209174] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1561209774] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1561209774] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561209906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561209937] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9747644 kB) free!
[1561210026] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561210058] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13217676 kB) free.
[1561210176] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561210626] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561210657] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9312620 kB) free!
[1561210746] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561210777] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13217804 kB) free.
[1561211165] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (2619360 kB) free!
[1561211347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561211377] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8863408 kB) free!
[1561211467] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561211497] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13208280 kB) free.
[1561211554] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1561211674] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1561211765] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192156 kB) free!
[1561211794] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1561212067] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561212097] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9454312 kB) free!
[1561212187] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561212217] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13217932 kB) free.
[1561212787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561212817] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9494304 kB) free!
[1561212907] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561212937] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13208804 kB) free.
[1561213507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561213537] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9447400 kB) free!
[1561213627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561213657] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13215048 kB) free.
[1561214228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561214257] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9332980 kB) free!
[1561214348] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561214378] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13216188 kB) free.
[1561214948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561214977] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9235940 kB) free!
[1561215068] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561215097] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13207744 kB) free.
[1561215668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561215697] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9077068 kB) free!
[1561215774] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1561215788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561215817] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13214088 kB) free.
[1561215894] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561216101] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 114.01, 81.94, 38.31
[1561216221] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 27.73, 64.40, 37.55
[1561216306] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2724192 kB) free!
[1561216341] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.85, 43.15, 33.02
[1561216388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561216417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8934044 kB) free!
[1561216441] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1561216509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561216537] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13205860 kB) free.
[1561216594] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1561216739] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561216941] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 6.06, 17.45
[1561217108] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1561217137] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8703904 kB) free!
[1561217228] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561217257] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13214156 kB) free.
[1561217356] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 203.27, 119.22, 50.92
[1561217376] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 26.90, 6.85, 2.58
[1561217476] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 205.59, 147.89, 69.70
[1561217496] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 62.21, 26.99, 10.47
[1561217596] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 205.63, 166.99, 86.18
[1561217616] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 66.63, 40.37, 17.38
[1561217828] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561217857] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8653776 kB) free!
[1561217939] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2782564 kB) free!
[1561217948] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561217977] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13217980 kB) free.
[1561217995] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561218106] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2308164 kB) free!
[1561218216] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.49, 10.47, 14.39
[1561218548] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561218668] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561218742] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.20, 1.12, 3.75
[1561219268] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561219389] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561219396] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.59, 2.30, 17.06
[1561219416] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.23, 0.49, 4.18
[1561219988] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561220108] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561220708] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561220828] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561220939] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2441436 kB) free!
[1561221196] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.98, 1.81, 3.98
[1561221294] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561221428] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561221548] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561222749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561222777] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9674732 kB) free!
[1561222869] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561222897] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13206652 kB) free.
[1561223469] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561223497] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9522324 kB) free!
[1561223589] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561223616] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 132.49, 38.35, 13.46
[1561223617] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13212964 kB) free.
[1561223694] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1561223736] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 199.28, 95.02, 37.25
[1561223765] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2537072 kB) free!
[1561223814] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561223856] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1561223856] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 208.01, 132.87, 58.13
[1561224189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561224217] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9216732 kB) free!
[1561224309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561224337] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13205660 kB) free.
[1561224539] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.4% (3049660 kB) free!
[1561224909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561224937] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8951328 kB) free!
[1561225029] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561225058] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13211308 kB) free.
[1561225630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561225656] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 0.83, 11.90
[1561225657] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8797780 kB) free!
[1561225750] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561225777] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13213708 kB) free.
[1561226350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561226470] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561226814] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561226856] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 0.45, 3.56
[1561226934] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561227070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561227190] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes