[1561315042] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561315063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8367400 kB) free!
[1561315162] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561315183] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12908808 kB) free.
[1561315740] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4940816 kB) free!
[1561315762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561315860] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4930596 kB) free!
[1561315882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561315980] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4921512 kB) free!
[1561316384] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9768124 kB) free!
[1561316482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561316504] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12909448 kB) free.
[1561316602] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561316617] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 180.95, 69.44, 25.71
[1561316737] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 205.43, 115.72, 47.97
[1561316857] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 208.77, 146.70, 67.54
[1561317104] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9697216 kB) free!
[1561317181] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 23.6% (5831800 kB) free.
[1561317203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561317224] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12912648 kB) free.
[1561317323] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561317824] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9671516 kB) free!
[1561317923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1561317944] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12903472 kB) free.
[1561318043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561318335] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.0% (4169200 kB) free!
[1561318395] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4171388 kB) free!
[1561318455] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 17.0% (4171808 kB) free!
[1561318515] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.0% (4175020 kB) free!
[1561318544] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9606388 kB) free!
[1561318643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561318657] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 0.86, 11.61
[1561318664] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12901924 kB) free.
[1561318763] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561319236] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561319264] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9594568 kB) free!
[1561319363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561319384] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12903744 kB) free.
[1561319483] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561319857] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.41, 0.41, 3.48
[1561319984] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9498196 kB) free!
[1561320105] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12899992 kB) free.
[1561320457] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 24.39, 5.74, 3.78
[1561320577] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 98.17, 40.09, 16.60
[1561320683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561320697] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1561320697] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 107.80, 63.03, 27.86
[1561320704] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9416092 kB) free!
[1561320781] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (4035560 kB) free!
[1561320803] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561320824] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12901824 kB) free.
[1561320901] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 25.6% (6310944 kB) free.
[1561321403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561321424] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9309976 kB) free!
[1561321523] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561321544] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12890004 kB) free.
[1561321897] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.18, 2.74, 13.30
[1561322123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561322144] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9169944 kB) free!
[1561322243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561322264] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12899364 kB) free.
[1561322309] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (263576 kB) free!
[1561322429] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (250252 kB) free!
[1561322549] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (265948 kB) free!
[1561322843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561322864] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9094912 kB) free!
[1561322964] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561322984] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12902820 kB) free.
[1561323097] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.63, 0.48, 3.93
[1561323563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561323584] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9008024 kB) free!
[1561323683] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561323704] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12902116 kB) free.
[1561324283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561324304] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8923508 kB) free!
[1561324403] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561324424] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12902080 kB) free.
[1561325003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561325025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8794812 kB) free!
[1561325123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561325145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12889628 kB) free.
[1561325544] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 95.95, 47.25, 18.72
[1561325664] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 99.63, 64.78, 28.60
[1561325723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1561325745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8546556 kB) free!
[1561325784] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 60.83, 66.97, 34.12
[1561325807] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 583 processes
[1561325844] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561325865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12898256 kB) free.
[1561325927] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 582 processes
[1561326047] SERVICE ALERT: sservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 584 processes
[1561326384] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.17, 9.15, 17.97
[1561326443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561326465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8372416 kB) free!
[1561326563] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561326585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12900756 kB) free.
[1561326647] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 322 processes
[1561327163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561327283] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561327785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9665240 kB) free!
[1561327883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561327905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12900380 kB) free.
[1561328003] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561328184] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.14, 0.19, 2.74
[1561328505] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9596624 kB) free!
[1561328625] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12900200 kB) free.
[1561329204] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561329225] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9421180 kB) free!
[1561329324] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561329346] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12891620 kB) free.
[1561329924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561329945] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9252340 kB) free!
[1561330044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561330065] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12899132 kB) free.
[1561330644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561330665] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8992564 kB) free!
[1561330764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561330785] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12900376 kB) free.
[1561331364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561331385] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8695200 kB) free!
[1561331484] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561331505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12898360 kB) free.
[1561332085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561332105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8349136 kB) free!
[1561332205] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561332225] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12899064 kB) free.
[1561332805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561332925] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561333525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561333645] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561333765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561333897] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561334236] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 203.22, 115.46, 48.80
[1561334356] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 206.49, 145.75, 67.98
[1561334365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561334476] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 207.11, 166.04, 84.84
[1561334526] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (3168272 kB) free!
[1561334646] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 22.9% (5612548 kB) free.
[1561334966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561335086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561335206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561335225] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9020796 kB) free!
[1561335345] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12855880 kB) free.
[1561335806] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1561335806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561335846] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4826628 kB) free!
[1561335901] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4831704 kB) free!
[1561335966] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4735632 kB) free!
[1561336021] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4818228 kB) free!
[1561336086] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (4673104 kB) free!
[1561336141] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4806704 kB) free!
[1561336276] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.98, 2.73, 18.50
[1561336406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561336526] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561336566] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4506496 kB) free!
[1561336646] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561336686] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4705648 kB) free!
[1561336806] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (4549048 kB) free!
[1561336949] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 61.6% (10013196 kB) free.
[1561337145] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8577368 kB) free!
[1561337265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12479796 kB) free.
[1561337476] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.92, 23.15, 14.12
[1561337887] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2462148 kB) free!
[1561338315] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 21.7% (5325844 kB) free.
[1561338465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9723024 kB) free!
[1561338487] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (4834304 kB) free!
[1561338585] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8114208 kB) free!
[1561338676] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1561338676] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.44, 7.00, 18.62
[1561338705] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12056632 kB) free.
[1561339206] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.6% (5043852 kB) free.
[1561339905] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9214004 kB) free!
[1561340026] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12826360 kB) free.
[1561340287] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1561340287] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.4% (5006752 kB) free.
[1561340476] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.61, 1.72, 4.13
[1561341446] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561341487] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4895760 kB) free!
[1561341607] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4856920 kB) free!
[1561341727] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4853116 kB) free!
[1561342046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561342167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561342287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561342425] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9674016 kB) free!
[1561342545] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12813928 kB) free.
[1561342897] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 74.14, 25.84, 9.56
[1561343017] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 90.44, 48.06, 19.68
[1561343137] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 92.86, 63.03, 28.60
[1561343145] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9613224 kB) free!
[1561343265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12818256 kB) free.
[1561343487] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561343865] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9529412 kB) free!
[1561343985] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12820336 kB) free.
[1561344337] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.32, 2.07, 10.74
[1561344585] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9479020 kB) free!
[1561344706] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12815988 kB) free.
[1561345287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561345305] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9378860 kB) free!
[1561345407] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561345425] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12818116 kB) free.
[1561345537] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.95, 0.92, 3.49
[1561346008] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561346025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9219820 kB) free!
[1561346128] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561346145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12809124 kB) free.
[1561346728] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561346745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9054340 kB) free!
[1561346848] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561346865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12812788 kB) free.
[1561347448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561347465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8847772 kB) free!
[1561347568] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561347585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12811476 kB) free.
[1561348168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561348185] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8795804 kB) free!
[1561348288] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561348305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12810440 kB) free.
[1561348806] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4870352 kB) free!
[1561348889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561348905] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8695544 kB) free!
[1561348926] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4849212 kB) free!
[1561349009] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561349026] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12812720 kB) free.
[1561349046] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4429808 kB) free!
[1561349138] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 146.24, 45.84, 16.49
[1561349257] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 196.44, 98.29, 39.27
[1561349341] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1657576 kB) free!
[1561349377] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1561349377] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 203.36, 133.39, 59.29
[1561349609] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561349625] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8588004 kB) free!
[1561349729] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561349745] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12805172 kB) free.
[1561349941] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.1% (3489008 kB) free!
[1561350329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561350345] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8421952 kB) free!
[1561350449] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561350465] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12820168 kB) free.
[1561350676] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561350727] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437316 kB) free!
[1561350775] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1561350895] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561351049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561351065] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8261968 kB) free!
[1561351169] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561351177] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 1.05, 12.48
[1561351185] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12821648 kB) free.
[1561351741] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2308568 kB) free!
[1561351770] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561351785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8530580 kB) free!
[1561351890] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561351905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12818632 kB) free.
[1561352150] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.6% (2693364 kB) free!
[1561352270] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.8% (2569900 kB) free!
[1561352377] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.03, 0.79, 3.96
[1561352391] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.2% (2474364 kB) free!
[1561352490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561352505] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8854828 kB) free!
[1561352549] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (182796 kB) free!
[1561352610] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561352625] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12823672 kB) free.
[1561352669] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171468 kB) free!
[1561352789] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156496 kB) free!
[1561353076] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.65, 6.80, 6.67
[1561353106] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.92, 5.67, 5.03
[1561353127] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (2864024 kB) free!
[1561353196] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.62, 5.80, 6.32
[1561353210] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561353225] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8836588 kB) free!
[1561353226] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.90, 7.68, 5.85
[1561353316] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.12, 4.92, 5.93
[1561353330] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561353346] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12814964 kB) free.
[1561353346] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.36, 8.29, 6.30
[1561353591] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (322712 kB) free!
[1561353727] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2069984 kB) free!
[1561353916] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.18, 2.43, 4.09
[1561353930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561353945] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8919660 kB) free!
[1561354050] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561354065] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12823080 kB) free.
[1561354584] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 58.40, 52.27, 22.05
[1561354651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561354665] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8539572 kB) free!
[1561354704] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8.05, 35.04, 19.40
[1561354771] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561354785] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12820012 kB) free.
[1561354824] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.95, 23.72, 17.15
[1561355095] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561355184] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1561355215] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561355304] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561355371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561355385] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8979848 kB) free!
[1561355491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561355505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12821968 kB) free.
[1561355655] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9452084 kB) free!
[1561355775] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9451608 kB) free!
[1561355895] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (9451836 kB) free!
[1561356091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561356105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9125432 kB) free!
[1561356211] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561356211] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561356225] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12821228 kB) free.
[1561356234] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1561356354] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1561356474] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1561356624] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 0.43, 2.82
[1561356811] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561356825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8459000 kB) free!
[1561356931] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561356945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12809832 kB) free.
[1561357007] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (212856 kB) free!
[1561357127] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (120212 kB) free!
[1561357247] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (129024 kB) free!
[1561357532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561357652] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561357848] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.2% (6783220 kB) free.
[1561358144] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.5% (23078356 kB) free.
[1561358146] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9571800 kB) free!
[1561358189] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 37.2% (6040644 kB) free.
[1561358205] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.6% (22455752 kB) free.
[1561358252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561358262] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.9% (14443016 kB) free.
[1561358266] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12816384 kB) free.
[1561358295] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238392 kB) free!
[1561358319] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.8% (38810892 kB) free.
[1561358369] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 69.7% (17089200 kB) free.
[1561358372] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561358789] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (426856 kB) free!
[1561358805] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (505304 kB) free!
[1561358909] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (160600 kB) free!
[1561358925] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165360 kB) free!
[1561359029] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (170880 kB) free!
[1561359045] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164908 kB) free!
[1561359572] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561359692] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561359812] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561360015] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1561360127] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244772 kB) free!
[1561360135] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561360247] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233244 kB) free!
[1561360367] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1561360367] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240012 kB) free!
[1561360516] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.77, 5.02, 5.11
[1561360636] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.84, 3.94, 4.71
[1561360638] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.4% (2498488 kB) free!
[1561360759] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.4% (2498844 kB) free!
[1561360879] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (2439544 kB) free!
[1561361151] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4760796 kB) free!
[1561361271] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4622060 kB) free!
[1561361377] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 19.61, 5.92, 3.58
[1561361391] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (4466188 kB) free!
[1561361424] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.28, 5.73, 2.96
[1561361479] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167412 kB) free!
[1561361497] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.47, 4.24, 3.25
[1561361544] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.12, 4.06, 2.68
[1561362029] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1561362029] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 44.1% (7166752 kB) free.
[1561362436] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.42, 5.13, 5.08
[1561362557] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.23, 6.01, 5.49
[1561362674] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1561362677] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.53, 4.68, 5.06
[1561362950] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (3760228 kB) free!
[1561363068] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (3033928 kB) free!
[1561363070] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.7% (3634920 kB) free!
[1561363188] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.0% (2760240 kB) free!
[1561363190] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.6% (3594828 kB) free!
[1561363277] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1561363277] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.97, 2.39, 3.67
[1561363297] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561363308] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.3% (2484592 kB) free!
[1561363412] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561363908] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1561363908] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1388712 kB) free!
[1561364012] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561364132] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561364391] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2444480 kB) free!
[1561364448] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (675624 kB) free!
[1561364455] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1561364515] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 64 users currently logged in
[1561364568] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (446588 kB) free!
[1561364575] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1561364635] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1561364688] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (240392 kB) free!
[1561364732] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561364853] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561364935] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1561364990] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195648 kB) free!
[1561365055] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1561365452] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561365466] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9368112 kB) free!
[1561365572] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561365586] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12795464 kB) free.
[1561365655] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561365674] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1561365775] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1561365895] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1561366172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561366186] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9477836 kB) free!
[1561366292] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561366306] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12791820 kB) free.
[1561366495] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1561366615] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1561366735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1561366892] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561366906] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9305396 kB) free!
[1561367012] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561367026] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12792096 kB) free.
[1561367335] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561367477] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 19.82, 8.23, 5.11
[1561367597] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.95, 6.02, 4.69
[1561367612] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561367626] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8797444 kB) free!
[1561367733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561367746] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12782792 kB) free.
[1561368332] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561368346] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.61, 4.12, 4.64
[1561368346] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8176124 kB) free!
[1561368452] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561368466] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12789936 kB) free.
[1561368888] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.5% (6804308 kB) free.
[1561369652] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561369666] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9166824 kB) free!
[1561369772] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561369787] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12790364 kB) free.
[1561369967] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1561370335] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 67 users currently logged in
[1561370372] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561370386] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9406972 kB) free!
[1561370445] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561370492] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561370506] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12789196 kB) free.
[1561370635] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1561371029] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1561371074] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 8 users currently logged in
[1561371093] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561371106] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9443188 kB) free!
[1561371213] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561371226] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12789240 kB) free.
[1561371235] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 52 users currently logged in
[1561371333] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561371346] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.40, 6.28, 5.26
[1561371466] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.19, 6.15, 5.33
[1561371586] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.25, 5.49, 5.19
[1561371826] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9438884 kB) free!
[1561371946] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12785628 kB) free.
[1561372186] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.79, 4.53, 4.82
[1561372229] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (170200 kB) free!
[1561372274] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1561372349] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180564 kB) free!
[1561372394] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1561372469] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180764 kB) free!
[1561372500] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1561372514] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1561372514] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1561372533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561372546] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9319696 kB) free!
[1561372621] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1561372653] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561372666] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12787156 kB) free.
[1561373114] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1561373171] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.20, 7.20, 5.63
[1561373253] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561373266] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9491804 kB) free!
[1561373291] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.74, 6.57, 5.58
[1561373309] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.0% (23446564 kB) free.
[1561373374] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561373376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 36.0% (17758860 kB) free.
[1561373411] SERVICE ALERT: webservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.85, 6.43, 5.66
[1561373445] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 30.2% (7411644 kB) free.
[1561373502] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.9% (39380488 kB) free.
[1561373553] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.2% (14495952 kB) free.
[1561373597] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.15, 12.82, 8.38
[1561373629] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.6% (22449168 kB) free.
[1561373691] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.4% (14207208 kB) free.
[1561373714] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1561373717] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.92, 9.34, 7.65
[1561373754] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 57.7% (9385052 kB) free.
[1561373820] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1561373834] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1561373837] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.25, 6.96, 6.98
[1561373940] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1561373987] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.12, 5.92, 5.13
[1561374060] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1561374098] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.91, 5.81, 5.38
[1561374107] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.11, 6.13, 5.30
[1561374107] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (7164244 kB) free!
[1561374144] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.73, 5.20, 5.16
[1561374218] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 14.95, 9.02, 6.55
[1561374227] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1561374227] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.46, 5.99, 5.36
[1561374227] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (6037320 kB) free!
[1561374264] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.01, 6.26, 5.47
[1561374298] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171696 kB) free!
[1561374338] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.78, 6.35, 5.87
[1561374347] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3641840 kB) free!
[1561374358] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (170192 kB) free!
[1561374384] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.57, 4.24, 4.82
[1561374419] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (170292 kB) free!
[1561374437] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.26, 2.87, 4.65
[1561374478] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171724 kB) free!
[1561374539] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (170512 kB) free!
[1561374573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561374598] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157640 kB) free!
[1561374611] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.17, 4.01, 4.69
[1561374660] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1561374693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561374780] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 369 processes
[1561374813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561374901] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1561375037] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.24, 5.70, 6.40
[1561375157] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.00, 4.48, 5.87
[1561375211] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.94, 9.17, 6.91
[1561375277] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.28, 3.79, 5.45
[1561375331] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.31, 9.90, 7.46
[1561375452] SERVICE ALERT: webservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.56, 10.11, 7.83
[1561375501] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1561375538] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.22, 2.16, 4.78
[1561375787] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8584696 kB) free!
[1561375877] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.48, 2.61, 3.98
[1561375907] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12476988 kB) free.
[1561376088] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (1358120 kB) free!
[1561376101] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 368 processes
[1561376208] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (1203872 kB) free!
[1561376221] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1561376328] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.3% (1051236 kB) free!
[1561376928] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (199836 kB) free!
[1561376935] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1561376995] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1561377055] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1561377115] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1561377213] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561377252] SERVICE FLAPPING ALERT: webservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1561377252] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.43, 2.66, 4.94
[1561377434] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1561377554] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1561377601] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.9% (2913388 kB) free!
[1561377674] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1561377707] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8736944 kB) free!
[1561377721] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2845520 kB) free!
[1561377827] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13109800 kB) free.
[1561377841] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.2% (2787932 kB) free!
[1561378021] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 369 processes
[1561378141] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 368 processes
[1561378261] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1561378315] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 59 users currently logged in
[1561378413] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561378427] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.31, 3.73, 4.56
[1561378427] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9052316 kB) free!
[1561378533] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561378547] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13100216 kB) free.
[1561378961] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1561379133] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561379147] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8860848 kB) free!
[1561379253] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561379267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13108468 kB) free.
[1561379315] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182340 kB) free!
[1561379435] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189304 kB) free!
[1561379453] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171928 kB) free!
[1561379555] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (185076 kB) free!
[1561379573] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185192 kB) free!
[1561379641] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.5% (13577056 kB) free.
[1561379693] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178748 kB) free!
[1561379853] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561379868] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8898332 kB) free!
[1561379929] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1561379929] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6834200 kB) free.
[1561379974] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561379988] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13108644 kB) free.
[1561380573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561380588] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8682436 kB) free!
[1561380693] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561380708] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13104364 kB) free.
[1561381192] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.5% (18148336 kB) free.
[1561381237] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.8% (22267180 kB) free.
[1561381274] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 8 users currently logged in
[1561381293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561381304] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 93.4% (46010120 kB) free.
[1561381308] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8544792 kB) free!
[1561381340] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.0% (14958096 kB) free.
[1561381399] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 36.0% (17723372 kB) free.
[1561381413] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561381449] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 69.2% (11245584 kB) free.
[1561381797] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196728 kB) free!
[1561381848] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.2% (3225020 kB) free!
[1561381917] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186888 kB) free!
[1561381968] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196072 kB) free!
[1561382037] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1561382037] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191812 kB) free!
[1561382088] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1561382088] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174616 kB) free!
[1561382613] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561382733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561383061] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1561383061] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 362 processes
[1561383209] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8887472 kB) free!
[1561383329] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12672736 kB) free.
[1561383334] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561383454] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561383574] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561383761] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.7% (2711848 kB) free!
[1561383827] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.20, 7.01, 5.88
[1561383881] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.5% (2513176 kB) free!
[1561383947] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.51, 6.50, 5.84
[1561384001] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.2% (2311492 kB) free!
[1561384067] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.53, 5.95, 5.73
[1561384601] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1219436 kB) free!
[1561384735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1561384855] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 351 processes
[1561384975] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1561385129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8470428 kB) free!
[1561385249] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12382532 kB) free.
[1561385575] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561386529] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.2% (973604 kB) free!
[1561386650] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.1% (6774828 kB) free.
[1561386837] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 26.5% (6533780 kB) free.
[1561387277] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561387437] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (204860 kB) free!
[1561387452] SERVICE FLAPPING ALERT: webservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561387557] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (361488 kB) free!
[1561387593] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 384 processes
[1561387677] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197396 kB) free!
[1561387713] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 383 processes
[1561387774] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561388267] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.21, 3.48, 4.74
[1561388461] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1561388849] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9592248 kB) free!
[1561388970] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13111612 kB) free.
[1561389049] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1561389569] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9495016 kB) free!
[1561389574] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561389674] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1561389689] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13103276 kB) free.
[1561389694] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561390289] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9403088 kB) free!
[1561390294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561390409] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13112740 kB) free.
[1561390414] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561391009] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8986668 kB) free!
[1561391014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561391129] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13099644 kB) free.
[1561391135] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561391274] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1561391729] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9299840 kB) free!
[1561391734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561391738] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 230.89, 73.70, 27.58
[1561391849] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13109768 kB) free.
[1561391854] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561391859] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 310.07, 155.99, 63.32
[1561391979] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 320.64, 210.97, 94.70
[1561392288] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561392449] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9229236 kB) free!
[1561392454] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561392569] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13110572 kB) free.
[1561392574] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561393170] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9097812 kB) free!
[1561393174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561393290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13103972 kB) free.
[1561393294] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561393375] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561393495] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1561393615] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1561393890] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8922500 kB) free!
[1561393894] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561394010] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13112272 kB) free.
[1561394015] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561394215] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1561394215] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561394610] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8695224 kB) free!
[1561394614] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561394730] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13100692 kB) free.
[1561394734] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561394979] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.51, 2.53, 14.78
[1561395061] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1561395330] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8453804 kB) free!
[1561395334] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561395450] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13110012 kB) free.
[1561395454] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561396067] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1561396650] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9766020 kB) free!
[1561396654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561396770] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13112840 kB) free.
[1561396774] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561396779] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.78, 2.58, 4.24
[1561397370] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9577636 kB) free!
[1561397375] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561397491] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13103696 kB) free.
[1561397495] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561398090] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9408320 kB) free!
[1561398095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561398210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13111512 kB) free.
[1561398215] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561398810] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9141796 kB) free!
[1561398815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561398930] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13099956 kB) free.
[1561398935] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561399077] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561399530] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8868392 kB) free!
[1561399535] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561399650] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13107380 kB) free.
[1561399655] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561400184] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 291.55, 188.38, 82.82
[1561400250] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8627012 kB) free!
[1561400255] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561400304] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 211.81, 204.57, 102.25
[1561400370] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13108800 kB) free.
[1561400376] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561400424] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.87, 138.91, 90.59
[1561400815] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1561400935] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes