[1588271609] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10308304 kB) free!
[1588271730] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6694640 kB) free!
[1588271850] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29833536 kB) free.
[1588271899] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3735984 kB) free.
[1588272033] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 361 processes
[1588272450] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9050036 kB) free!
[1588272499] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2604708 kB) free!
[1588272570] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19390028 kB) free!
[1588272619] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3299296 kB) free.
[1588272690] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27813496 kB) free.
[1588273290] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7545840 kB) free!
[1588273410] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30342776 kB) free.
[1588273535] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195724 kB) free!
[1588273596] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194408 kB) free!
[1588273656] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (187228 kB) free!
[1588273719] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (165712 kB) free!
[1588273819] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3012592 kB) free!
[1588273939] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (158072 kB) free!
[1588274010] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9932148 kB) free!
[1588274059] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2156268 kB) free!
[1588274130] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6399456 kB) free!
[1588274250] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29034784 kB) free.
[1588274319] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 76.7% (18790648 kB) free.
[1588274850] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8891352 kB) free!
[1588274970] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30228604 kB) free.
[1588275259] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1604864 kB) free!
[1588275570] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10486204 kB) free!
[1588275690] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6958904 kB) free!
[1588275810] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29041860 kB) free.
[1588275859] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.4% (2183320 kB) free!
[1588276411] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7785836 kB) free!
[1588276531] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30038580 kB) free.
[1588276744] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 221571 MB (20% inode=99%):
[1588276855] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (377992 kB) free!
[1588276864] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 211552 MB (19% inode=99%):
[1588276935] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239524 kB) free!
[1588276975] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (376400 kB) free!
[1588276984] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 201959 MB (18% inode=99%):
[1588277055] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (220348 kB) free!
[1588277059] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (462356 kB) free!
[1588277095] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (373544 kB) free!
[1588277131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8595904 kB) free!
[1588277175] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242344 kB) free!
[1588277251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (5012732 kB) free!
[1588277371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27632680 kB) free.
[1588277971] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6320460 kB) free!
[1588278091] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29799960 kB) free.
[1588278691] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7086748 kB) free!
[1588278811] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29761384 kB) free.
[1588278975] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 52.7% (25960772 kB) free.
[1588279165] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1588279286] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1588279411] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9092804 kB) free!
[1588279451] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 34.4% (8421208 kB) free.
[1588279459] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2165660 kB) free!
[1588279531] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5588016 kB) free!
[1588279651] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26633720 kB) free.
[1588280051] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189680 kB) free!
[1588280059] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1358812 kB) free!
[1588280171] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193144 kB) free!
[1588280251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6441836 kB) free!
[1588280291] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188536 kB) free!
[1588280371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27053176 kB) free.
[1588280659] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2164956 kB) free!
[1588280971] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6756428 kB) free!
[1588281091] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26594364 kB) free.
[1588281129] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 92.2% (45417212 kB) free.
[1588281545] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.7% (46188024 kB) free.
[1588282233] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1588282291] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10173444 kB) free!
[1588282397] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(241) < 103.16.101.72 [01/May/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1588282411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29678648 kB) free.
[1588282450] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2414) < 203.212.70.170 [01/May/2020:03:02:32 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1588282517] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588282570] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588282644] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(302) < 2020/05/01 03:01:32 [error] 22122#0: *86433814 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588282764] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588283011] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9086464 kB) free!
[1588283131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27831740 kB) free.
[1588284331] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23567796 kB) free!
[1588284451] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27206328 kB) free.
[1588284491] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.9% (10263212 kB) free.
[1588285495] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.2% (13447404 kB) free!
[1588286251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11776656 kB) free!
[1588286292] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (228992 kB) free!
[1588286371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28820016 kB) free.
[1588286412] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (229304 kB) free!
[1588286532] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1588286532] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196720 kB) free!
[1588286659] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.2% (4255120 kB) free.
[1588287571] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11578884 kB) free!
[1588287691] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27865548 kB) free.
[1588287732] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.2% (9613112 kB) free.
[1588288459] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2996664 kB) free!
[1588288579] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.8% (4512964 kB) free.
[1588288891] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9479048 kB) free!
[1588289011] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 20.0% (26332904 kB) free!
[1588289131] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21523056 kB) free!
[1588289532] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171732 kB) free!
[1588289652] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174468 kB) free!
[1588289772] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195368 kB) free!
[1588290986] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.6% (20486884 kB) free.
[1588291072] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 56.8% (14026344 kB) free.
[1588291200] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 65.7% (16106196 kB) free.
[1588292131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27590948 kB) free.
[1588292779] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1588293295] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4540888 kB) free!
[1588294984] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 242841 MB (22% inode=99%):
[1588295095] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 49.0% (64567416 kB) free.
[1588295131] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26326968 kB) free!
[1588295251] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22501588 kB) free!
[1588295371] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18712196 kB) free!
[1588295872] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (1279416 kB) free!
[1588295992] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (1317656 kB) free!
[1588296112] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1313940 kB) free!
[1588296523] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4836364 kB) free!
[1588296583] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 34.3% (8402328 kB) free.
[1588298371] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1588298371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30069240 kB) free.
[1588299972] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588300771] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25846232 kB) free!
[1588300891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20961724 kB) free!
[1588301011] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16313756 kB) free!
[1588301272] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.94, 8.95, 5.98
[1588301392] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.70, 9.73, 6.62
[1588301512] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.79, 9.89, 7.07
[1588302600] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4452304 kB) free!
[1588302720] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4372560 kB) free!
[1588302840] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4298728 kB) free!
[1588303312] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.30, 3.16, 4.86
[1588304011] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12954824 kB) free!
[1588305240] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2110896 kB) free!
[1588305633] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1588305753] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1588305873] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1588305971] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(883) < 103.255.146.164 [01/May/2020:09:30:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588306091] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588306312] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.60, 6.51, 5.31
[1588306432] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.39, 6.04, 5.28
[1588306552] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.74, 6.27, 5.45
[1588306885] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 333 processes
[1588307005] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1588307011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29295232 kB) free.
[1588308211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22543660 kB) free!
[1588308331] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18209056 kB) free!
[1588308451] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43544944 kB) free.
[1588309651] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19562452 kB) free!
[1588309692] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(494) < 103.255.146.164 [01/May/2020:10:32:37 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588309771] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40493820 kB) free.
[1588309812] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588310605] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1588310725] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1588311571] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19400504 kB) free!
[1588311692] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39148608 kB) free.
[1588311979] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3207792 kB) free!
[1588312099] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (1288444 kB) free!
[1588312126] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 10.5% (2355992 kB) free!
[1588312219] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3075640 kB) free!
[1588312246] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1654100 kB) free!
[1588312291] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14729140 kB) free!
[1588312366] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (970176 kB) free!
[1588312411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36497540 kB) free.
[1588313011] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13502824 kB) free!
[1588313131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34345636 kB) free.
[1588313731] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10924576 kB) free!
[1588313851] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33655400 kB) free.
[1588314012] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(340) < 103.255.146.164 [01/May/2020:11:48:58 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588314132] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588314273] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1588314451] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10023532 kB) free!
[1588314571] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32178796 kB) free.
[1588314767] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.9% (15923588 kB) free.
[1588314883] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3820276 kB) free!
[1588314943] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4837688 kB) free!
[1588315003] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 17.6% (4303332 kB) free!
[1588315063] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.9% (4147648 kB) free!
[1588315171] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9268988 kB) free!
[1588315291] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30147412 kB) free.
[1588315363] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 22.2% (5439780 kB) free.
[1588315891] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6983900 kB) free!
[1588316011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26870652 kB) free.
[1588316533] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(767) < 103.255.146.164 [01/May/2020:12:31:09 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588316653] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588317211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18830624 kB) free!
[1588317331] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1588317331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38566636 kB) free.
[1588317932] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16314068 kB) free!
[1588317986] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4625584 kB) free!
[1588318052] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39723772 kB) free.
[1588318063] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.6% (2345568 kB) free!
[1588318106] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4517768 kB) free!
[1588318226] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (4405192 kB) free!
[1588318453] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(544) < 103.255.146.164 [01/May/2020:13:01:18 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588318573] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588318652] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16397748 kB) free!
[1588318772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38532512 kB) free.
[1588318819] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (949372 kB) free!
[1588319372] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14940844 kB) free!
[1588319419] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (2511156 kB) free!
[1588319492] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15248984 kB) free!
[1588319612] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34298164 kB) free.
[1588319774] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(583) < 103.255.146.164 [01/May/2020:13:16:38 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588319894] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588320019] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1588320019] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (148204 kB) free!
[1588320026] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2368280 kB) free!
[1588320212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11922276 kB) free!
[1588320332] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34181988 kB) free.
[1588320619] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3120768 kB) free!
[1588320932] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11057780 kB) free!
[1588321052] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34424664 kB) free.
[1588321145] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8742560 kB) free!
[1588321265] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8485512 kB) free!
[1588321385] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8257580 kB) free!
[1588321652] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11107576 kB) free!
[1588321772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33148184 kB) free.
[1588321967] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2796096 kB) free!
[1588322087] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.7% (2392264 kB) free!
[1588322207] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2046416 kB) free!
[1588322372] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11893760 kB) free!
[1588322492] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7781760 kB) free!
[1588322613] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31756624 kB) free.
[1588322895] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(294) < 103.16.101.72 [01/May/2020:14:16:04 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1588323015] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588323185] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4731524 kB) free!
[1588323212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8525772 kB) free!
[1588323332] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4411644 kB) free!
[1588323452] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28433224 kB) free.
[1588324052] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6749812 kB) free!
[1588324172] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29467368 kB) free.
[1588324772] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8111528 kB) free!
[1588324819] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1519864 kB) free!
[1588324892] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4429144 kB) free!
[1588325012] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27010596 kB) free.
[1588325152] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.81, 2.91, 4.49
[1588325415] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(341) < 185.255.8.58 [01/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1588325536] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(187) < 103.255.146.164 [01/May/2020:15:01:09 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588325612] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5731448 kB) free!
[1588325656] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(181) < 103.255.146.164 [01/May/2020:15:03:39 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588325732] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28061828 kB) free.
[1588325808] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.3% (15998640 kB) free.
[1588326019] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (2706948 kB) free!
[1588326256] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588326332] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5871576 kB) free!
[1588326352] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.90, 7.38, 6.33
[1588326452] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2118696 kB) free!
[1588326472] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.90, 7.44, 6.47
[1588326572] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25854448 kB) free!
[1588326592] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.23, 7.28, 6.52
[1588326856] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(615) < 103.255.146.164 [01/May/2020:15:22:37 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588326976] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588327172] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3910072 kB) free!
[1588327364] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4897544 kB) free!
[1588327424] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4896964 kB) free!
[1588327484] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 20.0% (4897916 kB) free!
[1588327525] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1588327544] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 20.0% (4893928 kB) free!
[1588327645] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 329 processes
[1588327765] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1588328373] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17753384 kB) free!
[1588329130] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9414888 kB) free!
[1588329250] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (9269316 kB) free!
[1588329370] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (9158168 kB) free!
[1588329376] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(127) < 103.255.146.164 [01/May/2020:16:01:22 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588329496] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588329573] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1588329573] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393960 kB) free!
[1588330173] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13549216 kB) free!
[1588330219] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (952372 kB) free!
[1588330696] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(652) < 103.255.146.164 [01/May/2020:16:19:24 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588330792] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.01, 3.60, 4.75
[1588330816] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588330819] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3202668 kB) free!
[1588331373] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390208 kB) free!
[1588331992] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.57, 5.72, 5.13
[1588332113] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.42, 5.79, 5.25
[1588332233] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1588332233] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.28, 4.94, 5.01
[1588332573] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (24056252 kB) free!
[1588332616] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(464) < 103.255.146.164 [01/May/2020:17:00:04 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588332736] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588333173] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411132 kB) free!
[1588333219] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (670536 kB) free!
[1588333819] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (2016620 kB) free!
[1588333936] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(448) < 103.255.146.164 [01/May/2020:17:18:56 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588334033] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.43, 3.76, 4.73
[1588334057] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588334373] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21902420 kB) free!
[1588334973] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32777740 kB) free.
[1588335019] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1071584 kB) free!
[1588335233] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.25, 8.82, 6.80
[1588335353] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.19, 8.22, 6.84
[1588335370] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4420888 kB) free!
[1588335408] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (3729236 kB) free!
[1588335473] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.78, 8.32, 7.03
[1588335528] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3197672 kB) free!
[1588335573] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7887292 kB) free!
[1588335619] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (2414612 kB) free!
[1588335648] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.5% (2575592 kB) free!
[1588335693] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3452792 kB) free!
[1588335813] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423480 kB) free!
[1588336166] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1588336249] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.1% (15974708 kB) free.
[1588336286] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1588336414] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20477616 kB) free!
[1588336820] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1067956 kB) free!
[1588337014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32846852 kB) free.
[1588337420] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3073244 kB) free!
[1588337614] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7352168 kB) free!
[1588337734] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38131368 kB) free.
[1588338334] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12477436 kB) free!
[1588338454] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8240852 kB) free!
[1588338574] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38305276 kB) free.
[1588338620] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (291144 kB) free!
[1588339174] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12824932 kB) free!
[1588339220] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (2752236 kB) free!
[1588339294] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8602844 kB) free!
[1588339414] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37641228 kB) free.
[1588340014] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12440256 kB) free!
[1588340134] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8121452 kB) free!
[1588340254] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30227928 kB) free.
[1588340420] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (559292 kB) free!
[1588340854] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14216772 kB) free!
[1588340874] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.07, 4.66, 4.88
[1588340974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9703980 kB) free!
[1588341020] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (2720924 kB) free!
[1588341095] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5348136 kB) free!
[1588341620] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1593088 kB) free!
[1588341695] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14806100 kB) free!
[1588342220] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2577608 kB) free!
[1588342243] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1588342363] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1588342483] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1588342895] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35582300 kB) free.
[1588343083] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1588343420] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (732140 kB) free!
[1588343495] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8783684 kB) free!
[1588343615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4483348 kB) free!
[1588343735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35122224 kB) free.
[1588344020] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (2481048 kB) free!
[1588344335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9467284 kB) free!
[1588344340] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 72.5% (17776716 kB) free.
[1588344455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (5000996 kB) free!
[1588344575] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (910512 kB) free!
[1588344620] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3961628 kB) free.
[1588345220] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2398712 kB) free!
[1588345341] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2609768 kB) free!
[1588345461] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3068804 kB) free!
[1588345775] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22665540 kB) free!
[1588346061] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3751164 kB) free.
[1588346375] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33976904 kB) free.
[1588346661] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2212748 kB) free!
[1588346781] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2424964 kB) free!
[1588346901] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (2685440 kB) free!
[1588346975] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9391552 kB) free!
[1588347095] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4883404 kB) free!
[1588347215] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (803496 kB) free!
[1588347474] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1588347501] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3569568 kB) free.
[1588348415] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20692788 kB) free!
[1588348544] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4914428 kB) free.
[1588348844] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4907204 kB) free!
[1588348904] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4912076 kB) free!
[1588348965] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4914988 kB) free.
[1588349015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33042920 kB) free.
[1588349301] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (1574224 kB) free!
[1588349421] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (1585692 kB) free!
[1588349541] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (1700056 kB) free!
[1588349615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7894792 kB) free!
[1588349735] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4194828 kB) free!
[1588349855] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32851508 kB) free.
[1588349864] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4907712 kB) free!
[1588349924] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4913624 kB) free!
[1588349984] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4915176 kB) free.
[1588350284] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4909092 kB) free!
[1588350332] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1693 processes
[1588350344] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4906128 kB) free!
[1588350404] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 20.0% (4905532 kB) free!
[1588350455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8796056 kB) free!
[1588350464] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 20.0% (4904624 kB) free!
[1588350575] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4279916 kB) free!
[1588350695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32571540 kB) free.
[1588350741] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3265688 kB) free.
[1588350932] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1792 processes
[1588351295] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4370264 kB) free!
[1588351342] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3154384 kB) free!
[1588351365] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4933688 kB) free.
[1588351415] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (759972 kB) free!
[1588351462] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3126248 kB) free!
[1588351535] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34320612 kB) free.
[1588351582] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3125492 kB) free!
[1588352135] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9986176 kB) free!
[1588352255] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6113856 kB) free!
[1588352375] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34084272 kB) free.
[1588352976] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10870768 kB) free!
[1588353096] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6985524 kB) free!
[1588353216] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33906752 kB) free.
[1588353382] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1508236 kB) free!
[1588353816] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10647756 kB) free!
[1588353936] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6872748 kB) free!
[1588353982] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.8% (5172128 kB) free.
[1588354056] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33785700 kB) free.
[1588354656] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9390076 kB) free!
[1588354776] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5695852 kB) free!
[1588354896] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32996360 kB) free.
[1588355496] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8406028 kB) free!
[1588355616] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4751772 kB) free!
[1588355736] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31172692 kB) free.
[1588355865] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4913960 kB) free!
[1588355925] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4914956 kB) free.
[1588356173] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 26303 MB (20% inode=99%):
[1588356294] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 26303 MB (20% inode=99%):
[1588356336] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8307160 kB) free!
[1588356414] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 26303 MB (20% inode=99%):
[1588356456] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33212292 kB) free.
[1588356825] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4907440 kB) free!
[1588356885] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4909992 kB) free!
[1588356945] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 20.0% (4914092 kB) free!
[1588357005] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 20.0% (4906356 kB) free!
[1588357056] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10841840 kB) free!
[1588357176] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7075376 kB) free!
[1588357297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33080100 kB) free.