[1552761014] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.83, 5.45, 5.31
[1552761262] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11918984 kB) free.
[1552761614] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.21, 4.34, 4.75
[1552761862] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4076220 kB) free!
[1552761919] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12019 MB (5% inode=99%):
[1552761982] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11897168 kB) free.
[1552762163] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552762582] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6474292 kB) free!
[1552762702] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5192332 kB) free!
[1552762763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552762823] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11913832 kB) free.
[1552762883] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552763119] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25739 MB (12% inode=99%):
[1552763363] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 47.7% (3505392 kB) free.
[1552763422] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7235008 kB) free!
[1552763436] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 334 processes
[1552763483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552763542] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5954620 kB) free!
[1552763603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552763662] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4417948 kB) free!
[1552763723] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552764262] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.0% (7882352 kB) free!
[1552764319] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17568 MB (8% inode=99%):
[1552764323] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552764443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552764564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552764862] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11918408 kB) free.
[1552764919] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 28901 MB (13% inode=99%):
[1552765214] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.30, 5.74, 5.07
[1552765334] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.66, 5.74, 5.16
[1552765454] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.36, 5.58, 5.16
[1552765462] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4282804 kB) free!
[1552765519] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17637 MB (8% inode=99%):
[1552765582] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11898600 kB) free.
[1552765764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552766182] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6451488 kB) free!
[1552766302] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5172476 kB) free!
[1552766364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552766422] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3570260 kB) free!
[1552766484] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552766719] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27562 MB (13% inode=99%):
[1552767084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552767204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552767319] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17664 MB (8% inode=99%):
[1552767623] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.3% (8040660 kB) free!
[1552767804] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552767854] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.53, 4.53, 4.95
[1552767925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552768045] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552768223] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11906708 kB) free.
[1552768519] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27653 MB (13% inode=99%):
[1552768645] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552768765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552768823] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4745968 kB) free!
[1552768885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552768943] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11890352 kB) free.
[1552769054] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.86, 7.06, 5.78
[1552769174] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.48, 7.81, 6.20
[1552769294] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.49, 8.86, 6.79
[1552769543] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3407516 kB) free!
[1552769663] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11443716 kB) free.
[1552769719] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20741 MB (9% inode=99%):
[1552769856] SERVICE ALERT: sig4-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (1439972 kB) free!
[1552769976] SERVICE ALERT: sig4-vm2;Memory;WARNING;SOFT;2;WARNING - 19.4% (1427944 kB) free!
[1552770097] SERVICE ALERT: sig4-vm2;Memory;WARNING;HARD;3;WARNING - 19.2% (1413064 kB) free!
[1552770863] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4069496 kB) free!
[1552770983] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11871540 kB) free.
[1552771519] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26928 MB (12% inode=99%):
[1552771583] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6426420 kB) free!
[1552771703] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5144372 kB) free!
[1552771824] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11856444 kB) free.
[1552771886] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552772119] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 16197 MB (7% inode=99%):
[1552772423] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7154020 kB) free!
[1552772486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552772543] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5874424 kB) free!
[1552772606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552772663] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4369284 kB) free!
[1552772726] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1552773106] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 505.04, 325.70, 143.42
[1552773226] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 508.08, 386.20, 187.69
[1552773263] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7933572 kB) free!
[1552773320] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26466 MB (12% inode=99%):
[1552773326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552773346] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 159.08, 320.75, 189.77
[1552773446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552773566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552773632] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5378 MB (10% inode=97%):
[1552773863] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11863808 kB) free.
[1552773920] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17312 MB (8% inode=99%):
[1552774094] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.53, 4.66, 4.92
[1552774463] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5103932 kB) free!
[1552774583] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11878736 kB) free.
[1552774767] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552775120] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 28561 MB (13% inode=99%):
[1552775183] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7208784 kB) free!
[1552775294] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.42, 5.94, 5.45
[1552775303] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5968688 kB) free!
[1552775367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552775414] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.02, 6.02, 5.54
[1552775423] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4533440 kB) free!
[1552775487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552775534] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.98, 6.05, 5.61
[1552775607] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552775720] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17457 MB (8% inode=99%):
[1552775746] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.46, 1.61, 15.77
[1552776023] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.4% (8079384 kB) free!
[1552776207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552776320] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26576 MB (12% inode=99%):
[1552776327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552776447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552776624] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11874328 kB) free.
[1552776734] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.15, 4.50, 4.99
[1552776920] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17396 MB (8% inode=99%):
[1552777047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552777168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552777224] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5174304 kB) free!
[1552777288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552777344] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11863608 kB) free.
[1552777520] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27030 MB (12% inode=99%):
[1552777546] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.62, 1.48, 3.52
[1552777944] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7267120 kB) free!
[1552778064] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (5997548 kB) free!
[1552778184] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4569804 kB) free!
[1552778488] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552778721] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20932 MB (10% inode=99%):
[1552778784] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.4% (8089856 kB) free!
[1552779088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552779208] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552779321] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29158 MB (13% inode=99%):
[1552779384] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11860380 kB) free.
[1552779808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552779921] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20966 MB (10% inode=99%):
[1552779928] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552779984] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5018696 kB) free!
[1552780049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552780104] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11880264 kB) free.
[1552780334] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.57, 6.10, 5.39
[1552780454] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.77, 6.31, 5.55
[1552780521] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27145 MB (13% inode=99%):
[1552780574] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.93, 6.40, 5.67
[1552780704] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7007324 kB) free!
[1552780825] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.6% (5734296 kB) free!
[1552780945] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4133056 kB) free!
[1552781249] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552781545] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7732700 kB) free!
[1552781849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552781969] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552782145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11877176 kB) free.
[1552782374] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 5.02, 4.58, 4.81
[1552782569] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552782689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552782745] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4429412 kB) free!
[1552782809] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552782865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11862860 kB) free.
[1552782921] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20999 MB (10% inode=99%):
[1552783410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552783465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6671704 kB) free!
[1552783521] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25116 MB (12% inode=99%):
[1552783530] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552783585] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.0% (5400848 kB) free!
[1552783650] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552783705] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3523876 kB) free!
[1552784305] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7431212 kB) free!
[1552784721] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21037 MB (10% inode=99%):
[1552784850] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552784905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11861016 kB) free.
[1552785321] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29246 MB (14% inode=99%):
[1552785374] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.29, 6.07, 5.52
[1552785450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552785495] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.94, 6.02, 5.57
[1552785506] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3921448 kB) free!
[1552785570] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552785615] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.77, 5.99, 5.62
[1552785626] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11872024 kB) free.
[1552785922] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17982 MB (8% inode=99%):
[1552786170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552786226] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6380936 kB) free!
[1552786291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552786346] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5111872 kB) free!
[1552786411] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552786466] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11862448 kB) free.
[1552786522] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25165 MB (12% inode=99%):
[1552787066] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7109156 kB) free!
[1552787186] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.8% (5815136 kB) free!
[1552787306] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4253848 kB) free!
[1552787415] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.84, 4.21, 4.72
[1552787611] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552787722] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19366 MB (9% inode=99%):
[1552787906] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7845144 kB) free!
[1552788211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552788322] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25223 MB (12% inode=99%):
[1552788331] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1552788506] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11860480 kB) free.
[1552788931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552789051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552789106] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4426656 kB) free!
[1552789171] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552789226] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11869276 kB) free.
[1552789523] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17029 MB (8% inode=99%):
[1552789772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552789815] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.62, 5.54, 5.11
[1552789827] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6592808 kB) free!
[1552789892] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552789935] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.08, 5.67, 5.21
[1552789947] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.8% (5309024 kB) free!
[1552790012] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552790055] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.29, 5.88, 5.34
[1552790067] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3362140 kB) free!
[1552790667] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7265220 kB) free!
[1552790723] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 28326 MB (13% inode=99%):
[1552791212] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552791267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11872624 kB) free.
[1552791323] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17062 MB (8% inode=99%):
[1552791812] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552791867] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3660392 kB) free!
[1552791932] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552791987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11852904 kB) free.
[1552792532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552792587] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6235512 kB) free!
[1552792653] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552792707] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4941432 kB) free!
[1552792773] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552792827] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11866732 kB) free.
[1552793055] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.70, 4.83, 4.97
[1552793123] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25232 MB (12% inode=99%):
[1552793427] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6986716 kB) free!
[1552793548] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.5% (5691164 kB) free!
[1552793668] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4041520 kB) free!
[1552794255] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.42, 5.93, 5.48
[1552794268] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5100904 kB) free!
[1552794323] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17048 MB (8% inode=99%):
[1552794375] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.35, 5.70, 5.46
[1552794495] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.42, 5.70, 5.49
[1552795523] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26717 MB (12% inode=99%):
[1552796723] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18971 MB (9% inode=99%):
[1552797268] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10448464 kB) free.
[1552797323] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24751 MB (11% inode=99%):
[1552798134] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (703540 kB) free!
[1552798468] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3137140 kB) free!
[1552798524] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 16339 MB (7% inode=99%):
[1552798588] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11241504 kB) free.
[1552798780] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (691812 kB) free!
[1552799724] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24179 MB (11% inode=99%):
[1552799788] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3726256 kB) free!
[1552799908] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604464 kB) free.
[1552800924] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15739 MB (7% inode=99%):
[1552801108] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4055196 kB) free!
[1552801228] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11760584 kB) free.
[1552801613] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47174 MB (20% inode=99%):
[1552801733] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47156 MB (20% inode=99%):
[1552801828] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3145676 kB) free!
[1552801853] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47110 MB (20% inode=99%):
[1552801948] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11014428 kB) free.
[1552802124] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25525 MB (12% inode=99%):
[1552803148] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3501852 kB) free!
[1552803252] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 20.0% (1470860 kB) free!
[1552803268] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11223504 kB) free.
[1552803324] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15685 MB (7% inode=99%):
[1552803373] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 20.0% (1470132 kB) free!
[1552803493] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 20.0% (1468272 kB) free!
[1552803573] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552804469] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3636724 kB) free!
[1552804524] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24963 MB (11% inode=99%):
[1552804589] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11324912 kB) free.
[1552805295] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552805725] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14682 MB (7% inode=99%):
[1552805789] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3820284 kB) free!
[1552805909] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11462568 kB) free.
[1552807109] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4035468 kB) free!
[1552807174] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552807229] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1552807229] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11588708 kB) free.
[1552807774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552807894] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552808014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552808429] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4133924 kB) free!
[1552808549] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11638544 kB) free.
[1552808725] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25147 MB (12% inode=99%):
[1552809150] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8705924 kB) free!
[1552809270] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10833896 kB) free.
[1552809814] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552809925] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17757 MB (8% inode=99%):
[1552810414] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1552810470] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3126016 kB) free!
[1552810534] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552810574] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181540 kB) free!
[1552810590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11725372 kB) free.
[1552811125] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23454 MB (11% inode=99%):
[1552811135] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552811190] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3983700 kB) free!
[1552811255] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1552811310] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11711536 kB) free.
[1552811855] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552811910] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6308824 kB) free!
[1552811975] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552812030] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.1% (4988548 kB) free!
[1552812095] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1552812095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552812150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11725320 kB) free.
[1552812325] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15262 MB (7% inode=99%):
[1552812750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5816616 kB) free!
[1552812871] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4342956 kB) free!
[1552812991] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11705508 kB) free.
[1552813591] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5307204 kB) free!
[1552813711] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (3425836 kB) free!
[1552813831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11717292 kB) free.
[1552813895] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552814125] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25506 MB (12% inode=99%):
[1552814431] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5919132 kB) free!
[1552814496] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552814551] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4503908 kB) free!
[1552814616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552814671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11706168 kB) free.
[1552814736] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552815271] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6100820 kB) free!
[1552815325] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15267 MB (7% inode=99%):
[1552815336] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552815391] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4779032 kB) free!
[1552815456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552815511] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11718724 kB) free.
[1552815576] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552816111] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.7% (5291464 kB) free!
[1552816176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552816232] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3383128 kB) free!
[1552816296] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552816352] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11021068 kB) free.
[1552816525] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23461 MB (11% inode=99%):
[1552817496] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552817552] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6047008 kB) free!
[1552817616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552817672] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (4649288 kB) free!
[1552817725] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15260 MB (7% inode=99%):
[1552817736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552817792] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11696980 kB) free.
[1552818392] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3342488 kB) free!
[1552818512] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11178396 kB) free.
[1552819712] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4237848 kB) free!
[1552819832] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11705104 kB) free.
[1552820433] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3771048 kB) free!
[1552820553] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11694644 kB) free.
[1552820725] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26107 MB (12% inode=99%):
[1552821153] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5101184 kB) free!
[1552821273] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (3188140 kB) free!
[1552821336] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552821393] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11702456 kB) free.
[1552821936] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552821993] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5871436 kB) free!
[1552822056] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552822113] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4476712 kB) free!
[1552822176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552822233] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11683552 kB) free.
[1552822525] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18939 MB (9% inode=99%):
[1552822833] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3141452 kB) free!
[1552822953] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10844080 kB) free.
[1552824154] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5559760 kB) free!
[1552824274] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (3915660 kB) free!
[1552824394] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11484056 kB) free.
[1552825525] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26120 MB (12% inode=99%):
[1552825594] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4122744 kB) free!
[1552825714] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11693768 kB) free.
[1552825776] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552826314] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4563304 kB) free!
[1552826376] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552826434] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11668596 kB) free.
[1552826496] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552826616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552826725] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18951 MB (9% inode=99%):
[1552827034] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3597432 kB) free!
[1552827154] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11399724 kB) free.
[1552827754] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3038328 kB) free!
[1552827816] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552827874] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10934864 kB) free.
[1552827926] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26123 MB (12% inode=99%):
[1552828416] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552828537] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552828657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552829074] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3000504 kB) free!
[1552829194] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10771180 kB) free.
[1552829726] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18955 MB (9% inode=99%):
[1552830994] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3601252 kB) free!
[1552831057] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552831114] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11296724 kB) free.
[1552831657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552831734] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 18.2% (1336576 kB) free!
[1552831777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552831854] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (388796 kB) free!
[1552831898] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552831974] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (387588 kB) free!
[1552832314] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3269680 kB) free!
[1552832434] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11116276 kB) free.
[1552833326] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26119 MB (12% inode=99%):
[1552833634] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3851824 kB) free!
[1552833698] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552833755] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11639136 kB) free.
[1552833926] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18951 MB (9% inode=99%):
[1552834298] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552834354] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4150172 kB) free!
[1552834418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552834474] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11624216 kB) free.
[1552834526] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24582 MB (11% inode=99%):
[1552834538] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552835074] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5005380 kB) free!
[1552835126] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18438 MB (8% inode=99%):
[1552835138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552835194] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (3133292 kB) free!
[1552835258] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552835314] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11637808 kB) free.
[1552835859] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552835914] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5916232 kB) free!
[1552835979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552836034] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4703440 kB) free!
[1552836099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552836154] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11621424 kB) free.
[1552836754] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6674512 kB) free!
[1552836874] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5370280 kB) free!
[1552836994] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3758328 kB) free!
[1552837299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552837595] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.6% (7664228 kB) free!
[1552837899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552838019] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552838195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11615304 kB) free.
[1552838619] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552838726] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26738 MB (12% inode=99%):
[1552838740] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552838795] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4316156 kB) free!
[1552838860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552838896] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.53, 4.84, 4.98
[1552838915] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11634348 kB) free.
[1552839326] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14703 MB (7% inode=99%):
[1552839496] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.51, 5.27, 5.12
[1552839515] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6362280 kB) free!
[1552839616] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.19, 4.88, 5.00
[1552839635] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.2% (5044184 kB) free!
[1552839755] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1552839755] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.5% (6648104 kB) free!
[1552840216] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.69, 6.29, 5.65
[1552840336] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.02, 6.54, 5.81
[1552840456] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.42, 6.46, 5.87
[1552840526] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27012 MB (12% inode=99%):
[1552840660] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552840955] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11506540 kB) free.
[1552841126] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14715 MB (7% inode=99%):
[1552841260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552841380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552841500] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552842155] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6843580 kB) free!
[1552842276] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5507972 kB) free!
[1552842396] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3780852 kB) free!
[1552842700] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552842856] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.33, 4.64, 5.00
[1552842996] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7428232 kB) free!
[1552843300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552843420] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552843596] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11609100 kB) free.
[1552844020] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552844140] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552844196] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3709988 kB) free!
[1552844260] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552844316] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11628064 kB) free.
[1552844656] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.29, 5.22, 5.08
[1552844776] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.06, 5.40, 5.16
[1552844860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552844896] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1552844896] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.67, 5.89, 5.37
[1552844916] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5976052 kB) free!
[1552844981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552845036] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4594740 kB) free!
[1552845101] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552845156] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11613244 kB) free.
[1552845326] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26479 MB (12% inode=99%):
[1552845756] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6479932 kB) free!
[1552845876] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5150684 kB) free!
[1552845926] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 16014 MB (7% inode=99%):
[1552845997] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3203120 kB) free!
[1552846301] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552846597] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7043244 kB) free!
[1552846901] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552847021] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552847197] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604680 kB) free.