[1568226713] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.15, 5.31, 5.13
[1568226733] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1568226833] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.89, 4.76, 4.95
[1568226944] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568227433] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.24, 5.66, 5.31
[1568227553] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.86, 6.06, 5.49
[1568227673] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.36, 6.48, 5.71
[1568227787] SERVICE FLAPPING ALERT: secondary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568229348] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 655 processes
[1568229464] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 149.81, 118.46, 59.56
[1568229468] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 414 processes
[1568229584] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 149.72, 128.79, 70.49
[1568229704] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 21.27, 86.60, 62.10
[1568229936] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2669240 kB) free!
[1568230536] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2200204 kB) free!
[1568230906] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.04, 2.28, 17.53
[1568232706] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.66, 0.60, 3.11
[1568235663] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233152 kB) free!
[1568235783] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242536 kB) free!
[1568235903] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (400224 kB) free!
[1568236307] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1568236428] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1568236504] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.2% (42482292 kB) free.
[1568238343] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 46.9% (11487996 kB) free.
[1568238843] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 60.1% (14735864 kB) free.
[1568241473] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.07, 4.41, 4.97
[1568241827] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568241947] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1568242073] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.86, 5.01, 5.02
[1568242193] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.36, 4.58, 4.85
[1568243393] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.83, 5.37, 5.24
[1568243513] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.99, 5.75, 5.40
[1568243633] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.73, 5.63, 5.39
[1568247833] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.23, 4.39, 4.75
[1568250233] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.40, 11.57, 8.08
[1568250353] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.45, 9.04, 7.57
[1568250473] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1568250473] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.47, 6.80, 6.93
[1568250844] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182320 kB) free!
[1568250944] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190052 kB) free!
[1568250966] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187764 kB) free!
[1568251064] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189964 kB) free!
[1568251087] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172284 kB) free!
[1568251184] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176224 kB) free!
[1568251673] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.31, 2.90, 4.17
[1568252397] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 453 processes
[1568252457] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 455 processes
[1568260547] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568260667] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568261744] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4885644 kB) free!
[1568261864] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.8% (4856232 kB) free!
[1568261985] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.7% (4830872 kB) free!
[1568262467] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1568262473] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568262587] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1568264506] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 40.16, 35.03, 19.91
[1568264626] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 39.65, 36.52, 22.29
[1568264631] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568264746] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.53, 34.52, 23.40
[1568264751] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1568264871] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1568265346] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.91, 5.35, 12.65
[1568266071] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1568266296] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 26.8% (6623532 kB) free.
[1568266333] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (9467780 kB) free!
[1568266546] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 0.98, 4.19
[1568266896] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.4% (3564076 kB) free!
[1568267016] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (774604 kB) free!
[1568267136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186932 kB) free!
[1568267533] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224096 kB) free!
[1568267633] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.9% (4879772 kB) free!
[1568267753] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.9% (4877948 kB) free!
[1568267873] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.9% (4876216 kB) free!
[1568267873] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 166.79, 132.38, 67.47
[1568267993] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 168.38, 144.28, 79.69
[1568268114] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 24.79, 97.22, 70.27
[1568269071] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568269191] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1568269311] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1568270181] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 34.96, 19.06, 8.67
[1568270301] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.95, 24.71, 12.00
[1568270421] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.31, 28.51, 14.93
[1568270514] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.76, 4.43, 17.44
[1568270887] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (2516344 kB) free!
[1568270947] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (201196 kB) free!
[1568271021] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.47, 9.86, 13.13
[1568271487] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (1062100 kB) free!
[1568272221] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.83, 1.95, 4.96
[1568272314] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.80, 1.83, 4.43
[1568273385] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2399316 kB) free!
[1568273704] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8787736 kB) free!
[1568273824] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8537772 kB) free!
[1568273944] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.9% (8325108 kB) free!
[1568274357] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 525 processes
[1568274417] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 527 processes
[1568275744] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4425232 kB) free!
[1568276502] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1568276623] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1568277114] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.96, 5.50, 5.10
[1568277234] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.50, 5.03, 4.97
[1568277835] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.67, 6.20, 5.62
[1568277955] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.66, 6.46, 5.79
[1568278075] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1568278075] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.25, 6.70, 5.95
[1568279511] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1568279519] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 93.9% (23013144 kB) free.
[1568279556] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21978864 kB) free.
[1568279587] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.3% (21155676 kB) free.
[1568279629] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 61.0% (30052408 kB) free.
[1568279658] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 29.9% (14716000 kB) free.
[1568279693] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.4% (20675592 kB) free.
[1568279719] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.5% (14545636 kB) free.
[1568279875] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.23, 3.54, 4.62
[1568281362] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (4074152 kB) free!
[1568281482] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.2% (2508124 kB) free!
[1568281602] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1636272 kB) free!
[1568281787] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1568281908] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1568282633] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4757632 kB) free!
[1568282753] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (2396628 kB) free!
[1568282873] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.7% (2386844 kB) free!
[1568282874] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1309180 kB) free!
[1568287742] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.0% (18005092 kB) free.
[1568287773] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.4% (35200672 kB) free.
[1568287802] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 93.5% (22923436 kB) free.
[1568289546] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4881396 kB) free!
[1568289666] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (686752 kB) free!
[1568289787] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164580 kB) free!
[1568290075] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1568291973] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8182228 kB) free!
[1568292093] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.3% (7056264 kB) free!
[1568292214] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.0% (6879736 kB) free!
[1568293414] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1568293414] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3531232 kB) free!
[1568293467] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8899716 kB) free!
[1568293494] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (4308852 kB) free!
[1568293548] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1568293587] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8577228 kB) free!
[1568293614] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (4167984 kB) free!
[1568293707] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8271212 kB) free!
[1568293734] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (4038196 kB) free!
[1568294014] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23396580 kB) free!
[1568294134] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 34.5% (45521408 kB) free.
[1568294595] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4649828 kB) free!
[1568294715] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4542148 kB) free!
[1568294734] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25308256 kB) free!
[1568294835] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4429152 kB) free!
[1568294854] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 16.1% (21209416 kB) free!
[1568294974] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 32.9% (43408216 kB) free.
[1568295507] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3901508 kB) free!
[1568295534] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2135740 kB) free!
[1568295574] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22375492 kB) free!
[1568295694] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 33.8% (44543956 kB) free.
[1568296294] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23693716 kB) free!
[1568296414] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19437172 kB) free!
[1568296534] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.8% (41921148 kB) free.
[1568296907] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568297027] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568297134] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20954140 kB) free!
[1568297235] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2021288 kB) free!
[1568297254] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 33.0% (43542872 kB) free.
[1568297854] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22626248 kB) free!
[1568297974] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18425992 kB) free!
[1568298094] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.3% (41214652 kB) free.
[1568298695] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20744320 kB) free!
[1568298815] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (16696412 kB) free!
[1568298935] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 30.1% (39661836 kB) free.
[1568299535] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18404708 kB) free!
[1568299655] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 31.6% (41717580 kB) free.
[1568300255] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20254940 kB) free!
[1568300375] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15942840 kB) free!
[1568300495] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 29.8% (39242368 kB) free.
[1568300747] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45630 MB (20% inode=99%):
[1568300867] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45415 MB (20% inode=99%):
[1568300987] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44639 MB (19% inode=99%):
[1568301095] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18198628 kB) free!
[1568301215] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13924464 kB) free!
[1568301335] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 28.4% (37490744 kB) free.
[1568301667] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.6% (35269200 kB) free.
[1568301708] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.1% (22824572 kB) free.
[1568301742] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.2% (14866688 kB) free.
[1568301773] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 79.8% (19689452 kB) free.
[1568301798] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 83.7% (20523232 kB) free.
[1568301935] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16562584 kB) free!
[1568302055] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 30.1% (39709220 kB) free.
[1568302198] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1568302221] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 73.33, 18.97, 7.34
[1568302318] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1568302341] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 169.08, 73.60, 28.77
[1568302461] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 181.72, 110.08, 47.58
[1568302655] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19722484 kB) free!
[1568302775] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15706236 kB) free!
[1568302895] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 29.9% (39352296 kB) free.
[1568303496] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20013100 kB) free!
[1568303573] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.8% (3642944 kB) free!
[1568303616] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16034588 kB) free!
[1568303693] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (1168948 kB) free!
[1568303736] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 23.7% (31209664 kB) free.
[1568303813] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174864 kB) free!
[1568304261] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.08, 2.32, 16.87
[1568304336] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20595720 kB) free!
[1568304346] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 139.52, 95.46, 43.75
[1568304456] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (16647332 kB) free!
[1568304466] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 139.80, 110.16, 55.41
[1568304576] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12709424 kB) free!
[1568304586] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 139.69, 119.89, 65.61
[1568305176] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21274624 kB) free!
[1568305776] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 22.6% (29804656 kB) free.
[1568305876] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8464264 kB) free!
[1568305996] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.4% (8074536 kB) free!
[1568306061] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.05, 1.23, 3.53
[1568306116] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.5% (7629024 kB) free!
[1568306386] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.77, 1.30, 11.30
[1568306976] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18505784 kB) free!
[1568307096] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (14558568 kB) free!
[1568307216] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 29.4% (38769100 kB) free.
[1568307227] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568307347] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568307586] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.32, 0.97, 3.74
[1568307816] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19362528 kB) free!
[1568307916] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4374676 kB) free!
[1568307936] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15360564 kB) free!
[1568308056] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 24.1% (31781440 kB) free.
[1568308656] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20169420 kB) free!
[1568308777] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16245888 kB) free!
[1568308897] SERVICE FLAPPING ALERT: rq-2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1568308897] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12330192 kB) free!
[1568309497] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21250800 kB) free!
[1568310097] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 22.8% (30061724 kB) free.
[1568311297] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18830456 kB) free!
[1568311417] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14856140 kB) free!
[1568311537] SERVICE ALERT: rq-2;Memory;OK;HARD;1;OK - 29.2% (38452480 kB) free.
[1568312137] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19754564 kB) free!
[1568312257] SERVICE ALERT: rq-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15825508 kB) free!
[1568312377] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11875072 kB) free!
[1568312977] SERVICE ALERT: rq-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20777064 kB) free!