[1558551670] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8605392 kB) free!
[1558551791] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.3% (13954308 kB) free.
[1558552628] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8963812 kB) free!
[1558552748] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.2% (14401132 kB) free.
[1558554190] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8600064 kB) free!
[1558554310] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.3% (13923836 kB) free.
[1558555209] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1620596 kB) free!
[1558556110] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8960340 kB) free!
[1558556230] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.7% (14647276 kB) free.
[1558556659] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 141.63, 37.76, 13.95
[1558556779] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 269.01, 120.82, 47.27
[1558556846] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 233.68, 84.13, 30.70
[1558556899] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 287.04, 176.70, 76.66
[1558556966] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 276.14, 149.81, 61.25
[1558557086] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 281.47, 193.65, 88.05
[1558558148] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9132748 kB) free!
[1558558248] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558558268] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.1% (14360744 kB) free.
[1558558630] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9345896 kB) free!
[1558558751] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.5% (14516316 kB) free.
[1558558886] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.55, 1.01, 15.25
[1558559299] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.90, 2.16, 16.16
[1558560048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558560068] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7757192 kB) free!
[1558560087] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.63, 0.49, 4.54
[1558560168] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558560188] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.1% (14360196 kB) free.
[1558560534] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1558560654] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1558560769] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558560788] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8764936 kB) free!
[1558560889] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558560908] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.1% (14348732 kB) free.
[1558561099] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.77, 1.82, 3.89
[1558561150] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8972928 kB) free!
[1558561270] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.6% (14116700 kB) free.
[1558561489] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558561508] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9682420 kB) free!
[1558561609] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558561629] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.9% (14244740 kB) free.
[1558561729] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558562329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558562449] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1558563049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1558563070] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9736748 kB) free!
[1558563170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558563190] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.6% (14573108 kB) free.
[1558563290] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558564490] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558565591] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9424252 kB) free!
[1558565690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558565711] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.2% (14397372 kB) free.
[1558565810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558565930] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558566531] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558566651] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1558567251] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1558567371] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558567971] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558568091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558568211] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558568812] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558568932] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558569052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1558570511] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8417480 kB) free!
[1558570631] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.3% (13967912 kB) free.
[1558571452] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558571828] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9237368 kB) free!
[1558571948] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (6994816 kB) free!
[1558571973] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 597 processes
[1558572033] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 596 processes
[1558572052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1558572075] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.6% (13605816 kB) free.
[1558572173] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558572260] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1631688 kB) free!
[1558572293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558572432] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9216088 kB) free!
[1558572552] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.4% (13983288 kB) free.
[1558574152] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1558574272] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 362 processes
[1558574393] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1558574693] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1558574710] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1558574716] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1558574899] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 228.83, 77.77, 28.83
[1558574952] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9024100 kB) free!
[1558575019] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 285.29, 149.28, 60.97
[1558575072] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.3% (13969448 kB) free.
[1558575075] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9203368 kB) free!
[1558575139] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 292.99, 197.23, 89.25
[1558575195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.6% (14070848 kB) free.
[1558575293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558575414] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558575534] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558576872] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9252840 kB) free!
[1558576992] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.3% (13948212 kB) free.
[1558577174] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1558577294] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1558577334] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1558577492] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15.52, 109.12, 76.83
[1558577539] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.56, 1.73, 15.14
[1558577612] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.42, 73.14, 67.56
[1558577660] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2784812 kB) free!
[1558577732] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.51, 49.03, 59.40
[1558577934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558578054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558578174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558578793] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9782396 kB) free!
[1558578795] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9060864 kB) free!
[1558578913] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8046700 kB) free!
[1558578916] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.5% (14038152 kB) free.
[1558578932] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.18, 1.12, 16.53
[1558579033] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.0% (13301260 kB) free.
[1558579339] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.81, 1.84, 3.76
[1558579460] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461868 kB) free!
[1558580132] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.45, 4.85
[1558580574] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558580833] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8149040 kB) free!
[1558580953] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.4% (13493712 kB) free.
[1558581174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558581294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558581414] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558582515] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9004156 kB) free!
[1558582635] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6691392 kB) free!
[1558582753] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8388276 kB) free!
[1558582755] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13420384 kB) free.
[1558582873] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.7% (13637448 kB) free.
[1558582926] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 112.31, 121.15, 54.06
[1558583046] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.34, 81.15, 47.54
[1558583167] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.28, 54.41, 41.82
[1558583214] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558583814] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558583907] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4901184 kB) free!
[1558583934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558584027] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4798068 kB) free!
[1558584054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1558584147] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (4692152 kB) free!
[1558584367] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 1.34, 11.73
[1558584674] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8117048 kB) free!
[1558584794] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6293112 kB) free!
[1558584914] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.1% (12381160 kB) free.
[1558585156] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9157336 kB) free!
[1558585276] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6952188 kB) free!
[1558585396] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.7% (13654012 kB) free.
[1558585567] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.62, 3.59
[1558586114] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8062648 kB) free!
[1558586234] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6086384 kB) free!
[1558586354] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 36.4% (17939096 kB) free.
[1558586455] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558586856] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46974 MB (20% inode=99%):
[1558586977] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46942 MB (20% inode=99%):
[1558587055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1558587097] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46869 MB (20% inode=99%):
[1558587147] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198436 kB) free!
[1558587175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558587295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558587646] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.1% (14804932 kB) free.
[1558587684] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.3% (22877632 kB) free.
[1558587738] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.8% (14913900 kB) free.
[1558587789] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.9% (23172736 kB) free.
[1558588057] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178832 kB) free!
[1558588177] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172884 kB) free!
[1558588297] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186972 kB) free!
[1558588347] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (194856 kB) free!
[1558588396] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9781908 kB) free!
[1558588467] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (193016 kB) free!
[1558588516] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8238308 kB) free!
[1558588588] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (192700 kB) free!
[1558588636] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.9% (5863456 kB) free!
[1558589236] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.3% (13460076 kB) free.
[1558589354] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8451140 kB) free!
[1558589397] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 363 processes
[1558589474] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 32.2% (15852752 kB) free.
[1558589497] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 96844 MB (43% inode=99%):
[1558589517] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1558590295] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558590717] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1558590739] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2825044 kB) free!
[1558590837] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558590859] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.7% (2557464 kB) free!
[1558590979] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.5% (2358968 kB) free!
[1558591181] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.37, 5.22, 5.20
[1558591301] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.18, 4.57, 4.97
[1558591495] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558591516] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558591579] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1445548 kB) free!
[1558591615] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558591636] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558591735] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558591756] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558591889] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182616 kB) free!
[1558591995] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191980 kB) free!
[1558592010] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170452 kB) free!
[1558592115] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194328 kB) free!
[1558592130] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1558592130] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196736 kB) free!
[1558592235] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187640 kB) free!
[1558592236] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9440864 kB) free!
[1558592356] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7682228 kB) free!
[1558592474] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9017024 kB) free!
[1558592476] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5288004 kB) free!
[1558592501] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.30, 5.23, 5.10
[1558592594] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6945152 kB) free!
[1558592621] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.86, 5.26, 5.14
[1558592714] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.4% (13996708 kB) free.
[1558592741] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.22, 5.24, 5.14
[1558593050] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (2997268 kB) free!
[1558593076] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1558593076] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12933924 kB) free.
[1558593170] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (2797808 kB) free!
[1558593290] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.0% (2597312 kB) free!
[1558593342] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.56, 4.68, 4.89
[1558593380] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 80.5% (13084012 kB) free.
[1558593556] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558593890] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1311460 kB) free!
[1558593910] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1558593932] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 19.49, 4.86, 2.11
[1558593942] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.52, 5.53, 5.04
[1558594052] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.37, 6.37, 3.14
[1558594062] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.25, 5.38, 5.05
[1558594182] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.10, 5.38, 5.10
[1558594437] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 5 crit, 0 warn out of 359 processes
[1558594557] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558594735] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558595114] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8961196 kB) free!
[1558595234] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6854020 kB) free!
[1558595335] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558595354] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4871228 kB) free!
[1558595455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558595476] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9838816 kB) free!
[1558595575] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558595596] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8120152 kB) free!
[1558595710] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1558595716] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6233032 kB) free!
[1558595955] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (9280012 kB) free!
[1558596316] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4718064 kB) free!
[1558596555] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.7% (11690684 kB) free.
[1558596917] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12345856 kB) free.
[1558597557] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 362 processes
[1558597677] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558598117] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7218656 kB) free!
[1558598237] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.0% (5430344 kB) free!
[1558598357] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3687332 kB) free!
[1558599270] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (594728 kB) free!
[1558599330] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1558599390] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (312024 kB) free!
[1558599510] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (124892 kB) free!
[1558600757] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.1% (9425408 kB) free!
[1558601180] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 12.3% (2005688 kB) free!
[1558601300] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 10.4% (1684492 kB) free!
[1558601420] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1369816 kB) free!
[1558601877] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 361 processes
[1558601997] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558602333] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 587 processes
[1558602339] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.0% (4177980 kB) free!
[1558602393] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 583 processes
[1558602399] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.7% (4090368 kB) free!
[1558602459] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.7% (4092700 kB) free!
[1558602519] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.7% (4086436 kB) free!
[1558603797] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 366 processes
[1558603917] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1558603939] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.50, 15.13, 9.36
[1558604059] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.24, 10.81, 8.48
[1558604179] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.09, 7.95, 7.72
[1558604310] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.5% (6876380 kB) free.
[1558604356] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9605140 kB) free!
[1558604476] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7480884 kB) free!
[1558604596] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (5430788 kB) free!
[1558604779] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.37, 2.58, 4.91
[1558605557] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4471604 kB) free!
[1558605582] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.19, 4.12, 4.58
[1558605717] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 364 processes
[1558605837] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1558605957] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1558605979] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.08, 6.71, 8.01
[1558606099] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.43, 5.25, 7.31
[1558606182] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.97, 5.73, 5.23
[1558606219] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.40, 4.33, 6.73
[1558606302] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.88, 6.24, 5.48
[1558606375] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558606396] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1558606396] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4197832 kB) free!
[1558606422] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.12, 6.95, 5.83
[1558606819] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1558606819] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.86, 2.38, 4.56
[1558606996] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.4% (10044924 kB) free.
[1558607575] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558607696] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558607816] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558608357] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 369 processes
[1558608477] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1558608796] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9839068 kB) free!
[1558608822] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.35, 4.16, 4.91
[1558608916] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7839640 kB) free!
[1558609036] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 12.1% (5951284 kB) free!
[1558609078] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 374 processes
[1558609198] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 3 warn out of 381 processes
[1558609297] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46869 MB (20% inode=99%):
[1558609318] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 2 warn out of 378 processes
[1558609417] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46699 MB (20% inode=99%):
[1558609537] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46171 MB (20% inode=99%):
[1558609918] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 375 processes
[1558609993] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.17, 6.94, 5.04
[1558610113] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.23, 6.68, 5.18
[1558610233] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.44, 6.34, 5.24
[1558610518] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 370 processes
[1558610833] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.50, 3.16, 3.90
[1558611118] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1558611118] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1558611557] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.9% (8805816 kB) free!
[1558612110] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.9% (1092616 kB) free!
[1558612156] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1558612231] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.5% (993996 kB) free!
[1558612277] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1558612318] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 369 processes
[1558612351] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.4% (911320 kB) free!
[1558612397] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1558612438] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1558612951] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (323212 kB) free!
[1558613639] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 365 processes
[1558613759] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 366 processes
[1558613879] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 366 processes
[1558614479] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1558614620] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 77.2% (12545340 kB) free.
[1558614751] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1558614751] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.9% (6906844 kB) free.
[1558616084] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.2% (23492984 kB) free.
[1558616125] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.9% (23252784 kB) free.
[1558616163] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.6% (14893680 kB) free.
[1558616200] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.7% (22470436 kB) free.
[1558616216] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558616237] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.8% (14918552 kB) free.
[1558616850] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (177884 kB) free!
[1558616879] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 362 processes
[1558616970] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161652 kB) free!
[1558616999] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558617090] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159864 kB) free!
[1558617416] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558617536] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558617557] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1558617599] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 361 processes
[1558617656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558617720] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1558618219] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558619096] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.8% (14919952 kB) free.
[1558619236] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1558619412] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558619519] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1558619532] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558619639] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1558619653] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558619701] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173508 kB) free!
[1558619710] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1558619759] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1558619768] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169056 kB) free!
[1558619821] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169072 kB) free!
[1558619888] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (170868 kB) free!
[1558619941] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1558619941] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173008 kB) free!
[1558620008] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (183888 kB) free!
[1558620478] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558620959] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558622016] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (553192 kB) free!
[1558622136] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (685376 kB) free!
[1558622208] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (3726224 kB) free!
[1558622256] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (796904 kB) free!
[1558622328] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.1% (3702312 kB) free!
[1558622448] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (3647228 kB) free!
[1558622767] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 284.03, 186.15, 82.39
[1558622887] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 285.28, 219.02, 107.02
[1558622957] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4421900 kB) free!
[1558623007] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 42.27, 149.93, 95.45
[1558623290] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189732 kB) free!
[1558623332] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187796 kB) free!
[1558623410] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194384 kB) free!
[1558623452] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197440 kB) free!
[1558623530] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177720 kB) free!
[1558623572] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194668 kB) free!
[1558623648] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172680 kB) free!
[1558624807] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.43, 0.93, 14.29
[1558624951] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1558625836] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243352 kB) free!
[1558626607] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.19, 3.06, 3.95
[1558626656] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558628456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558628576] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558628697] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558628957] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8792236 kB) free!
[1558630057] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (2474908 kB) free!
[1558630741] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558631278] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1558631398] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1558631518] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1558631759] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558634697] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558635119] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558635897] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558636017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558636137] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes