[1583260295] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34720200 kB) free.
[1583260313] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583260347] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.8% (29448292 kB) free.
[1583260659] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1583260692] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2020/03/04 00:01:35 [error] 12257#0: *64405824 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583260779] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1583260812] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583260845] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(19) < 172.105.59.173 [04/Mar/2020:00:03:33 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583260896] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12379900 kB) free!
[1583260965] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583261016] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34410812 kB) free.
[1583261372] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1583261616] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12178300 kB) free!
[1583261736] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33642104 kB) free.
[1583262200] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (381080 kB) free!
[1583262260] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163616 kB) free!
[1583262321] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (180540 kB) free!
[1583262381] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (176328 kB) free!
[1583262572] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583262579] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1583262664] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.7% (3125576 kB) free!
[1583262681] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.1% (16932444 kB) free.
[1583262692] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1583262699] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1583262765] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(41) < 103.16.101.72 [04/Mar/2020:00:36:10 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1583262885] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583262936] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13061316 kB) free!
[1583263056] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33881804 kB) free.
[1583263264] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194228 kB) free!
[1583263292] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1583263334] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1583263412] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583263532] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1583263864] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1583263864] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 36.6% (9026488 kB) free.
[1583264031] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4050088 kB) free!
[1583264256] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13739344 kB) free!
[1583264376] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34269036 kB) free.
[1583264464] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (382344 kB) free!
[1583264499] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1583264584] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182572 kB) free!
[1583264620] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1583264705] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194856 kB) free!
[1583264976] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26314848 kB) free!
[1583265096] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31700168 kB) free.
[1583266896] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14542696 kB) free!
[1583267016] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33674932 kB) free.
[1583268713] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.4% (10137912 kB) free.
[1583268816] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14618312 kB) free!
[1583268936] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33052664 kB) free.
[1583269211] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (225444 kB) free!
[1583269331] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217608 kB) free!
[1583269451] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234868 kB) free!
[1583270051] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.9% (41824284 kB) free.
[1583270136] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (4970876 kB) free!
[1583270256] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30531560 kB) free.
[1583271012] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(66) < 2020/03/04 03:00:12 [error] 12256#0: *64470160 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583271071] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(256) < 103.16.101.72 [04/Mar/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583271132] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(237) < 2020/03/04 03:01:32 [error] 12247#0: *64470739 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583271193] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583271252] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583271285] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4445) < 203.212.70.170 [04/Mar/2020:03:02:32 +0530] TCP 502 0 0 120.001 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1583271405] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583271456] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5997852 kB) free!
[1583271576] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30195500 kB) free.
[1583271932] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583273376] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13181304 kB) free!
[1583273496] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33420032 kB) free.
[1583276105] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583278313] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192820 kB) free!
[1583278419] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1583278434] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (159640 kB) free!
[1583278539] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1583278554] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (161200 kB) free!
[1583279533] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 46.1% (11305952 kB) free.
[1583279706] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.9% (12065480 kB) free.
[1583282929] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (26658532 kB) free.
[1583284897] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25294924 kB) free!
[1583284916] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22003508 kB) free.
[1583284943] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 46.5% (22924128 kB) free.
[1583285017] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21546108 kB) free!
[1583285137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35309476 kB) free.
[1583286906] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1633272 kB) free!
[1583287026] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (1515176 kB) free!
[1583287146] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1514840 kB) free!
[1583289337] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25497724 kB) free!
[1583289405] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [04/Mar/2020:08:03:07 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1583289457] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.4% (21609440 kB) free!
[1583289525] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583289577] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35868868 kB) free.
[1583289939] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1583290059] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1583290452] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2020/03/04 08:24:06 [error] 12243#0: *64563685 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:1735/1697, bytes from/to upstream:1697/1735
[1583290572] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(89) < 2020/03/04 08:25:35 [error] 12255#0: *64583841 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583290692] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583290697] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21858116 kB) free.
[1583290726] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(156) < 172.105.59.173 [04/Mar/2020:08:27:36 +0530] TCP 502 0 0 180.001 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.001, 0.000"
[1583290846] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583291583] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4910748 kB) free!
[1583291703] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.8% (4852044 kB) free!
[1583291823] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.6% (4802604 kB) free!
[1583292394] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1583295424] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2122984 kB) free!
[1583296659] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1583296779] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1583298133] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4538132 kB) free!
[1583298253] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4385692 kB) free!
[1583298374] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (4256512 kB) free!
[1583298502] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (4351660 kB) free!
[1583298622] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (4207372 kB) free!
[1583298742] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (4182724 kB) free!
[1583299246] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(27) < 103.255.146.164 [04/Mar/2020:10:41:42 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583299325] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168008 kB) free!
[1583299366] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583299445] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172164 kB) free!
[1583299450] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 707 processes
[1583299565] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184688 kB) free!
[1583299571] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 707 processes
[1583299691] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 708 processes
[1583299942] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173808 kB) free!
[1583299961] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (211380 kB) free!
[1583300081] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (188824 kB) free!
[1583300132] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1583300174] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1978808 kB) free!
[1583300201] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (193100 kB) free!
[1583300252] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1583300292] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/04 11:02:22 [error] 12257#0: *64648868 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1583300372] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1583300377] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25899820 kB) free!
[1583300412] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583300497] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21362676 kB) free!
[1583300617] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40505868 kB) free.
[1583300794] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1583300914] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1583301034] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1583301728] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4014664 kB) free!
[1583301766] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 103.255.146.164 [04/Mar/2020:11:24:30 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583301848] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (3734452 kB) free!
[1583301886] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583301968] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.3% (3427604 kB) free!
[1583302091] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1583302417] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18515992 kB) free!
[1583302537] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36880248 kB) free.
[1583302568] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1777276 kB) free!
[1583302772] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1583303086] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1080) < 103.255.146.164 [04/Mar/2020:11:51:19 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583303137] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12384552 kB) free!
[1583303206] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583303257] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43190512 kB) free.
[1583303292] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 707 processes
[1583303412] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 711 processes
[1583303532] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 708 processes
[1583303769] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.6% (15637008 kB) free.
[1583303807] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(30) < 49.50.86.152 [04/Mar/2020:12:01:55 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583303927] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583303972] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1583304012] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(38) < 2020/03/04 12:01:47 [error] 12247#0: *64682313 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583304132] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583304457] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21947032 kB) free!
[1583304572] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1583304577] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39471412 kB) free.
[1583304692] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 357 processes
[1583304812] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1583304812] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1583305127] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1066) < 103.255.146.164 [04/Mar/2020:12:20:13 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583305137] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1583305177] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11350700 kB) free!
[1583305247] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583305257] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1583305297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40055164 kB) free.
[1583305332] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1583305332] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1583305352] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8096352 kB) free!
[1583305412] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1583305472] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.4% (8091380 kB) free!
[1583305592] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (8080060 kB) free!
[1583305898] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10044688 kB) free!
[1583306018] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40665244 kB) free.
[1583306447] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(999) < 103.255.146.164 [04/Mar/2020:12:44:03 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583306568] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583306618] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10925780 kB) free!
[1583306738] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41224804 kB) free.
[1583307338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11311488 kB) free!
[1583307458] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40193652 kB) free.
[1583307634] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1583307767] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(13) < 103.255.146.164 [04/Mar/2020:13:11:13 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583307849] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 21 users currently logged in
[1583307887] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583307969] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 21 users currently logged in
[1583308058] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14324248 kB) free!
[1583308089] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 21 users currently logged in
[1583308178] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38166224 kB) free.
[1583308487] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(12) < 103.255.146.164 [04/Mar/2020:13:22:56 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583308608] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583308778] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13896336 kB) free!
[1583308898] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37271244 kB) free.
[1583309192] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4426068 kB) free!
[1583309498] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12475364 kB) free!
[1583309618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35348884 kB) free.
[1583309807] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1350) < 103.255.146.164 [04/Mar/2020:13:44:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583309927] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583310818] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13724384 kB) free!
[1583310938] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21379468 kB) free!
[1583311059] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32923252 kB) free.
[1583311128] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1111) < 103.255.146.164 [04/Mar/2020:14:00:56 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583311248] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583311451] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9110456 kB) free!
[1583311498] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21663536 kB) free.
[1583311550] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39227440 kB) free.
[1583311569] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.2% (12304060 kB) free.
[1583311571] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8898404 kB) free!
[1583311598] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.9% (21542084 kB) free.
[1583311626] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.1% (14155916 kB) free.
[1583311658] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8664552 kB) free!
[1583311660] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.7% (21749964 kB) free.
[1583311691] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (8667956 kB) free!
[1583311747] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 55.8% (73559032 kB) free.
[1583312703] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.9% (2683188 kB) free!
[1583312823] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2360776 kB) free!
[1583312943] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1441464 kB) free!
[1583313048] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1011) < 103.255.146.164 [04/Mar/2020:14:31:06 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583313169] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(696) < 103.255.146.164 [04/Mar/2020:14:42:49 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583313289] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(498) < 103.255.146.164 [04/Mar/2020:14:42:51 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583314691] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4096964 kB) free!
[1583315089] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583315457] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1583315533] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583315578] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1583315690] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(990) < 103.255.146.164 [04/Mar/2020:15:20:33 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583315810] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583315889] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 17 users currently logged in
[1583316551] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21515496 kB) free!
[1583316671] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17246128 kB) free!
[1583316791] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42553960 kB) free.
[1583316812] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583317391] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16375960 kB) free!
[1583317511] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42625444 kB) free.
[1583317610] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 103.255.146.164 [04/Mar/2020:15:53:38 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583317730] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583317932] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(178) < 2020/03/04 16:01:52 [error] 12251#0: *64814785 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1583318052] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(15) < 2020/03/04 16:02:36 [error] 12253#0: *64815134 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583318111] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16189764 kB) free!
[1583318173] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583318231] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41657088 kB) free.
[1583318331] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 172.105.59.173 [04/Mar/2020:16:02:36 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 0.000, 60.000"
[1583318451] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583318612] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1583318732] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1583318831] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15524440 kB) free!
[1583318843] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6353664 kB) free!
[1583318852] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1583318952] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36717188 kB) free.
[1583318963] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (3013672 kB) free!
[1583319083] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225080 kB) free!
[1583319452] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1583319491] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9026704 kB) free!
[1583319551] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17474480 kB) free!
[1583319651] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1011) < 103.255.146.164 [04/Mar/2020:16:22:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583319671] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13453472 kB) free!
[1583319771] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583319791] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38446540 kB) free.
[1583320283] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 12.5% (6184064 kB) free!
[1583320391] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12491784 kB) free!
[1583320511] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39189512 kB) free.
[1583320972] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1013) < 103.255.146.164 [04/Mar/2020:16:45:59 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583321092] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583321111] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12628680 kB) free!
[1583321231] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38429592 kB) free.
[1583321692] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1110) < 103.255.146.164 [04/Mar/2020:16:59:05 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583321812] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583321831] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13379832 kB) free!
[1583321951] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9250908 kB) free!
[1583322071] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35484708 kB) free.
[1583322083] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4817444 kB) free!
[1583322412] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 103.255.146.164 [04/Mar/2020:17:13:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583322533] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583322671] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10742988 kB) free!
[1583322683] SERVICE FLAPPING ALERT: qservercluster2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1583322683] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 41.9% (20639264 kB) free.
[1583322791] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36402576 kB) free.
[1583323132] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 103.255.146.164 [04/Mar/2020:17:23:29 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583323252] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583323391] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11180400 kB) free!
[1583323511] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36506768 kB) free.
[1583323691] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4502272 kB) free!
[1583324112] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11327912 kB) free!
[1583324232] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35842868 kB) free.
[1583324452] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1066) < 103.255.146.164 [04/Mar/2020:17:44:10 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583324572] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583324832] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12177848 kB) free!
[1583324952] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7857816 kB) free!
[1583325072] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34264880 kB) free.
[1583325672] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8982300 kB) free!
[1583325792] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34459268 kB) free.
[1583326372] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(832) < 103.255.146.164 [04/Mar/2020:18:19:59 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583326392] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10734812 kB) free!
[1583326492] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583326512] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6433376 kB) free!
[1583326571] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4881308 kB) free!
[1583326632] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32771188 kB) free.
[1583326691] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4823672 kB) free!
[1583326811] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (4718128 kB) free!
[1583327092] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1052) < 103.255.146.164 [04/Mar/2020:18:30:59 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583327153] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8982100 kB) free!
[1583327213] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583327232] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7269684 kB) free!
[1583327273] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8975212 kB) free!
[1583327352] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33264292 kB) free.
[1583327393] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (8971508 kB) free!
[1583327812] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1053) < 103.255.146.164 [04/Mar/2020:18:38:22 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583327932] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(70) < 103.255.146.164 [04/Mar/2020:18:48:43 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583327952] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8050928 kB) free!
[1583328052] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583328072] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4022116 kB) free!
[1583328192] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30942864 kB) free.
[1583328652] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 103.255.146.164 [04/Mar/2020:18:58:52 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583328773] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583328792] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6286912 kB) free!
[1583328913] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32180160 kB) free.
[1583329512] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6843948 kB) free!
[1583329632] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (18997612 kB) free!
[1583329752] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29107012 kB) free.
[1583329811] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2382480 kB) free!
[1583330352] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4265848 kB) free!
[1583330472] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30492888 kB) free.
[1583331072] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3743532 kB) free!
[1583331192] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30729156 kB) free.
[1583331462] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.8% (4370648 kB) free!
[1583331582] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.4% (4259428 kB) free!
[1583331702] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.6% (4303412 kB) free!
[1583331792] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4667348 kB) free!
[1583331912] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (666328 kB) free!
[1583332032] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31374064 kB) free.
[1583332632] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4374532 kB) free!
[1583332752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389868 kB) free!
[1583332872] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28667484 kB) free.
[1583332972] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 49.50.86.152 [04/Mar/2020:20:03:00 +0530] TCP 502 0 0 91.068 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "31.068, 60.000, 0.000"
[1583333092] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583333472] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3298792 kB) free!
[1583333593] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30847096 kB) free.
[1583334083] SERVICE FLAPPING ALERT: qservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583334192] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3272912 kB) free!
[1583334312] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (398244 kB) free!
[1583334432] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30156744 kB) free.
[1583334702] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2157044 kB) free!
[1583335032] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5456740 kB) free!
[1583335152] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31890564 kB) free.
[1583335752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7343204 kB) free!
[1583335872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3333520 kB) free!
[1583335992] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29484780 kB) free.
[1583336592] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4873900 kB) free!
[1583336712] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28363996 kB) free.
[1583336993] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4487100 kB) free!
[1583337312] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3055724 kB) free!
[1583337432] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39577716 kB) free.
[1583338032] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15123000 kB) free!
[1583338152] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42584648 kB) free.
[1583338752] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16278192 kB) free!
[1583338873] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (41999752 kB) free.
[1583339472] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20163924 kB) free!
[1583339592] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16508700 kB) free!
[1583339712] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40144268 kB) free.
[1583340312] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18002360 kB) free!
[1583340432] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41669476 kB) free.
[1583341032] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19330412 kB) free!
[1583341152] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41800264 kB) free.
[1583341752] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20128380 kB) free!
[1583341872] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41576348 kB) free.
[1583342472] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20776736 kB) free!
[1583342592] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41729200 kB) free.
[1583343192] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21708964 kB) free!
[1583343312] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17963228 kB) free!
[1583343432] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39551984 kB) free.
[1583344032] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17827260 kB) free!
[1583344153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38576976 kB) free.
[1583345331] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.3% (21479272 kB) free!
[1583345352] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21082668 kB) free!
[1583345451] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.4% (13698536 kB) free!
[1583345472] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41623576 kB) free.
[1583345547] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231520 kB) free!
[1583345571] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5624932 kB) free!
[1583345667] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237136 kB) free!
[1583345787] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235964 kB) free!
[1583345980] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1583346072] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21161544 kB) free!
[1583346100] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1583346192] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41415752 kB) free.