[1566499331] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 324 processes
[1566499402] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197132 kB) free!
[1566499451] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1566499501] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3624332 kB) free!
[1566499523] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187920 kB) free!
[1566499531] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (3694612 kB) free!
[1566499621] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3616328 kB) free!
[1566499643] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180172 kB) free!
[1566499652] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175516 kB) free!
[1566499741] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.8% (3617728 kB) free!
[1566499772] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190248 kB) free!
[1566500283] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.79, 10.77, 8.91
[1566500403] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.46, 8.00, 8.12
[1566500523] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.54, 6.23, 7.46
[1566500739] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 231.60, 175.91, 85.61
[1566500843] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (2574332 kB) free!
[1566500860] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 119.32, 165.45, 93.59
[1566500980] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16.70, 110.96, 82.34
[1566501723] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.72, 2.68, 3.94
[1566502172] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.3% (10686332 kB) free.
[1566502216] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 565 processes
[1566502276] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 565 processes
[1566502337] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 559 processes
[1566502643] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2250576 kB) free!
[1566502741] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2108288 kB) free!
[1566502780] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.55, 1.01, 14.10
[1566503941] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2759004 kB) free!
[1566503980] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.38, 0.58, 4.33
[1566504743] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1566505323] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.77, 7.47, 8.03
[1566505443] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.63, 5.86, 7.37
[1566505563] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.03, 4.94, 6.85
[1566506163] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.64, 2.97, 4.86
[1566506380] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.84, 13.43, 8.13
[1566506500] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.63, 9.07, 7.17
[1566506620] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 6.31, 6.39
[1566506906] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1566506972] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166712 kB) free!
[1566507027] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1566507092] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192708 kB) free!
[1566507212] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172748 kB) free!
[1566507220] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 1.20, 3.55
[1566509742] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.43, 15.73, 9.87
[1566509764] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 15.59, 23.80, 12.53
[1566509862] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.12, 11.38, 8.98
[1566509884] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.93, 17.04, 11.42
[1566509941] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 45.8% (11235660 kB) free.
[1566509982] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.98, 7.92, 8.01
[1566510004] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.09, 12.32, 10.37
[1566510443] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.6% (10687576 kB) free.
[1566510582] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 1.27, 4.31
[1566511204] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.56, 2.80, 4.80
[1566511421] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.08, 7.65, 6.55
[1566511541] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.73, 5.37, 5.85
[1566511661] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 3.84, 5.23
[1566512261] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 1.04, 3.03
[1566514826] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566514946] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1566515861] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 56.39, 24.60, 9.44
[1566515981] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.89, 16.57, 8.33
[1566516101] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.46, 11.23, 7.37
[1566516701] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 1.94, 4.12
[1566519604] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.43, 12.17, 9.25
[1566519724] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.00, 9.10, 8.48
[1566519844] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.98, 7.04, 7.80
[1566520901] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.20, 16.24, 8.32
[1566521022] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.38, 11.02, 7.37
[1566521044] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.20, 2.87, 4.10
[1566521142] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 7.51, 6.52
[1566521742] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.49, 1.56, 3.72
[1566522136] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 573 processes
[1566522196] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 572 processes
[1566522256] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 568 processes
[1566522746] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566522812] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 62.4% (15396464 kB) free.
[1566522866] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566522942] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 101.15, 149.49, 85.97
[1566523062] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.01, 100.18, 75.60
[1566523182] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.35, 67.19, 66.49
[1566524382] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.87, 1.99, 18.83
[1566525173] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1566525293] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1566525593] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8731264 kB) free!
[1566525636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566525713] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7458408 kB) free!
[1566525756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566525833] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11397104 kB) free.
[1566525845] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 246.47, 160.90, 72.21
[1566525876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566525965] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 247.74, 189.70, 93.52
[1566526085] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 164.09, 189.24, 105.71
[1566526182] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.59, 3.32
[1566527034] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9660276 kB) free!
[1566527154] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8386940 kB) free!
[1566527274] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7077624 kB) free!
[1566527874] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11381476 kB) free.
[1566527885] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.95, 3.53, 17.91
[1566528876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566529074] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9396764 kB) free!
[1566529085] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.13, 23.79, 15.47
[1566529194] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8118640 kB) free!
[1566529314] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6576604 kB) free!
[1566529685] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.87, 5.64, 9.43
[1566530076] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566530196] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566530796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566530885] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.49, 2.68, 4.53
[1566530917] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566530983] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.55, 4.84, 5.74
[1566531103] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.58, 3.42, 5.11
[1566531114] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1566531114] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11371364 kB) free.
[1566531223] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.53, 2.49, 4.57
[1566531266] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566531386] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566533612] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (363136 kB) free!
[1566533732] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179692 kB) free!
[1566533852] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164268 kB) free!
[1566533885] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 140.25, 36.48, 14.17
[1566533941] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (265816 kB) free!
[1566533956] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 573 processes
[1566534005] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 20.81, 25.16, 12.73
[1566534017] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 576 processes
[1566534061] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167296 kB) free!
[1566534125] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.51, 17.82, 11.54
[1566534182] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194484 kB) free!
[1566534443] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195416 kB) free!
[1566534563] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195996 kB) free!
[1566534683] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190132 kB) free!
[1566535925] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.09, 3.83, 4.79
[1566536317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566536437] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566536557] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566536623] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1566537653] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566537714] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9254824 kB) free!
[1566537773] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566537834] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8039568 kB) free!
[1566537893] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566537954] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6382088 kB) free!
[1566538554] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11155872 kB) free.
[1566538925] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.18, 18.49, 11.96
[1566539045] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.69, 13.43, 10.89
[1566539165] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.98, 10.00, 9.95
[1566539754] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8723632 kB) free!
[1566539874] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7426464 kB) free!
[1566539994] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11348260 kB) free.
[1566540158] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1566540224] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.61, 19.55, 10.34
[1566540344] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.96, 13.39, 9.19
[1566540365] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.70, 3.20, 4.97
[1566540464] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.41, 9.34, 8.22
[1566540594] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7674444 kB) free!
[1566540604] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 64.5% (15812424 kB) free.
[1566540714] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11343820 kB) free.
[1566540887] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.72, 5.84, 5.29
[1566541007] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.28, 5.42, 5.19
[1566541064] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.17, 2.11, 4.75
[1566541127] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.14, 5.35, 5.20
[1566541314] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8022872 kB) free!
[1566541358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566541434] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6380268 kB) free!
[1566541478] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566541554] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11336364 kB) free.
[1566541727] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.01, 4.19, 4.53
[1566542355] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.3% (14520036 kB) free.
[1566542408] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.9% (21685444 kB) free.
[1566542422] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.1% (21842732 kB) free.
[1566542454] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.2% (40023108 kB) free.
[1566542497] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.3% (20910876 kB) free.
[1566542928] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.09, 6.00, 5.63
[1566543048] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.09, 6.26, 5.76
[1566543168] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.39, 5.96, 5.71
[1566543175] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1566543295] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1566543895] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1566544015] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1566544664] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 240.15, 195.31, 100.83
[1566544784] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 59.28, 151.00, 96.83
[1566544904] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 66.38, 128.65, 95.38
[1566545364] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (182100 kB) free!
[1566545484] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165528 kB) free!
[1566545605] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179184 kB) free!
[1566545982] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 270.39, 163.14, 69.71
[1566546102] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 272.97, 199.65, 94.42
[1566546222] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 273.53, 224.12, 116.15
[1566546704] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1566546704] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 1.19, 14.55
[1566547904] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.18, 1.10, 4.82
[1566548022] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.59, 1.25, 18.96
[1566548186] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566548306] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566548414] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178096 kB) free!
[1566548434] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.3% (3261396 kB) free!
[1566548534] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171120 kB) free!
[1566548555] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201404 kB) free!
[1566548654] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183588 kB) free!
[1566548675] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164300 kB) free!
[1566549102] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193440 kB) free!
[1566549222] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176316 kB) free!
[1566549342] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183896 kB) free!
[1566549823] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.59, 3.31
[1566550265] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8754116 kB) free!
[1566550385] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8753908 kB) free!
[1566550505] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (8654596 kB) free!
[1566550555] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566551165] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1566551197] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.2% (14492796 kB) free.
[1566551237] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.0% (21332740 kB) free.
[1566551291] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.3% (22372528 kB) free.
[1566551323] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.0% (22454004 kB) free.
[1566551705] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2060240 kB) free!
[1566552415] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1566552535] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1566553001] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3168224 kB) free!
[1566553121] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (3163888 kB) free!
[1566553135] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1566553241] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.4% (2016892 kB) free!
[1566553255] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1566553368] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.95, 4.30, 4.75
[1566553841] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1566553841] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176948 kB) free!
[1566553855] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1566553975] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1566554576] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1566554693] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566554696] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1566554841] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.2% (3000348 kB) free!
[1566554894] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171840 kB) free!
[1566554935] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178584 kB) free!
[1566554961] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185168 kB) free!
[1566555015] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191700 kB) free!
[1566555055] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198640 kB) free!
[1566555082] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163556 kB) free!
[1566555135] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181888 kB) free!
[1566555175] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180112 kB) free!
[1566555296] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 357 processes
[1566555416] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1566556493] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566556614] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566556734] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566557678] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566557798] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566558354] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7867636 kB) free!
[1566558399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566558416] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1566558474] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11814996 kB) free.
[1566558519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566558537] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1566558704] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1566559744] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9291528 kB) free!
[1566559864] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9051668 kB) free!
[1566559984] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.8% (8781948 kB) free!
[1566560568] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.18, 6.24, 5.17
[1566560688] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.99, 5.83, 5.16
[1566560808] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.08, 5.13, 4.98
[1566560919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566560934] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566561039] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1566561159] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1566561747] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21993668 kB) free.
[1566561760] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1566561784] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4397504 kB) free!
[1566561790] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.9% (22035272 kB) free.
[1566561829] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.6% (38722884 kB) free.
[1566561863] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 44.8% (22058104 kB) free.
[1566561898] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.5% (20945140 kB) free.
[1566561936] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.3% (14516072 kB) free.
[1566562008] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.41, 6.35, 5.41
[1566562128] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.27, 6.60, 5.61
[1566562248] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.66, 6.61, 5.73
[1566562423] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 25.04, 93.26, 58.07
[1566562543] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.25, 62.77, 51.16
[1566562663] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.63, 42.33, 45.08
[1566563143] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (159852 kB) free!
[1566563263] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165036 kB) free!
[1566563383] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 50.9% (8272336 kB) free.
[1566563863] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.09, 1.73, 13.06
[1566563983] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (192516 kB) free!
[1566564103] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167176 kB) free!
[1566564223] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (153096 kB) free!
[1566564823] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 30.4% (4949368 kB) free.
[1566565063] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.56, 0.57, 4.06
[1566565423] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (147872 kB) free!
[1566565543] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (170436 kB) free!
[1566565663] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174588 kB) free!
[1566565848] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.99, 4.04, 4.61
[1566566504] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 274.68, 172.32, 75.04
[1566566624] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 276.14, 206.72, 99.44
[1566566633] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7558568 kB) free!
[1566566744] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 276.35, 229.79, 120.91
[1566566753] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6797864 kB) free!
[1566566873] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.3% (6081260 kB) free!
[1566567152] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (1042268 kB) free!
[1566567272] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186984 kB) free!
[1566567392] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195752 kB) free!
[1566567473] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2261020 kB) free!
[1566568544] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.80, 1.45, 19.11
[1566570344] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.66, 0.94, 3.60
[1566573100] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.3% (14508892 kB) free.
[1566573134] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 68.7% (33879224 kB) free.
[1566573176] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.2% (21276160 kB) free.
[1566573783] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177396 kB) free!
[1566573851] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 320 processes
[1566573904] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174356 kB) free!
[1566573971] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1566574024] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178616 kB) free!
[1566575504] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (1381540 kB) free!
[1566575624] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (355144 kB) free!
[1566575744] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (193992 kB) free!
[1566577130] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.2% (14497996 kB) free.
[1566577136] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 353 processes
[1566577168] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1566577168] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.7% (22129844 kB) free.
[1566577256] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1566579743] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8731016 kB) free!
[1566579864] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (8033544 kB) free!
[1566579984] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (7540648 kB) free!
[1566580134] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (857696 kB) free!
[1566580255] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165368 kB) free!
[1566580375] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157156 kB) free!
[1566580584] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4621572 kB) free!
[1566581307] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1566581427] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1566584972] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)