[1567451556] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 63.0% (15544720 kB) free.
[1567452934] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.5% (21721264 kB) free!
[1567452948] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 293 processes
[1567453054] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12267084 kB) free!
[1567453174] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2769084 kB) free!
[1567453416] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (231748 kB) free!
[1567453537] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (241988 kB) free!
[1567453657] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (755312 kB) free!
[1567453956] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173084 kB) free!
[1567454076] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186212 kB) free!
[1567454148] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 289 processes
[1567454196] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175984 kB) free!
[1567460157] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182828 kB) free!
[1567460277] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191056 kB) free!
[1567460394] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191608 kB) free!
[1567460397] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199660 kB) free!
[1567460514] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193984 kB) free!
[1567460634] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181692 kB) free!
[1567461235] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 56.8% (13911532 kB) free.
[1567461597] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 51.9% (12728608 kB) free.
[1567471058] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.9% (33583732 kB) free.
[1567472974] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 48.7% (64235648 kB) free.
[1567474596] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.7% (11032656 kB) free.
[1567475627] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 168.80, 94.97, 39.84
[1567475747] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 170.98, 120.23, 55.78
[1567475796] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187224 kB) free!
[1567475867] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 82.07, 113.55, 61.82
[1567475916] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (2149380 kB) free!
[1567475997] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4859580 kB) free!
[1567476036] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1849324 kB) free!
[1567476117] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4813180 kB) free!
[1567476237] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (4786864 kB) free!
[1567477067] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.52, 2.71, 17.77
[1567477437] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 22.1% (5426472 kB) free.
[1567478867] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.96, 1.02, 3.38
[1567481637] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4874328 kB) free!
[1567481757] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4852812 kB) free!
[1567481877] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (4830004 kB) free!
[1567486579] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22953504 kB) free.
[1567486632] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40723128 kB) free.
[1567487137] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1567487257] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1567489077] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2085160 kB) free!
[1567490185] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (1238036 kB) free!
[1567490305] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (1217716 kB) free!
[1567490425] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (1013908 kB) free!
[1567493035] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4764636 kB) free!
[1567493146] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1567493155] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4679916 kB) free!
[1567493266] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1567493275] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4595900 kB) free!
[1567493386] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1567493832] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2651208 kB) free!
[1567493953] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230660 kB) free!
[1567494073] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241844 kB) free!
[1567494457] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1567494577] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1567494823] SERVICE ALERT: rq-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1567497467] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.72, 5.95, 5.50
[1567497586] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1567497587] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.98, 5.94, 5.55
[1567497707] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.43, 5.46, 5.43
[1567498675] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2193832 kB) free!
[1567498907] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.10, 4.78, 4.91
[1567499507] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.33, 5.82, 5.35
[1567499627] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.42, 5.58, 5.31
[1567499747] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.04, 5.68, 5.40
[1567502144] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 403 processes
[1567502147] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1567502147] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.45, 4.71, 4.92
[1567502264] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 403 processes
[1567502385] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567504547] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.66, 5.36, 5.10
[1567504667] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.32, 5.75, 5.28
[1567504787] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.50, 5.65, 5.30
[1567505668] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.4% (20826228 kB) free.
[1567505683] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.4% (22893832 kB) free.
[1567505694] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.9% (39874272 kB) free.
[1567505774] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.7% (23219096 kB) free.
[1567505986] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1567505987] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.37, 3.90, 4.48
[1567506106] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1567506226] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1567506273] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (3860032 kB) free!
[1567506393] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (439112 kB) free!
[1567506513] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192712 kB) free!
[1567506903] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (223204 kB) free!
[1567507023] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213960 kB) free!
[1567507143] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (265548 kB) free!
[1567508984] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 403 processes
[1567508987] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.41, 5.60, 5.34
[1567509104] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567509108] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.79, 5.38, 5.30
[1567509228] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.40, 4.86, 5.10
[1567510177] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9517932 kB) free!
[1567510297] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9315996 kB) free!
[1567510394] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4881020 kB) free!
[1567510418] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (9144924 kB) free!
[1567510424] SERVICE ALERT: rq-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1567510428] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.65, 3.74, 4.32
[1567510514] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.9% (4873604 kB) free!
[1567510544] SERVICE ALERT: rq-2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1567510634] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.9% (4866104 kB) free!
[1567510664] SERVICE ALERT: rq-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1567511135] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1567511255] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1567511853] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: /home 47135 MB (20% inode=99%):
[1567511855] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 317 processes
[1567511973] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: /home 47135 MB (20% inode=99%):
[1567511975] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1567512093] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: /home 47135 MB (20% inode=99%):
[1567513177] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 317 processes
[1567513297] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1567513418] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4161376 kB) free!
[1567513547] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 304 processes
[1567513667] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 306 processes
[1567513785] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 577 processes
[1567513787] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 311 processes
[1567513845] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 573 processes
[1567513905] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 1 warn out of 576 processes
[1567513965] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 1 warn out of 576 processes
[1567514265] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 578 processes
[1567514987] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1567515228] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.93, 6.66, 5.18
[1567515348] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.06, 7.10, 5.52
[1567515468] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.79, 6.94, 5.65
[1567515962] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.0% (21714172 kB) free.
[1567515975] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39841284 kB) free.
[1567515997] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.7% (45181232 kB) free.
[1567516426] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1567516562] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (2305120 kB) free!
[1567516682] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182208 kB) free!
[1567516802] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1567516802] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170140 kB) free!
[1567517183] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2703516 kB) free!
[1567517304] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (1853124 kB) free!
[1567517424] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (1101028 kB) free!
[1567517868] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.89, 2.53, 3.86
[1567521983] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4771856 kB) free!
[1567522034] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2402452 kB) free!
[1567522104] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4673944 kB) free!
[1567522224] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4584584 kB) free!
[1567525824] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437748 kB) free!
[1567526264] SERVICE ALERT: rq-2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1567526504] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567526624] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567527468] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1567528202] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)