[1564597817] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6581844 kB) free!
[1564597903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564597937] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11869772 kB) free.
[1564598503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1564598537] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7132720 kB) free!
[1564598623] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564598631] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181228 kB) free!
[1564598657] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11871412 kB) free.
[1564598751] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200564 kB) free!
[1564598871] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (208988 kB) free!
[1564599223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564599257] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7553184 kB) free!
[1564599344] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1564599377] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11864196 kB) free.
[1564599943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1564599977] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7901228 kB) free!
[1564600063] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1564600097] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11869804 kB) free.
[1564600663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564600698] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8140672 kB) free!
[1564600783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1564600818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11864088 kB) free.
[1564600903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564600971] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.7% (10031592 kB) free.
[1564601418] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8518736 kB) free!
[1564601503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564601538] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6923072 kB) free!
[1564601623] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564601658] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11861492 kB) free.
[1564602224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564602258] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7263264 kB) free!
[1564602344] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1564602378] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11855560 kB) free.
[1564602490] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1564602610] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1564602944] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1564602978] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7631704 kB) free!
[1564603064] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564603098] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11860716 kB) free.
[1564603664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1564603698] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7991000 kB) free!
[1564603784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564603818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11850032 kB) free.
[1564603904] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564604419] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8258560 kB) free!
[1564604504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564604539] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11762528 kB) free.
[1564604624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564605139] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8496504 kB) free!
[1564605225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564605259] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6802220 kB) free!
[1564605345] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1564605379] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11850740 kB) free.
[1564605945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1564605979] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7271544 kB) free!
[1564606065] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564606099] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11859668 kB) free.
[1564606371] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185264 kB) free!
[1564606387] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 233.35, 730.26, 455.83
[1564606491] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201136 kB) free!
[1564606507] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 33.14, 489.21, 400.77
[1564606611] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189632 kB) free!
[1564606627] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.55, 328.10, 352.45
[1564606665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1564606699] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7739232 kB) free!
[1564606785] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1564606819] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11852396 kB) free.
[1564607301] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.5% (13405244 kB) free.
[1564607385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564607419] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8063284 kB) free!
[1564607505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564607539] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11858620 kB) free.
[1564607625] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564608139] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8361748 kB) free!
[1564608226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564608260] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6560412 kB) free!
[1564608346] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1564608380] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11852948 kB) free.
[1564608946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1564608980] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7005416 kB) free!
[1564609066] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1564609100] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11859068 kB) free.
[1564609627] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.74, 0.89, 14.69
[1564609666] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564609700] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7369012 kB) free!
[1564609786] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564609820] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11851828 kB) free.
[1564610386] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1564610420] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7778144 kB) free!
[1564610506] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564610540] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11859228 kB) free.
[1564610827] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.70, 0.61, 4.46
[1564611107] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564611140] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8008996 kB) free!
[1564611227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564611260] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11860196 kB) free.
[1564611347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1564611860] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8263344 kB) free!
[1564611947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1564611980] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10362616 kB) free.
[1564612067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564612581] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8510268 kB) free!
[1564612667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564612701] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6820472 kB) free!
[1564612787] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564612821] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11844312 kB) free.
[1564613387] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564613421] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7338588 kB) free!
[1564613508] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564613541] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11850552 kB) free.
[1564614107] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564614141] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7790104 kB) free!
[1564614227] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564614261] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11850708 kB) free.
[1564614827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564614861] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8088508 kB) free!
[1564614947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1564614981] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11845028 kB) free.
[1564615067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1564615581] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6843564 kB) free!
[1564615701] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11608480 kB) free.
[1564618667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564618787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564618907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564620107] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564620501] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8912416 kB) free!
[1564620621] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7927552 kB) free!
[1564620707] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564620741] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6562816 kB) free!
[1564620828] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564621427] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564621547] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1564621667] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564621810] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1564621930] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1564621941] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10981732 kB) free.
[1564622267] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564622387] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564622442] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1564622507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564622531] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1564622562] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1564622651] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1564623707] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1564623707] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564623762] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1564623882] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1564624907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1564625027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564625147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1564625541] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6655468 kB) free!
[1564625661] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11413584 kB) free.
[1564626947] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1564627461] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8789184 kB) free!
[1564627547] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1564627581] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7768216 kB) free!
[1564627668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564627701] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6363184 kB) free!
[1564627788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564628388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564628508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1564628628] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564628671] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4902112 kB) free!
[1564628791] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4892256 kB) free!
[1564628902] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1564628902] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10994756 kB) free.
[1564628912] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4884196 kB) free!
[1564629228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1564629348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1564629468] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564630069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1564630102] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7550724 kB) free!
[1564630189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1564630222] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11636276 kB) free.
[1564630309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564630822] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8814044 kB) free!
[1564630909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564630942] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7851428 kB) free!
[1564631029] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564631062] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6577912 kB) free!
[1564631149] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1564631749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1564631869] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1564631990] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564632262] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10107392 kB) free.
[1564632589] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564632709] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564632829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564632862] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9326972 kB) free!
[1564632982] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7998484 kB) free!
[1564633102] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6257060 kB) free!
[1564633703] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11275264 kB) free.
[1564634901] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.3% (2971288 kB) free!
[1564634903] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6498060 kB) free!
[1564635021] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.8% (2731100 kB) free!
[1564635023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11040576 kB) free.
[1564635141] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.6% (2527568 kB) free!
[1564636341] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1024092 kB) free!
[1564636619] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 378 processes
[1564636739] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 381 processes
[1564637629] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564638023] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8448364 kB) free!
[1564638143] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6795632 kB) free!
[1564638263] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11338984 kB) free.
[1564638829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564638949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1564639069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564639112] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2261372 kB) free!
[1564640063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7044892 kB) free!
[1564640183] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11586176 kB) free.
[1564640424] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.9% (1171728 kB) free!
[1564640544] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (1049688 kB) free!
[1564640664] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.1% (891704 kB) free!
[1564640809] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.1% (2616384 kB) free!
[1564640811] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.3% (21296824 kB) free.
[1564640862] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 74.9% (12173964 kB) free.
[1564640872] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.0% (22313836 kB) free.
[1564640893] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40321296 kB) free.
[1564640894] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.1% (6770352 kB) free.
[1564640912] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.8% (21761244 kB) free.
[1564640914] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.7% (14085528 kB) free.
[1564640929] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.8% (2404312 kB) free!
[1564640932] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.8% (2397900 kB) free!
[1564640980] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.0% (14147120 kB) free.
[1564641383] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6316084 kB) free!
[1564641411] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (609088 kB) free!
[1564641462] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (190888 kB) free!
[1564641469] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1564641503] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11591916 kB) free.
[1564641531] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181932 kB) free!
[1564641582] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174112 kB) free!
[1564641651] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181916 kB) free!
[1564641702] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176244 kB) free!
[1564642069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564642082] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193908 kB) free!
[1564642093] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227652 kB) free!
[1564642103] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7313252 kB) free!
[1564642189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564642202] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193240 kB) free!
[1564642213] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215332 kB) free!
[1564642223] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11583024 kB) free.
[1564642252] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 32.8% (8086092 kB) free.
[1564642309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564642322] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192492 kB) free!
[1564642333] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237320 kB) free!
[1564642627] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.69, 5.47, 5.03
[1564642747] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.36, 4.65, 4.78
[1564642823] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7983596 kB) free!
[1564642909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1564642943] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6711616 kB) free!
[1564643030] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564643063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11590740 kB) free.
[1564643452] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (3700896 kB) free!
[1564643522] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 74.2% (18184760 kB) free.
[1564643572] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163192 kB) free!
[1564643629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564643663] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8057504 kB) free!
[1564643674] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22334792 kB) free.
[1564643695] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.3% (38104884 kB) free.
[1564643739] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 34.5% (17020984 kB) free.
[1564643749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 344 processes
[1564643760] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 45.6% (7416836 kB) free.
[1564643869] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564644282] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.1% (3239928 kB) free!
[1564644304] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8431020 kB) free!
[1564644365] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169124 kB) free!
[1564644403] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191624 kB) free!
[1564644425] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3513128 kB) free!
[1564644469] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564644485] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (150208 kB) free!
[1564644523] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1564644523] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198076 kB) free!
[1564644545] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1564644545] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (259180 kB) free!
[1564644589] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564644605] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1564644605] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162272 kB) free!
[1564644709] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1564645147] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.23, 5.30, 5.03
[1564645191] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1564645267] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.40, 4.40, 4.74
[1564645310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564645430] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1564645543] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1564646030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564646150] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1564646525] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182320 kB) free!
[1564646646] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182164 kB) free!
[1564646743] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8143236 kB) free!
[1564646750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1564646766] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1564646766] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164980 kB) free!
[1564646863] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6749256 kB) free!
[1564646870] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564646912] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (3484448 kB) free!
[1564646923] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.7% (2393380 kB) free!
[1564646983] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11595520 kB) free.
[1564646990] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564647032] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.1% (3461460 kB) free!
[1564647043] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 13.2% (2151792 kB) free!
[1564647153] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.0% (3438072 kB) free!
[1564647163] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.9% (1938464 kB) free!
[1564647590] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564647763] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (931708 kB) free!
[1564648183] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6860772 kB) free!
[1564648191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1564648303] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11894748 kB) free.
[1564648311] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1564649503] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7512036 kB) free!
[1564649511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564649623] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11892440 kB) free.
[1564649631] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564649902] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (1302976 kB) free!
[1564650022] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.2% (1194212 kB) free!
[1564650121] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.4% (22305472 kB) free.
[1564650142] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (1075312 kB) free!
[1564650164] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.6% (22202684 kB) free.
[1564650201] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 66.9% (32982092 kB) free.
[1564650224] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7707468 kB) free!
[1564650231] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564650237] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.1% (14315828 kB) free.
[1564650270] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1564650270] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.2% (14342008 kB) free.
[1564650344] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11887476 kB) free.
[1564650351] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1564650742] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1564650742] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (458296 kB) free!
[1564650841] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (185044 kB) free!
[1564650944] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8095088 kB) free!
[1564650952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1564650961] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164912 kB) free!
[1564651064] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11894936 kB) free.
[1564651072] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564651081] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179652 kB) free!
[1564651353] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2378600 kB) free!
[1564651401] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (931500 kB) free!
[1564651521] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214144 kB) free!
[1564651642] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240300 kB) free!
[1564651664] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8286808 kB) free!
[1564651672] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564651747] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 27.12, 50.28, 24.27
[1564651773] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.66, 50.62, 29.83
[1564651784] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11772988 kB) free.
[1564651792] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564651867] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 6.34, 25.88, 20.27
[1564651867] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.29, 33.93, 21.44
[1564651893] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.44, 34.94, 26.60
[1564651942] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.8% (6824392 kB) free.
[1564651987] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.59, 19.16, 18.49
[1564651987] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.07, 22.90, 18.91
[1564652013] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.21, 24.34, 23.73
[1564652107] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.94, 14.49, 16.87
[1564652384] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8344724 kB) free!
[1564652392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564652504] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11777992 kB) free.
[1564652512] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564652521] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185700 kB) free!
[1564652613] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.13, 5.73, 13.77
[1564652641] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191808 kB) free!
[1564652761] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186484 kB) free!
[1564653104] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8414096 kB) free!
[1564653112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564653224] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6584308 kB) free!
[1564653232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564653311] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.28, 6.13, 5.37
[1564653344] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11880380 kB) free.
[1564653353] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1564653432] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.20, 6.52, 5.60
[1564653552] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.97, 6.35, 5.66
[1564653787] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.15, 1.02, 3.44
[1564653944] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7154004 kB) free!
[1564653952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1564654064] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11887620 kB) free.
[1564654072] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564654184] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 14.2% (2300508 kB) free!
[1564654304] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 13.1% (2136308 kB) free!
[1564654372] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179012 kB) free!
[1564654414] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.01, 3.19, 4.65
[1564654424] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.0% (1951080 kB) free!
[1564654492] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166988 kB) free!
[1564654507] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.40, 1.79, 3.81
[1564654612] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190120 kB) free!
[1564655024] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (974368 kB) free!
[1564655272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564655352] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.69, 4.67, 4.97
[1564655392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564655512] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564655673] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3123120 kB) free!
[1564655793] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (2922072 kB) free!
[1564655913] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.7% (2718104 kB) free!
[1564656552] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.48, 5.57, 5.24
[1564656672] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.38, 5.62, 5.30
[1564656792] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.32, 5.54, 5.32
[1564657113] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (778676 kB) free!
[1564657322] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.5% (14060572 kB) free.
[1564657363] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.4% (21658612 kB) free.
[1564657392] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1564657392] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.67, 4.26, 4.84
[1564657424] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22828024 kB) free.
[1564657474] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.9% (22519604 kB) free.
[1564657511] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 51.7% (25460780 kB) free.
[1564657550] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.9% (14287208 kB) free.
[1564657586] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 86.7% (14090708 kB) free.
[1564657932] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (154560 kB) free!
[1564658052] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167820 kB) free!
[1564658172] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165808 kB) free!
[1564658264] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9395240 kB) free!
[1564658384] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8062484 kB) free!
[1564658504] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11872756 kB) free.
[1564658512] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564658715] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3513840 kB) free!
[1564658835] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2948648 kB) free!
[1564658955] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2384824 kB) free!
[1564659104] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8328592 kB) free!
[1564659112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564659224] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6437340 kB) free!
[1564659232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564659344] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11123612 kB) free.
[1564659827] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (2148264 kB) free!
[1564659832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564659947] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178068 kB) free!
[1564659952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1564660067] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181880 kB) free!
[1564660073] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564661052] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1564661172] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1564661542] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.3% (1344088 kB) free!
[1564661662] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (1186484 kB) free!
[1564661782] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (1070304 kB) free!
[1564662382] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (655028 kB) free!
[1564662983] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.5% (6872888 kB) free.
[1564663797] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1564665283] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166836 kB) free!
[1564665354] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.3% (2812908 kB) free!
[1564665403] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169900 kB) free!
[1564665474] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.5% (2675764 kB) free!
[1564665523] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178584 kB) free!
[1564665594] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.7% (2546408 kB) free!
[1564666197] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1564666371] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (2309124 kB) free!
[1564666491] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (2239616 kB) free!
[1564666611] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2180732 kB) free!
[1564666794] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1094468 kB) free!
[1564668373] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1564668792] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1564668944] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9257308 kB) free!
[1564669064] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7925568 kB) free!
[1564669067] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1564669156] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1564669184] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11832832 kB) free.
[1564669387] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14.73, 40.67, 21.24
[1564669414] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.43, 40.30, 22.81
[1564669507] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.90, 27.56, 18.79
[1564669507] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.87, 28.54, 21.16
[1564669534] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.32, 27.93, 20.40
[1564669627] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.72, 18.58, 16.56
[1564669627] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.98, 19.35, 18.68
[1564669654] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.85, 19.59, 18.26
[1564669674] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1564669747] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.74, 13.15, 16.49
[1564670874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564670994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564671427] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.00, 1.07, 3.31
[1564671454] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.36, 2.55, 4.86
[1564671523] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1564671547] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.30, 1.25, 3.42
[1564671584] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8120948 kB) free!
[1564671594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564671704] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11826356 kB) free.
[1564671714] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564672305] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8374532 kB) free!
[1564672314] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564672425] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6678068 kB) free!
[1564672434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564672545] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11817008 kB) free.
[1564672554] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1564672572] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1564672574] SERVICE ALERT: sig6-vm1;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig6-vm1' was out of bounds)
[1564672612] SERVICE ALERT: sig6-vm1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1564672615] HOST ALERT: sig6-vm1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1564672670] SERVICE ALERT: sig6-vm1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6-vm1' was out of bounds)
[1564672678] SERVICE ALERT: sig6-vm2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1564672681] HOST ALERT: sig6-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.193)
[1564672692] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 398 processes
[1564672697] SERVICE ALERT: sig6-vm1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig6-vm1' was out of bounds)
[1564672735] SERVICE ALERT: sig6-vm1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6-vm1' was out of bounds)
[1564672735] SERVICE ALERT: sig6-vm1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig6-vm1' was out of bounds)
[1564672801] SERVICE ALERT: sig6-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig6-vm2' was out of bounds)
[1564672801] SERVICE ALERT: sig6-vm1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig6-vm1' was out of bounds)
[1564672812] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1564672866] SERVICE ALERT: sig3-vm1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1564672867] SERVICE ALERT: sig6-vm1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig6-vm1' was out of bounds)
[1564672869] HOST ALERT: sig3-vm1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.60)
[1564672912] HOST ALERT: sig3-vm2;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1564672916] SERVICE ALERT: sig3-vm1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig3-vm1' was out of bounds)
[1564672989] SERVICE ALERT: sig3-vm1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig3-vm1' was out of bounds)
[1564673027] SERVICE ALERT: sig6-vm2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig6-vm2' was out of bounds)
[1564673042] SERVICE ALERT: sig3-vm2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig3-vm2' was out of bounds)
[1564673044] SERVICE ALERT: sig6-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6-vm2' was out of bounds)
[1564673069] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig3-vm1' was out of bounds)
[1564673104] SERVICE ALERT: sig6-vm1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig6-vm1' was out of bounds)
[1564673106] SERVICE ALERT: sig6-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6-vm2' was out of bounds)
[1564673139] SERVICE ALERT: sig3-vm1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig3-vm1' was out of bounds)
[1564673145] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7101780 kB) free!
[1564673153] SERVICE ALERT: sig6-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig6-vm2' was out of bounds)
[1564673153] SERVICE ALERT: sig6-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig6-vm2' was out of bounds)
[1564673154] SERVICE ALERT: sig3-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig3-vm2' was out of bounds)
[1564673155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1564673164] SERVICE ALERT: sig3-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig3-vm2' was out of bounds)
[1564673167] SERVICE ALERT: sig3-vm1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig3-vm1' was out of bounds)
[1564673177] SERVICE ALERT: sig3-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig3-vm2' was out of bounds)
[1564673179] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig6-vm2' was out of bounds)
[1564673182] HOST ALERT: sig6-vm1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.58)
[1564673186] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig3-vm2' was out of bounds)
[1564673205] SERVICE ALERT: sig3-vm1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig3-vm1' was out of bounds)
[1564673228] SERVICE ALERT: sig3-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig3-vm2' was out of bounds)
[1564673245] HOST ALERT: sig6-vm2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.193)
[1564673265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11823292 kB) free.
[1564673271] SERVICE ALERT: sig3-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig3-vm2' was out of bounds)
[1564673275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1564673394] SERVICE ALERT: sig3-vm1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig3-vm1' was out of bounds)
[1564673437] HOST ALERT: sig3-vm1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.60)
[1564673455] SERVICE ALERT: sig6-vm1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1564673459] HOST ALERT: sig6-vm1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1564673480] HOST ALERT: sig3-vm2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.181)
[1564673521] SERVICE ALERT: sig6-vm1;Memory;OK;HARD;1;OK - 98.1% (23071076 kB) free.
[1564673549] HOST ALERT: sig6-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1564673587] SERVICE ALERT: sig6-vm1;Total Processes;OK;HARD;1;PROCS OK: 260 processes
[1564673657] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'sig5-vm2' was out of bounds)
[1564673673] HOST ALERT: sig5-vm2;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1564673747] SERVICE ALERT: sig6-vm2;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.10, 0.07
[1564673764] SERVICE ALERT: sig6-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 73184 MB (61% inode=96%):
[1564673780] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig5-vm2' was out of bounds)
[1564673796] SERVICE ALERT: sig5-vm2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig5-vm2' was out of bounds)
[1564673824] SERVICE ALERT: sig6-vm1;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.07, 0.06
[1564673827] SERVICE ALERT: sig6-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1564673837] SERVICE ALERT: sig5-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig5-vm2' was out of bounds)
[1564673854] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 24.85, 15.21, 7.66
[1564673865] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7510728 kB) free!
[1564673865] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig5-vm2' was out of bounds)
[1564673873] SERVICE ALERT: sig6-vm2;Total Processes;OK;HARD;1;PROCS OK: 262 processes
[1564673873] SERVICE ALERT: sig6-vm2;CPU_Procs;OK;HARD;1;CPU OK: 262 processes
[1564673875] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564673888] SERVICE ALERT: sig5-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig5-vm2' was out of bounds)
[1564673894] SERVICE ALERT: sig6-vm1;CPU_Procs;OK;HARD;1;CPU OK: 260 processes
[1564673899] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 98.1% (23070216 kB) free.
[1564673914] SERVICE ALERT: sig5-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig5-vm2' was out of bounds)
[1564673917] SERVICE ALERT: sig5-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig5-vm2' was out of bounds)
[1564673932] SERVICE ALERT: sig6-vm1;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1564673974] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.45, 11.00, 7.03
[1564673985] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11826704 kB) free.
[1564673990] SERVICE ALERT: sig6-vm1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 72029 MB (60% inode=96%):
[1564673995] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564673998] SERVICE ALERT: sig6-vm2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1564674094] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.92, 8.20, 6.49
[1564674240] HOST ALERT: sig5-vm2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.192)
[1564674585] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7857984 kB) free!
[1564674595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564674609] SERVICE ALERT: sig5-vm2;Total Processes;OK;HARD;1;PROCS OK: 211 processes
[1564674613] HOST ALERT: sig5-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1564674694] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1564674694] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.25, 3.15, 4.55
[1564674705] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11824632 kB) free.
[1564674715] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564674978] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 97.3% (34084588 kB) free.
[1564674993] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (179740 kB) free!
[1564675010] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1564675040] HOST ALERT: sig4-vm2;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1564675114] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180688 kB) free!
[1564675116] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.06, 0.05
[1564675133] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1564675157] SERVICE ALERT: sig5-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1564675185] SERVICE ALERT: sig5-vm2;CPU_Procs;OK;HARD;1;CPU OK: 197 processes
[1564675199] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig4-vm2' was out of bounds)
[1564675234] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (854956 kB) free!
[1564675234] SERVICE ALERT: sig5-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 245368 MB (97% inode=99%):
[1564675237] SERVICE ALERT: sig5-vm2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1564675305] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7879400 kB) free!
[1564675315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564675425] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11824620 kB) free.
[1564675435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564675466] SERVICE ALERT: sig4-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig4-vm2' was out of bounds)
[1564675473] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig4-vm2' was out of bounds)
[1564675476] SERVICE ALERT: sig4-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4-vm2' was out of bounds)
[1564675504] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4-vm2' was out of bounds)
[1564675554] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4-vm2' was out of bounds)
[1564675607] HOST ALERT: sig4-vm2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.191)
[1564675911] HOST ALERT: sig4-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1564675919] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 214 processes
[1564676025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7881440 kB) free!
[1564676036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564676146] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11816776 kB) free.
[1564676156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1564676186] SERVICE ALERT: sig4-vm2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1564676193] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.01, 0.10, 0.08
[1564676196] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 115391 MB (89% inode=99%):
[1564676225] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 94.8% (33032252 kB) free.
[1564676227] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 19.10, 33.40, 16.13
[1564676274] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 227 processes
[1564676330] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1564676347] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.73, 18.15, 13.20
[1564676347] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.07, 22.58, 14.26
[1564676456] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 86.3% (14023536 kB) free.
[1564676467] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.64, 12.29, 11.65
[1564676467] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.82, 15.27, 12.59
[1564676547] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.4% (14370720 kB) free.
[1564676548] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1564676587] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.55, 8.40, 10.31
[1564676599] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.9% (21304280 kB) free.
[1564676637] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.3% (14357584 kB) free.
[1564676669] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1564676672] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 64.2% (31661492 kB) free.
[1564676712] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22426672 kB) free.
[1564676741] HOST ALERT: sig3-vm1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1564676745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7901672 kB) free!
[1564676751] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.3% (23273264 kB) free.
[1564676756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1564676762] SERVICE ALERT: sig3-vm2;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.30, 0.11
[1564676766] HOST ALERT: sig3-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1564676790] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 97.2% (22896884 kB) free.
[1564676859] SERVICE ALERT: sig3-vm1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1564676865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11822784 kB) free.
[1564676874] SERVICE ALERT: sig3-vm2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1564676876] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1564676876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564676884] SERVICE ALERT: sig3-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 90352 MB (75% inode=96%):
[1564676887] SERVICE ALERT: sig3-vm1;CPU Load;OK;HARD;1;OK - load average: 0.06, 0.26, 0.14
[1564676897] SERVICE ALERT: sig3-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1564676906] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.6% (22776508 kB) free.
[1564676925] SERVICE ALERT: sig3-vm1;Total Processes;OK;HARD;1;PROCS OK: 290 processes
[1564676948] SERVICE ALERT: sig3-vm2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564676992] SERVICE ALERT: sig3-vm2;Total Processes;OK;HARD;1;PROCS OK: 336 processes
[1564677114] SERVICE ALERT: sig3-vm1;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1564677186] SERVICE ALERT: sig3-vm1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1564677236] SERVICE ALERT: sig3-vm1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 60359 MB (50% inode=96%):
[1564677465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7920084 kB) free!
[1564677476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564677585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11823432 kB) free.
[1564677596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1564677668] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1564677668] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.58, 0.87, 3.88
[1564677759] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (2302120 kB) free!
[1564677777] SERVICE ALERT: sig2-vm2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1564677780] HOST ALERT: sig2-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.180)
[1564677787] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1564677787] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.68, 0.88, 3.35
[1564677879] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (1579656 kB) free!
[1564677900] SERVICE ALERT: sig2-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig2-vm2' was out of bounds)
[1564677910] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig2-vm2' was out of bounds)
[1564677910] SERVICE ALERT: sig2-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig2-vm2' was out of bounds)
[1564677999] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (182852 kB) free!
[1564678023] SERVICE ALERT: sig2-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2-vm2' was out of bounds)
[1564678072] SERVICE ALERT: sig2-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig2-vm2' was out of bounds)
[1564678130] SERVICE ALERT: sig2-vm2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig2-vm2' was out of bounds)
[1564678185] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7899472 kB) free!
[1564678196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1564678253] SERVICE ALERT: sig2-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig2-vm2' was out of bounds)
[1564678305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11821336 kB) free.
[1564678316] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564678347] HOST ALERT: sig2-vm2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.180)
[1564678743] SERVICE ALERT: sig2-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1564678747] HOST ALERT: sig2-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1564678747] SERVICE ALERT: smscapp1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1564678764] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'smscapp1' was out of bounds)
[1564678777] HOST ALERT: smscapp1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1564678792] SERVICE ALERT: sig2-vm2;CPU_Procs;OK;HARD;1;CPU OK: 223 processes
[1564678850] SERVICE ALERT: sig2-vm2;CPU Load;OK;HARD;1;OK - load average: 0.25, 0.30, 0.13
[1564678870] SERVICE ALERT: smscapp1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'smscapp1' was out of bounds)
[1564678887] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'smscapp1' was out of bounds)
[1564678905] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7820100 kB) free!
[1564678916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1564678973] SERVICE ALERT: sig2-vm2;Total Processes;OK;HARD;1;PROCS OK: 214 processes
[1564679025] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11820876 kB) free.
[1564679037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1564679046] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'smscapp1' was out of bounds)
[1564679081] HOST ALERT: smscapp1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1564679082] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (257376 kB) free!
[1564679097] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 120633 MB (87% inode=99%):
[1564679166] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 97.9% (48245332 kB) free.
[1564679202] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (266552 kB) free!
[1564679205] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1564679230] SERVICE ALERT: sig2-vm2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1564679230] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 97.0% (33794828 kB) free.
[1564679322] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250912 kB) free!
[1564679325] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1564679423] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1564679426] HOST ALERT: smscapp2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.54)
[1564679444] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'smscapp2' was out of bounds)
[1564679445] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1564679467] SERVICE ALERT: smscapp1;Sensors Status;OK;HARD;1;SENSORS OK
[1564679470] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'smscapp2' was out of bounds)
[1564679470] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'smscapp2' was out of bounds)
[1564679479] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'smscapp2' was out of bounds)
[1564679484] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 127551 MB (56% inode=99%):
[1564679546] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'smscapp2' was out of bounds)
[1564679546] SERVICE ALERT: smscapp2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'smscapp2' was out of bounds)
[1564679625] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7671000 kB) free!
[1564679636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564679642] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'smscapp2' was out of bounds)
[1564679712] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'smscapp2' was out of bounds)
[1564679742] HOST ALERT: smscapp2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1564679745] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11813208 kB) free.
[1564679756] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1564679762] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38902 MB (76% inode=98%):
[1564679833] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1564680143] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 0.38, 0.27, 0.13
[1564680164] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 293 processes
[1564680190] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 159758 MB (71% inode=99%):
[1564680190] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1564680199] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1564680199] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 97.8% (48179776 kB) free.
[1564680266] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1564680356] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1564680377] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 293 processes
[1564680476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1564680497] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1564680596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1564681123] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40746080 kB) free.
[1564681294] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8.55, 28.40, 16.70
[1564681414] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.46, 19.87, 14.99
[1564681534] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.78, 14.17, 13.50
[1564681600] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 54.6% (8875068 kB) free.
[1564682145] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8037544 kB) free!
[1564682200] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167456 kB) free!
[1564682265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11820180 kB) free.
[1564682320] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174768 kB) free!
[1564682323] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (212024 kB) free!
[1564682397] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1564682440] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1564682440] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157972 kB) free!
[1564682443] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241908 kB) free!
[1564682563] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1564682563] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237772 kB) free!
[1564682587] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 40.88, 31.36, 13.28
[1564682708] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.83, 21.09, 11.70
[1564682828] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.55, 14.40, 10.39
[1564682865] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8300820 kB) free!
[1564682985] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6499348 kB) free!
[1564683105] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11816380 kB) free.
[1564683334] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.48, 3.02, 4.36
[1564683597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1564683705] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6999284 kB) free!
[1564683717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1564683825] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11814852 kB) free.
[1564683837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1564684028] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.90, 1.08, 3.42