[1560969029] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560969106] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560969149] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1560969226] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560969270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1560969295] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232648 kB) free!
[1560969346] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560969388] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2760052 kB) free!
[1560969466] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560969495] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.46, 2.93
[1560969508] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.4% (11517816 kB) free.
[1560969543] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 319.09, 209.73, 93.67
[1560969586] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560969663] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 320.61, 246.49, 121.21
[1560969706] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560969783] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 321.03, 271.15, 145.44
[1560969826] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560969946] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560969951] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2766468 kB) free!
[1560970053] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.25, 2.03, 3.89
[1560970066] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560970186] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560970306] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560970426] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560970489] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2453344 kB) free!
[1560970546] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560970666] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560970708] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4696584 kB) free!
[1560970786] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560970828] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2950284 kB) free!
[1560970906] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560970948] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.5% (11603540 kB) free.
[1560971026] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560971146] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560971267] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560971387] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560971507] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560971627] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560971671] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1560971695] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13258032 kB) free.
[1560971747] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560971867] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560971987] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560972107] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560972148] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4468604 kB) free!
[1560972183] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1560972183] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.82, 1.83, 13.48
[1560972227] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560972268] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.0% (12318444 kB) free.
[1560972270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1560972295] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (1124392 kB) free!
[1560972347] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560972390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560972415] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (323660 kB) free!
[1560972467] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560972510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1560972535] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231740 kB) free!
[1560972587] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560972707] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560972827] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560972868] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4399852 kB) free!
[1560972947] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560972988] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.0% (12304548 kB) free.
[1560973067] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560973187] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560973307] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560973427] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560973547] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560973589] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4081504 kB) free!
[1560973667] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560973709] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.0% (12314216 kB) free.
[1560973787] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560973907] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560973983] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.71, 1.63, 3.41
[1560974027] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560974147] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560974151] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2432396 kB) free!
[1560974267] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560974309] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3925408 kB) free!
[1560974310] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1560974335] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12054196 kB) free.
[1560974387] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560974429] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.8% (12236404 kB) free.
[1560974507] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560974627] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560974747] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560974867] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560974910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1560974935] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (228604 kB) free!
[1560974987] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560975029] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3655188 kB) free!
[1560975055] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236740 kB) free!
[1560975107] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560975149] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.2% (11937000 kB) free.
[1560975175] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226388 kB) free!
[1560975202] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560975227] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560975347] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560975467] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560975587] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560975707] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560975749] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3532992 kB) free!
[1560975827] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560975869] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.8% (11749572 kB) free.
[1560975947] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560976067] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560976187] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560976307] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560976427] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560976469] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3478636 kB) free!
[1560976547] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560976589] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.7% (11675880 kB) free.
[1560976667] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560976711] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1560976787] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560976907] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560976975] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7723236 kB) free!
[1560977027] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560977147] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560977189] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3360916 kB) free!
[1560977267] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560977309] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.4% (11535492 kB) free.
[1560977311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1560977387] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560977431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560977507] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560977551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1560977575] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (279916 kB) free!
[1560977627] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560977747] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560977867] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560977987] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560978107] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560978183] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 106.38, 76.04, 36.12
[1560978227] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560978286] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5372 MB (10% inode=97%):
[1560978289] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1560978303] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 106.48, 86.04, 44.62
[1560978347] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560978423] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 89.94, 89.16, 50.93
[1560978467] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560978510] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4469456 kB) free!
[1560978587] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560978630] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.0% (12320236 kB) free.
[1560978707] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560978751] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1560978776] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10704120 kB) free.
[1560978827] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560978886] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7020 MB (13% inode=98%):
[1560978947] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560978951] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (3047908 kB) free!
[1560979067] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560979095] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 64.30, 47.49, 22.64
[1560979187] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560979215] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.77, 37.25, 22.11
[1560979230] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4119472 kB) free!
[1560979307] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560979335] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.37, 25.01, 19.46
[1560979350] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.7% (12149712 kB) free.
[1560979351] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1560979365] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 591 processes
[1560979376] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (370240 kB) free!
[1560979428] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560979485] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 321 processes
[1560979489] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2696552 kB) free!
[1560979496] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (295064 kB) free!
[1560979548] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560979616] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217764 kB) free!
[1560979623] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.60, 3.34, 15.29
[1560979668] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560979788] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560979908] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560979936] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.07, 3.77, 10.38
[1560979950] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3620044 kB) free!
[1560980028] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560980070] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.7% (11656688 kB) free.
[1560980148] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560980216] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11516512 kB) free.
[1560980253] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560980268] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560980388] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560980508] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560980552] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1560980628] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560980748] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560980816] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5172188 kB) free!
[1560980868] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560980936] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (4429304 kB) free!
[1560980988] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560981056] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3595100 kB) free!
[1560981108] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560981136] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.28, 0.38, 3.07
[1560981152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1560981228] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560981270] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4548292 kB) free!
[1560981348] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560981390] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.1% (12346820 kB) free.
[1560981424] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.91, 1.74, 3.69
[1560981468] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560981588] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560981708] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560981828] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560981948] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560981990] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4133640 kB) free!
[1560982068] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560982110] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.5% (12083152 kB) free.
[1560982188] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560982308] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560982428] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560982489] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2401156 kB) free!
[1560982548] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560982668] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560982710] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3605896 kB) free!
[1560982788] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560982830] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.3% (11503732 kB) free.
[1560982856] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11904760 kB) free.
[1560982908] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560982952] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1560983028] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560983148] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560983268] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560983388] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560983456] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232092 kB) free!
[1560983508] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560983552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1560983576] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226044 kB) free!
[1560983628] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560983672] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1560983696] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (221456 kB) free!
[1560983748] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560983751] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2440472 kB) free!
[1560983792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560983868] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560983988] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560984030] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4464940 kB) free!
[1560984108] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560984150] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.9% (12274136 kB) free.
[1560984228] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560984348] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560984468] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560984588] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560984708] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560984750] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3916508 kB) free!
[1560984828] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560984870] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.9% (11778240 kB) free.
[1560984948] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560984992] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1560985068] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560985188] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560985308] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560985428] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560985496] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6049052 kB) free!
[1560985548] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560985592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560985668] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560985712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1560985788] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560985832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560985908] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560986028] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560986070] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4898816 kB) free!
[1560986097] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214740 kB) free!
[1560986148] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560986190] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.9% (12289332 kB) free.
[1560986268] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560986388] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560986508] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560986628] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560986748] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560986790] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4344656 kB) free!
[1560986853] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.18, 33.48, 17.36
[1560986868] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560986910] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.5% (12093124 kB) free.
[1560986973] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.18, 23.56, 15.70
[1560986988] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560987032] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1560987093] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.98, 16.48, 14.06
[1560987108] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560987228] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560987298] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7724044 kB) free!
[1560987348] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560987468] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560987510] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4016620 kB) free!
[1560987588] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560987630] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.0% (11812320 kB) free.
[1560987632] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560987709] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560987753] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1560987829] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560987873] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1560987898] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241968 kB) free!
[1560987949] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560988069] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560988189] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560988309] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560988429] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560988549] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560988669] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560988789] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560988831] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4619896 kB) free!
[1560988893] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.16, 2.18, 3.86
[1560988909] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560988951] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.0% (12300512 kB) free.
[1560989029] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560989073] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1560989098] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7507420 kB) free!
[1560989149] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560989269] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560989389] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560989509] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560989551] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4186812 kB) free!
[1560989629] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560989671] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.1% (11881620 kB) free.
[1560989673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1560989698] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233100 kB) free!
[1560989749] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560989869] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560989989] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560990109] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560990229] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560990349] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560990469] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560990589] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560990709] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560990829] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560990871] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4476500 kB) free!
[1560990898] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.8% (7774236 kB) free!
[1560990949] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560990991] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.9% (12286928 kB) free.
[1560991069] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560991189] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560991309] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560991429] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560991498] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212936 kB) free!
[1560991549] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560991591] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3462408 kB) free!
[1560991592] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 384 processes
[1560991624] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560991669] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560991711] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.5% (12089604 kB) free.
[1560991713] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 383 processes
[1560991789] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560991893] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 289.47, 222.94, 110.10
[1560991909] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560991936] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 262.30, 232.42, 122.80
[1560991991] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 43.4% (10627576 kB) free.
[1560992013] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 288.76, 244.71, 131.76
[1560992029] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560992057] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.46, 155.56, 107.93
[1560992085] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 692 processes
[1560992133] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 286.27, 258.49, 150.52
[1560992149] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560992177] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.16, 104.24, 94.90
[1560992190] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 56.5% (9178168 kB) free.
[1560992206] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 691 processes
[1560992269] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560992311] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2748356 kB) free!
[1560992326] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 325 processes
[1560992389] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560992431] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.5% (11587160 kB) free.
[1560992509] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560992591] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4738016 kB) free!
[1560992629] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560992651] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4735912 kB) free!
[1560992711] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.3% (4738628 kB) free!
[1560992749] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560992771] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.3% (4731980 kB) free!
[1560992869] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560992989] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560993109] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560993229] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560993349] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560993469] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560993589] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560993631] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3209152 kB) free!
[1560993709] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560993751] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.9% (11765044 kB) free.
[1560993829] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560993874] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 589 processes
[1560993934] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 597 processes
[1560993949] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560993977] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.35, 0.83, 14.03
[1560994069] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560994189] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560994234] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 590 processes
[1560994295] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 587 processes
[1560994309] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560994352] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2507604 kB) free!
[1560994429] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560994472] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.5% (11101452 kB) free.
[1560994474] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1560994498] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.0% (9373172 kB) free!
[1560994533] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1560994533] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.21, 2.57, 18.28
[1560994549] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560994669] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560994789] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560994909] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560995029] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560995068] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.7% (2069568 kB) free!
[1560995074] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1560995098] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238084 kB) free!
[1560995149] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560995177] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.09, 0.24, 4.03
[1560995269] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560995389] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560995509] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560995629] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560995672] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3754224 kB) free!
[1560995674] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1560995699] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10711860 kB) free.
[1560995749] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560995792] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.6% (11650680 kB) free.
[1560995869] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560995989] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560996109] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560996229] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560996277] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1560996299] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (225236 kB) free!
[1560996334] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.05, 2.14, 4.51
[1560996350] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560996392] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3003836 kB) free!
[1560996397] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560996419] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234128 kB) free!
[1560996470] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560996512] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.5% (11074808 kB) free.
[1560996517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1560996539] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214768 kB) free!
[1560996590] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560996710] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560996830] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560996950] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560997070] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560997118] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 334 processes
[1560997139] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5899740 kB) free!
[1560997190] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560997310] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560997430] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560997550] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560997670] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560997712] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3445356 kB) free!
[1560997718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1560997739] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (304516 kB) free!
[1560997790] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560997832] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.0% (11310932 kB) free.
[1560997910] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560998030] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560998068] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1620272 kB) free!
[1560998150] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560998224] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 140.12, 66.75, 27.06
[1560998270] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560998318] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1560998340] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10519028 kB) free.
[1560998344] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 146.91, 93.68, 41.75
[1560998390] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560998432] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2602492 kB) free!
[1560998464] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 147.64, 111.63, 54.63
[1560998510] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560998552] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.1% (10871280 kB) free.
[1560998630] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560998727] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 724 processes
[1560998750] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560998847] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 707 processes
[1560998870] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560998918] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1560998942] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (258812 kB) free!
[1560998967] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 361 processes
[1560998990] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560999038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560999062] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244616 kB) free!
[1560999110] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560999158] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1560999182] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236880 kB) free!
[1560999230] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560999350] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560999470] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560999590] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560999710] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560999753] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3125036 kB) free!
[1560999758] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1560999782] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11047052 kB) free.
[1560999830] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1560999873] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.9% (11290812 kB) free.
[1560999950] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561000070] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561000112] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 379 processes
[1561000190] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561000233] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1561000310] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561000359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561000382] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (641012 kB) free!
[1561000430] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561000473] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2169860 kB) free!
[1561000479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561000502] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (380088 kB) free!
[1561000550] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561000593] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 21.4% (10533256 kB) free.
[1561000599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561000622] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (285840 kB) free!
[1561000670] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561000790] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561000910] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561001030] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561001150] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561001151] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (2730648 kB) free!
[1561001270] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561001390] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561001510] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561001630] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561001750] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561001793] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2371580 kB) free!
[1561001799] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1561001823] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.2% (7993416 kB) free!
[1561001870] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561001913] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.7% (11172052 kB) free.
[1561001990] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561002110] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561002230] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561002350] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561002351] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1925576 kB) free!
[1561002399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561002423] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226116 kB) free!
[1561002470] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561002513] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (901256 kB) free!
[1561002590] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561002633] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 21.4% (10569280 kB) free.
[1561002664] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.77, 2.18, 17.49
[1561002710] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561002830] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561002950] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561003070] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561003190] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561003233] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (293500 kB) free!
[1561003310] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561003353] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9403088 kB) free!
[1561003430] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561003473] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 14.4% (7074340 kB) free!
[1561003550] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561003670] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561003790] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561003910] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561004030] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561004073] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 21.7% (10672972 kB) free.
[1561004150] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561004199] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 339 processes
[1561004270] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561004390] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561004464] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.00, 2.07, 4.39
[1561004510] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561004514] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1561004630] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561004634] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1561004673] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (661620 kB) free!
[1561004751] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561004754] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1561004794] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 21.4% (10568544 kB) free.
[1561004871] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561004991] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561005111] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561005231] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561005351] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561005393] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240904 kB) free!
[1561005471] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561005513] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.7% (11667336 kB) free.
[1561005591] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561005711] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561005831] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561005951] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561005977] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 25.86, 14.73, 6.37
[1561006071] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561006097] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 25.41, 18.25, 8.68
[1561006113] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (1170208 kB) free!
[1561006126] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 580 processes
[1561006191] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561006217] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.50, 17.95, 9.82
[1561006233] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.5% (11596324 kB) free.
[1561006246] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 582 processes
[1561006253] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.3% (14836068 kB) free.
[1561006299] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.1% (22338088 kB) free.
[1561006311] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561006346] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.8% (22751212 kB) free.
[1561006366] SERVICE ALERT: sservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 582 processes
[1561006422] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 69.8% (17235456 kB) free.
[1561006431] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561006485] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.0% (14622988 kB) free.
[1561006534] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561006551] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561006594] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 606 processes
[1561006599] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561006623] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11080096 kB) free.
[1561006654] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 605 processes
[1561006658] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.83, 6.22, 5.62
[1561006671] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561006778] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.28, 5.63, 5.48
[1561006791] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561006833] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (999524 kB) free!
[1561006871] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.9% (2415632 kB) free!
[1561006898] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.59, 5.40, 5.42
[1561006911] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561006953] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.7% (11700084 kB) free.
[1561006966] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 326 processes
[1561007002] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9556900 kB) free!
[1561007027] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184984 kB) free!
[1561007031] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561007122] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 20.1% (9925712 kB) free.
[1561007147] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (148172 kB) free!
[1561007151] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561007199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561007223] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (1617628 kB) free!
[1561007267] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196656 kB) free!
[1561007271] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561007319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1561007344] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (1312968 kB) free!
[1561007391] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561007417] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.68, 0.98, 3.16
[1561007440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561007464] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228344 kB) free!
[1561007471] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 13.2% (3226612 kB) free!
[1561007511] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561007553] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (1210208 kB) free!
[1561007631] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561007673] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244152 kB) free!
[1561007722] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3046076 kB) free!
[1561007751] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561007790] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (152876 kB) free!
[1561007793] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.5% (11074120 kB) free.
[1561007842] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (2919328 kB) free!
[1561007871] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561007910] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176176 kB) free!
[1561007962] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (959540 kB) free!
[1561007985] SERVICE ALERT: redis-1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1561007991] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561008030] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178276 kB) free!
[1561008104] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168548 kB) free!
[1561008111] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561008224] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 80.4% (19695072 kB) free.
[1561008231] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561008351] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561008393] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (598656 kB) free!
[1561008461] SERVICE ALERT: redis-1;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.01, 0.05
[1561008471] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561008513] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.0% (11356916 kB) free.
[1561008591] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561008640] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1561008711] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561008831] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561008951] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561009071] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561009113] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (348548 kB) free!
[1561009173] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 97.3% (63984300 kB) free.
[1561009191] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561009233] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.7% (12170444 kB) free.
[1561009240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561009311] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561009386] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 1069171 MB (99% inode=99%):
[1561009431] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561009485] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (170644 kB) free!
[1561009494] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561009508] SERVICE ALERT: redis-1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45722 MB (89% inode=99%):
[1561009540] SERVICE ALERT: redis-1;Sensors Status;OK;HARD;1;SENSORS OK
[1561009551] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561009605] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173376 kB) free!
[1561009671] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561009693] SERVICE ALERT: redis-1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 373128 MB (99% inode=99%):
[1561009702] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1561009717] SERVICE ALERT: redis-1;Total Processes;OK;HARD;1;PROCS OK: 339 processes
[1561009725] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174776 kB) free!
[1561009791] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561009834] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (1670544 kB) free!
[1561009911] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561009954] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234664 kB) free!
[1561010031] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561010074] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 21.1% (10402208 kB) free.
[1561010151] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561010271] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561010391] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561010511] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561010631] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1561010674] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (242264 kB) free!
[1561010751] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;NRPE: Command 'check_cpu_procs' not defined
[1561010794] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (8692904 kB) free!
[1561010871] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;3;NRPE: Command 'check_cpu_procs' not defined
[1561010914] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 12.9% (6371964 kB) free!
[1561010925] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 57.5% (9348232 kB) free.
[1561010936] SERVICE ALERT: redis-2;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.04, 0.07
[1561010953] SERVICE ALERT: redis-2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1561010959] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 96.2% (63259208 kB) free.
[1561010991] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1561011015] SERVICE ALERT: redis-2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 373128 MB (99% inode=99%):
[1561011074] SERVICE ALERT: redis-2;Sensors Status;OK;HARD;1;SENSORS OK
[1561011077] SERVICE ALERT: redis-2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45541 MB (88% inode=99%):
[1561011146] SERVICE ALERT: redis-2;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 1069171 MB (99% inode=99%):
[1561011221] SERVICE ALERT: redis-2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1561011243] SERVICE ALERT: redis-2;Total Processes;OK;HARD;1;PROCS OK: 369 processes
[1561011641] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1561012033] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1561012114] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1561012114] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.8% (12245708 kB) free.
[1561012153] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1561012259] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (160172 kB) free!
[1561012272] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 8.2% (2010764 kB) free!
[1561012273] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1561012379] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (179676 kB) free!
[1561012492] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 369 processes
[1561012499] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166560 kB) free!
[1561012572] SERVICE FLAPPING ALERT: webserver2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1561012572] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 18.3% (4485504 kB) free!
[1561012612] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1561012714] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (1351492 kB) free!
[1561012834] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217124 kB) free!
[1561012955] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.7% (12645716 kB) free.
[1561013212] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 370 processes
[1561013332] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 369 processes
[1561013441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561013452] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1561013464] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5019112 kB) free!
[1561013554] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215948 kB) free!
[1561013674] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.5% (13039272 kB) free.
[1561014064] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243300 kB) free!
[1561014098] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.92, 4.12, 4.99
[1561014274] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (1869152 kB) free!
[1561014394] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238396 kB) free!
[1561014514] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.6% (12132800 kB) free.
[1561014525] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (3146772 kB) free!
[1561014645] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.2% (2965776 kB) free!
[1561014652] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 369 processes
[1561014672] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 21.6% (5296328 kB) free.
[1561014699] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.58, 4.98, 5.24
[1561014765] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1561014765] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.9% (2742244 kB) free!
[1561014773] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1561014819] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.50, 5.47, 5.40
[1561014939] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.30, 5.07, 5.27
[1561015114] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (294368 kB) free!
[1561015234] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.9% (13276204 kB) free.
[1561015356] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195960 kB) free!
[1561015372] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 371 processes
[1561015424] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4813480 kB) free!
[1561015476] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166852 kB) free!
[1561015492] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 371 processes
[1561015544] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193468 kB) free!
[1561015596] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179824 kB) free!
[1561015612] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 370 processes
[1561015664] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183944 kB) free!
[1561015834] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (1286924 kB) free!
[1561015954] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.1% (13349564 kB) free.
[1561015966] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (755280 kB) free!
[1561016212] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1561016554] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1753024 kB) free!
[1561016675] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (221976 kB) free!
[1561016795] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.4% (13524100 kB) free.
[1561016872] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 63 users currently logged in
[1561016932] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 63 users currently logged in
[1561016992] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 64 users currently logged in
[1561017052] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1561017064] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11420772 kB) free.
[1561017073] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1561017395] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (1156648 kB) free!
[1561017515] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232668 kB) free!
[1561017635] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.3% (12492272 kB) free.
[1561017641] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561018012] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 376 processes
[1561018133] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 376 processes
[1561018217] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561018235] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (493600 kB) free!
[1561018253] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 376 processes
[1561018355] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.1% (13861716 kB) free.
[1561018572] SERVICE FLAPPING ALERT: webserver2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1561018841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561018848] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (1470768 kB) free!
[1561018853] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1561018853] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1561018955] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (2216032 kB) free!
[1561018961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561018968] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (1303104 kB) free!
[1561019075] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (221072 kB) free!
[1561019081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561019088] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 15.1% (1107548 kB) free!
[1561019195] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.3% (13478292 kB) free.
[1561019452] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1561019689] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (228496 kB) free!
[1561019752] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1561019795] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (1535296 kB) free!
[1561019915] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.0% (13777668 kB) free.
[1561020064] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9631648 kB) free!
[1561020184] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7955312 kB) free!
[1561020305] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6301840 kB) free!
[1561020515] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2657124 kB) free!
[1561020635] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (537800 kB) free!
[1561020756] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.9% (13756880 kB) free.
[1561020939] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.34, 4.17, 4.48
[1561021252] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1561021252] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1561021273] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1561021355] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (1889152 kB) free!
[1561021475] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.7% (13635988 kB) free.
[1561022075] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (2311168 kB) free!
[1561022082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561022105] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11404616 kB) free.
[1561022195] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (218500 kB) free!
[1561022315] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.3% (12473188 kB) free.
[1561022417] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 214.36, 130.41, 55.81
[1561022473] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1561022473] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 16 users currently logged in
[1561022537] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 215.91, 158.79, 75.26
[1561022657] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 73.16, 135.56, 77.75
[1561022864] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2675296 kB) free!
[1561022915] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (826816 kB) free!
[1561023035] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.5% (13532244 kB) free.
[1561023282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561023289] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.3% (6790360 kB) free.
[1561023402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561023464] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2260144 kB) free!
[1561023522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561023630] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 42.2% (6858528 kB) free.
[1561023635] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (1922768 kB) free!
[1561023755] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234620 kB) free!
[1561023875] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.2% (12415796 kB) free.
[1561023934] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 219.47, 122.97, 51.99
[1561024054] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 224.66, 156.63, 72.91
[1561024174] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 224.29, 179.00, 91.26
[1561024253] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 385 processes
[1561024373] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 385 processes
[1561024457] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.37, 0.77, 11.61
[1561024475] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (741776 kB) free!
[1561024493] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 384 processes
[1561024505] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4273280 kB) free!
[1561024539] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.02, 5.85, 5.43
[1561024595] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.4% (13488100 kB) free.
[1561024625] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11067092 kB) free.
[1561024659] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.12, 6.61, 5.77
[1561024779] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.60, 5.98, 5.66
[1561024830] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (437432 kB) free!
[1561024864] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 118.20, 102.10, 53.32
[1561024950] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171924 kB) free!
[1561024954] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 623 processes
[1561024984] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 91.70, 101.38, 59.19
[1561025014] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 624 processes
[1561025070] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160044 kB) free!
[1561025093] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 381 processes
[1561025104] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 13.68, 68.37, 52.21
[1561025196] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2011040 kB) free!
[1561025196] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.8% (2645644 kB) free!
[1561025316] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233876 kB) free!
[1561025379] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1561025379] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.36, 4.27, 4.91
[1561025436] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.0% (12317788 kB) free.
[1561025658] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 62.14, 33.79, 17.04
[1561025752] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1561025801] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189384 kB) free!
[1561025825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9690580 kB) free!
[1561025945] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (8007760 kB) free!
[1561026036] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (857152 kB) free!
[1561026052] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1561026065] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6246980 kB) free!
[1561026112] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1561026156] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561026156] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.5% (13072864 kB) free.
[1561026172] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1561026232] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1561026258] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1561026258] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 14.46, 18.27
[1561026293] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 384 processes
[1561026304] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.16, 3.52, 16.15
[1561026414] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 384 processes
[1561026534] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 384 processes
[1561026574] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.99, 4.65, 12.88
[1561026756] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (630148 kB) free!
[1561026766] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561026876] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.1% (9921600 kB) free.
[1561027134] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 382 processes
[1561027179] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.13, 5.53, 5.18
[1561027299] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.80, 6.07, 5.43
[1561027419] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.48, 5.83, 5.42
[1561027477] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233920 kB) free!
[1561027597] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240372 kB) free!
[1561027605] SERVICE ALERT: sigqscluster1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23112 MB (10% inode=99%):
[1561027717] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (9774156 kB) free!
[1561027722] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561027865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11197720 kB) free.
[1561028020] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.11, 4.16, 4.85
[1561028058] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.93, 0.50, 2.99
[1561028104] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.67, 2.05, 4.10
[1561028205] SERVICE ALERT: sigqscluster1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25248 MB (11% inode=99%):
[1561028317] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232416 kB) free!
[1561028322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561028374] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.03, 3.19, 4.52
[1561028442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561028562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561028934] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 383 processes
[1561029054] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 384 processes
[1561029162] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561029405] SERVICE ALERT: sigqscluster1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23615 MB (10% inode=99%):
[1561029517] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (7550844 kB) free!
[1561029665] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4159568 kB) free!
[1561029785] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11357580 kB) free.
[1561030006] SERVICE FLAPPING ALERT: sigqscluster1;HDD_Home status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1561030006] SERVICE ALERT: sigqscluster1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25661 MB (11% inode=99%):
[1561030117] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1561030117] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.1% (10901680 kB) free.
[1561030254] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 389 processes
[1561030362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561030374] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 385 processes
[1561030385] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4750104 kB) free!
[1561030482] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561030495] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 383 processes
[1561030505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11341776 kB) free.
[1561030606] SERVICE ALERT: sigqscluster1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23072 MB (10% inode=99%):
[1561031082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561031105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5152888 kB) free!
[1561031202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561031225] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (3040076 kB) free!
[1561031317] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (1990992 kB) free!
[1561031322] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561031345] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10881580 kB) free.
[1561031437] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216960 kB) free!
[1561031557] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.0% (12313460 kB) free.
[1561031694] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 386 processes
[1561031814] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 386 processes
[1561031922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561031934] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 387 processes
[1561031974] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 70.37, 47.46, 22.60
[1561032033] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 380 processes
[1561032042] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561032094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 70.90, 55.23, 28.45
[1561032153] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1561032157] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (1246564 kB) free!
[1561032162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561032214] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1561032214] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.11, 49.89, 30.03
[1561032277] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.5% (13064384 kB) free.
[1561032289] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.6% (1071192 kB) free!
[1561032409] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.5% (994360 kB) free!
[1561032529] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.3% (903576 kB) free!
[1561032534] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 381 processes
[1561032545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9511504 kB) free!
[1561032665] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7738388 kB) free!
[1561032785] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6084236 kB) free!
[1561032814] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.74, 8.89, 16.93
[1561032878] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (1642020 kB) free!
[1561032998] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.7% (13171984 kB) free.
[1561033129] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (454324 kB) free!
[1561033273] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561033598] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3712860 kB) free!
[1561033718] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (1841116 kB) free!
[1561033734] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 388 processes
[1561033838] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.8% (12232160 kB) free.
[1561033854] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 390 processes
[1561033962] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561033974] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 387 processes
[1561033985] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332628 kB) free.
[1561034330] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1561034330] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.4% (6865976 kB) free.
[1561034438] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (2239808 kB) free!
[1561034558] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.6% (12612924 kB) free.
[1561034575] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 382 processes
[1561034614] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.30, 2.45, 4.57
[1561035158] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2764016 kB) free!
[1561035162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561035185] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4225412 kB) free!
[1561035278] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.9% (12757672 kB) free.
[1561035282] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561035305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11325460 kB) free.
[1561035404] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.0% (18744104 kB) free.
[1561035445] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.1% (22580340 kB) free.
[1561035537] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 67.7% (33364084 kB) free.
[1561035766] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.4% (14699960 kB) free.
[1561035835] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.7% (21996192 kB) free.
[1561035878] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3037568 kB) free!
[1561035882] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561035905] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5680048 kB) free!
[1561035961] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1561035998] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.2% (12924984 kB) free.
[1561036002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561036007] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.8% (2674272 kB) free!
[1561036025] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3497988 kB) free!
[1561036081] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1561036123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561036128] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (954972 kB) free!
[1561036145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11325700 kB) free.
[1561036201] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1561036248] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203524 kB) free!
[1561036318] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.0% (18744552 kB) free.
[1561036365] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.8% (14915576 kB) free.
[1561036375] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 379 processes
[1561036435] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 67.6% (10995888 kB) free.
[1561036495] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1561036598] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3472548 kB) free!
[1561036718] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.1% (12875940 kB) free.
[1561036722] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561036732] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 59 users currently logged in
[1561036746] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4491676 kB) free!
[1561036801] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1561036842] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561036866] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11316596 kB) free.
[1561036918] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4571688 kB) free!
[1561037043] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177392 kB) free!
[1561037163] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1561037163] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195592 kB) free!
[1561037318] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3835060 kB) free!
[1561037439] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.1% (12857512 kB) free.
[1561037442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561037466] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5894240 kB) free!
[1561037562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561037570] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173220 kB) free!
[1561037586] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (3987896 kB) free!
[1561037682] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561037690] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (172236 kB) free!
[1561037706] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11310252 kB) free.
[1561037810] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166496 kB) free!
[1561037932] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1561037992] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1561038038] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2428084 kB) free!
[1561038052] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1561038112] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1561038158] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2946628 kB) free!
[1561038258] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561038278] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.4% (11544240 kB) free.
[1561038282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561038306] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5018924 kB) free!
[1561038402] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561038426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11327736 kB) free.
[1561038878] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (584988 kB) free!
[1561038998] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.5% (11564572 kB) free.
[1561039003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561039012] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 57 users currently logged in
[1561039026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5797652 kB) free!
[1561039123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561039146] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11331672 kB) free.
[1561039420] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561039598] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (760196 kB) free!
[1561039718] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.0% (11839132 kB) free.
[1561039723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561039746] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6600704 kB) free!
[1561039843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561039866] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4589268 kB) free!
[1561039963] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561039986] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332112 kB) free.
[1561040318] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (1695568 kB) free!
[1561040438] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.3% (11964328 kB) free.
[1561040563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561040587] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4252636 kB) free!
[1561040683] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561040707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11309664 kB) free.
[1561041038] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2146716 kB) free!
[1561041073] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1561041158] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.6% (12106640 kB) free.
[1561041283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561041307] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4083848 kB) free!
[1561041404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561041427] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11304260 kB) free.
[1561041758] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2180444 kB) free!
[1561041879] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.6% (12110604 kB) free.
[1561042003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561042006] SERVICE FLAPPING ALERT: sigqscluster1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561042027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5335064 kB) free!
[1561042123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1561042146] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9003000 kB) free!
[1561042147] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (3370972 kB) free!
[1561042243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561042266] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8864020 kB) free!
[1561042267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11293476 kB) free.
[1561042386] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 24.2% (11907728 kB) free.
[1561042478] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (1481612 kB) free!
[1561042598] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.7% (12192588 kB) free.
[1561042843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561042867] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4485468 kB) free!
[1561042963] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561042987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11315052 kB) free.
[1561043198] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2187884 kB) free!
[1561043318] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.6% (12141364 kB) free.
[1561043563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561043587] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5125476 kB) free!
[1561043683] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561043707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11305816 kB) free.
[1561043918] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2482724 kB) free!
[1561044038] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561044038] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.6% (12111024 kB) free.
[1561044283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561044307] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5584140 kB) free!
[1561044354] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1561044404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561044428] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11322368 kB) free.
[1561044638] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2752504 kB) free!
[1561044758] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.6% (12108488 kB) free.
[1561044814] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561044904] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 136.44, 60.70, 24.11
[1561045003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561045024] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 146.68, 89.60, 39.13
[1561045027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5899060 kB) free!
[1561045123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561045130] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561045144] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 147.79, 108.93, 52.34
[1561045147] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323760 kB) free.
[1561045358] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2603752 kB) free!
[1561045478] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.1% (11857320 kB) free.
[1561045723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561045747] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6359188 kB) free!
[1561045843] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561045867] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323628 kB) free.
[1561046078] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2360484 kB) free!
[1561046095] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561046198] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.2% (11417004 kB) free.
[1561046443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561046467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6268688 kB) free!
[1561046563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561046587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11326840 kB) free.
[1561046658] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 57.63, 166.82, 101.84
[1561046683] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561046683] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561046778] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.98, 111.68, 89.52
[1561046799] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2273960 kB) free!
[1561046898] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.40, 74.87, 78.71
[1561046919] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.0% (11314392 kB) free.
[1561046944] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.91, 2.34, 13.22
[1561047187] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7512632 kB) free!
[1561047284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1561047307] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5924760 kB) free!
[1561047404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561047427] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561047427] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11334472 kB) free.
[1561047954] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1561048004] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561048027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6365160 kB) free!
[1561048074] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1561048119] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3119208 kB) free!
[1561048124] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561048144] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.59, 1.79, 4.94
[1561048147] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11333044 kB) free.
[1561048194] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1561048239] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.0% (11814376 kB) free.
[1561048414] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 270.65, 119.60, 46.75
[1561048535] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 288.49, 176.49, 76.44
[1561048563] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561048655] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 290.37, 214.31, 102.43
[1561048699] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.14, 0.37, 11.50
[1561048724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561048748] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6279888 kB) free!
[1561048839] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2654044 kB) free!
[1561048844] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561048868] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11328240 kB) free.
[1561048959] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 22.9% (11272476 kB) free.
[1561049163] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2827672 kB) free!
[1561049231] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.0% (23202472 kB) free.
[1561049260] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.8% (14928008 kB) free.
[1561049444] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561049468] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6167432 kB) free!
[1561049564] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561049588] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11331208 kB) free.
[1561049899] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.33, 3.37
[1561050159] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3689812 kB) free!
[1561050165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561050188] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6547852 kB) free!
[1561050279] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.7% (12178456 kB) free.
[1561050285] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561050308] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11329800 kB) free.
[1561050455] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 1.06, 18.51
[1561050879] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3092064 kB) free!
[1561050885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561050908] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6993620 kB) free!
[1561050999] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.3% (11492696 kB) free.
[1561051005] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561051028] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11336296 kB) free.
[1561051125] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561051628] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7219892 kB) free!
[1561051725] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1561051748] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11341516 kB) free.
[1561051845] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561052200] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3632972 kB) free!
[1561052255] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.40, 3.00
[1561052320] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.7% (12191808 kB) free.
[1561052348] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7278040 kB) free!
[1561052445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561052468] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11328192 kB) free.
[1561052566] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561052920] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (3316540 kB) free!
[1561053040] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.4% (11534124 kB) free.
[1561053069] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7379648 kB) free!
[1561053165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561053189] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11341564 kB) free.
[1561053285] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561053499] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.44, 5.29, 5.03
[1561053619] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 3.58, 4.43
[1561053786] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9097448 kB) free!
[1561053789] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7468924 kB) free!
[1561053885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561053907] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8542148 kB) free!
[1561053909] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11329992 kB) free.
[1561054005] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561054027] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (8129352 kB) free!
[1561054115] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 595 processes
[1561054175] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 599 processes
[1561054240] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3914572 kB) free!
[1561054360] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.8% (12222008 kB) free.
[1561054509] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7520200 kB) free!
[1561054605] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561054629] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11341476 kB) free.
[1561054725] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561054960] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3157728 kB) free!
[1561055080] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.1% (11388928 kB) free.
[1561055229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7538528 kB) free!
[1561055325] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561055349] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11341600 kB) free.