[1544035002] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544035122] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544035242] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544037042] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544037162] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544037762] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544037882] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544041413] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1544042014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544042134] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544042254] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544046737] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1544046857] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1544053482] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544053602] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544053654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544053774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544053894] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1544056003] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544056123] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544056723] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544056843] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544056894] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544057014] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544057443] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544057563] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544059363] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544059483] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544060615] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544060735] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1544060855] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1544061283] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544061403] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544062003] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544062123] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1544063923] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544064043] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544064644] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544064764] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544065055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1544065175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1544065295] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544065299] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1544065420] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1544065540] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1544066564] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544066684] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544066804] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544067340] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 13 users currently logged in
[1544067404] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544067524] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544068125] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544068245] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544068365] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1544068895] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1544068965] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544069015] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1544069085] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544069205] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544069615] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1544069735] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1544069805] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544069925] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544070099] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.9% (23513432 kB) free.
[1544070335] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544070455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544070576] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1544071125] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544071245] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544071775] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544071798] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 298 processes
[1544071845] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544071895] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1544071918] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1544071965] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544072085] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544072495] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1544072615] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544072620] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.58, 13.32, 11.44
[1544072735] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544072740] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 14.29, 13.64, 11.79
[1544072860] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.61, 13.93, 12.11
[1544073285] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544073405] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544074005] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544074126] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544074318] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1544074439] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1544074956] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1544075260] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.52, 2.66, 7.53
[1544075556] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1544075676] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1544075735] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1544075796] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1544076239] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;1;USERS WARNING - 42 users currently logged in
[1544076299] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;2;USERS WARNING - 42 users currently logged in
[1544076303] SERVICE ALERT: localhost;Sensors Status;CRITICAL;SOFT;1;NRPE: Command 'check_sensors' not defined
[1544076336] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544076360] SERVICE ALERT: localhost;Current Users;WARNING;SOFT;3;USERS WARNING - 42 users currently logged in
[1544076363] SERVICE ALERT: localhost;Sensors Status;CRITICAL;SOFT;2;NRPE: Command 'check_sensors' not defined
[1544076419] SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 42 users currently logged in
[1544076423] SERVICE ALERT: localhost;Sensors Status;CRITICAL;SOFT;3;NRPE: Command 'check_sensors' not defined
[1544076456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544076483] SERVICE ALERT: localhost;Sensors Status;CRITICAL;HARD;4;NRPE: Command 'check_sensors' not defined
[1544076526] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544076548] SERVICE ALERT: localhost;Sensors Status;OK;HARD;1;SENSORS OK
[1544076576] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1544076597] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 2 warn out of 335 processes
[1544076646] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544076766] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544077308] SERVICE ALERT: localhost;Nginx_status;CRITICAL;SOFT;1;NRPE: Command 'check_nginx_status' not defined
[1544077368] SERVICE ALERT: localhost;Nginx_status;CRITICAL;SOFT;2;NRPE: Command 'check_nginx_status' not defined
[1544077428] SERVICE ALERT: localhost;Nginx_status;CRITICAL;SOFT;3;NRPE: Command 'check_nginx_status' not defined
[1544077438] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 303 processes
[1544077488] SERVICE ALERT: localhost;Nginx_status;CRITICAL;HARD;4;NRPE: Command 'check_nginx_status' not defined
[1544077558] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 303 processes
[1544077678] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 306 processes
[1544077776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1544077896] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1544077966] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544078086] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544078462] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 43 users currently logged in
[1544078522] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 43 users currently logged in
[1544078582] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 43 users currently logged in
[1544078593] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.4% (14533844 kB) free.
[1544078642] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 43 users currently logged in
[1544078642] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1544078644] SERVICE ALERT: webserver2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1544078713] SERVICE ALERT: webserver2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1544078723] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1544078731] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1544078777] SERVICE ALERT: webserver2;Sensors Status;OK;HARD;1;SENSORS OK
[1544078783] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 44.5% (10901956 kB) free.
[1544078791] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 498 processes
[1544078878] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 303 processes
[1544078997] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 337 processes
[1544079096] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544079216] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544079286] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544079337] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1544079406] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544079579] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 300 processes
[1544079598] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.9% (15427492 kB) free.
[1544079700] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 299 processes
[1544079820] SERVICE ALERT: sig1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 298 processes
[1544080079] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 302 processes
[1544080198] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 338 processes
[1544080420] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1544080536] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544080656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544080776] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1544080798] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1544080798] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 338 processes
[1544081279] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1544081279] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1544081998] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 1 warn out of 340 processes
[1544082576] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1544082697] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544082817] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544083006] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544083126] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544083417] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1544083417] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1544086337] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1544086592] SERVICE ALERT: sigqscluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1544086596] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1544087196] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1544087316] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1544087436] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1544087618] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544087738] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544087792] SERVICE ALERT: sigqscluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1544087879] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1544087912] SERVICE ALERT: sigqscluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1544087999] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1544088032] SERVICE ALERT: sigqscluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1544088119] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 302 processes
[1544088526] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544088646] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544088938] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1544089058] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544089178] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544089346] SERVICE ALERT: qservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 303 processes
[1544089420] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 299 processes
[1544089466] SERVICE ALERT: qservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 303 processes
[1544089540] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 297 processes
[1544089587] SERVICE ALERT: qservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 303 processes
[1544089660] SERVICE ALERT: sig1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 298 processes
[1544089938] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1544090058] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 338 processes
[1544090178] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1544090447] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544090567] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544090687] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544090905] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.65, 6.44, 6.61
[1544091025] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.75, 4.57, 5.91
[1544091146] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.05, 3.37, 5.30
[1544091746] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.71, 1.02, 3.09
[1544091887] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544092487] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544092503] SERVICE ALERT: primary-server;Nginx Status;CRITICAL;SOFT;1;NGINX CRITICAL - 404 Not Found
[1544092598] SERVICE ALERT: primary-server;Nginx Status;OK;HARD;1;NGINX OK -  0.084 sec. response time, Active: 0 (Writing: 0 Reading: 0 Waiting: 0) ReqPerSec: 0.000 ConnPerSec: 0.000 ReqPerConn: 0.000
[1544092607] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544092707] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4722428 kB) free!
[1544092727] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544092737] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3075384 kB) free!
[1544092827] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4682332 kB) free!
[1544092857] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (3044536 kB) free!
[1544092947] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4644448 kB) free!
[1544092977] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (2994804 kB) free!
[1544093178] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1544093378] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544094528] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1544094528] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544095179] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544095299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1544095346] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.03, 7.94, 8.21
[1544095466] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.82, 5.66, 7.35
[1544095586] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.17, 4.12, 6.58
[1544095947] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2139452 kB) free!
[1544096186] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1544096186] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.72, 1.61, 4.06
[1544096499] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544096547] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.0% (3196520 kB) free!
[1544096619] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544096928] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544097048] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1544097219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1544097339] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544097459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544097777] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1572976 kB) free!
[1544098248] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544098347] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1544098347] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1391076 kB) free!
[1544098369] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544099186] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.45, 18.35, 9.92
[1544099306] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.16, 12.58, 8.82
[1544099426] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.44, 8.89, 7.93
[1544099568] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544099594] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3122916 kB) free!
[1544099688] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544099714] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3034424 kB) free!
[1544099834] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.2% (2964260 kB) free!
[1544100026] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.74, 2.09, 4.66
[1544100778] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.3% (14996732 kB) free.
[1544101488] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544101608] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544102209] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544102329] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544102449] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544103049] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544103119] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1544103242] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 40 users currently logged in
[1544103460] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1544104060] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544104180] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544104849] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544104969] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544105042] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1544105102] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1544105162] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1544105222] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 41 users currently logged in
[1544105569] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544105690] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544107060] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.30, 10.84, 10.10
[1544107180] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.87, 11.17, 10.31
[1544107300] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.65, 11.43, 10.52
[1544107660] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1544107900] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 11.87, 8.99, 9.42
[1544108089] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544108209] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544108547] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1544109100] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.63, 11.87, 10.50
[1544109220] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.82, 11.89, 10.68
[1544109340] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.96, 12.06, 10.89
[1544110540] SERVICE FLAPPING ALERT: qservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1544110540] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.83, 7.01, 8.51
[1544111427] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1544113609] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544113729] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544114929] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544115050] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1544115050] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544117998] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544118598] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 340 processes
[1544118718] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 340 processes
[1544118838] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 338 processes