[1576262093] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18552456 kB) free!
[1576262213] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39950852 kB) free.
[1576262527] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199192 kB) free!
[1576262548] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194404 kB) free!
[1576262647] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186100 kB) free!
[1576262647] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185636 kB) free!
[1576262666] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182124 kB) free!
[1576262668] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183172 kB) free!
[1576262767] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191740 kB) free!
[1576262767] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187232 kB) free!
[1576262786] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183516 kB) free!
[1576262790] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190800 kB) free!
[1576262813] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19745992 kB) free!
[1576262887] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179628 kB) free!
[1576262906] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196888 kB) free!
[1576262933] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40412516 kB) free.
[1576263534] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20127432 kB) free!
[1576263654] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40388780 kB) free.
[1576263692] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (2200864 kB) free!
[1576263753] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171532 kB) free!
[1576263813] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (168620 kB) free!
[1576263873] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (177736 kB) free!
[1576264173] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 56.0% (13728512 kB) free.
[1576264254] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21077360 kB) free!
[1576264374] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40363244 kB) free.
[1576264974] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21466732 kB) free!
[1576265094] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40346812 kB) free.
[1576265694] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21616192 kB) free!
[1576265814] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40316752 kB) free.
[1576266366] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 597 processes
[1576266410] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576266414] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22093620 kB) free!
[1576266515] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3781424 kB) free!
[1576266530] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576266534] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40309332 kB) free.
[1576266966] SERVICE ALERT: webservercluster1;Total Processes;OK;HARD;1;PROCS OK: 511 processes
[1576267134] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21489016 kB) free!
[1576267254] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40345208 kB) free.
[1576267681] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(99) < 2019/12/14 01:32:08 [error] 20758#0: *16768275 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576267801] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576267854] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22896516 kB) free!
[1576267974] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39994968 kB) free.
[1576268574] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22076256 kB) free!
[1576268695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36533472 kB) free.
[1576270386] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (254692 kB) free!
[1576270461] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 368 processes
[1576270506] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230852 kB) free!
[1576270581] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1576270626] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228672 kB) free!
[1576271227] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.4% (42075832 kB) free.
[1576271767] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.9% (11829884 kB) free.
[1576275530] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1576275650] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576279801] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(94) < 2019/12/14 04:59:14 [error] 20754#0: *16800941 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576279921] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576281650] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576281770] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1576283767] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (141528 kB) free!
[1576283887] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (142396 kB) free!
[1576284007] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192068 kB) free!
[1576284487] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 82.1% (20116224 kB) free.
[1576284506] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.7% (21242872 kB) free.
[1576284522] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.5% (22579944 kB) free.
[1576284557] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.8% (21767268 kB) free.
[1576284597] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39102892 kB) free.
[1576285793] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 323 processes
[1576291370] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576291490] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576291922] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2019/12/14 08:21:56 [error] 20756#0: *16833303 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576292042] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/14 08:22:06 [error] 20760#0: *16833321 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576292162] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(24) < 2019/12/14 08:25:55 [error] 20760#0: *16833750 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576292690] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576292810] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576293363] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576294729] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4898000 kB) free!
[1576294849] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (4194676 kB) free!
[1576294969] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4180892 kB) free!
[1576295569] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1486300 kB) free!
[1576296176] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8689648 kB) free!
[1576296296] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236372 kB) free!
[1576296416] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232236 kB) free!
[1576297016] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 81.0% (39925240 kB) free.
[1576297610] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576297730] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576298330] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576298450] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576298599] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.6% (20128008 kB) free.
[1576298827] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4359564 kB) free!
[1576298947] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236892 kB) free!
[1576299067] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222436 kB) free!
[1576299209] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177324 kB) free!
[1576299329] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166292 kB) free!
[1576299449] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1576299449] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193508 kB) free!
[1576299667] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 61.5% (30297392 kB) free.
[1576301803] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.3% (22030756 kB) free.
[1576303169] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4591896 kB) free!
[1576303289] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.0% (3432804 kB) free!
[1576303409] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (4346384 kB) free!
[1576303510] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2781104 kB) free!
[1576303563] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8772) < 2019/12/14 11:36:03 [error] 20760#0: *16872972 no live upstreams while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "backend_3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576303591] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(11395) < 103.255.146.164 [14/Dec/2019:11:36:31 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576303631] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.9% (2454648 kB) free!
[1576303684] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(12697) < 2019/12/14 11:38:03 [error] 20758#0: *16886182 no live upstreams while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "backend_3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576303711] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(12781) < 103.255.146.164 [14/Dec/2019:11:38:31 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576303751] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2161192 kB) free!
[1576303804] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(14013) < 2019/12/14 11:40:04 [error] 20756#0: *16900699 no live upstreams while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "backend_3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576303831] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(12767) < 103.255.146.164 [14/Dec/2019:11:40:31 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576303973] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 79.2% (17784660 kB) free.
[1576304207] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.2% (3245460 kB) free!
[1576304328] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.6% (3118524 kB) free!
[1576304448] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (3006672 kB) free!
[1576304609] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185308 kB) free!
[1576305048] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2246296 kB) free!
[1576305202] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.9% (3917744 kB) free!
[1576305265] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(33) < 2019/12/14 12:00:52 [error] 28866#0: *20985735 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576305322] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.5% (3587388 kB) free!
[1576305385] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576305412] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.5% (22587932 kB) free.
[1576305442] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.2% (3247208 kB) free!
[1576305450] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.7% (22233048 kB) free.
[1576306042] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 94.0% (23208780 kB) free.
[1576306205] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576306231] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576306808] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9428284 kB) free!
[1576306928] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8375692 kB) free!
[1576307048] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (8319784 kB) free!
[1576307287] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4370824 kB) free!
[1576307407] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.3% (4239792 kB) free!
[1576307527] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4144212 kB) free!
[1576307818] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.1% (3243096 kB) free!
[1576307938] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.8% (3155872 kB) free!
[1576308058] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (3087264 kB) free!
[1576308894] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23977004 kB) free!
[1576309014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46613752 kB) free.
[1576309258] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (491812 kB) free!
[1576309327] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2289116 kB) free!
[1576309614] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23154488 kB) free!
[1576309706] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4912816 kB) free!
[1576309734] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45731220 kB) free.
[1576309826] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4817008 kB) free!
[1576309946] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (4733036 kB) free!
[1576310334] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22764000 kB) free!
[1576310454] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44924576 kB) free.
[1576310648] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4825064 kB) free!
[1576311054] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22323316 kB) free!
[1576311175] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18829112 kB) free!
[1576311295] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41451108 kB) free.
[1576311895] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18420428 kB) free!
[1576312015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38724688 kB) free.
[1576312267] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9200760 kB) free!
[1576312387] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9073596 kB) free!
[1576312507] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8930796 kB) free!
[1576312615] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16685616 kB) free!
[1576312735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.9% (51341568 kB) free.
[1576314146] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2441692 kB) free!
[1576315591] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.7% (22388752 kB) free.
[1576315786] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39066832 kB) free.
[1576315832] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.0% (45847324 kB) free.
[1576315869] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.3% (21638548 kB) free.
[1576315928] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.7% (21746112 kB) free.
[1576317397] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.9% (4171588 kB) free!
[1576317518] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4123476 kB) free!
[1576317606] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(100) < 2019/12/14 15:27:52 [error] 20758#0: *17077375 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576317638] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2649476 kB) free!
[1576317726] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(33) < 2019/12/14 15:32:03 [error] 20751#0: *17077846 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576317846] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576318238] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (818272 kB) free!
[1576319935] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25820420 kB) free!
[1576320055] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48663500 kB) free.
[1576320655] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24717704 kB) free!
[1576320775] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48645888 kB) free.
[1576321058] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4589524 kB) free!
[1576321178] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4462704 kB) free!
[1576321298] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4353104 kB) free!
[1576321375] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23505392 kB) free!
[1576321495] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.9% (47293892 kB) free.
[1576322018] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.4% (22318016 kB) free.
[1576322054] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.5% (21694292 kB) free.
[1576322095] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23918884 kB) free!
[1576322215] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.0% (48785564 kB) free.
[1576322815] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24581516 kB) free!
[1576322935] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.7% (48368404 kB) free.
[1576323535] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24668852 kB) free!
[1576323650] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576323655] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (47994568 kB) free.
[1576323770] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576324256] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23934252 kB) free!
[1576324376] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (47544952 kB) free.
[1576324976] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25336456 kB) free!
[1576325096] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21438204 kB) free!
[1576325216] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (44059412 kB) free.
[1576325816] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20471996 kB) free!
[1576325936] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44302944 kB) free.
[1576326536] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19781416 kB) free!
[1576326656] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44616228 kB) free.
[1576326827] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (226244 kB) free!
[1576326947] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181288 kB) free!
[1576327067] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193628 kB) free!
[1576327256] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18831332 kB) free!
[1576327277] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171724 kB) free!
[1576327334] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (208820 kB) free!
[1576327376] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.4% (45368412 kB) free.
[1576327397] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175412 kB) free!
[1576327454] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206380 kB) free!
[1576327457] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198852 kB) free!
[1576327517] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172360 kB) free!
[1576327574] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (220884 kB) free!
[1576327577] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165208 kB) free!
[1576327697] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191268 kB) free!
[1576327976] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21044536 kB) free!
[1576328096] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45538592 kB) free.
[1576328696] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21118980 kB) free!
[1576328816] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45230224 kB) free.
[1576329417] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20262720 kB) free!
[1576329537] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44889164 kB) free.
[1576329847] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(69) < 2019/12/14 18:53:26 [error] 20756#0: *17110539 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576329967] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(34) < 2019/12/14 18:56:01 [error] 20756#0: *17110857 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576330088] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2019/12/14 18:56:12 [error] 20758#0: *17110871 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576330137] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20051136 kB) free!
[1576330257] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44408700 kB) free.
[1576330857] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20366204 kB) free!
[1576330977] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43804744 kB) free.
[1576331216] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9493448 kB) free!
[1576331288] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576331336] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9412728 kB) free!
[1576331457] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (9335412 kB) free!
[1576331577] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20140524 kB) free!
[1576331697] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30819200 kB) free.
[1576332297] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20859956 kB) free!
[1576332417] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16914876 kB) free!
[1576332537] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40481552 kB) free.
[1576333067] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1576333137] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17859352 kB) free!
[1576333257] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40690868 kB) free.
[1576333688] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2019/12/14 19:50:50 [error] 20756#0: *17120026 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576333808] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576333857] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18273984 kB) free!
[1576333977] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40853396 kB) free.
[1576334578] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19109872 kB) free!
[1576334698] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40801900 kB) free.
[1576335298] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18798680 kB) free!
[1576335418] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41514364 kB) free.
[1576336018] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17267728 kB) free!
[1576336138] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40280156 kB) free.
[1576336738] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18401916 kB) free!
[1576336858] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39224220 kB) free.
[1576337182] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1576337302] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1576338058] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20859892 kB) free!
[1576338178] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41908492 kB) free.
[1576338778] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17324656 kB) free!
[1576338898] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39964188 kB) free.
[1576340098] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21427032 kB) free!
[1576340173] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.9% (22186928 kB) free.
[1576340200] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.1% (21599328 kB) free.
[1576340219] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41553616 kB) free.
[1576340236] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 93.3% (45977044 kB) free.
[1576340272] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.7% (22712676 kB) free.
[1576340301] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.6% (22690504 kB) free.
[1576340408] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2019/12/14 21:42:34 [error] 20758#0: *17138475 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576340529] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576340818] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22418436 kB) free!
[1576340938] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41428856 kB) free.
[1576342138] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19773280 kB) free!
[1576342198] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9703592 kB) free!
[1576342258] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.2% (53044508 kB) free.
[1576342318] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9703776 kB) free!
[1576342329] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2019/12/14 22:13:54 [error] 20759#0: *17143235 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576342438] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (9703396 kB) free!
[1576342449] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576343050] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2019/12/14 22:29:10 [error] 20756#0: *17145598 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576343170] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found