[1582223525] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.4% (28775756 kB) free.
[1582223659] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 203.212.70.170 [21/Feb/2020:00:03:09 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1582223756] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2020/02/21 00:01:09 [error] 12255#0: *57265071 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582223780] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582223832] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16560044 kB) free!
[1582223877] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582223952] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12689028 kB) free!
[1582224072] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9150280 kB) free!
[1582224672] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21799152 kB) free!
[1582225259] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (13669568 kB) free!
[1582225278] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37654812 kB) free.
[1582225846] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582225878] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14169788 kB) free!
[1582225998] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10436436 kB) free!
[1582226025] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 64.0% (15695068 kB) free.
[1582226118] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14497112 kB) free!
[1582226459] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12910096 kB) free!
[1582227059] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1582227059] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.9% (14357232 kB) free!
[1582227318] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29570452 kB) free.
[1582227918] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9494328 kB) free!
[1582228038] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38941516 kB) free.
[1582228638] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15414748 kB) free!
[1582228758] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11785100 kB) free!
[1582228878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40736360 kB) free.
[1582229478] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17973544 kB) free!
[1582229599] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14292248 kB) free!
[1582229719] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (10889900 kB) free!
[1582230319] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20485120 kB) free!
[1582230919] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31007356 kB) free.
[1582231519] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12093832 kB) free!
[1582231639] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39376604 kB) free.
[1582232095] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238076 kB) free!
[1582232215] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227596 kB) free!
[1582232239] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16613332 kB) free!
[1582232276] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2020/02/21 02:27:54 [error] 12246#0: *57291788 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1582232326] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 103.16.101.72 [21/Feb/2020:02:28:10 +0530] TCP 502 0 0 60.000 "10.147.212.38:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1582232335] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229780 kB) free!
[1582232359] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13068540 kB) free!
[1582232397] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(14) < 2020/02/21 02:29:09 [error] 12251#0: *57291779 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1582232447] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582232479] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40652880 kB) free.
[1582232517] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582232779] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4022) < 103.255.146.164 [21/Feb/2020:02:27:38 +0530] TCP 502 0 0 0.003 "10.147.212.142:3000" "0" "0" "0.003"
[1582232900] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582232935] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 77.9% (38391744 kB) free.
[1582233079] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17393576 kB) free!
[1582233199] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13716156 kB) free!
[1582233319] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40426236 kB) free.
[1582233343] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 46.0% (11268840 kB) free.
[1582233708] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.5% (42118652 kB) free.
[1582233919] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17057072 kB) free!
[1582234039] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13617532 kB) free!
[1582234159] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39598520 kB) free.
[1582234248] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(208) < 103.16.101.72 [21/Feb/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582234317] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(252) < 2020/02/21 03:01:54 [error] 12244#0: *57300346 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1582234370] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582234437] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582234699] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2282) < 203.212.70.170 [21/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582234760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16732672 kB) free!
[1582234819] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582234880] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13330832 kB) free!
[1582235000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36610524 kB) free.
[1582235600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9349728 kB) free!
[1582235720] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3997060 kB) free!
[1582235840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37652428 kB) free.
[1582236440] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10129216 kB) free!
[1582236560] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4804212 kB) free!
[1582236680] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37021380 kB) free.
[1582237280] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11840792 kB) free!
[1582237400] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6692796 kB) free!
[1582237520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37799452 kB) free.
[1582238037] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2020/02/21 04:00:41 [error] 12257#0: *57309201 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582238120] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12968408 kB) free!
[1582238157] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582238240] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13338180 kB) free!
[1582238360] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38716152 kB) free.
[1582238459] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582238961] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16985276 kB) free!
[1582239081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41890700 kB) free.
[1582239681] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21868664 kB) free!
[1582239801] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18172508 kB) free!
[1582239921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42743224 kB) free.
[1582240521] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25016352 kB) free!
[1582240641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44185044 kB) free.
[1582241743] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176324 kB) free!
[1582241864] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172092 kB) free!
[1582241984] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181788 kB) free!
[1582242680] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 68.7% (16833292 kB) free.
[1582243217] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 64.1% (15811400 kB) free.
[1582244841] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582246169] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1582246259] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 41.8% (55160224 kB) free.
[1582246289] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1582246642] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23161396 kB) free!
[1582246762] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41899296 kB) free.
[1582247362] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21094068 kB) free!
[1582247482] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43866472 kB) free.
[1582248082] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23174520 kB) free!
[1582248089] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1582248202] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43863140 kB) free.
[1582248209] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1582248507] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 64.1% (31581204 kB) free.
[1582248533] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21934136 kB) free.
[1582248577] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.4% (14038084 kB) free.
[1582248802] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23668024 kB) free!
[1582248922] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44188388 kB) free.
[1582249817] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.4% (3066828 kB) free!
[1582249937] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.4% (3067840 kB) free!
[1582250057] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181592 kB) free!
[1582252522] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22730084 kB) free!
[1582252642] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42305140 kB) free.
[1582253157] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(19) < 2020/02/21 08:11:18 [error] 12253#0: *57353943 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582253243] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21707788 kB) free!
[1582253278] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582253363] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42001808 kB) free.
[1582253419] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 172.105.59.173 [21/Feb/2020:08:12:22 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.000, 0.000"
[1582253539] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582253963] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23473532 kB) free!
[1582254083] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42231592 kB) free.
[1582254683] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23719596 kB) free!
[1582254803] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42819560 kB) free.
[1582255403] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24190408 kB) free!
[1582255523] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42587356 kB) free.
[1582256123] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14946532 kB) free!
[1582256243] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41528136 kB) free.
[1582256843] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21445452 kB) free!
[1582256963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41665620 kB) free.
[1582257287] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21941504 kB) free.
[1582257563] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20251552 kB) free!
[1582257683] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41040104 kB) free.
[1582257945] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.9% (4626452 kB) free!
[1582258065] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.3% (4477284 kB) free!
[1582258185] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.8% (4371268 kB) free!
[1582258284] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17873120 kB) free!
[1582258404] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39955464 kB) free.
[1582259004] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17186180 kB) free!
[1582259124] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13109456 kB) free!
[1582259244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34733256 kB) free.
[1582259844] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11495428 kB) free!
[1582259964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39472580 kB) free.
[1582259985] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2068496 kB) free!
[1582260085] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.8% (4864932 kB) free!
[1582260205] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.7% (4834420 kB) free!
[1582260325] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.6% (4812988 kB) free!
[1582260564] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16424912 kB) free!
[1582260579] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (218920 kB) free!
[1582260684] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40892368 kB) free.
[1582260699] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (211284 kB) free!
[1582260819] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (210236 kB) free!
[1582260890] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185656 kB) free!
[1582261010] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187440 kB) free!
[1582261130] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183912 kB) free!
[1582261284] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16481220 kB) free!
[1582261404] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12223064 kB) free!
[1582261524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33247824 kB) free.
[1582261846] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1582261966] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1582262087] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1582262124] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5756164 kB) free!
[1582262244] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1639680 kB) free!
[1582262364] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33787168 kB) free.
[1582262943] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (2214260 kB) free!
[1582262965] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9563856 kB) free!
[1582263063] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165988 kB) free!
[1582263085] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5522320 kB) free!
[1582263183] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183444 kB) free!
[1582263205] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34514064 kB) free.
[1582263805] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9369048 kB) free!
[1582263925] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5405204 kB) free!
[1582264045] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33545464 kB) free.
[1582264111] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8878060 kB) free!
[1582264184] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4266388 kB) free!
[1582264231] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8879164 kB) free!
[1582264280] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4857700 kB) free!
[1582264304] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.6% (3955972 kB) free!
[1582264351] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (8877012 kB) free!
[1582264384] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.2% (21762680 kB) free.
[1582264400] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4648092 kB) free!
[1582264417] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.1% (21583112 kB) free.
[1582264424] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.9% (3579660 kB) free!
[1582264446] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.2% (21616100 kB) free.
[1582264470] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.5% (14061440 kB) free.
[1582264487] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582264520] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4470764 kB) free!
[1582264525] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.1% (22087520 kB) free.
[1582264645] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8439444 kB) free!
[1582264676] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31785812 kB) free.
[1582264706] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.7% (16088580 kB) free.
[1582264734] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 70.0% (34519880 kB) free.
[1582264988] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4807024 kB) free!
[1582265109] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (966360 kB) free!
[1582265229] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1582265229] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (473972 kB) free!
[1582266125] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9125636 kB) free!
[1582266245] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.5% (8603372 kB) free!
[1582266320] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1610368 kB) free!
[1582266366] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.4% (8096900 kB) free!
[1582266680] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4416468 kB) free!
[1582266800] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3973844 kB) free!
[1582266887] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1582266920] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3467404 kB) free!
[1582266971] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.0% (30539916 kB) free.
[1582267077] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2020/02/21 12:00:46 [error] 12247#0: *57409716 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582267197] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582267339] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [21/Feb/2020:12:02:37 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1582267459] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582267488] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582267520] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.5% (22838428 kB) free.
[1582267608] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582267728] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1582267728] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582268279] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24605884 kB) free!
[1582268400] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20379364 kB) free!
[1582268520] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16329492 kB) free!
[1582268649] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184336 kB) free!
[1582268769] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169544 kB) free!
[1582268889] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178208 kB) free!
[1582268926] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1839352 kB) free!
[1582268941] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9202796 kB) free!
[1582268986] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (1838880 kB) free!
[1582269046] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 7.5% (1843304 kB) free!
[1582269061] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8460884 kB) free!
[1582269106] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 7.5% (1844336 kB) free!
[1582269120] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26786688 kB) free.
[1582269181] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (7401836 kB) free!
[1582269708] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8654704 kB) free!
[1582269829] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.8% (8295536 kB) free!
[1582269949] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.1% (7929032 kB) free!
[1582270006] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 15.0% (3665648 kB) free!
[1582270320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13922460 kB) free!
[1582270409] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.3044.2179> (10.147.212.53:46240 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1582270440] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9786556 kB) free!
[1582270529] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1582270560] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5856732 kB) free!
[1582270728] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1582270981] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1582270981] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (1871448 kB) free!
[1582271149] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4712008 kB) free!
[1582271760] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1582271760] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23677024 kB) free!
[1582272360] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (830576 kB) free!
[1582272960] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14226164 kB) free!
[1582273128] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1582273420] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.1% (4425188 kB) free!
[1582273540] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.5% (4291624 kB) free!
[1582273560] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30085720 kB) free.
[1582273660] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.9% (4139460 kB) free!
[1582274160] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5231036 kB) free!
[1582274280] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (990648 kB) free!
[1582274400] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36235000 kB) free.
[1582274581] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (5979524 kB) free!
[1582274806] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.5% (2328376 kB) free!
[1582274829] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 20.7% (5101352 kB) free.
[1582274860] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1962284 kB) free!
[1582275001] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10795712 kB) free!
[1582275106] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 34.0% (8325808 kB) free.
[1582275121] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6367048 kB) free!
[1582275241] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2162952 kB) free!
[1582275429] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (1324360 kB) free!
[1582275511] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4071140 kB) free!
[1582275549] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (504740 kB) free!
[1582275571] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3192052 kB) free!
[1582275631] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.1% (3830908 kB) free!
[1582275669] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178192 kB) free!
[1582275692] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2891924 kB) free!
[1582275751] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.8% (3553568 kB) free!
[1582275781] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2178812 kB) free!
[1582275812] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2807992 kB) free!
[1582275841] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19632696 kB) free!
[1582276136] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9518064 kB) free!
[1582276256] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9397556 kB) free!
[1582276376] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (9266712 kB) free!
[1582276412] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.1% (5373420 kB) free.
[1582276441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35984560 kB) free.
[1582276951] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1571984 kB) free!
[1582277041] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9300876 kB) free!
[1582277161] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (5014888 kB) free!
[1582277260] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 36.4% (8910680 kB) free.
[1582277281] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (794664 kB) free!
[1582277551] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16086388 kB) free.
[1582277659] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(236) < 92.42.104.166 [21/Feb/2020:15:00:13 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582277779] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582277881] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14565516 kB) free!
[1582278481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34171140 kB) free.
[1582279081] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1229236 kB) free!
[1582279201] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400992 kB) free!
[1582279321] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408148 kB) free!
[1582280012] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2635484 kB) free!
[1582280132] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (2782736 kB) free!
[1582280252] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.2% (4414908 kB) free.
[1582280521] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38910560 kB) free.
[1582280643] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582280763] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582280861] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (239360 kB) free!
[1582280883] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582280943] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582280981] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175916 kB) free!
[1582281063] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582281101] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1582281101] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179040 kB) free!
[1582281122] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6102480 kB) free!
[1582281183] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1582281242] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1208916 kB) free!
[1582281362] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375152 kB) free!
[1582281379] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 49.50.86.152 [21/Feb/2020:16:02:18 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1582281483] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1582281499] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582281528] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1582281597] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(107) < 2020/02/21 16:02:13 [error] 12254#0: *57484094 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582281648] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582281717] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582281769] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582281776] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4515476 kB) free!
[1582281962] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18007976 kB) free!
[1582282512] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46376 MB (20% inode=99%):
[1582282562] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393620 kB) free!
[1582282632] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46376 MB (20% inode=99%):
[1582282653] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2969980 kB) free!
[1582282683] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582282752] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46376 MB (20% inode=99%):
[1582282773] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3134324 kB) free!
[1582282803] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1582282893] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3419232 kB) free.
[1582282923] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1582283162] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14426600 kB) free!
[1582283762] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44253164 kB) free.
[1582284102] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 65.0% (15927664 kB) free.
[1582284183] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1582284362] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10013744 kB) free!
[1582284482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5239640 kB) free!
[1582284552] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 85085 MB (37% inode=99%):
[1582284602] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1075932 kB) free!
[1582284693] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2411336 kB) free!
[1582284783] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582284813] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2580216 kB) free!
[1582284903] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582284933] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2816556 kB) free!
[1582285023] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1582285202] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20210208 kB) free!
[1582285302] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174708 kB) free!
[1582285422] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173128 kB) free!
[1582285533] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1582285533] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.9% (4858948 kB) free.
[1582285542] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177448 kB) free!
[1582285802] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384316 kB) free!
[1582286133] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2553108 kB) free!
[1582286253] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2766668 kB) free!
[1582286373] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (3090216 kB) free!
[1582286402] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19509956 kB) free!
[1582286581] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582286973] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.2% (4425432 kB) free.
[1582287002] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38282120 kB) free.
[1582287069] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582287170] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1582287602] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11548672 kB) free!
[1582287723] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7277164 kB) free!
[1582287843] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3234516 kB) free!
[1582288443] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16156004 kB) free!
[1582288542] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 62.8% (15403816 kB) free.
[1582289043] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33821896 kB) free.
[1582289223] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1582289223] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1582289643] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8165952 kB) free!
[1582289763] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4126356 kB) free!
[1582289823] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582289883] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27018440 kB) free.
[1582289943] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582290063] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582290170] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1582290342] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (968276 kB) free!
[1582290462] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (892520 kB) free!
[1582290483] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (570304 kB) free!
[1582290582] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (828552 kB) free!
[1582290603] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (491652 kB) free!
[1582290663] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1582290723] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378440 kB) free!
[1582291782] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 47.1% (11552132 kB) free.
[1582291863] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1582291923] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34064212 kB) free.
[1582291971] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1582291983] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1582292091] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582292103] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582292211] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1582292382] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174416 kB) free!
[1582292502] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172096 kB) free!
[1582292523] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5577536 kB) free!
[1582292622] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194124 kB) free!
[1582292643] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1113748 kB) free!
[1582292703] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1582292763] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392364 kB) free!
[1582293363] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14810144 kB) free!
[1582293822] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 75.1% (18403484 kB) free.
[1582293964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36880984 kB) free.
[1582294564] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8063052 kB) free!
[1582294684] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3790960 kB) free!
[1582294804] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387732 kB) free!
[1582295404] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16787296 kB) free!
[1582295517] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(72) < 2020/02/21 20:01:49 [error] 12249#0: *57567167 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582295637] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582295899] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(457) < 203.212.70.170 [21/Feb/2020:20:02:10 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582296004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36259524 kB) free.
[1582296019] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582296604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11077800 kB) free!
[1582296724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7094876 kB) free!
[1582296844] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3199176 kB) free!
[1582297444] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15640988 kB) free!
[1582298044] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30650496 kB) free.
[1582298373] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582298644] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6609260 kB) free!
[1582298764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2604408 kB) free!
[1582298884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38559352 kB) free.
[1582299411] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1582299484] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13210556 kB) free!
[1582299583] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21981456 kB) free.
[1582299605] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9175128 kB) free!
[1582299612] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 30.7% (15128556 kB) free.
[1582299652] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 44.7% (58911840 kB) free.
[1582299684] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39058608 kB) free.
[1582299724] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 54.4% (26802856 kB) free.
[1582299762] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22021408 kB) free.
[1582299888] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.5% (21940392 kB) free.
[1582300410] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582300530] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582300789] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (4436892 kB) free!
[1582300909] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173044 kB) free!
[1582301029] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189272 kB) free!
[1582301461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25633224 kB) free!
[1582301581] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21401116 kB) free!
[1582301701] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17433144 kB) free!
[1582302301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29107536 kB) free.
[1582303502] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20054044 kB) free!
[1582303622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40735380 kB) free.
[1582304103] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582304622] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582305422] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19327492 kB) free!
[1582305542] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.3% (50429572 kB) free.
[1582306142] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25996852 kB) free!
[1582306262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.9% (22300716 kB) free!
[1582306382] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18846068 kB) free!
[1582306530] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582306650] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582306982] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28975840 kB) free.
[1582307850] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582307971] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1582308182] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21343372 kB) free!
[1582308302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.4% (50672040 kB) free.
[1582308377] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 11.7% (5744988 kB) free!
[1582308498] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240568 kB) free!
[1582308618] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230356 kB) free!
[1582309259] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (621452 kB) free!
[1582309379] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (426932 kB) free!
[1582309499] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415584 kB) free!