[1578421889] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14667212 kB) free!
[1578422009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36637404 kB) free.
[1578422111] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.1% (27628712 kB) free.
[1578422560] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1578422609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14494216 kB) free!
[1578422729] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35012264 kB) free.
[1578423761] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578423881] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578423929] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15715948 kB) free!
[1578424001] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578424049] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36560484 kB) free.
[1578424601] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1578425201] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578425249] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15171248 kB) free!
[1578425369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36426432 kB) free.
[1578425801] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1578426569] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16710932 kB) free!
[1578426689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36558204 kB) free.
[1578427001] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1578427122] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578427242] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578427842] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1578427889] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17244052 kB) free!
[1578428009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36565864 kB) free.
[1578428096] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(614) < 183.82.105.122 [08/Jan/2020:01:44:32 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1578428216] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578428387] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(627) < 2020/01/08 01:44:32 [error] 21675#0: *23371987 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578428507] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578429042] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578429162] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578429209] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16720652 kB) free!
[1578429282] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578429329] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34622220 kB) free.
[1578429707] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(46) < 2020/01/08 02:11:40 [error] 21682#0: *23373250 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:44/0, bytes from/to upstream:0/44
[1578429827] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578429882] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1578430529] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16419556 kB) free!
[1578430649] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34615320 kB) free.
[1578430768] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 66.5% (32768204 kB) free.
[1578430951] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.5% (10906152 kB) free.
[1578431082] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578431202] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1578431775] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.2% (42955988 kB) free.
[1578433002] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578433049] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20038132 kB) free!
[1578433122] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578433169] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31676796 kB) free.
[1578433242] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1578435043] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578435163] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578435283] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1578436483] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578436603] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1578437203] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578437323] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578437444] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578438416] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(724) < 183.82.105.122 [08/Jan/2020:04:35:54 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578438536] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578438828] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(737) < 2020/01/08 04:35:54 [error] 21682#0: *23399119 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578438948] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578439844] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1578440444] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1578440564] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578440684] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578441284] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1578441884] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578442004] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578442125] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578443325] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1578443551] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181072 kB) free!
[1578443671] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172400 kB) free!
[1578443791] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170524 kB) free!
[1578444188] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 21.2% (5186416 kB) free.
[1578444525] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1578444645] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578444756] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.8% (22010040 kB) free.
[1578444765] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1578444789] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177084 kB) free!
[1578444800] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.6% (21377192 kB) free.
[1578444848] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.1% (38499268 kB) free.
[1578444900] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.9% (21785592 kB) free.
[1578444909] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169544 kB) free!
[1578444936] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.5% (21680020 kB) free.
[1578445365] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578445407] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189252 kB) free!
[1578445457] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (257708 kB) free!
[1578445485] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578445527] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197080 kB) free!
[1578445547] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181104 kB) free!
[1578445577] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242324 kB) free!
[1578445594] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1578445605] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578445647] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195172 kB) free!
[1578445667] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190240 kB) free!
[1578445697] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228360 kB) free!
[1578445714] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 330 processes
[1578445787] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185444 kB) free!
[1578445835] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1578446108] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (569880 kB) free!
[1578446206] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1578446228] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (559964 kB) free!
[1578446348] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (555020 kB) free!
[1578446369] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24895788 kB) free!
[1578446421] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 646 processes
[1578446489] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38506148 kB) free.
[1578446542] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 408 processes
[1578446806] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578446926] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1578447506] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.6% (22359032 kB) free.
[1578447526] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578447550] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39536596 kB) free.
[1578447586] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1578447586] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.5% (22183268 kB) free.
[1578447628] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.5% (20477248 kB) free.
[1578447646] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1578448118] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173156 kB) free!
[1578448238] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200768 kB) free!
[1578448238] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182776 kB) free!
[1578448246] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578448358] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197032 kB) free!
[1578448358] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1578448358] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177220 kB) free!
[1578448366] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1578448478] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1578448478] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173268 kB) free!
[1578448754] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (1719820 kB) free!
[1578448874] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219696 kB) free!
[1578448966] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1578448994] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1578448994] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226264 kB) free!
[1578449086] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578449148] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(693) < 2020/01/08 07:31:33 [error] 21677#0: *23429741 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578449206] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1578449268] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578449336] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(680) < 183.82.105.122 [08/Jan/2020:07:31:33 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1578449457] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578449700] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2020/01/08 07:40:02 [error] 18360#0: *34915644 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:217/17, bytes from/to upstream:17/217
[1578449820] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578449869] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1461) < 2020/01/08 07:46:08 [error] 21681#0: *23435277 upstream timed out (110: Connection timed out) while connecting to upstream, client: 182.18.144.246, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578449989] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578450056] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(954) < 172.105.59.173 [08/Jan/2020:07:46:02 +0530] TCP 502 0 0 120.001 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1578450176] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578450406] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1578450526] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1578450646] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578451037] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.1% (21847564 kB) free.
[1578451077] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.1% (22719396 kB) free.
[1578451122] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.9% (38885768 kB) free.
[1578451246] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1578451648] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (863056 kB) free!
[1578451686] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193636 kB) free!
[1578451726] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238024 kB) free!
[1578451768] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180416 kB) free!
[1578451806] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193608 kB) free!
[1578451846] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215120 kB) free!
[1578451888] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173280 kB) free!
[1578451926] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165916 kB) free!
[1578451966] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (245576 kB) free!
[1578452446] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1578452566] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1578452577] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(138) < 92.42.104.166 [08/Jan/2020:08:30:52 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1578452686] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1578452697] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(856) < 183.82.105.122 [08/Jan/2020:08:33:38 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1578452817] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578452989] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1040) < 2020/01/08 08:33:38 [error] 21675#0: *23445755 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578453109] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578454617] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1121) < 183.82.105.122 [08/Jan/2020:09:03:08 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578454738] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578454909] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1121) < 2020/01/08 09:03:08 [error] 21683#0: *23452996 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578455029] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578455086] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578455206] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1578455326] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1578456594] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1578457558] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.8% (22907948 kB) free.
[1578457609] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39508496 kB) free.
[1578457650] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.8% (23230212 kB) free.
[1578458164] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189300 kB) free!
[1578458217] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235660 kB) free!
[1578458284] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183880 kB) free!
[1578458337] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (252236 kB) free!
[1578458404] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184308 kB) free!
[1578458457] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226600 kB) free!
[1578458937] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2924) < 103.255.146.164 [08/Jan/2020:10:17:27 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578459057] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578459229] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2925) < 2020/01/08 10:17:27 [error] 21677#0: *23471530 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578459349] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578461749] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3755) < 2020/01/08 11:05:49 [error] 21677#0: *23487159 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578461869] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(619) < 2020/01/08 11:05:53 [error] 21682#0: *23488431 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578461989] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578462057] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4346) < 103.255.146.164 [08/Jan/2020:11:05:53 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1578462177] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578462254] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.1% (22731184 kB) free.
[1578462332] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 73.2% (36094124 kB) free.
[1578462458] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1578462590] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1626) < 2020/01/08 11:16:12 [error] 21670#0: *23493450 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578462710] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578462777] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1624) < 103.255.146.164 [08/Jan/2020:11:16:12 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1578462857] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198768 kB) free!
[1578462881] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578462897] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578462954] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (242928 kB) free!
[1578462977] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193396 kB) free!
[1578463001] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578463074] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (220100 kB) free!
[1578463097] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165148 kB) free!
[1578463121] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578463194] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231420 kB) free!
[1578463289] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22985008 kB) free!
[1578463310] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1346) < 2020/01/08 11:31:21 [error] 21682#0: *23498947 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578463409] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42866908 kB) free.
[1578463427] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 105.91, 69.79, 38.35
[1578463430] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578463498] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1834) < 103.255.146.164 [08/Jan/2020:11:34:57 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578463547] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 32.34, 53.80, 36.31
[1578463558] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4374040 kB) free!
[1578463618] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(5518) < 103.255.146.164 [08/Jan/2020:11:36:38 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1578463667] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 20.00, 41.63, 33.95
[1578463678] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4073176 kB) free!
[1578463726] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1578463738] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(4647) < 103.255.146.164 [08/Jan/2020:11:38:47 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578463798] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.8% (3761364 kB) free!
[1578464031] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13061) < 2020/01/08 11:43:45 [error] 21676#0: *23527214 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:44/0, bytes from/to upstream:0/44
[1578464151] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(10) < 2020/01/08 11:45:42 [error] 21682#0: *23527306 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1578464267] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 16.52, 19.47, 25.39
[1578464271] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(35) < 2020/01/08 11:47:45 [error] 21681#0: *23527541 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1578464939] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578464998] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1899480 kB) free!
[1578465471] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578465599] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.0% (16836888 kB) free.
[1578466721] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578467258] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176564 kB) free!
[1578467271] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3529) < 2020/01/08 12:30:25 [error] 21683#0: *23541624 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578467340] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3529) < 103.255.146.164 [08/Jan/2020:12:30:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578467378] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172752 kB) free!
[1578467391] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578467460] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578467499] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189088 kB) free!
[1578468297] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.7% (22129832 kB) free.
[1578468507] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.3% (38099384 kB) free.
[1578468585] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.9% (20807804 kB) free.
[1578469121] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578469188] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183116 kB) free!
[1578469241] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578469308] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187684 kB) free!
[1578469362] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578469429] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195872 kB) free!
[1578469507] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184952 kB) free!
[1578469627] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170468 kB) free!
[1578469718] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (211424 kB) free!
[1578469747] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166320 kB) free!
[1578469838] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214944 kB) free!
[1578469958] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216956 kB) free!
[1578471162] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1578471162] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578471591] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/08 13:47:09 [error] 21677#0: *23554837 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3395/104, bytes from/to upstream:104/3395
[1578471711] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578471762] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578471766] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.6% (4311568 kB) free!
[1578471809] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21843000 kB) free!
[1578471882] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578471886] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.2% (4215444 kB) free!
[1578471929] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.4% (45410704 kB) free.
[1578472002] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578472006] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.9% (4152328 kB) free!
[1578472234] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1578472311] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/08 13:54:04 [error] 21677#0: *23556103 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:1400/104, bytes from/to upstream:104/1400
[1578472354] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1578472431] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578472529] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19910020 kB) free!
[1578472564] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1578472564] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 54.6% (13392148 kB) free.
[1578472628] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.1% (38509160 kB) free.
[1578472649] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42638756 kB) free.
[1578472677] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22106668 kB) free.
[1578472794] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4565944 kB) free!
[1578472914] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4434568 kB) free!
[1578473034] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4330432 kB) free!
[1578473168] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172116 kB) free!
[1578473235] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (212332 kB) free!
[1578473249] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18060460 kB) free!
[1578473288] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186332 kB) free!
[1578473333] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168188 kB) free!
[1578473356] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 66.8% (32936680 kB) free.
[1578473369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41046084 kB) free.
[1578473408] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191968 kB) free!
[1578473453] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195760 kB) free!
[1578473460] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.2% (20158496 kB) free.
[1578473518] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.7% (22123148 kB) free.
[1578474068] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189276 kB) free!
[1578474132] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186208 kB) free!
[1578474188] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192840 kB) free!
[1578474252] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193972 kB) free!
[1578474308] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191480 kB) free!
[1578474372] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194196 kB) free!
[1578474376] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8955448 kB) free!
[1578474406] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2339544 kB) free!
[1578474496] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8794404 kB) free!
[1578474569] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17471900 kB) free!
[1578474583] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224816 kB) free!
[1578474616] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8630996 kB) free!
[1578474689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41915988 kB) free.
[1578474703] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (223344 kB) free!
[1578474823] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240804 kB) free!
[1578474834] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2203068 kB) free!
[1578475289] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16756080 kB) free!
[1578475409] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40436224 kB) free.
[1578475425] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 389 processes
[1578475546] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 385 processes
[1578475608] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 18.1% (4427792 kB) free!
[1578475666] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 385 processes
[1578475729] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 0.7% (171308 kB) free!
[1578475841] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.7% (21987608 kB) free.
[1578475849] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 0.8% (197132 kB) free!
[1578475886] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.9% (21686648 kB) free.
[1578475922] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.4% (22156664 kB) free.
[1578475954] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 30.1% (14854064 kB) free.
[1578475984] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 73.4% (17983980 kB) free.
[1578476009] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15047232 kB) free!
[1578476022] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 40.3% (9883568 kB) free.
[1578476130] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38610548 kB) free.
[1578476496] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182552 kB) free!
[1578476556] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (260744 kB) free!
[1578476616] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188080 kB) free!
[1578476626] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194940 kB) free!
[1578476677] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219704 kB) free!
[1578476729] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13416084 kB) free!
[1578476736] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176848 kB) free!
[1578476746] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199796 kB) free!
[1578476797] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (252972 kB) free!
[1578476849] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37610928 kB) free.
[1578476866] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183624 kB) free!
[1578476866] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 382 processes
[1578477124] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22078208 kB) free.
[1578477169] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.8% (38331556 kB) free.
[1578477229] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.6% (20736960 kB) free.
[1578477283] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.4% (44048500 kB) free.
[1578477449] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11467836 kB) free!
[1578477466] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 385 processes
[1578477569] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33889336 kB) free.
[1578477586] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 383 processes
[1578477728] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174836 kB) free!
[1578477771] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (251904 kB) free!
[1578477831] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/08 15:28:03 [error] 21682#0: *23572591 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3996/121, bytes from/to upstream:121/3996
[1578477837] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174852 kB) free!
[1578477848] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164352 kB) free!
[1578477891] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234044 kB) free!
[1578477951] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578477957] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184400 kB) free!
[1578477968] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186648 kB) free!
[1578477978] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.3% (18342548 kB) free.
[1578478011] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214732 kB) free!
[1578478030] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 43.4% (21373216 kB) free.
[1578478064] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 33.4% (8180140 kB) free.
[1578478168] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9676932 kB) free!
[1578478288] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9579136 kB) free!
[1578478408] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (9475512 kB) free!
[1578478587] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202604 kB) free!
[1578478620] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/08 15:46:48 [error] 18365#0: *35233499 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/view/ HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/view/", host: "altsmsc.tesync.net"
[1578478636] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (213668 kB) free!
[1578478667] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191132 kB) free!
[1578478707] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183264 kB) free!
[1578478740] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578478756] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241792 kB) free!
[1578478769] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25812020 kB) free!
[1578478787] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194684 kB) free!
[1578478827] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173460 kB) free!
[1578478877] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (247964 kB) free!
[1578478889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.4% (50676520 kB) free.
[1578478907] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195100 kB) free!
[1578479076] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 72.6% (17786136 kB) free.
[1578479115] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.7% (38791252 kB) free.
[1578479148] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.7% (21883184 kB) free.
[1578479152] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/08 15:54:46 [error] 21673#0: *23577225 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:7175/87, bytes from/to upstream:87/7175
[1578479273] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578479489] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23391968 kB) free!
[1578479609] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.2% (50376096 kB) free.
[1578479684] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.0% (2937192 kB) free!
[1578479716] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236276 kB) free!
[1578479764] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178684 kB) free!
[1578479805] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (2075104 kB) free!
[1578479836] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247468 kB) free!
[1578479858] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.7% (21893540 kB) free.
[1578479899] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.0% (38937280 kB) free.
[1578479925] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1905628 kB) free!
[1578479941] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.0% (21578080 kB) free.
[1578479987] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.6% (45160468 kB) free.
[1578480209] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23737932 kB) free!
[1578480329] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19680372 kB) free!
[1578480449] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46311440 kB) free.
[1578480465] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180972 kB) free!
[1578480507] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (251844 kB) free!
[1578480574] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.5% (21840372 kB) free.
[1578480606] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.8% (38857564 kB) free.
[1578481144] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178920 kB) free!
[1578481177] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191512 kB) free!
[1578481219] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215536 kB) free!
[1578481264] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188256 kB) free!
[1578481297] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191156 kB) free!
[1578481324] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21852864 kB) free.
[1578481339] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241532 kB) free!
[1578481354] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1578481360] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39063572 kB) free.
[1578481384] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193816 kB) free!
[1578481404] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.9% (18841116 kB) free.
[1578481475] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1578481952] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188548 kB) free!
[1578481976] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3802796 kB) free!
[1578482006] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168372 kB) free!
[1578482072] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189864 kB) free!
[1578482096] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239076 kB) free!
[1578482117] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.2% (21754904 kB) free.
[1578482126] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173056 kB) free!
[1578482184] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242888 kB) free!
[1578482230] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.4% (39107416 kB) free.
[1578482246] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163172 kB) free!
[1578482276] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.3% (20177832 kB) free.
[1578482762] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193012 kB) free!
[1578482837] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1586028 kB) free!
[1578482843] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.2% (20040860 kB) free.
[1578482882] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174632 kB) free!
[1578482914] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.4% (36673284 kB) free.
[1578482988] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.5% (17521324 kB) free.
[1578483402] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578483449] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188132 kB) free!
[1578483472] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/08 17:00:32 [error] 21676#0: *23588578 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:9355/87, bytes from/to upstream:87/9355
[1578483519] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2514428 kB) free!
[1578483569] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171708 kB) free!
[1578483592] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578483597] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190492 kB) free!
[1578483604] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.9% (21695812 kB) free.
[1578483639] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247664 kB) free!
[1578483658] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39271192 kB) free.
[1578483703] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.5% (20214968 kB) free.
[1578484192] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/08 17:17:13 [error] 21682#0: *23591462 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:18688/87, bytes from/to upstream:87/18688
[1578484209] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191520 kB) free!
[1578484266] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220848 kB) free!
[1578484307] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173176 kB) free!
[1578484312] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578484329] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197508 kB) free!
[1578484386] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227456 kB) free!
[1578484427] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188472 kB) free!
[1578484449] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169584 kB) free!
[1578484506] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (252348 kB) free!
[1578484547] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180520 kB) free!
[1578485020] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.4% (21567688 kB) free.
[1578485056] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.7% (38766608 kB) free.
[1578485074] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1578485117] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 67.5% (16554680 kB) free.
[1578485194] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1578485399] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.5% (3924704 kB) free!
[1578485512] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/08 17:33:33 [error] 21679#0: *23594243 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:2555/104, bytes from/to upstream:104/2555
[1578485519] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3649156 kB) free!
[1578485621] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197636 kB) free!
[1578485632] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578485639] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.2% (3406448 kB) free!
[1578485662] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (257020 kB) free!
[1578485741] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185184 kB) free!
[1578485763] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164600 kB) free!
[1578485782] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234808 kB) free!
[1578485861] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22098936 kB) free.
[1578485883] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163944 kB) free!
[1578485902] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230084 kB) free!
[1578485928] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.4% (38640420 kB) free.
[1578486003] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170348 kB) free!
[1578486232] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/08 17:50:08 [error] 21683#0: *23597081 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:18924/104, bytes from/to upstream:104/18924
[1578486239] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2099764 kB) free!
[1578486311] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9537904 kB) free!
[1578486352] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578486431] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.7% (9236540 kB) free!
[1578486470] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187548 kB) free!
[1578486536] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18122264 kB) free!
[1578486536] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216596 kB) free!
[1578486551] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.1% (8904996 kB) free!
[1578486591] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196640 kB) free!
[1578486656] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.1% (52834860 kB) free.
[1578486656] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228648 kB) free!
[1578486707] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241560 kB) free!
[1578486711] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186156 kB) free!
[1578486761] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.5% (38697720 kB) free.
[1578486811] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.5% (20107448 kB) free.
[1578486839] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.1% (16873400 kB) free.
[1578486908] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.6% (21958516 kB) free.
[1578486943] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.5% (27822128 kB) free.
[1578486953] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/08 18:05:06 [error] 21677#0: *23599662 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:5293/87, bytes from/to upstream:87/5293
[1578487073] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578487256] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20942476 kB) free!
[1578487366] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236512 kB) free!
[1578487376] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16885700 kB) free!
[1578487418] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184016 kB) free!
[1578487486] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (229004 kB) free!
[1578487496] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13261720 kB) free!
[1578487509] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.6% (21618996 kB) free.
[1578487515] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179664 kB) free!
[1578487590] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 48.0% (11755892 kB) free.
[1578487606] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 75.4% (37154080 kB) free.
[1578487667] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 52.0% (68602484 kB) free.
[1578488116] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194108 kB) free!
[1578488194] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182472 kB) free!
[1578488236] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195568 kB) free!
[1578488254] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (267076 kB) free!
[1578488314] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192084 kB) free!
[1578488356] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193024 kB) free!
[1578488374] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232052 kB) free!
[1578488434] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.3% (18460764 kB) free.
[1578488460] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 67.8% (16733736 kB) free.
[1578488473] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.8% (35402156 kB) free.
[1578489034] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (212112 kB) free!
[1578489065] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191824 kB) free!
[1578489090] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240156 kB) free!
[1578489154] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182204 kB) free!
[1578489185] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177088 kB) free!
[1578489187] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4498276 kB) free!
[1578489210] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (245344 kB) free!
[1578489274] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164304 kB) free!
[1578489306] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200504 kB) free!
[1578489307] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4412100 kB) free!
[1578489326] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.9% (21191872 kB) free.
[1578489331] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.7% (37777288 kB) free.
[1578489407] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.9% (22033852 kB) free.
[1578489419] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 45.3% (11113720 kB) free.
[1578489930] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (223768 kB) free!
[1578489933] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197016 kB) free!
[1578490032] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4215916 kB) free!
[1578490040] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.11476.2190> (10.147.212.158:32978 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1578490050] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239792 kB) free!
[1578490053] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182608 kB) free!
[1578490152] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (4133096 kB) free!
[1578490160] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1578490170] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (265336 kB) free!
[1578490173] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174592 kB) free!
[1578490272] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179132 kB) free!
[1578490679] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14472348 kB) free!
[1578490737] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1578490799] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10520816 kB) free!
[1578490858] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1578490920] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41507148 kB) free.
[1578491519] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13246604 kB) free!
[1578491639] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43660652 kB) free.
[1578492239] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14418096 kB) free!
[1578492359] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10249832 kB) free!
[1578492479] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42513868 kB) free.
[1578493079] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14479752 kB) free!
[1578493199] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10334192 kB) free!
[1578493319] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40063448 kB) free.
[1578493919] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12713916 kB) free!
[1578494039] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41192308 kB) free.
[1578494639] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15400404 kB) free!
[1578494759] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11042276 kB) free!
[1578494879] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40085164 kB) free.
[1578495479] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14750372 kB) free!
[1578495600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40388756 kB) free.
[1578496170] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1578496199] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17261376 kB) free!
[1578496319] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13036716 kB) free!
[1578496439] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40134652 kB) free.
[1578497039] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15267716 kB) free!
[1578497159] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13470844 kB) free!
[1578497279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37497912 kB) free.
[1578497340] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2020/01/08 20:52:18 [error] 18359#0: *35381170 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:5631452/1445319, bytes from/to upstream:1445319/5631452
[1578497460] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578497879] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14347436 kB) free!
[1578497999] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39363780 kB) free.
[1578498402] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578498573] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1578498599] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14883412 kB) free!
[1578498719] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39602156 kB) free.
[1578498967] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 55.6% (13724280 kB) free.
[1578499022] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39345104 kB) free.
[1578499053] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.5% (21695992 kB) free.
[1578499126] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1578499319] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16163948 kB) free!
[1578499439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12069248 kB) free!
[1578499559] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37762756 kB) free.
[1578499577] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191336 kB) free!
[1578499602] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578499625] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (218664 kB) free!
[1578499656] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (707020 kB) free!
[1578499697] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196892 kB) free!
[1578499722] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578499746] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233244 kB) free!
[1578499776] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (682356 kB) free!
[1578499817] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194364 kB) free!
[1578499842] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578499866] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246700 kB) free!
[1578499896] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (664288 kB) free!
[1578500159] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14858768 kB) free!
[1578500280] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39322120 kB) free.
[1578500879] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15279720 kB) free!
[1578500999] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39244356 kB) free.
[1578501526] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1578501599] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15767436 kB) free!
[1578501647] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1578501719] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (11969112 kB) free!
[1578501839] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36542912 kB) free.
[1578502439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12495740 kB) free!
[1578502559] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38334996 kB) free.
[1578502842] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578502846] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578502966] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1578503159] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13410032 kB) free!
[1578503279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38747156 kB) free.
[1578503879] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13181888 kB) free!
[1578503999] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38534324 kB) free.
[1578504166] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1578504286] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1578504599] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15071400 kB) free!
[1578504719] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11351784 kB) free!
[1578504839] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35604524 kB) free.
[1578505439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12442884 kB) free!
[1578505486] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1578505560] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36520672 kB) free.
[1578505606] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 388 processes
[1578505673] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/08 23:17:51 [error] 21682#0: *23651696 connect() failed (111: Connection refused) while connecting to upstream, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "10.147.212.143:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578505793] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578506159] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13118780 kB) free!
[1578506279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (36993632 kB) free.
[1578506879] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14111324 kB) free!
[1578506999] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37707480 kB) free.
[1578507599] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14594384 kB) free!
[1578507719] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (38014260 kB) free.
[1578507946] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (241040 kB) free!
[1578508066] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216724 kB) free!
[1578508186] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216944 kB) free!