[1578508319] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15130644 kB) free!
[1578508439] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37954248 kB) free.
[1578508606] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166280 kB) free!
[1578508727] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163656 kB) free!
[1578508787] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.0% (27578852 kB) free.
[1578508847] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196384 kB) free!
[1578509039] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15633152 kB) free!
[1578509159] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37852244 kB) free.
[1578509759] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16196924 kB) free!
[1578509879] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37836796 kB) free.
[1578510479] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16391860 kB) free!
[1578510599] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37802760 kB) free.
[1578510696] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578511200] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16469784 kB) free!
[1578511320] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37809028 kB) free.
[1578511920] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14647840 kB) free!
[1578512040] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37701244 kB) free.
[1578512446] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4891184 kB) free!
[1578512566] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.9% (4881344 kB) free!
[1578512640] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13634540 kB) free!
[1578512686] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.9% (4872456 kB) free!
[1578512760] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34942940 kB) free.
[1578513960] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14024956 kB) free!
[1578514080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34814260 kB) free.
[1578514464] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1305.00, 841.65, 370.14
[1578514584] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1312.59, 997.89, 484.57
[1578514704] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 295.03, 761.52, 464.06
[1578514793] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2020/01/09 01:47:35 [error] 21679#0: *23674922 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1578514913] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578515280] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15081868 kB) free!
[1578515305] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.29, 103.69, 244.09
[1578515400] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37218648 kB) free.
[1578516647] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.1% (10325032 kB) free.
[1578516657] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1578516777] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1578516883] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230560 kB) free!
[1578517003] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (250256 kB) free!
[1578517123] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229628 kB) free!
[1578517200] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15156568 kB) free!
[1578517321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34693744 kB) free.
[1578517723] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.1% (41427160 kB) free.
[1578518514] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2020/01/09 02:44:01 [error] 21678#0: *23682329 connect() failed (111: Connection refused) while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578518634] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(5) < 2020/01/09 02:53:44 [error] 21682#0: *23682799 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1578518754] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(7) < 2020/01/09 02:55:54 [error] 21675#0: *23682984 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1578518806] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1578518926] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 480 processes
[1578519096] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (2601264 kB) free!
[1578519660] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1199) < 92.42.104.166 [09/Jan/2020:03:10:12 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1578519720] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8434296 kB) free!
[1578519780] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(163) < 172.105.59.173 [09/Jan/2020:03:12:11 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1578519840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32842000 kB) free.
[1578519901] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578520260] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/09 03:11:12 [error] 18358#0: *35641917 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578520381] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578520555] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578522096] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2333092 kB) free!
[1578525177] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1578525297] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1578527640] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24943420 kB) free!
[1578527760] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38843200 kB) free.
[1578529847] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167784 kB) free!
[1578529967] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190252 kB) free!
[1578530087] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165956 kB) free!
[1578531096] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.0% (20830436 kB) free.
[1578531960] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26129864 kB) free!
[1578532080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39470788 kB) free.
[1578533541] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 623 processes
[1578533661] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 618 processes
[1578533686] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2365724 kB) free!
[1578533697] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1578533781] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 356 processes
[1578533817] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1578536280] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25795284 kB) free!
[1578536400] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (21737616 kB) free!
[1578536520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36333904 kB) free.
[1578538583] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 70.4% (17260112 kB) free.
[1578538602] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.5% (10745676 kB) free.
[1578538624] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (7140580 kB) free!
[1578538649] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 62.1% (15216376 kB) free.
[1578539208] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190960 kB) free!
[1578539329] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200100 kB) free!
[1578539449] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192884 kB) free!
[1578539520] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22918392 kB) free!
[1578539640] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18894024 kB) free!
[1578539760] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.9% (55270460 kB) free.
[1578539826] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (258112 kB) free!
[1578540300] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(160) < 92.42.104.166 [09/Jan/2020:08:48:19 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1578540355] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(237) < 2020/01/09 08:48:19 [error] 21681#0: *23746586 no live upstreams while connecting to upstream, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "backend_13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578540421] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578540475] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578541017] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1578541137] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1578546475] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2020/01/09 10:30:23 [error] 21677#0: *23688161 recv() failed (104: Connection reset by peer) while proxying connection, client: 85.10.207.167, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:3566850/985305, bytes from/to upstream:985305/3566850
[1578546595] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578547604] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 53.9% (13300452 kB) free.
[1578547700] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.8% (37834928 kB) free.
[1578548210] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191032 kB) free!
[1578548327] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1578548330] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166920 kB) free!
[1578548447] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1578548450] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193260 kB) free!
[1578548496] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191384 kB) free!
[1578548616] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181304 kB) free!
[1578548637] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1578548637] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.7% (22143996 kB) free.
[1578548668] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.7% (21742124 kB) free.
[1578548900] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (246728 kB) free!
[1578548995] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/09 11:17:08 [error] 21675#0: *23771551 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:10054/87, bytes from/to upstream:87/10054
[1578549005] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.1% (38990396 kB) free.
[1578549116] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578549245] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171444 kB) free!
[1578549269] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186256 kB) free!
[1578549360] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17799512 kB) free!
[1578549365] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174460 kB) free!
[1578549389] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166256 kB) free!
[1578549421] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.7% (21636568 kB) free.
[1578549446] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.5% (21925416 kB) free.
[1578549480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.2% (50299772 kB) free.
[1578549609] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1050412 kB) free!
[1578549729] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225932 kB) free!
[1578549849] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248592 kB) free!
[1578549858] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2870548 kB) free!
[1578549884] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1578549884] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 75.2% (37073496 kB) free.
[1578549925] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.1% (16865072 kB) free.
[1578549995] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1578550029] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169580 kB) free!
[1578550057] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181320 kB) free!
[1578550115] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1578550143] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22101088 kB) free.
[1578550162] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.5% (22178696 kB) free.
[1578550187] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8761568 kB) free!
[1578550307] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8091948 kB) free!
[1578550308] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.4% (8084436 kB) free!
[1578550315] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2020/01/09 11:39:42 [error] 21682#0: *23775403 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:12661/87, bytes from/to upstream:87/12661
[1578550404] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 57.6% (28401768 kB) free.
[1578550435] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578550484] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215840 kB) free!
[1578550604] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236544 kB) free!
[1578550724] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (254856 kB) free!
[1578550746] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173828 kB) free!
[1578550767] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184500 kB) free!
[1578550785] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39344380 kB) free.
[1578550801] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.7% (22371820 kB) free.
[1578550819] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.8% (22007136 kB) free.
[1578550842] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578550962] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578551035] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2020/01/09 11:53:40 [error] 21683#0: *23777791 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3811/104, bytes from/to upstream:104/3811
[1578551082] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578551155] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(4) < 2020/01/09 11:54:17 [error] 21675#0: *23777904 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3064/87, bytes from/to upstream:87/3064
[1578551235] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4237004 kB) free!
[1578551275] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578551315] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1578551355] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.5% (3836736 kB) free!
[1578551385] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3372652 kB) free!
[1578551404] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165580 kB) free!
[1578551429] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173652 kB) free!
[1578551435] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1578551446] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.9% (22679036 kB) free.
[1578551472] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.4% (40109220 kB) free.
[1578551475] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.7% (3389672 kB) free!
[1578551490] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.9% (20812228 kB) free.
[1578551855] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199688 kB) free!
[1578551876] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(53) < 2020/01/09 12:02:00 [error] 21682#0: *23779324 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:9475/176, bytes from/to upstream:176/9475
[1578551975] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191408 kB) free!
[1578551996] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578552048] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1578552052] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177864 kB) free!
[1578552075] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (922112 kB) free!
[1578552096] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180084 kB) free!
[1578552097] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173352 kB) free!
[1578552168] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1578552172] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181340 kB) free!
[1578552217] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191192 kB) free!
[1578552288] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1578552293] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184072 kB) free!
[1578552337] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194548 kB) free!
[1578552351] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.0% (22211856 kB) free.
[1578552390] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.3% (22614684 kB) free.
[1578552416] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.5% (22185788 kB) free.
[1578552523] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8383744 kB) free!
[1578552539] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.5% (23078704 kB) free.
[1578552596] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/01/09 12:19:26 [error] 21680#0: *23782439 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:20539/104, bytes from/to upstream:104/20539
[1578552643] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.1% (7927308 kB) free!
[1578552649] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.0% (7906168 kB) free!
[1578552675] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (213820 kB) free!
[1578552716] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578552735] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.5% (46062440 kB) free.
[1578552766] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39527992 kB) free.
[1578552956] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163760 kB) free!
[1578552986] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.2% (4464332 kB) free!
[1578553016] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169556 kB) free!
[1578553076] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169976 kB) free!
[1578553106] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.4% (4265304 kB) free!
[1578553136] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193856 kB) free!
[1578553146] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.6% (21966956 kB) free.
[1578553158] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.0% (20478380 kB) free.
[1578553187] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.6% (21960632 kB) free.
[1578553316] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2020/01/09 12:31:31 [error] 21681#0: *23784559 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3474/70, bytes from/to upstream:70/3474
[1578553376] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7413756 kB) free!
[1578553436] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578553488] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1578553497] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (249520 kB) free!
[1578553608] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1578553617] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243436 kB) free!
[1578553686] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 75.5% (37189580 kB) free.
[1578553728] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1578553772] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184312 kB) free!
[1578553797] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (952672 kB) free!
[1578553892] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186448 kB) free!
[1578553917] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (764056 kB) free!
[1578554012] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189216 kB) free!
[1578554037] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2020/01/09 12:40:55 [error] 21677#0: *23786213 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3760/87, bytes from/to upstream:87/3760
[1578554037] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188244 kB) free!
[1578554157] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(7) < 2020/01/09 12:45:05 [error] 21677#0: *23786969 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:26683/70, bytes from/to upstream:70/26683
[1578554159] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.2% (22250676 kB) free.
[1578554208] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1578554208] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.7% (22473408 kB) free.
[1578554277] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578554296] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239696 kB) free!
[1578554387] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.4% (37662844 kB) free.
[1578554766] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172108 kB) free!
[1578554797] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4638080 kB) free!
[1578554817] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187908 kB) free!
[1578554877] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/09 12:53:32 [error] 21681#0: *23788463 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:8656/104, bytes from/to upstream:104/8656
[1578554886] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.3% (3533088 kB) free!
[1578554892] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22322356 kB) free.
[1578554913] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.3% (21884892 kB) free.
[1578554917] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (4389812 kB) free!
[1578554928] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1578554945] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.9% (22527700 kB) free.
[1578554992] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (257004 kB) free!
[1578554997] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578555112] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (218592 kB) free!
[1578555232] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (244656 kB) free!
[1578555278] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.2% (38039436 kB) free.
[1578555496] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189176 kB) free!
[1578555516] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182328 kB) free!
[1578555616] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195236 kB) free!
[1578555636] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177540 kB) free!
[1578555662] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.6% (21126232 kB) free.
[1578555680] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.6% (22202824 kB) free.
[1578555885] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215708 kB) free!
[1578556005] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (222256 kB) free!
[1578556125] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (7360148 kB) free!
[1578556129] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1578556134] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39526612 kB) free.
[1578556249] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1578556262] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177984 kB) free!
[1578556287] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1849544 kB) free!
[1578556369] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1578556382] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174644 kB) free!
[1578556407] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195024 kB) free!
[1578556443] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21858660 kB) free.
[1578556465] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.6% (21955544 kB) free.
[1578556737] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.3% (7025872 kB) free!
[1578556857] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (2455104 kB) free!
[1578556898] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39731404 kB) free.
[1578557048] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183104 kB) free!
[1578557068] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (1256736 kB) free!
[1578557169] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187652 kB) free!
[1578557188] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (1097844 kB) free!
[1578557289] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185896 kB) free!
[1578557308] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183732 kB) free!
[1578557435] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1578557508] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (212676 kB) free!
[1578557555] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1578557628] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (211728 kB) free!
[1578557748] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236200 kB) free!
[1578557912] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22072792 kB) free!
[1578557986] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.8% (22407768 kB) free.
[1578557997] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/09 13:47:21 [error] 21681#0: *23797952 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:2639/87, bytes from/to upstream:87/2639
[1578558021] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.5% (22429312 kB) free.
[1578558025] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (18547308 kB) free!
[1578558046] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39619656 kB) free.
[1578558089] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 54.1% (71354008 kB) free.
[1578558117] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578558169] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1578558289] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1578558599] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193340 kB) free!
[1578558621] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.1% (3452328 kB) free!
[1578558647] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (1482328 kB) free!
[1578558668] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.7% (21154124 kB) free.
[1578558695] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.5% (39168884 kB) free.
[1578558710] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.6% (22195080 kB) free.
[1578558717] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/09 13:52:36 [error] 21680#0: *23798883 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:2665/70, bytes from/to upstream:70/2665
[1578558837] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578558889] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1578558925] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4168328 kB) free!
[1578559009] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1578559045] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.8% (3763288 kB) free!
[1578559129] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1578559165] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.0% (3361680 kB) free!
[1578559274] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197996 kB) free!
[1578559295] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238680 kB) free!
[1578559319] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187668 kB) free!
[1578559394] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193312 kB) free!
[1578559415] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240188 kB) free!
[1578559437] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/09 14:11:57 [error] 21680#0: *23802305 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:7940/104, bytes from/to upstream:104/7940
[1578559439] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166044 kB) free!
[1578559514] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195932 kB) free!
[1578559535] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239148 kB) free!
[1578559557] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578559559] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193164 kB) free!
[1578559700] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21951892 kB) free.
[1578559726] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39779332 kB) free.
[1578559765] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1436092 kB) free!
[1578559848] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.1% (16851668 kB) free.
[1578559870] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.5% (21947040 kB) free.
[1578560009] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7786812 kB) free!
[1578560130] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7206976 kB) free!
[1578560250] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.3% (6551748 kB) free!
[1578560300] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (225612 kB) free!
[1578560420] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188812 kB) free!
[1578560472] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194612 kB) free!
[1578560540] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192788 kB) free!
[1578560592] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182132 kB) free!
[1578560712] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194324 kB) free!
[1578560850] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3287104 kB) free!
[1578560947] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (4289088 kB) free!
[1578560964] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236280 kB) free!
[1578561067] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.7% (3875992 kB) free!
[1578561084] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (248608 kB) free!
[1578561187] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3476384 kB) free!
[1578561204] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233824 kB) free!
[1578561357] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/09 14:36:37 [error] 21677#0: *23806697 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:9351/87, bytes from/to upstream:87/9351
[1578561477] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578561545] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (4184712 kB) free!
[1578561665] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (3993036 kB) free!
[1578561785] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.4% (3782604 kB) free!
[1578561787] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1578561787] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1569144 kB) free!
[1578561910] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.8% (20781592 kB) free.
[1578561956] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.3% (22369832 kB) free.
[1578562012] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.7% (23131960 kB) free.
[1578562034] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.8% (22169624 kB) free.
[1578562055] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39612224 kB) free.
[1578562077] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2020/01/09 14:56:43 [error] 21681#0: *23810177 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:19747/121, bytes from/to upstream:121/19747
[1578562180] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 55.0% (27104232 kB) free.
[1578562197] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578562514] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174716 kB) free!
[1578562634] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170160 kB) free!
[1578562636] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176276 kB) free!
[1578562655] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (245080 kB) free!
[1578562729] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1578562729] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1578562754] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172468 kB) free!
[1578562756] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174908 kB) free!
[1578562775] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238152 kB) free!
[1578562797] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/09 15:04:57 [error] 21680#0: *23811643 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:6169/104, bytes from/to upstream:104/6169
[1578562876] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164140 kB) free!
[1578562895] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233972 kB) free!
[1578562918] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578563329] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1578563449] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1578563555] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 326 processes
[1578563676] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1578563765] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166404 kB) free!
[1578563847] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2100388 kB) free!
[1578563885] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193172 kB) free!
[1578563893] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.5% (21830272 kB) free.
[1578563922] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.9% (39865088 kB) free.
[1578563954] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21947004 kB) free.
[1578563967] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.7% (2388668 kB) free!
[1578563973] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.9% (20812956 kB) free.
[1578564001] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.2% (4261792 kB) free.
[1578564049] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1578564170] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1578564290] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1578564498] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196152 kB) free!
[1578564527] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (226096 kB) free!
[1578564550] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.4% (4760076 kB) free!
[1578564557] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171132 kB) free!
[1578564576] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164884 kB) free!
[1578564604] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2875904 kB) free!
[1578564618] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194856 kB) free!
[1578564648] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244460 kB) free!
[1578564671] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.8% (4617648 kB) free!
[1578564677] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183232 kB) free!
[1578564696] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (225612 kB) free!
[1578564717] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(147) < 2020/01/09 15:41:56 [error] 21678#0: *23818167 no live upstreams while connecting to upstream, client: 139.59.21.208, server: 0.0.0.0:13000, upstream: "backend_13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578564724] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3157208 kB) free!
[1578564738] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191912 kB) free!
[1578564768] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (252084 kB) free!
[1578564791] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.3% (4479996 kB) free!
[1578564797] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167464 kB) free!
[1578564816] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186864 kB) free!
[1578564837] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(8) < 2020/01/09 15:43:08 [error] 21682#0: *23818262 recv() failed (104: Connection reset by peer) while proxying connection, client: 139.59.21.208, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:89/17, bytes from/to upstream:17/89
[1578564844] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3701344 kB) free.
[1578564957] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578565021] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(81) < 103.16.101.72 [09/Jan/2020:15:41:57 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1578565141] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578565335] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9224044 kB) free!
[1578565444] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.1% (1801352 kB) free!
[1578565455] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8920156 kB) free!
[1578565490] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1578565564] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.7% (2221964 kB) free!
[1578565575] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8604200 kB) free!
[1578565684] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (2192100 kB) free!
[1578565892] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23439412 kB) free!
[1578566012] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19339128 kB) free!
[1578566132] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15467780 kB) free!
[1578566284] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.4% (4285324 kB) free.
[1578566885] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2722468 kB) free!
[1578567005] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3229424 kB) free!
[1578567125] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3800052 kB) free.
[1578567191] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1728336 kB) free!
[1578567333] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33257904 kB) free.
[1578567375] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4356956 kB) free!
[1578567725] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2557932 kB) free!
[1578567845] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3407240 kB) free.
[1578567882] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1578567890] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1578567933] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10715380 kB) free!
[1578568010] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1578568053] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6862224 kB) free!
[1578568130] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1578568173] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34094244 kB) free.
[1578568445] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2026676 kB) free!
[1578568565] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.8% (2570056 kB) free!
[1578568628] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.5% (21945908 kB) free.
[1578568656] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21928344 kB) free.
[1578568682] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.3% (22142780 kB) free.
[1578568685] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3030048 kB) free!
[1578568773] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9442196 kB) free!
[1578568893] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5436592 kB) free!
[1578569013] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1834908 kB) free!
[1578569116] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1578569116] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 58.4% (77000084 kB) free.
[1578569147] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.3% (45463780 kB) free.
[1578569282] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (2031616 kB) free!
[1578569285] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191860 kB) free!
[1578569402] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167352 kB) free!
[1578569405] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194688 kB) free!
[1578569522] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165448 kB) free!
[1578569526] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170500 kB) free!
[1578569885] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1578569885] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.0% (4548604 kB) free.
[1578570357] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/09 17:11:18 [error] 21677#0: *23833648 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:8140/70, bytes from/to upstream:70/8140
[1578570477] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578570655] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (3716364 kB) free!
[1578570775] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.1% (3383040 kB) free!
[1578570895] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3137668 kB) free!
[1578570926] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23092668 kB) free!
[1578571046] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.1% (54207972 kB) free.
[1578571495] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1583036 kB) free!
[1578572388] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9560748 kB) free!
[1578572508] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.8% (9251124 kB) free!
[1578572628] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.1% (8943468 kB) free!
[1578572682] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1578572816] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578572846] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20383348 kB) free!
[1578572886] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (2727724 kB) free!
[1578572966] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.4% (50559184 kB) free.
[1578573006] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3399480 kB) free.
[1578573229] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.4% (30767552 kB) free.
[1578573282] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1578573296] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 74.9% (16824156 kB) free.
[1578573416] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.0% (3707592 kB) free!
[1578573478] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(72) < 2020/01/09 18:03:35 [error] 21683#0: *23842595 no live upstreams while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "backend_13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578573536] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3371152 kB) free!
[1578573541] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(45) < 103.209.99.7 [09/Jan/2020:18:03:35 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1578573566] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26396256 kB) free!
[1578573598] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578573656] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.5% (3085188 kB) free!
[1578573661] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578573686] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.9% (22233828 kB) free!
[1578573806] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18345096 kB) free!
[1578574256] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.3% (23039788 kB) free.
[1578574406] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29701940 kB) free.
[1578574806] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2710364 kB) free!
[1578574926] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3392928 kB) free.
[1578574969] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1578575082] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1578575082] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1578575398] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(189) < 2020/01/09 18:34:04 [error] 21681#0: *23848254 no live upstreams while connecting to upstream, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "backend_13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578575461] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(131) < 92.42.104.166 [09/Jan/2020:18:34:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1578575518] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578575526] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3127772 kB) free!
[1578575581] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578575606] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19802700 kB) free!
[1578575646] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3917892 kB) free.
[1578575726] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15902184 kB) free!
[1578575847] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (12098120 kB) free!
[1578576118] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(49) < 2020/01/09 18:51:53 [error] 21669#0: *23851625 connect() failed (111: Connection refused) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578576181] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [09/Jan/2020:18:52:55 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1578576182] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(292) < 2020/01/09 18:48:33 [error] 18360#0: *36144679 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578576226] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(54) < 103.16.101.72 [09/Jan/2020:18:47:23 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1578576239] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(458) < 2020/01/09 18:53:58 [error] 21678#0: *23852144 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.27.86.30, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578576246] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3224824 kB) free!
[1578576302] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(10) < 203.212.70.170 [09/Jan/2020:18:54:54 +0530] TCP 502 0 0 180.000 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1578576302] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578576347] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578576359] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(38) < 2020/01/09 18:54:54 [error] 21681#0: *23851862 no live upstreams while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "backend_3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578576366] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.6% (3997576 kB) free.
[1578576422] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578576447] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23622088 kB) free!
[1578576530] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578576966] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3187800 kB) free!
[1578577047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39418764 kB) free.
[1578577086] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.3% (4281520 kB) free.
[1578577559] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578577647] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15112412 kB) free!
[1578577767] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11139212 kB) free!
[1578577887] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (7245768 kB) free!
[1578577926] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1578578286] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2538376 kB) free!
[1578578406] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3234696 kB) free!
[1578578487] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19132016 kB) free!
[1578578526] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3731424 kB) free.
[1578578759] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(773) < 2020/01/09 19:34:16 [error] 21680#0: *23861934 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578578823] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(761) < 103.255.146.164 [09/Jan/2020:19:34:16 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578578879] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578578943] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578579087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35578404 kB) free.
[1578579530] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578579650] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1578579687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11675308 kB) free!
[1578579807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7570720 kB) free!
[1578579927] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3635884 kB) free!
[1578580137] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1578580257] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1578580527] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16152268 kB) free!
[1578580922] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578580926] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2953080 kB) free!
[1578581046] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3667628 kB) free.
[1578581127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32152156 kB) free.
[1578581646] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2768896 kB) free!
[1578581727] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7593588 kB) free!
[1578581766] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578581766] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3335936 kB) free.
[1578581848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3504852 kB) free!
[1578581968] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395988 kB) free!
[1578582499] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.66, 6.01, 5.28
[1578582619] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.27, 5.71, 5.26
[1578582739] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.52, 5.34, 5.18
[1578583168] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30726936 kB) free.
[1578583322] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (2627236 kB) free!
[1578583339] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.60, 4.07, 4.62
[1578583566] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3027596 kB) free!
[1578583686] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3508620 kB) free.
[1578583768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7002692 kB) free!
[1578583888] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2887348 kB) free!
[1578583922] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2164748 kB) free!
[1578584008] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383700 kB) free!
[1578584608] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13516672 kB) free!
[1578585208] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31699436 kB) free.
[1578585282] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1578585739] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.19, 5.69, 5.10
[1578585808] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6110240 kB) free!
[1578585859] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.37, 5.63, 5.16
[1578585882] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1578585928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2119232 kB) free!
[1578585979] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.44, 5.24, 5.07
[1578586048] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394096 kB) free!
[1578586087] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1847824 kB) free!
[1578586207] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.3% (2002500 kB) free!
[1578586327] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.9% (4704176 kB) free.
[1578586579] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1578586579] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.96, 4.15, 4.55
[1578587248] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21030100 kB) free!
[1578587848] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407988 kB) free!
[1578589048] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38591400 kB) free.
[1578589648] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12895528 kB) free!
[1578589768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8602980 kB) free!
[1578589888] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4465772 kB) free!
[1578590488] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20313804 kB) free!
[1578591088] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38419896 kB) free.
[1578591688] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12323576 kB) free!
[1578591808] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8255768 kB) free!
[1578591928] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4260820 kB) free!
[1578592528] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18663712 kB) free!
[1578593128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33904024 kB) free.
[1578593728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10055412 kB) free!
[1578593848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6216816 kB) free!
[1578593968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30349436 kB) free.
[1578594229] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (242652 kB) free!
[1578594349] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (224184 kB) free!
[1578594434] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.8% (4865020 kB) free!
[1578594469] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238200 kB) free!
[1578594554] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.8% (4846592 kB) free!
[1578594568] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14549012 kB) free!