[1579631447] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32622628 kB) free.
[1579631738] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198532 kB) free!
[1579631876] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579632177] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197628 kB) free!
[1579632297] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189608 kB) free!
[1579632417] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189076 kB) free!
[1579632647] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14658740 kB) free!
[1579632767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34138892 kB) free.
[1579633316] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1579633367] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13158956 kB) free!
[1579633436] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1579633461] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185028 kB) free!
[1579633487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32120824 kB) free.
[1579633522] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194340 kB) free!
[1579633582] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 68.5% (16779424 kB) free.
[1579634147] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.0% (9871132 kB) free.
[1579634687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14487336 kB) free!
[1579634808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32408112 kB) free.
[1579635305] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579635347] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4634596 kB) free!
[1579635426] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579635467] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.8% (3159264 kB) free!
[1579635588] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1643808 kB) free!
[1579636007] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14737036 kB) free!
[1579636041] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(705) < 49.50.86.152 [22/Jan/2020:01:17:16 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579636127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32738584 kB) free.
[1579636161] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(32) < 203.212.70.170 [22/Jan/2020:01:18:12 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1579636246] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2096) < 2020/01/22 01:19:12 [error] 21679#0: *25803242 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579636282] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579636367] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579636939] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.2% (5511940 kB) free!
[1579637059] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.2% (5501720 kB) free!
[1579637179] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (5501196 kB) free!
[1579638417] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.6% (10941788 kB) free.
[1579639727] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15005960 kB) free!
[1579639807] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2255392 kB) free!
[1579639847] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31413668 kB) free.
[1579639928] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243548 kB) free!
[1579640048] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227772 kB) free!
[1579640075] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3539096 kB) free!
[1579640195] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3537584 kB) free!
[1579640315] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 80.8% (39800652 kB) free.
[1579641248] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.4% (42093896 kB) free.
[1579642282] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3194) < 203.212.70.170 [22/Jan/2020:03:00:41 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579642366] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3886) < 2020/01/22 03:02:42 [error] 21677#0: *25820869 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579642402] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(235) < 203.212.70.170 [22/Jan/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579642435] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(257) < 103.16.101.72 [22/Jan/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1579642487] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(14) < 2020/01/22 03:03:26 [error] 21678#0: *25820740 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:87/0, bytes from/to upstream:0/87
[1579642522] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579642556] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579642607] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579642760] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(330) < 2020/01/22 03:01:32 [error] 18370#0: *42327443 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579642880] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579644680] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/22 03:37:26 [error] 18370#0: *42330425 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1579644800] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579646388] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579646447] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25165500 kB) free!
[1579646567] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.2% (20087632 kB) free!
[1579646687] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15839660 kB) free!
[1579647887] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33256492 kB) free.
[1579649817] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199208 kB) free!
[1579649938] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164684 kB) free!
[1579650058] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (1076360 kB) free!
[1579650338] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.2% (21385100 kB) free.
[1579650588] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.6% (11753060 kB) free.
[1579650887] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24183172 kB) free!
[1579651008] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35394020 kB) free.
[1579651522] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 172.105.59.173 [22/Jan/2020:05:34:10 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579651579] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4949648 kB) free!
[1579651607] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(610) < 2020/01/22 05:34:10 [error] 21680#0: *25837959 no live upstreams while connecting to upstream, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "backend_3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579651642] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579651727] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579653407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24816392 kB) free!
[1579653527] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35408348 kB) free.
[1579655927] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24730420 kB) free!
[1579656047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35163920 kB) free.
[1579658225] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579658345] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579658447] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25027516 kB) free!
[1579658568] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34874768 kB) free.
[1579659527] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(441) < 2020/01/22 07:48:38 [error] 21678#0: *25852837 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579659647] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/22 07:49:00 [error] 21682#0: *25852715 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:70/17, bytes from/to upstream:17/70
[1579659767] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579660042] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(15) < 203.212.70.170 [22/Jan/2020:07:48:03 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579660162] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579660367] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26413692 kB) free!
[1579660487] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22561752 kB) free!
[1579660607] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33907132 kB) free.
[1579662407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22688364 kB) free!
[1579662527] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34042792 kB) free.
[1579664276] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 59.66, 101.08, 51.37
[1579664327] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23008804 kB) free!
[1579664397] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16.25, 70.46, 46.38
[1579664447] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (19019256 kB) free!
[1579664520] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.96, 50.98, 42.15
[1579664567] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32091440 kB) free.
[1579665563] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [22/Jan/2020:09:24:11 +0530] TCP 502 0 0 63.004 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "3.004, 60.000, 0.000"
[1579665683] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579665767] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(677) < 2020/01/22 09:26:22 [error] 21677#0: *25864668 upstream timed out (110: Connection timed out) while connecting to upstream, client: 182.18.144.246, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579665887] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579666321] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.36, 11.97, 36.46
[1579666921] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 10.03, 11.17, 24.59
[1579667109] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.3% (4731472 kB) free!
[1579667229] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.2% (4698648 kB) free!
[1579667349] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.0% (4661004 kB) free!
[1579667521] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 195.95, 50.03, 30.17
[1579667641] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 639.00, 270.83, 113.59
[1579667761] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1579667761] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 497.95, 366.52, 169.28
[1579669189] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197636 kB) free!
[1579669283] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 203.212.70.170 [22/Jan/2020:10:30:11 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1579669309] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182380 kB) free!
[1579669403] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579669429] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193116 kB) free!
[1579669488] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(725) < 2020/01/22 10:32:08 [error] 21681#0: *25871156 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579669608] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579669749] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2398232 kB) free!
[1579670208] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/01/22 10:45:17 [error] 21677#0: *25873464 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579670328] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579670485] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21783964 kB) free.
[1579670513] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.4% (22561616 kB) free.
[1579670565] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40315544 kB) free.
[1579670598] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.8% (22979160 kB) free.
[1579670600] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/22 10:48:45 [error] 18365#0: *42384080 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:223/17, bytes from/to upstream:17/223
[1579670720] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579671961] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 18.91, 18.86, 37.93
[1579672174] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.3% (14351068 kB) free.
[1579672317] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.5% (2825940 kB) free!
[1579672409] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.1% (22478188 kB) free.
[1579672446] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4307792 kB) free!
[1579672561] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 24.86, 21.83, 29.99
[1579672566] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.3% (3879824 kB) free!
[1579672687] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.5% (3489408 kB) free!
[1579673287] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1713872 kB) free!
[1579673589] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 74.2% (16650300 kB) free.
[1579675022] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (4035976 kB) free!
[1579675142] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3684852 kB) free!
[1579675262] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.4% (3318904 kB) free!
[1579675862] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.2% (22997484 kB) free.
[1579676616] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (1000040 kB) free!
[1579676736] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (836828 kB) free!
[1579676856] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (681532 kB) free!
[1579677127] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1579677127] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.6% (22354324 kB) free.
[1579678128] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2020/01/22 12:58:40 [error] 6237#0: *25879895 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1579678248] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(13) < 2020/01/22 13:00:40 [error] 6243#0: *25879977 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1579678368] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(14) < 2020/01/22 13:02:40 [error] 6243#0: *25880087 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1579678745] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579678865] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1579679567] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17263892 kB) free!
[1579679687] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35819380 kB) free.
[1579680887] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14482160 kB) free!
[1579680938] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4770160 kB) free!
[1579681008] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32951808 kB) free.
[1579681058] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4655092 kB) free!
[1579681178] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (4565436 kB) free!
[1579681327] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4902952 kB) free!
[1579681447] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4742264 kB) free!
[1579681568] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4592624 kB) free!
[1579682016] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (4009396 kB) free!
[1579682136] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (3877636 kB) free!
[1579682207] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13485688 kB) free!
[1579682256] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.1% (3700520 kB) free!
[1579682327] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32908068 kB) free.
[1579682568] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (7709696 kB) free!
[1579682688] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.6% (7698764 kB) free!
[1579682768] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (363552 kB) free!
[1579682808] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (8239764 kB) free!
[1579682927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11081924 kB) free!
[1579683047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42206676 kB) free.
[1579683456] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1953696 kB) free!
[1579683578] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2462696 kB) free!
[1579683647] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20957292 kB) free!
[1579683767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40516172 kB) free.
[1579683800] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.1% (3840260 kB) free!
[1579683920] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3513524 kB) free!
[1579683961] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579684041] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3166624 kB) free!
[1579684367] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18941700 kB) free!
[1579684487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39714476 kB) free.
[1579684641] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1737164 kB) free!
[1579685087] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17812428 kB) free!
[1579685207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38802588 kB) free.
[1579685488] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.3% (4234084 kB) free!
[1579685608] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.7% (4083648 kB) free!
[1579685728] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.1% (3943832 kB) free!
[1579685807] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16975332 kB) free!
[1579685841] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 74.6% (16751060 kB) free.
[1579685927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38556640 kB) free.
[1579686126] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21665132 kB) free.
[1579686170] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.0% (21333888 kB) free.
[1579686223] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.4% (21924748 kB) free.
[1579686254] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.6% (40199172 kB) free.
[1579686287] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21788228 kB) free.
[1579686527] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17225252 kB) free!
[1579686648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37578836 kB) free.
[1579686727] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188592 kB) free!
[1579686789] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.4% (20095156 kB) free.
[1579687247] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16181816 kB) free!
[1579687367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36719308 kB) free.
[1579687400] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4723868 kB) free!
[1579687521] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (640520 kB) free!
[1579687641] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190932 kB) free!
[1579687842] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1579687962] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1579687967] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16234408 kB) free!
[1579688087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36067976 kB) free.
[1579688593] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4825036 kB) free!
[1579688687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14867024 kB) free!
[1579688713] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4733640 kB) free!
[1579688807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35281072 kB) free.
[1579688833] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1579688833] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4618028 kB) free!
[1579689407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14188428 kB) free!
[1579689527] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34296452 kB) free.
[1579690127] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13298628 kB) free!
[1579690186] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21784588 kB) free.
[1579690247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9495152 kB) free!
[1579690367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27878604 kB) free.
[1579690634] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1982616 kB) free!
[1579690962] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1579690967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (383016 kB) free!
[1579691048] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9453092 kB) free!
[1579691082] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 342 processes
[1579691088] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (379092 kB) free!
[1579691131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 52.6% (69381324 kB) free.
[1579691159] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.0% (22309348 kB) free.
[1579691168] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9235716 kB) free!
[1579691202] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1579691204] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.4% (45512280 kB) free.
[1579691915] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9796256 kB) free!
[1579692035] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9643744 kB) free!
[1579692155] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (9510196 kB) free!
[1579692803] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(369) < 103.255.146.164 [22/Jan/2020:17:02:18 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579692923] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579693538] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16579140 kB) free!
[1579693659] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.2% (50420876 kB) free.
[1579694258] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22072712 kB) free!
[1579694378] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17411316 kB) free!
[1579694498] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.5% (49449968 kB) free.
[1579694657] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8361464 kB) free!
[1579694778] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8354068 kB) free!
[1579694898] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (8337032 kB) free!
[1579695098] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20877072 kB) free!
[1579695218] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16406708 kB) free!
[1579695338] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45531128 kB) free.
[1579695938] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16146608 kB) free!
[1579696041] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (3860748 kB) free!
[1579696058] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.5% (25697016 kB) free!
[1579696161] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3649120 kB) free!
[1579696178] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30668348 kB) free.
[1579696281] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 74.6% (16743356 kB) free.
[1579696778] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8363700 kB) free!
[1579696898] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34758572 kB) free.
[1579697389] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (204728 kB) free!
[1579697498] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10944356 kB) free!
[1579697509] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189700 kB) free!
[1579697555] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4921460 kB) free!
[1579697618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31254576 kB) free.
[1579697629] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166524 kB) free!
[1579698218] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13258964 kB) free!
[1579698339] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9248500 kB) free!
[1579698459] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34642956 kB) free.
[1579698921] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1180_n"
[1579698967] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167680 kB) free!
[1579699042] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1579699059] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10713376 kB) free!
[1579699087] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206012 kB) free!
[1579699179] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6404988 kB) free!
[1579699207] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177820 kB) free!
[1579699299] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33652416 kB) free.
[1579699427] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193832 kB) free!
[1579699547] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177260 kB) free!
[1579699667] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169400 kB) free!
[1579699899] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10361800 kB) free!
[1579700019] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5976476 kB) free!
[1579700139] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31198712 kB) free.
[1579700739] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5997284 kB) free!
[1579700859] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32690672 kB) free.
[1579700898] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4452172 kB) free!
[1579701459] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7498996 kB) free!
[1579701579] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3342788 kB) free!
[1579701699] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31643868 kB) free.
[1579702299] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6903104 kB) free!
[1579702419] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2574632 kB) free!
[1579702540] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30865552 kB) free.
[1579702663] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.6% (22605084 kB) free.
[1579702695] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 26.6% (13088348 kB) free.
[1579702734] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.0% (38924892 kB) free.
[1579702765] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 73.1% (17922240 kB) free.
[1579702793] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.5% (22428488 kB) free.
[1579703139] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4213908 kB) free!
[1579703259] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408480 kB) free!
[1579703379] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30908092 kB) free.
[1579703979] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6249904 kB) free!
[1579704099] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2391236 kB) free!
[1579704219] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29334968 kB) free.
[1579704468] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (739008 kB) free!
[1579704588] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (715940 kB) free!
[1579704708] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (698388 kB) free!
[1579704819] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6979292 kB) free!
[1579704939] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3164108 kB) free!
[1579705059] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28852804 kB) free.
[1579705659] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6078572 kB) free!
[1579705779] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30499740 kB) free.
[1579706379] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8887924 kB) free!
[1579706499] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5165952 kB) free!
[1579706619] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30347288 kB) free.
[1579707220] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7644492 kB) free!
[1579707340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (4006024 kB) free!
[1579707460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28530956 kB) free.
[1579708060] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6338696 kB) free!
[1579708180] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30121676 kB) free.
[1579708780] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4690872 kB) free!
[1579708900] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1210720 kB) free!
[1579709020] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29424716 kB) free.
[1579709030] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.0% (22449520 kB) free.
[1579709620] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5626144 kB) free!
[1579709740] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (19089872 kB) free!
[1579709860] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26429444 kB) free!
[1579710460] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4538316 kB) free!
[1579710568] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1579711660] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17175152 kB) free!
[1579712860] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1579712860] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3301512 kB) free!
[1579713065] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579713186] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579714060] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13769044 kB) free!
[1579714439] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1579715260] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27120800 kB) free.
[1579715860] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5942692 kB) free!
[1579715980] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29209852 kB) free.
[1579716580] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7218196 kB) free!
[1579716605] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (223128 kB) free!
[1579716680] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26359968 kB) free!
[1579716700] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29222584 kB) free.
[1579716725] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235620 kB) free!
[1579716800] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.4% (20292536 kB) free!
[1579716845] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248420 kB) free!
[1579716920] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12545520 kB) free!
[1579717300] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8240152 kB) free!
[1579717385] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579717420] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29193532 kB) free.
[1579717505] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes