[1556649067] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12869032 kB) free.
[1556649876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1556650476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1556650597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556650717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556651317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1556651437] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556651467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9621008 kB) free!
[1556651557] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1556651587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13358884 kB) free.
[1556651943] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 109.92, 170.12, 91.98
[1556652063] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.13, 113.93, 80.87
[1556652097] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2807232 kB) free!
[1556652183] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.12, 76.29, 71.08
[1556652757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556652878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556652998] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556653323] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 258.14, 195.56, 95.11
[1556653383] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 1.68, 19.74
[1556653443] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 258.66, 216.50, 114.97
[1556653563] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 81.27, 176.70, 113.71
[1556654798] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556655183] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 0.40, 3.15
[1556655363] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.83, 2.24, 17.82
[1556655998] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556656118] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556657017] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 20.16, 102.48, 67.68
[1556657122] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2767240 kB) free!
[1556657137] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.91, 68.64, 59.50
[1556657163] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.86, 1.85, 4.06
[1556657257] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.50, 45.98, 52.31
[1556657918] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556658038] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556658457] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.16, 1.02, 14.52
[1556659238] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556659358] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556659657] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.31, 4.19
[1556660559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556660587] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (8997024 kB) free!
[1556660679] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556660707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13394792 kB) free.
[1556660722] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2377052 kB) free!
[1556661322] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (2487120 kB) free!
[1556661879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1556661907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8915420 kB) free!
[1556661922] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1556661922] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2454072 kB) free!
[1556661999] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1556662027] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13408000 kB) free.
[1556662297] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2355404 kB) free!
[1556662599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556662719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1556663228] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9250884 kB) free!
[1556663348] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13412184 kB) free.
[1556664519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556664548] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9166120 kB) free!
[1556664639] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556664668] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13408964 kB) free.
[1556665839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556665868] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9028112 kB) free!
[1556665959] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1556665959] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556665988] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13421464 kB) free.
[1556667160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556667188] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8892492 kB) free!
[1556667280] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556667308] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13417104 kB) free.
[1556667963] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 253.46, 70.46, 25.19
[1556668083] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 420.36, 194.79, 76.23
[1556668203] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 442.58, 277.92, 121.07
[1556668480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1556668600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556668720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1556669708] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9296688 kB) free!
[1556669828] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13209260 kB) free.
[1556671028] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8665388 kB) free!
[1556671148] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12702204 kB) free.
[1556671203] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.53, 2.26, 13.39
[1556671917] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.9% (20811104 kB) free.
[1556671990] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.1% (22971440 kB) free.
[1556672038] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.9% (22047468 kB) free.
[1556672100] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556672159] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.5% (14545508 kB) free.
[1556672213] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.9% (21779444 kB) free.
[1556672320] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1556672920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1556673003] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.62, 1.70, 3.50
[1556673040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556673160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1556673548] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9139124 kB) free!
[1556673668] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12893192 kB) free.
[1556674203] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 256.85, 150.30, 64.33
[1556674323] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 260.00, 186.80, 88.13
[1556674443] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1556674443] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 261.03, 211.39, 109.10
[1556676068] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9444628 kB) free!
[1556676160] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1556676160] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1556676188] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13155064 kB) free.
[1556676243] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.58, 2.25, 18.54
[1556676760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1556676880] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1556677000] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1556677600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556677721] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1556677841] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556678043] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.68, 1.54, 3.99
[1556678441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556678561] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1556678588] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9316204 kB) free!
[1556678681] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1556678708] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13020240 kB) free.
[1556679127] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (329788 kB) free!
[1556679183] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 189.17, 62.18, 22.34
[1556679195] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (1015956 kB) free!
[1556679247] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (275520 kB) free!
[1556679281] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1556679303] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 241.46, 124.23, 49.91
[1556679315] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (517712 kB) free!
[1556679367] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (215060 kB) free!
[1556679401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556679417] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.2% (3962704 kB) free!
[1556679423] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 248.49, 165.69, 74.12
[1556679436] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165904 kB) free!
[1556679521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1556679537] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.1% (3950296 kB) free!
[1556679657] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.1% (3936824 kB) free!
[1556680122] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1556680242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556680362] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1556680657] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 196.42, 76.91, 28.56
[1556680777] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 220.49, 125.68, 52.30
[1556680897] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 223.86, 158.32, 73.15
[1556680962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1556681082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1556681108] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9741672 kB) free!
[1556681202] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1556681223] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 1.23, 16.08
[1556681228] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13280212 kB) free.
[1556681457] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (213456 kB) free!
[1556681802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1556681922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1556682043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556682428] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8761372 kB) free!
[1556682548] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.6% (12609648 kB) free.
[1556682642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556682697] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.18, 0.97, 13.32
[1556682762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1556682882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1556683023] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.24, 1.77, 3.36
[1556683482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1556683602] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556683722] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1556683857] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (2961248 kB) free!
[1556683897] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 17.09, 123.60, 92.00
[1556684322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1556684442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556684457] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1556684457] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174812 kB) free!
[1556684544] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 580 processes
[1556684563] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1556684604] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 572 processes
[1556684948] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8686672 kB) free!
[1556685068] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12738868 kB) free.
[1556685162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556685282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556685402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556685698] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1556685698] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.31, 1.42, 14.11
[1556686268] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8639076 kB) free!
[1556686388] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12708768 kB) free.
[1556686898] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 0.76, 4.54
[1556687202] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1556687767] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (164012 kB) free!
[1556687802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556687887] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165180 kB) free!
[1556687923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1556688007] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171176 kB) free!
[1556688043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1556688243] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1556688607] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 80.7% (13124468 kB) free.
[1556688643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556688763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556688788] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9239844 kB) free!
[1556688883] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556688908] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12811692 kB) free.
[1556689483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556689603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556689724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556689780] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.1% (23477164 kB) free.
[1556689838] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.3% (19441600 kB) free.
[1556690043] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 261.00, 76.15, 27.25
[1556690108] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8730300 kB) free!
[1556690163] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 36.84, 51.60, 24.18
[1556690228] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12766096 kB) free.
[1556690274] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1556690283] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.69, 35.13, 21.47
[1556690389] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192536 kB) free!
[1556690394] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1556690509] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165764 kB) free!
[1556690514] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1556690629] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192116 kB) free!
[1556690883] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.68, 6.29, 12.13
[1556691046] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (670068 kB) free!
[1556691166] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (601568 kB) free!
[1556691286] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (718568 kB) free!
[1556691428] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7063056 kB) free!
[1556691548] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11267500 kB) free.
[1556692083] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.79, 1.69, 4.50
[1556692148] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9260052 kB) free!
[1556692268] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7692468 kB) free!
[1556692388] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11704112 kB) free.
[1556692988] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9715548 kB) free!
[1556693108] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8228708 kB) free!
[1556693228] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12028936 kB) free.
[1556694428] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7918288 kB) free!
[1556694524] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556694548] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11977588 kB) free.
[1556695125] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556695149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9679300 kB) free!
[1556695245] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556695269] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8196844 kB) free!
[1556695298] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 38.03, 19.16, 7.52
[1556695365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1556695389] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11950452 kB) free.
[1556695418] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.14, 13.25, 6.77
[1556695538] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.97, 9.29, 6.10
[1556695857] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556695965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556696085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556696138] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.64, 3.44, 4.37
[1556696205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556696589] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7700272 kB) free!
[1556696709] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11748552 kB) free.
[1556696883] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 271.56, 217.54, 110.73
[1556697003] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 249.88, 230.92, 128.78
[1556697123] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1556697123] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.88, 154.96, 113.33
[1556697309] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9735132 kB) free!
[1556697429] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8301720 kB) free!
[1556697549] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11867572 kB) free.
[1556698006] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556698606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556698726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556698749] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8170732 kB) free!
[1556698846] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556698869] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11810980 kB) free.
[1556698923] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.81, 2.05, 17.73
[1556699446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556699566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556699686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556700070] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8183556 kB) free!
[1556700190] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11786384 kB) free.
[1556701625] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 239.11, 136.39, 56.82
[1556701746] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 245.12, 172.67, 79.76
[1556701866] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 244.82, 196.66, 99.82
[1556701923] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.38, 1.77, 4.47
[1556701990] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7957604 kB) free!
[1556702086] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556702110] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11706864 kB) free.
[1556702686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556702738] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 202.25, 71.26, 26.16
[1556702806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556702858] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 243.32, 130.31, 53.27
[1556702926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556702978] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 248.62, 169.69, 77.04
[1556703310] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7937092 kB) free!
[1556703430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11653436 kB) free.
[1556703607] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.0% (2923300 kB) free!
[1556703666] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 1.08, 15.45
[1556703727] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.5% (2839768 kB) free!
[1556703847] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.9% (2754700 kB) free!
[1556704630] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8109784 kB) free!
[1556704750] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604148 kB) free.
[1556704778] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.35, 3.53, 16.39
[1556704866] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.56, 0.55, 4.62
[1556705647] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1415608 kB) free!
[1556705951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7545420 kB) free!
[1556706071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11540748 kB) free.
[1556706526] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1556706578] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.20, 2.30, 3.79
[1556706671] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9602628 kB) free!
[1556706791] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8090596 kB) free!
[1556706911] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6282516 kB) free!
[1556707126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556707246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556707366] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1556707511] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9949200 kB) free.
[1556707966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556708086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556708111] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7672144 kB) free!
[1556708207] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1556708231] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11457096 kB) free.
[1556708806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556708831] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9466180 kB) free!
[1556708926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556708951] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7999808 kB) free!
[1556709046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556709071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10554032 kB) free.
[1556709671] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9795956 kB) free!
[1556709792] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8253180 kB) free!
[1556709912] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6371216 kB) free!
[1556710923] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1556711712] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1556711712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11020192 kB) free.
[1556712312] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9001880 kB) free!
[1556712432] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7524512 kB) free!
[1556712552] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11299200 kB) free.
[1556713152] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9438648 kB) free!
[1556713272] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7993640 kB) free!
[1556713392] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6224460 kB) free!
[1556713846] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1556714446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556714566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556714686] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1556715192] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11069976 kB) free.
[1556715286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556715406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556715527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556715792] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8967396 kB) free!
[1556715912] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7372556 kB) free!
[1556716032] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11199536 kB) free.
[1556716632] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9146352 kB) free!
[1556716752] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7720224 kB) free!
[1556716778] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1556716872] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11169672 kB) free.
[1556717327] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1556717472] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9594980 kB) free!
[1556717592] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8069500 kB) free!
[1556717712] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6287580 kB) free!
[1556717927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1556717978] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 366.89, 134.86, 49.59
[1556718047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556718098] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 428.44, 235.16, 96.68
[1556718167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556718218] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 435.91, 302.00, 137.97
[1556718312] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9885988 kB) free.
[1556718913] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7656384 kB) free!
[1556719033] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11109580 kB) free.
[1556719633] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9367492 kB) free!
[1556719753] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7882712 kB) free!
[1556719873] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10490396 kB) free.
[1556719923] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 280.48, 136.46, 54.40
[1556720043] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 292.20, 188.58, 83.44
[1556720163] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 293.79, 223.52, 109.00
[1556720473] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9644692 kB) free!
[1556720593] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8120088 kB) free!
[1556720618] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 0.70, 12.61
[1556720713] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6391308 kB) free!
[1556721066] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 56.31, 12.87, 4.97
[1556721167] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1556721186] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 225.32, 91.88, 34.88
[1556721306] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 247.59, 144.58, 61.11
[1556721767] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556721818] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.46, 3.76
[1556721847] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.1% (3694392 kB) free!
[1556721887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556721968] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3650972 kB) free!
[1556722007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556722088] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.8% (3622540 kB) free!
[1556722122] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1556722242] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1556722513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11028852 kB) free.
[1556722563] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.54, 1.68, 13.56
[1556723106] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.36, 3.54, 19.26
[1556723114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8848228 kB) free!
[1556723234] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7342664 kB) free!
[1556723354] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11025056 kB) free.
[1556723763] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.79, 1.70, 4.92
[1556723807] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1556723954] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9102324 kB) free!
[1556724042] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1556724074] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7553464 kB) free!
[1556724163] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1556724194] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11005376 kB) free.
[1556724407] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556724527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556724647] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556724794] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9409212 kB) free!
[1556724906] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.44, 3.20
[1556724914] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7978652 kB) free!
[1556725034] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6386896 kB) free!
[1556725247] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556725362] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1556725367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556725482] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1556725487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556725634] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10108752 kB) free.
[1556726234] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7980504 kB) free!
[1556726354] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6236968 kB) free!
[1556726475] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10261240 kB) free.
[1556727074] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8320764 kB) free!
[1556727194] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6831760 kB) free!
[1556727287] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1556727314] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10779260 kB) free.
[1556727887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556727914] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8676544 kB) free!
[1556728007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556728034] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7168768 kB) free!
[1556728127] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1556728154] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10959752 kB) free.
[1556728727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1556728754] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8897780 kB) free!
[1556728847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1556728874] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7400156 kB) free!
[1556728968] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1556728994] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10949752 kB) free.
[1556729567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556729594] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9216012 kB) free!
[1556729687] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556729714] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7806484 kB) free!
[1556729807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556729835] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6130988 kB) free!
[1556730882] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1556731002] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1556731607] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1556731635] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10610872 kB) free.
[1556732208] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1556732235] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8576860 kB) free!
[1556732328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1556732355] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7057460 kB) free!
[1556732448] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556732475] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10917836 kB) free.
[1556733048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556733075] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8991760 kB) free!
[1556733168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556733195] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7551736 kB) free!
[1556733288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1556733315] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10913912 kB) free.
[1556733915] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9323328 kB) free!
[1556734036] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7868556 kB) free!
[1556734156] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6148432 kB) free!