[1553106875] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553107177] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.6% (7678816 kB) free!
[1553107475] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553107595] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553107778] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11982172 kB) free.
[1553108195] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553108315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553108378] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4328616 kB) free!
[1553108435] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553108498] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11971088 kB) free.
[1553108985] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1553109098] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6624016 kB) free!
[1553109218] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.8% (5308948 kB) free!
[1553109338] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3457744 kB) free!
[1553109636] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553109938] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7458876 kB) free!
[1553110236] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553110356] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553110538] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11967784 kB) free.
[1553110956] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553111076] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553111138] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4167724 kB) free!
[1553111196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553111258] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11978940 kB) free.
[1553111858] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6509176 kB) free!
[1553111978] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5201200 kB) free!
[1553112098] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3209948 kB) free!
[1553112396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553112698] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7322944 kB) free!
[1553112997] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553113117] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553113298] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11977188 kB) free.
[1553113717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553113837] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553113898] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3571500 kB) free!
[1553113957] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553114018] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11969308 kB) free.
[1553114557] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553114618] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6184644 kB) free!
[1553114677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553114738] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (4867592 kB) free!
[1553114797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553114858] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11978272 kB) free.
[1553115458] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7003976 kB) free!
[1553115578] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.5% (5691480 kB) free!
[1553115698] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4092632 kB) free!
[1553115998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553116017] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 48.3% (23783132 kB) free.
[1553116299] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7838900 kB) free!
[1553116598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553116718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553116899] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11971240 kB) free.
[1553117318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553117438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553117499] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4626520 kB) free!
[1553117558] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553117619] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11965932 kB) free.
[1553118219] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6884572 kB) free!
[1553118339] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5569676 kB) free!
[1553118459] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3911000 kB) free!
[1553118758] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553119059] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7816948 kB) free!
[1553119359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553119479] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553119659] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11955296 kB) free.
[1553120079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553120199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553120259] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4856596 kB) free!
[1553120319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553120323] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1525728 kB) free!
[1553120379] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11965156 kB) free.
[1553120980] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7080492 kB) free!
[1553121100] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5774728 kB) free!
[1553121220] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4282532 kB) free!
[1553121519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553121820] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.3% (8051444 kB) free!
[1553122119] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553122240] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553122420] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11963112 kB) free.
[1553122839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553122959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553123020] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4945880 kB) free!
[1553123079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553123140] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11950988 kB) free.
[1553123740] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7071732 kB) free!
[1553123860] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.6% (5730212 kB) free!
[1553123980] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4164332 kB) free!
[1553124279] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553124580] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.0% (7869596 kB) free!
[1553124879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553124999] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553125180] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11952892 kB) free.
[1553125599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553125720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553125781] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4697596 kB) free!
[1553125840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553125901] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11956932 kB) free.
[1553126501] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6885740 kB) free!
[1553126621] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5589144 kB) free!
[1553126741] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3939144 kB) free!
[1553127040] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553127341] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7745524 kB) free!
[1553127640] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553127760] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553127941] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11955528 kB) free.
[1553128360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553128480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553128541] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4300444 kB) free!
[1553128601] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553128661] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11951860 kB) free.
[1553129261] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6573544 kB) free!
[1553129381] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5236360 kB) free!
[1553129502] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3332104 kB) free!
[1553129801] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553130102] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7308648 kB) free!
[1553130401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553130521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553130702] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11948188 kB) free.
[1553131121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553131241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553131302] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3636244 kB) free!
[1553131361] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553131422] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11957036 kB) free.
[1553131893] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56088 MB (24% inode=99%):
[1553131962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553132022] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6271056 kB) free!
[1553132082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553132142] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4940196 kB) free!
[1553132202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553132262] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11950160 kB) free.
[1553132303] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 77.0% (18871480 kB) free.
[1553132862] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7114932 kB) free!
[1553132982] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5771636 kB) free!
[1553133102] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4244332 kB) free!
[1553133402] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553133702] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7930020 kB) free!
[1553134002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553134122] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553134303] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11946132 kB) free.
[1553134722] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553134843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553134903] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4643900 kB) free!
[1553134963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553135023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11955740 kB) free.
[1553135623] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6826432 kB) free!
[1553135743] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5520344 kB) free!
[1553135863] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3839556 kB) free!
[1553136463] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.1% (6961688 kB) free!
[1553136763] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553137063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11845328 kB) free.
[1553137363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553137483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553137603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553137663] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3250568 kB) free!
[1553137783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11943040 kB) free.
[1553138383] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5980100 kB) free!
[1553138503] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (4555064 kB) free!
[1553138624] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11948856 kB) free.
[1553138803] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553139223] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6811188 kB) free!
[1553139343] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5499220 kB) free!
[1553139403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553139463] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3790364 kB) free!
[1553139523] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553140063] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4981092 kB) free!
[1553140123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553140243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553140843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553140964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553141084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553143063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10468660 kB) free.
[1553144264] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4368320 kB) free!
[1553144384] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11895020 kB) free.
[1553144984] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4201968 kB) free!
[1553145104] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11870552 kB) free.
[1553145221] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3156384 kB) free!
[1553145341] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3135152 kB) free!
[1553145461] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.2% (3118124 kB) free!
[1553145523] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (14929912 kB) free.
[1553145704] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3553292 kB) free!
[1553145824] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11592932 kB) free.
[1553146087] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (158692 kB) free!
[1553146207] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (150896 kB) free!
[1553146327] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (134900 kB) free!
[1553146421] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 281 processes
[1553146541] SERVICE ALERT: sig5-vm2;CPU_Procs;OK;HARD;1;CPU OK: 276 processes
[1553146646] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (159044 kB) free!
[1553146766] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (167912 kB) free!
[1553146781] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (156516 kB) free!
[1553146886] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (161704 kB) free!
[1553146901] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (139676 kB) free!
[1553147021] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (131128 kB) free!
[1553147024] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3657096 kB) free!
[1553147144] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11833232 kB) free.
[1553147744] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3002312 kB) free!
[1553147864] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11233636 kB) free.
[1553149064] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3689860 kB) free!
[1553149184] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11624712 kB) free.
[1553149784] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2974456 kB) free!
[1553149904] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11033892 kB) free.
[1553151104] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3440788 kB) free!
[1553151224] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11448992 kB) free.
[1553151284] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1553152424] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3819960 kB) free!
[1553152484] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553152544] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11672816 kB) free.
[1553153084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553153144] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3010148 kB) free!
[1553153204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553153265] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1553153265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11030816 kB) free.
[1553153804] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553153924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553154045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553154461] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.8% (14763968 kB) free.
[1553154464] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3849460 kB) free!
[1553154584] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11669664 kB) free.
[1553154686] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.3% (7009960 kB) free.
[1553154821] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 95.8% (7047328 kB) free.
[1553155184] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3063268 kB) free!
[1553155245] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553155304] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11788072 kB) free.
[1553155845] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553155904] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4794124 kB) free!
[1553155965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553156024] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11784624 kB) free.
[1553156085] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553156624] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6228712 kB) free!
[1553156685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553156744] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4913504 kB) free!
[1553156805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553156864] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2933788 kB) free!
[1553156926] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553157465] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5837932 kB) free!
[1553157525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553157645] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553157765] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1553157765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553158365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553158965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553159085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553159205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553159265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10736192 kB) free.
[1553160465] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4734716 kB) free!
[1553160585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11773632 kB) free.
[1553161185] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3830792 kB) free!
[1553161305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11778240 kB) free.
[1553161905] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4868516 kB) free!
[1553162025] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2877008 kB) free!
[1553162146] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11135728 kB) free.
[1553162805] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553163345] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.7% (5748924 kB) free!
[1553163405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553163465] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (4424560 kB) free!
[1553163525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553163585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11769120 kB) free.
[1553163645] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553164185] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.7% (5759976 kB) free!
[1553164246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553164305] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4469396 kB) free!
[1553164366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553164425] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11775196 kB) free.
[1553164486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553165025] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3616412 kB) free!
[1553165086] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553165145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11528260 kB) free.
[1553165686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553165745] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3425452 kB) free!
[1553165806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553165865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11579516 kB) free.
[1553166406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553166466] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2941840 kB) free!
[1553166526] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553166586] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10946500 kB) free.
[1553167127] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553167247] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553167367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553167786] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3460792 kB) free!
[1553167906] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11744008 kB) free.
[1553168506] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3161944 kB) free!
[1553168567] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553168626] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11672396 kB) free.
[1553169167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553169226] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3090784 kB) free!
[1553169287] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553169346] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11738300 kB) free.
[1553169887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553169946] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5903268 kB) free!
[1553170008] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553170066] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4605256 kB) free!
[1553170128] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553170187] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11747360 kB) free.
[1553170786] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6721396 kB) free!
[1553170906] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.0% (5404772 kB) free!
[1553171026] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3768204 kB) free!
[1553171028] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45986 MB (20% inode=99%):
[1553171148] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57543 MB (25% inode=99%):
[1553171626] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5373600 kB) free!
[1553172528] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553173128] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553173248] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553173368] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553174568] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553174626] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11719480 kB) free.
[1553175168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553175226] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5432296 kB) free!
[1553175288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553175346] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (3888440 kB) free!
[1553175408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553175466] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11712392 kB) free.
[1553176066] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6339208 kB) free!
[1553176186] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.2% (5031576 kB) free!
[1553176306] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1553176306] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3202756 kB) free!
[1553176608] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553176906] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7231764 kB) free!
[1553177208] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553177329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553177449] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553178049] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553178106] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11307000 kB) free.
[1553178649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553178769] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553178889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553179306] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4665840 kB) free!
[1553179426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11712220 kB) free.
[1553179489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553180026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7020884 kB) free!
[1553180089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553180147] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.8% (5815968 kB) free!
[1553180210] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553180267] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4596996 kB) free!
[1553180330] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553180867] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.3% (8523272 kB) free!
[1553181467] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11714004 kB) free.
[1553181530] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553182067] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5474072 kB) free!
[1553182130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553182187] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (3996004 kB) free!
[1553182250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553182307] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698604 kB) free.
[1553182370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553182907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6826356 kB) free!
[1553182970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553183027] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5555220 kB) free!
[1553183090] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553183147] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4157076 kB) free!
[1553183211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553183747] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.6% (8157216 kB) free!
[1553184347] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11695464 kB) free.
[1553184411] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553184948] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4965272 kB) free!
[1553185011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553185068] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (3102152 kB) free!
[1553185131] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553185188] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11703260 kB) free.
[1553185731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553185788] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5816840 kB) free!
[1553185851] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553185908] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (4528308 kB) free!
[1553185971] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553186028] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698036 kB) free.
[1553186486] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 11.2% (825236 kB) free!
[1553186606] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 11.2% (825244 kB) free!
[1553186621] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (170280 kB) free!
[1553186628] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6707956 kB) free!
[1553186726] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 11.2% (824228 kB) free!
[1553186741] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (170568 kB) free!
[1553186748] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.0% (5418708 kB) free!
[1553186861] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (170180 kB) free!
[1553186868] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3876464 kB) free!
[1553187171] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553187468] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7551676 kB) free!
[1553187771] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553187891] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553188068] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11695660 kB) free.
[1553188491] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553188611] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553188669] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4583700 kB) free!
[1553188731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553188789] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11702480 kB) free.
[1553189389] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6671876 kB) free!
[1553189509] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5361180 kB) free!
[1553189629] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3764180 kB) free!
[1553189932] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553190229] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7530040 kB) free!
[1553190532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553190652] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553190829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698792 kB) free.
[1553191252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553191372] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553191429] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4554572 kB) free!
[1553191492] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553191549] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11690320 kB) free.
[1553192149] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6630920 kB) free!
[1553192269] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.8% (5328704 kB) free!
[1553192389] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3660640 kB) free!
[1553192692] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553192990] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7390108 kB) free!