[1588098798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14113236 kB) free!
[1588098918] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36251156 kB) free.
[1588099518] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15055656 kB) free!
[1588099638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36345344 kB) free.
[1588100239] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16433772 kB) free!
[1588100359] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12946252 kB) free!
[1588100479] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33917928 kB) free.
[1588100553] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1588100666] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178272 kB) free!
[1588100673] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1588100726] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185168 kB) free!
[1588100786] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (196232 kB) free!
[1588100846] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (184648 kB) free!
[1588101079] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13543812 kB) free!
[1588101132] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (150876 kB) free!
[1588101146] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 72.8% (17841688 kB) free.
[1588101199] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35097448 kB) free.
[1588101252] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1161236 kB) free!
[1588101372] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1160892 kB) free!
[1588101799] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14286960 kB) free!
[1588101919] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35597288 kB) free.
[1588101973] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (2649332 kB) free!
[1588102519] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14478932 kB) free!
[1588102639] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35495268 kB) free.
[1588103239] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14242852 kB) free!
[1588103359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35074436 kB) free.
[1588103814] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9138352 kB) free!
[1588103934] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393916 kB) free!
[1588103959] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13889760 kB) free!
[1588104054] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395548 kB) free!
[1588104079] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36089900 kB) free.
[1588104254] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (218960 kB) free!
[1588104374] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217164 kB) free!
[1588104494] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215604 kB) free!
[1588104679] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14007116 kB) free!
[1588104799] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35547928 kB) free.
[1588105400] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14222556 kB) free!
[1588105520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34072960 kB) free.
[1588106294] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.1% (5973616 kB) free!
[1588106649] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4939824 kB) free!
[1588106720] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16858148 kB) free!
[1588106840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36121356 kB) free.
[1588107249] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 51.6% (25408568 kB) free.
[1588108040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17715120 kB) free!
[1588108160] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1588108160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35795504 kB) free.
[1588108265] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.7% (43214316 kB) free.
[1588108450] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 29.2% (7165456 kB) free.
[1588109360] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17928748 kB) free!
[1588109480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35235940 kB) free.
[1588109519] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3903972 kB) free!
[1588109529] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2102) < 172.105.59.173 [29/Apr/2020:03:00:23 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1588109649] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(32) < 203.212.70.170 [29/Apr/2020:03:03:22 +0530] TCP 502 0 0 180.001 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.001, 0.000"
[1588109650] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1658120 kB) free!
[1588109769] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588109770] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (606440 kB) free!
[1588109890] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185732 kB) free!
[1588109894] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.2% (44450896 kB) free.
[1588109956] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(213) < 103.16.101.72 [29/Apr/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1588109962] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(326) < 2020/04/29 03:01:22 [error] 22135#0: *86186670 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588110076] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588110082] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588110119] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1559092 kB) free!
[1588110680] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13601384 kB) free!
[1588110800] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35278328 kB) free.
[1588111573] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4034812 kB) free.
[1588112001] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15000556 kB) free!
[1588112121] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35372836 kB) free.
[1588112890] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.0% (10059644 kB) free.
[1588113321] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14819292 kB) free!
[1588113441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34528936 kB) free.
[1588113490] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (237216 kB) free!
[1588113610] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186824 kB) free!
[1588113730] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1588113730] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (242236 kB) free!
[1588114930] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 37.8% (9274864 kB) free.
[1588116441] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22752632 kB) free!
[1588116561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37245684 kB) free.
[1588116730] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195924 kB) free!
[1588116850] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175640 kB) free!
[1588116970] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (155828 kB) free!
[1588117573] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1588117705] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 78.5% (19241284 kB) free.
[1588117791] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 57.1% (14089056 kB) free.
[1588117920] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 54.2% (13289692 kB) free.
[1588121992] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (1651600 kB) free!
[1588122054] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 44.0% (57961652 kB) free.
[1588122112] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (1393336 kB) free!
[1588122232] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1348992 kB) free!
[1588123761] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22011452 kB) free!
[1588123881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35711620 kB) free.
[1588127170] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588131369] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(597) < 103.255.146.164 [29/Apr/2020:09:01:42 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588131490] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588131587] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.04, 8.30, 6.50
[1588131708] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.63, 7.79, 6.54
[1588131828] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.36, 8.06, 6.79
[1588132920] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4654540 kB) free!
[1588133040] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4543124 kB) free!
[1588133160] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4426748 kB) free!
[1588133773] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2711536 kB) free!
[1588133893] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3589356 kB) free.
[1588134820] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4794424 kB) free!
[1588134940] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.1% (4670716 kB) free!
[1588135060] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2363652 kB) free!
[1588135363] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1588135463] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(2) < Channel error on connection <0.5837.1614> (10.147.212.158:47450 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588135473] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1588135483] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1588135560] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1969908 kB) free!
[1588135583] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588135593] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1588135603] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1588135705] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (208072 kB) free!
[1588135713] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1588135825] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (217524 kB) free!
[1588135881] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9327920 kB) free!
[1588135945] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200360 kB) free!
[1588136001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27861916 kB) free.
[1588136203] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1588136243] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1588136363] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1588136436] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (208668 kB) free!
[1588136482] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/29 10:29:26 [error] 22135#0: *86220501 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1588136483] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1588136603] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588136893] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3221368 kB) free!
[1588137013] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3183444 kB) free!
[1588137133] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (3154804 kB) free!
[1588137684] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1588137801] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (875856 kB) free!
[1588137804] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1588137921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28971204 kB) free.
[1588138521] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (390156 kB) free!
[1588138641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29452916 kB) free.
[1588138660] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 60.2% (14758744 kB) free.
[1588138713] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1588138833] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1588138953] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1588139183] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.32007.1599> (10.147.212.147:45576 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588139241] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20495040 kB) free!
[1588139303] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588139361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.0% (25025928 kB) free!
[1588139481] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20435636 kB) free!
[1588140081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34948768 kB) free.
[1588140804] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1588140884] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (3732928 kB) free!
[1588140924] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1588141004] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3503448 kB) free!
[1588141124] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (3317560 kB) free!
[1588141953] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1588142324] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.1% (15958840 kB) free.
[1588143081] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19331484 kB) free!
[1588143202] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40718004 kB) free.
[1588143801] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17396440 kB) free!
[1588143921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38191688 kB) free.
[1588144428] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.29, 4.43, 4.74
[1588145121] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16975500 kB) free!
[1588145241] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38428720 kB) free.
[1588145841] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15017144 kB) free!
[1588145961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36523276 kB) free.
[1588147161] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15590080 kB) free!
[1588147281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37517704 kB) free.
[1588147881] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13988780 kB) free!
[1588147933] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1591952 kB) free!
[1588148001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36106784 kB) free.
[1588148303] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.2915.1505> (10.147.212.158:51198 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588148423] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588148533] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.0% (2270412 kB) free!
[1588148601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12400160 kB) free!
[1588148721] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34263676 kB) free.
[1588148860] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.2% (4713360 kB) free!
[1588148980] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.9% (4623012 kB) free!
[1588149100] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.1% (4432904 kB) free!
[1588149133] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1619524 kB) free!
[1588149733] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1588149733] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (1975680 kB) free!
[1588149921] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13691812 kB) free!
[1588150041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35369444 kB) free.
[1588150641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12414048 kB) free!
[1588150761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34798952 kB) free.
[1588151361] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10539732 kB) free!
[1588151481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33647340 kB) free.
[1588151500] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2144468 kB) free!
[1588151533] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1266396 kB) free!
[1588151628] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.57, 6.00, 5.47
[1588151748] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.89, 5.57, 5.38
[1588151868] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.53, 5.66, 5.44
[1588152081] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8921724 kB) free!
[1588152134] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.3% (2002700 kB) free!
[1588152201] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31953252 kB) free.
[1588152801] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6920984 kB) free!
[1588152921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31060380 kB) free.
[1588153089] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 185.255.8.58 [29/Apr/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1588153209] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588153521] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7120736 kB) free!
[1588153641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29578292 kB) free.
[1588153668] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.53, 4.05, 4.50
[1588153809] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(311) < 103.255.146.164 [29/Apr/2020:15:14:28 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588153929] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588154841] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8693300 kB) free!
[1588154961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31625564 kB) free.
[1588155468] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.81, 8.54, 6.52
[1588155562] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3948816 kB) free!
[1588155588] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.50, 7.70, 6.45
[1588155682] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30681964 kB) free.
[1588155708] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1588155708] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.20, 6.03, 5.99
[1588155729] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(230) < 103.255.146.164 [29/Apr/2020:15:44:02 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588155849] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588156282] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6871348 kB) free!
[1588156308] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.95, 4.15, 4.84
[1588156402] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30225704 kB) free.
[1588156934] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (662212 kB) free!
[1588157002] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4228720 kB) free!
[1588157122] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29686660 kB) free.
[1588157722] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3042880 kB) free!
[1588157842] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28461972 kB) free.
[1588158442] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (1959848 kB) free!
[1588158562] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27390996 kB) free.
[1588158665] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9711316 kB) free!
[1588158734] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (2114908 kB) free!
[1588158785] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9578704 kB) free!
[1588158905] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.1% (9430872 kB) free!
[1588159162] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (400020 kB) free!
[1588159282] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37306476 kB) free.
[1588159908] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.21, 6.94, 5.55
[1588159934] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1184400 kB) free!
[1588160028] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.73, 6.93, 5.72
[1588160148] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.61, 7.42, 6.03
[1588160482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10391152 kB) free!
[1588160535] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3001256 kB) free!
[1588160602] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39826256 kB) free.
[1588161135] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3789872 kB) free.
[1588161203] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10696604 kB) free!
[1588161305] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4729212 kB) free!
[1588161323] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5812628 kB) free!
[1588161443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37986148 kB) free.
[1588161735] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3241940 kB) free!
[1588161855] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2839688 kB) free!
[1588161975] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (2469608 kB) free!
[1588162043] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6772064 kB) free!
[1588162124] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 10.9% (2443716 kB) free!
[1588162163] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2102556 kB) free!
[1588162244] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (1963124 kB) free!
[1588162283] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33638968 kB) free.
[1588162364] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1559684 kB) free!
[1588162883] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4269844 kB) free!
[1588162953] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1588163003] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34538208 kB) free.
[1588163073] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1588163175] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3654268 kB) free.
[1588163193] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1588163564] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.1% (15954716 kB) free.
[1588163603] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6844004 kB) free!
[1588163649] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9093764 kB) free!
[1588163723] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2070384 kB) free!
[1588163769] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (8883156 kB) free!
[1588163843] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31882436 kB) free.
[1588163889] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (8750000 kB) free!
[1588164375] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3036792 kB) free!
[1588164443] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5057188 kB) free!
[1588164495] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2884028 kB) free!
[1588164563] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32605060 kB) free.
[1588164615] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (2803780 kB) free!
[1588165163] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6635664 kB) free!
[1588165283] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32117544 kB) free.
[1588165883] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5931948 kB) free!
[1588166004] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9414852 kB) free!
[1588166124] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28634684 kB) free.
[1588166415] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1527220 kB) free!
[1588166724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3041404 kB) free!
[1588166748] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.03, 4.63, 4.89
[1588166844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28922404 kB) free.
[1588167444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1298876 kB) free!
[1588167564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27677956 kB) free.
[1588167948] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.75, 5.21, 5.10
[1588168068] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.36, 5.37, 5.15
[1588168164] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1417784 kB) free!
[1588168188] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.96, 5.17, 5.11
[1588168216] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (1984304 kB) free!
[1588168284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28454528 kB) free.
[1588168884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1511452 kB) free!
[1588169004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29656464 kB) free.
[1588169289] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4737276 kB) free!
[1588169604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (396140 kB) free!
[1588169724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (392940 kB) free!
[1588169844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38441480 kB) free.
[1588170016] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1624816 kB) free!
[1588170444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5724020 kB) free!
[1588170564] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1181860 kB) free!
[1588170588] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.40, 3.53, 4.72
[1588170616] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (1730892 kB) free!
[1588170685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34566208 kB) free.
[1588171284] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2488800 kB) free!
[1588171404] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15307040 kB) free!
[1588171524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31689364 kB) free.
[1588172124] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (1963132 kB) free!
[1588172244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37789632 kB) free.
[1588172844] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6686112 kB) free!
[1588172964] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2439576 kB) free!
[1588173084] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35616248 kB) free.
[1588173684] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5212356 kB) free!
[1588173804] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35784228 kB) free.
[1588174404] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5672752 kB) free!
[1588174524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35536984 kB) free.
[1588175124] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6085952 kB) free!
[1588175244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37060908 kB) free.
[1588175844] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8539720 kB) free!
[1588175965] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4886188 kB) free!
[1588176016] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3587160 kB) free.
[1588176085] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35098348 kB) free.
[1588176685] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6935500 kB) free!
[1588176805] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37399240 kB) free.
[1588177405] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10727408 kB) free!
[1588177525] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37793744 kB) free.
[1588178125] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15909052 kB) free!
[1588178245] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11843452 kB) free!
[1588178365] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35207624 kB) free.
[1588178965] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12397344 kB) free!
[1588179016] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2984016 kB) free!
[1588179085] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35502332 kB) free.
[1588179136] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.4% (2829368 kB) free!
[1588179256] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (2685260 kB) free!
[1588179685] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12314232 kB) free!
[1588179805] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35237144 kB) free.
[1588180188] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1588180405] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12676984 kB) free!
[1588180525] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34968488 kB) free.
[1588181725] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16086896 kB) free!
[1588181845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37240032 kB) free.
[1588182445] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16172312 kB) free!
[1588182565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37232312 kB) free.
[1588183165] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15936652 kB) free!
[1588183285] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37168624 kB) free.
[1588183525] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1588183593] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1588183645] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1588183885] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15357100 kB) free!
[1588184005] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36302620 kB) free.