[1578767661] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (1056928 kB) free!
[1578767759] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 49.5% (24379484 kB) free.
[1578767781] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404120 kB) free!
[1578767821] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37737716 kB) free.
[1578767901] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378552 kB) free!
[1578768143] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 223283 MB (20% inode=99%):
[1578768263] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 223283 MB (20% inode=99%):
[1578768383] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 223283 MB (20% inode=99%):
[1578768421] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16550740 kB) free!
[1578768541] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12909572 kB) free!
[1578768661] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40916284 kB) free.
[1578769261] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20522564 kB) free!
[1578769378] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164072 kB) free!
[1578769381] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16905792 kB) free!
[1578769438] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177516 kB) free!
[1578769499] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (186476 kB) free!
[1578769501] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13335924 kB) free!
[1578769559] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (195060 kB) free!
[1578769859] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 56.9% (13943780 kB) free.
[1578770701] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35586772 kB) free.
[1578771301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16313228 kB) free!
[1578771421] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12868440 kB) free!
[1578771541] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38984260 kB) free.
[1578772141] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19029816 kB) free!
[1578772261] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15300160 kB) free!
[1578772382] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41079576 kB) free.
[1578772981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19283640 kB) free!
[1578773041] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/01/12 01:33:41 [error] 21675#0: *24393901 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.142:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578773101] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15262500 kB) free!
[1578773162] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578773221] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41110348 kB) free.
[1578773821] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19049672 kB) free!
[1578773941] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15253356 kB) free!
[1578774061] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37799456 kB) free.
[1578774661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19826956 kB) free!
[1578774781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15768252 kB) free!
[1578774901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41053672 kB) free.
[1578775501] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19795968 kB) free!
[1578775561] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(71) < 2020/01/12 02:15:51 [error] 21681#0: *24399276 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:35/0, bytes from/to upstream:0/35
[1578775621] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16075352 kB) free!
[1578775682] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578775741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41030420 kB) free.
[1578775925] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.7% (21996320 kB) free.
[1578776302] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 86.9% (114487348 kB) free.
[1578776327] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 64.3% (31667696 kB) free.
[1578776341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19404776 kB) free!
[1578776461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15719692 kB) free!
[1578776582] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40566316 kB) free.
[1578777181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19158208 kB) free!
[1578777301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15578588 kB) free!
[1578777377] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.8% (43744032 kB) free.
[1578777421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39752284 kB) free.
[1578777481] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/12 02:39:51 [error] 21683#0: *24402441 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:47/0, bytes from/to upstream:0/47
[1578777601] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578777806] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.9% (10507996 kB) free.
[1578778021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18878148 kB) free!
[1578778141] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15376548 kB) free!
[1578778261] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39456136 kB) free.
[1578778861] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19420924 kB) free!
[1578778981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15845512 kB) free!
[1578779101] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38779668 kB) free.
[1578779701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13359836 kB) free!
[1578779821] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9464576 kB) free!
[1578779833] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1578779941] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39628084 kB) free.
[1578779953] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1578780541] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13597780 kB) free!
[1578780661] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9827712 kB) free!
[1578780782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38986860 kB) free.
[1578781381] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14529508 kB) free!
[1578781501] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40842236 kB) free.
[1578782101] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21446756 kB) free!
[1578782221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17687896 kB) free!
[1578782341] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41935720 kB) free.
[1578782941] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21824128 kB) free!
[1578783061] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18222216 kB) free!
[1578783181] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40036612 kB) free.
[1578783781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21678024 kB) free!
[1578783901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43139556 kB) free.
[1578784501] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24802688 kB) free!
[1578784621] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21325544 kB) free!
[1578784741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41471684 kB) free.
[1578785341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23328960 kB) free!
[1578785461] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578785461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43342084 kB) free.
[1578785606] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (150884 kB) free!
[1578785727] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195472 kB) free!
[1578785847] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202432 kB) free!
[1578786062] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25144280 kB) free!
[1578786182] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (21805124 kB) free!
[1578786302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41865028 kB) free.
[1578786838] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.8% (22412328 kB) free.
[1578786902] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22684604 kB) free!
[1578786998] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.6% (20973076 kB) free.
[1578787022] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41809988 kB) free.
[1578787378] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1578787498] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1578787622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23266992 kB) free!
[1578787742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42955656 kB) free.
[1578788342] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25101420 kB) free!
[1578788462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43307628 kB) free.
[1578789062] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25284180 kB) free!
[1578789182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43303756 kB) free.
[1578789782] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25896800 kB) free!
[1578789902] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.8% (22149744 kB) free!
[1578790022] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42030568 kB) free.
[1578790273] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.5% (40656780 kB) free.
[1578790308] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.7% (22712648 kB) free.
[1578790351] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.0% (14145572 kB) free.
[1578790583] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 249384 MB (23% inode=99%):
[1578790622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22236372 kB) free!
[1578790743] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42803552 kB) free.
[1578790920] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185404 kB) free!
[1578791040] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186904 kB) free!
[1578791160] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190672 kB) free!
[1578791342] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23367988 kB) free!
[1578791462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42993884 kB) free.
[1578792062] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24173568 kB) free!
[1578792182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42928608 kB) free.
[1578792782] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24087028 kB) free!
[1578792902] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42600120 kB) free.
[1578792994] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1578792994] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.3% (22623576 kB) free.
[1578793502] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25052560 kB) free!
[1578793622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.1% (21203356 kB) free!
[1578793742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41707432 kB) free.
[1578794342] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22796996 kB) free!
[1578794462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42250196 kB) free.
[1578795062] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24262640 kB) free!
[1578795182] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.7% (20658840 kB) free!
[1578795302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40842516 kB) free.
[1578795902] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22407092 kB) free!
[1578796023] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41969116 kB) free.
[1578796622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21334384 kB) free!
[1578796742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41658892 kB) free.
[1578797038] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4734572 kB) free!
[1578797158] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4701988 kB) free!
[1578797278] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4657412 kB) free!
[1578797342] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21336680 kB) free!
[1578797462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41177372 kB) free.
[1578798062] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22657340 kB) free!
[1578798182] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18729596 kB) free!
[1578798302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39235008 kB) free.
[1578798902] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19151460 kB) free!
[1578799022] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40105200 kB) free.
[1578799622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20050556 kB) free!
[1578799742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40208764 kB) free.
[1578800342] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18799520 kB) free!
[1578800462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38338480 kB) free.
[1578800878] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2129124 kB) free!
[1578801063] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17906012 kB) free!
[1578801183] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37206012 kB) free.
[1578801783] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16252992 kB) free!
[1578801903] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38214808 kB) free.
[1578801998] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1578802503] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16962680 kB) free!
[1578802623] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28422752 kB) free.
[1578802789] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 728.63, 258.09, 96.98
[1578802909] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 769.12, 432.08, 180.99
[1578802988] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1578803030] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 778.19, 545.29, 252.76
[1578803108] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1578803223] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18250612 kB) free!
[1578803228] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1578803343] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14322368 kB) free!
[1578803463] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36122496 kB) free.
[1578803800] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (459224 kB) free!
[1578803920] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194720 kB) free!
[1578803954] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1578804001] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(144) < 2020/01/12 10:05:53 [error] 21672#0: *24456990 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:44/0, bytes from/to upstream:0/44
[1578804040] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164800 kB) free!
[1578804063] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12158504 kB) free!
[1578804066] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(78) < 94.237.68.73 [12/Jan/2020:10:03:58 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1578804074] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1578804122] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578804183] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8370540 kB) free!
[1578804183] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2020/01/12 10:03:49 [error] 18355#0: *37610810 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:1195/4567, bytes from/to upstream:4567/1195
[1578804186] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578804303] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34286616 kB) free.
[1578804304] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578804712] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22859044 kB) free.
[1578804749] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.2% (22343572 kB) free.
[1578804903] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12838600 kB) free!
[1578805023] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34860220 kB) free.
[1578805029] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1578805624] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13132840 kB) free!
[1578805744] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8663060 kB) free!
[1578805864] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32987132 kB) free.
[1578805877] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237272 kB) free!
[1578805919] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186660 kB) free!
[1578805956] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195512 kB) free!
[1578805997] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239344 kB) free!
[1578806039] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192964 kB) free!
[1578806076] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191908 kB) free!
[1578806117] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237852 kB) free!
[1578806159] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1578806159] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188376 kB) free!
[1578806196] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195756 kB) free!
[1578806229] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578806349] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578806464] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9930068 kB) free!
[1578806469] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578806584] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.8% (26143068 kB) free!
[1578806630] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.90, 13.08, 33.74
[1578806704] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27016228 kB) free.
[1578807069] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1578807069] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1578807230] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 14.25, 11.83, 23.10
[1578807304] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (390460 kB) free!
[1578807424] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35853200 kB) free.
[1578807607] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21927604 kB) free.
[1578807673] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 75.4% (37168812 kB) free.
[1578807705] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1578807705] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.5% (22432396 kB) free.
[1578808024] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9359744 kB) free!
[1578808144] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5104896 kB) free!
[1578808218] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195952 kB) free!
[1578808264] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35128044 kB) free.
[1578808268] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.8% (2414156 kB) free!
[1578808278] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (225424 kB) free!
[1578808309] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172024 kB) free!
[1578808338] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192448 kB) free!
[1578808398] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (372816 kB) free!
[1578808429] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182324 kB) free!
[1578808458] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181148 kB) free!
[1578808518] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214972 kB) free!
[1578808549] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172668 kB) free!
[1578808703] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1578808824] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1578808864] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13093272 kB) free!
[1578808869] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578808925] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184360 kB) free!
[1578808944] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1578808984] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8863064 kB) free!
[1578808990] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1578809045] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199288 kB) free!
[1578809083] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.7% (21987292 kB) free.
[1578809104] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34323612 kB) free.
[1578809110] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1578809119] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.1% (22732912 kB) free.
[1578809161] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1578809161] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39972284 kB) free.
[1578809236] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.4% (19709576 kB) free.
[1578809351] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 91.0% (22295528 kB) free.
[1578809544] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1578809704] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11546780 kB) free!
[1578809825] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7230200 kB) free!
[1578809839] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200916 kB) free!
[1578809945] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33382752 kB) free.
[1578809959] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196188 kB) free!
[1578810079] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178828 kB) free!
[1578810121] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2020/01/12 11:43:21 [error] 21675#0: *24468811 connect() failed (111: Connection refused) while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578810186] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 180.179.218.150 [12/Jan/2020:11:43:19 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1578810241] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578810306] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578810545] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9897884 kB) free!
[1578810665] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5652408 kB) free!
[1578810785] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32517420 kB) free.
[1578810929] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (2401844 kB) free!
[1578811049] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (2355088 kB) free!
[1578811095] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.7% (22618672 kB) free.
[1578811143] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.8% (21032924 kB) free.
[1578811385] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8247048 kB) free!
[1578811505] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4107208 kB) free!
[1578811602] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578811625] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30547880 kB) free.
[1578811723] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578811747] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181188 kB) free!
[1578811843] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578811867] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187908 kB) free!
[1578811987] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165292 kB) free!
[1578812110] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578812225] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6535316 kB) free!
[1578812231] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1578812345] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2288844 kB) free!
[1578812465] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28717624 kB) free.
[1578812767] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8299780 kB) free!
[1578812799] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200876 kB) free!
[1578812887] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236276 kB) free!
[1578812898] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192528 kB) free!
[1578812919] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201804 kB) free!
[1578813007] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241052 kB) free!
[1578813018] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184172 kB) free!
[1578813039] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199056 kB) free!
[1578813065] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4408380 kB) free!
[1578813138] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193044 kB) free!
[1578813185] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401500 kB) free!
[1578813305] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28272764 kB) free.
[1578813905] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3782600 kB) free!
[1578814026] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374192 kB) free!
[1578814146] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28511552 kB) free.
[1578814441] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/12 12:54:02 [error] 21682#0: *24477470 recv() failed (104: Connection reset by peer) while proxying connection, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "10.147.212.154:13001", bytes from/to client:41/0, bytes from/to upstream:0/41
[1578814561] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578814631] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578814734] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22661192 kB) free.
[1578814746] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3829604 kB) free!
[1578814751] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1578814776] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40700888 kB) free.
[1578814843] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578814857] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 44.1% (58118488 kB) free.
[1578814921] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.7% (22974288 kB) free.
[1578814973] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.5% (19494116 kB) free.
[1578815443] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578815563] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578815574] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179336 kB) free!
[1578815683] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578815694] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182964 kB) free!
[1578815814] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182816 kB) free!
[1578816540] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (1964532 kB) free!
[1578816660] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (1939352 kB) free!
[1578816780] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1909748 kB) free!
[1578817281] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24136236 kB) free!
[1578817402] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.8% (51146800 kB) free.
[1578818951] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578819801] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21431720 kB) free!
[1578819921] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17360752 kB) free!
[1578820041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29808688 kB) free.
[1578820151] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1578820271] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1578820641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5946476 kB) free!
[1578820761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33993716 kB) free.
[1578821361] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9575556 kB) free!
[1578821481] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5235436 kB) free!
[1578821601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33313736 kB) free.
[1578822126] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193432 kB) free!
[1578822201] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8115264 kB) free!
[1578822246] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194172 kB) free!
[1578822321] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3818104 kB) free!
[1578822366] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200972 kB) free!
[1578822441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32527560 kB) free.
[1578823042] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7516428 kB) free!
[1578823162] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3071956 kB) free!
[1578823282] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32637440 kB) free.
[1578823355] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1578823475] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 330 processes
[1578823596] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 331 processes
[1578823882] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7686908 kB) free!
[1578824002] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3443664 kB) free!
[1578824122] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404348 kB) free!
[1578824196] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1578824976] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1578825287] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.7% (4081436 kB) free!
[1578825322] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16620872 kB) free!
[1578825361] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/12 15:57:44 [error] 21680#0: *24500001 recv() failed (104: Connection reset by peer) while proxying connection, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.142:13001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1578825407] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.0% (3932348 kB) free!
[1578825481] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578825527] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.5% (3798240 kB) free!
[1578825922] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33100740 kB) free.
[1578826380] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1578826522] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7174564 kB) free!
[1578826577] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9477696 kB) free!
[1578826642] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2798464 kB) free!
[1578826698] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9321844 kB) free!
[1578826762] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33239060 kB) free.
[1578826818] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (9166088 kB) free!
[1578827214] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578827327] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1676496 kB) free!
[1578827362] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5387416 kB) free!
[1578827482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1063812 kB) free!
[1578827602] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34298372 kB) free.
[1578828203] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9210796 kB) free!
[1578828323] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4816180 kB) free!
[1578828443] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1578828443] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (818128 kB) free!
[1578829080] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.8% (22009000 kB) free.
[1578829122] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.1% (22728740 kB) free.
[1578829163] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.3% (57035308 kB) free.
[1578829205] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (45297564 kB) free.
[1578829232] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.7% (22464672 kB) free.
[1578829269] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.8% (23001032 kB) free.
[1578830479] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188360 kB) free!
[1578830599] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 44.9% (11004404 kB) free.
[1578830967] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18967256 kB) free!
[1578831087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.0% (51414004 kB) free.
[1578831199] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175428 kB) free!
[1578831319] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169964 kB) free!
[1578831440] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174140 kB) free!
[1578831687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23896636 kB) free!
[1578831807] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19506516 kB) free!
[1578831898] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1578831927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.4% (50667516 kB) free.
[1578832018] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1578832271] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1578832391] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578832511] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1578832527] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23266800 kB) free!
[1578832647] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18887380 kB) free!
[1578832767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.9% (49910664 kB) free.
[1578833368] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23425600 kB) free!
[1578833488] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.5% (19169356 kB) free!
[1578833608] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.4% (49247488 kB) free.
[1578833976] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9286608 kB) free!
[1578834096] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8758688 kB) free!
[1578834208] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23953084 kB) free!
[1578834216] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (8755816 kB) free!
[1578834328] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19715612 kB) free!
[1578834448] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15821584 kB) free!
[1578835648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34779912 kB) free.
[1578836848] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20914752 kB) free!
[1578836968] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16843380 kB) free!
[1578837088] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46750744 kB) free.
[1578837688] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20382412 kB) free!
[1578837808] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16287424 kB) free!
[1578837928] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45741384 kB) free.
[1578838529] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19683920 kB) free!
[1578838649] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15696732 kB) free!
[1578838769] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (44067788 kB) free.
[1578839369] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18495364 kB) free!
[1578839489] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.3% (46571156 kB) free.
[1578840089] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21339284 kB) free!
[1578840209] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (16967696 kB) free!
[1578840329] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.0% (46161772 kB) free.
[1578840929] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20449624 kB) free!
[1578841049] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16073288 kB) free!
[1578841169] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45702260 kB) free.
[1578841769] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19314424 kB) free!
[1578841889] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (15021464 kB) free!
[1578842009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45226908 kB) free.
[1578842609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17833668 kB) free!
[1578842730] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13457268 kB) free!
[1578842850] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44686828 kB) free.
[1578843450] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16035496 kB) free!
[1578843570] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11700012 kB) free!
[1578843690] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42736828 kB) free.
[1578843816] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4913592 kB) free!
[1578844290] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14456804 kB) free!
[1578844410] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10274736 kB) free!
[1578844530] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37079064 kB) free.
[1578845130] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11050284 kB) free!
[1578845250] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32570036 kB) free.
[1578845326] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4852968 kB) free!
[1578845446] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4806328 kB) free!
[1578845566] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (4760304 kB) free!
[1578845850] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6008124 kB) free!
[1578845970] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36530736 kB) free.
[1578846570] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11264912 kB) free!
[1578846690] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (7051980 kB) free!
[1578846810] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35570744 kB) free.
[1578847087] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.6% (4564360 kB) free!
[1578847207] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.5% (4542884 kB) free!
[1578847327] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.5% (4527212 kB) free!
[1578847410] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11360696 kB) free!
[1578847531] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7785976 kB) free!
[1578847651] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33679792 kB) free.
[1578848218] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1578848251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11619384 kB) free!
[1578848338] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1578848371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35951704 kB) free.
[1578848566] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2432024 kB) free!
[1578848938] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1578848971] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13852648 kB) free!
[1578849058] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1578849091] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10143300 kB) free!
[1578849211] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34075848 kB) free.
[1578849811] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12617824 kB) free!
[1578849931] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578849931] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35874160 kB) free.
[1578850313] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578850511] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1578850531] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15813660 kB) free!
[1578850651] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12147384 kB) free!
[1578850771] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35858148 kB) free.
[1578851111] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1578851371] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14632876 kB) free!
[1578851491] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11099908 kB) free!
[1578851611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33440460 kB) free.
[1578851708] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(840) < 103.255.146.164 [12/Jan/2020:23:19:29 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578851828] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(868) < 103.255.146.164 [12/Jan/2020:23:25:53 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578851881] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1708) < 2020/01/12 23:25:53 [error] 21683#0: *24558177 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578851948] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578852001] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578852212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12336260 kB) free!
[1578852311] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1578852332] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34654936 kB) free.
[1578852359] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (695416 kB) free!
[1578852479] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236196 kB) free!
[1578852599] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (253880 kB) free!
[1578852911] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1578852911] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1578852932] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13073488 kB) free!
[1578853052] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35916524 kB) free.
[1578853652] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15250316 kB) free!
[1578853772] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11749200 kB) free!