[1572633874] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30454396 kB) free.
[1572634474] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16816568 kB) free!
[1572634594] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33742556 kB) free.
[1572634673] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.0% (3200796 kB) free!
[1572635107] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192160 kB) free!
[1572635167] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 59.5% (14592880 kB) free.
[1572635195] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19200136 kB) free!
[1572635273] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185920 kB) free!
[1572635315] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16060876 kB) free!
[1572635435] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32824756 kB) free.
[1572636035] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16458120 kB) free!
[1572636155] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33105084 kB) free.
[1572636755] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17576964 kB) free!
[1572636875] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33697176 kB) free.
[1572637475] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18992312 kB) free!
[1572637595] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15925416 kB) free!
[1572637715] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32612008 kB) free.
[1572638315] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16388816 kB) free!
[1572638435] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32666292 kB) free.
[1572639035] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16450268 kB) free!
[1572639155] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32634476 kB) free.
[1572640355] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18301672 kB) free!
[1572640476] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33676300 kB) free.
[1572641075] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17077584 kB) free!
[1572641195] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32594324 kB) free.
[1572641506] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9865232 kB) free!
[1572641626] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (3347576 kB) free!
[1572641746] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242640 kB) free!
[1572642946] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 89.7% (44204668 kB) free.
[1572642995] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17637496 kB) free!
[1572643115] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33180968 kB) free.
[1572643447] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/02 02:45:50 [error] 24130#0: *1845398 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:55/17, bytes from/to upstream:17/55
[1572643567] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572644168] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/11/02 03:00:36 [error] 24130#0: *1847324 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572644288] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572644315] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18347784 kB) free!
[1572644435] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33661640 kB) free.
[1572644888] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(11) < 2019/11/02 03:16:32 [error] 24130#0: *1847889 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572645008] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2019/11/02 03:20:05 [error] 24130#0: *1848074 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572645128] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(23) < 2019/11/02 03:21:50 [error] 24130#0: *1848196 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572645451] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2759344 kB) free!
[1572645635] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17821080 kB) free!
[1572645755] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32994720 kB) free.
[1572646328] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572646955] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18061776 kB) free!
[1572647075] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33650352 kB) free.
[1572647528] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2019/11/02 04:00:51 [error] 24131#0: *1855144 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572647551] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1572647648] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572647671] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1572648275] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17265116 kB) free!
[1572648395] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32202908 kB) free.
[1572651995] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17587068 kB) free!
[1572652051] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2424392 kB) free!
[1572652115] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32595080 kB) free.
[1572653315] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17726280 kB) free!
[1572653435] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32787128 kB) free.
[1572654271] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1572654391] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1572654635] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17774972 kB) free!
[1572654756] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33569720 kB) free.
[1572655355] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19159708 kB) free!
[1572655475] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16103304 kB) free!
[1572655595] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33422068 kB) free.
[1572655752] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1572655872] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 374 processes
[1572655927] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.7% (21506028 kB) free.
[1572656049] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2019/11/02 06:19:18 [error] 24133#0: *1882531 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.142:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572656169] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572656188] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.6% (21237588 kB) free.
[1572656195] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18119848 kB) free!
[1572656251] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 75.9% (18613692 kB) free.
[1572656315] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14843120 kB) free!
[1572656435] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32217764 kB) free.
[1572656791] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1572656911] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572657035] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16247068 kB) free!
[1572657155] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32727508 kB) free.
[1572657755] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18216300 kB) free!
[1572657875] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32997404 kB) free.
[1572658475] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19467832 kB) free!
[1572658595] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16161648 kB) free!
[1572658715] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32532728 kB) free.
[1572659315] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17713124 kB) free!
[1572659391] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.9% (23180392 kB) free.
[1572659435] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32871624 kB) free.
[1572659440] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.3% (40560216 kB) free.
[1572660036] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18456800 kB) free!
[1572660156] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15620428 kB) free!
[1572660276] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31763804 kB) free.
[1572660876] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16636956 kB) free!
[1572660996] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32751652 kB) free.
[1572661596] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17436124 kB) free!
[1572661716] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32690736 kB) free.
[1572662316] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17905056 kB) free!
[1572662436] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (14966188 kB) free!
[1572662556] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31593480 kB) free.
[1572663156] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16298700 kB) free!
[1572663276] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32502236 kB) free.
[1572663876] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16841640 kB) free!
[1572663996] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32404824 kB) free.
[1572664546] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 304 processes
[1572664596] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16781328 kB) free!
[1572664666] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1572664717] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32283364 kB) free.
[1572665311] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1572665316] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16678816 kB) free!
[1572665431] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1572665436] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32172700 kB) free.
[1572666036] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15060908 kB) free!
[1572666156] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31927072 kB) free.
[1572666597] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4727104 kB) free!
[1572666717] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.8% (4642096 kB) free!
[1572666756] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11323340 kB) free!
[1572666837] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (4557632 kB) free!
[1572666876] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31312048 kB) free.
[1572667476] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (9006752 kB) free!
[1572667596] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30056124 kB) free.
[1572668037] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178484 kB) free!
[1572668196] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9834336 kB) free!
[1572668316] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29249432 kB) free.
[1572668916] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8175036 kB) free!
[1572669036] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28327760 kB) free.
[1572669451] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3997776 kB) free!
[1572669571] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.2% (3737212 kB) free!
[1572669636] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5740328 kB) free!
[1572669691] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (3496356 kB) free!
[1572669757] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27638172 kB) free.
[1572670291] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2364160 kB) free!
[1572670356] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6045568 kB) free!
[1572670476] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27421924 kB) free.
[1572671076] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4561380 kB) free!
[1572671196] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 20.0% (26365008 kB) free!
[1572671316] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21585152 kB) free!
[1572671333] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22829256 kB) free.
[1572671383] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.8% (22490224 kB) free.
[1572672516] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3367576 kB) free!
[1572672747] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 64.12, 55.45, 36.84
[1572672867] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 31.74, 45.81, 35.55
[1572672987] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 32.05, 41.01, 35.01
[1572673238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.4% (54537820 kB) free.
[1572674187] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.23, 77.98, 52.17
[1572674787] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 23.13, 31.12, 38.85
[1572675533] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195412 kB) free!
[1572675654] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188912 kB) free!
[1572675774] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1572675774] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168208 kB) free!
[1572675988] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1572675988] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.93, 38.46, 46.18
[1572676242] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12530740 kB) free!
[1572676362] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8678408 kB) free!
[1572676482] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31463672 kB) free.
[1572676546] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9235828 kB) free!
[1572676618] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 78.5% (19380108 kB) free.
[1572676666] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (8897484 kB) free!
[1572676786] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (8562640 kB) free!
[1572677082] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8612368 kB) free!
[1572677202] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32712324 kB) free.
[1572677223] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201256 kB) free!
[1572677343] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177124 kB) free!
[1572677383] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4492072 kB) free!
[1572677463] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175952 kB) free!
[1572677503] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (4210900 kB) free!
[1572677623] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1572677623] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3899020 kB) free!
[1572677802] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8016008 kB) free!
[1572677922] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10999356 kB) free!
[1572677986] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4305180 kB) free!
[1572678042] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28355908 kB) free.
[1572678388] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4890036 kB) free!
[1572678508] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4712352 kB) free!
[1572678628] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4552892 kB) free!
[1572678642] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5009432 kB) free!
[1572678763] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29781152 kB) free.
[1572678823] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1928408 kB) free!
[1572679362] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5994932 kB) free!
[1572679482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2121800 kB) free!
[1572679589] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 29.64, 30.14, 36.52
[1572679602] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27910576 kB) free.
[1572680111] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 79.2% (19536616 kB) free.
[1572680149] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22421224 kB) free.
[1572680202] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3141952 kB) free!
[1572680205] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 95.2% (46913608 kB) free.
[1572680258] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.8% (55061488 kB) free.
[1572680445] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9001756 kB) free!
[1572680565] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.5% (8117020 kB) free!
[1572680685] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (8105596 kB) free!
[1572680715] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176300 kB) free!
[1572680789] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 46.72, 87.94, 64.28
[1572680835] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201208 kB) free!
[1572680955] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200988 kB) free!
[1572681970] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/11/02 13:36:04 [error] 24132#0: *1964863 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572681989] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 27.04, 30.96, 39.69
[1572682090] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(26) < 2019/11/02 13:38:07 [error] 24130#0: *1965093 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572682210] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(19) < 2019/11/02 13:39:31 [error] 24130#0: *1965130 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572682662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20332212 kB) free!
[1572682782] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16212208 kB) free!
[1572682810] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572682828] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2231808 kB) free!
[1572682902] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (12261956 kB) free!
[1572682927] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4407764 kB) free!
[1572683047] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.99, 5.91, 5.80
[1572683047] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (4312532 kB) free!
[1572683167] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.83, 8.50, 6.78
[1572683168] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (4204172 kB) free!
[1572683244] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.1% (19031828 kB) free.
[1572683287] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.90, 7.88, 6.72
[1572683301] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (1282380 kB) free!
[1572683302] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1572683302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.5% (54651752 kB) free.
[1572683348] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.7% (22785544 kB) free.
[1572683365] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4074604 kB) free!
[1572683400] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.6% (22930636 kB) free.
[1572683485] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (3758352 kB) free!
[1572683605] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3278700 kB) free!
[1572683789] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 27.26, 26.97, 29.60
[1572683853] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184860 kB) free!
[1572683973] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186516 kB) free!
[1572684093] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165776 kB) free!
[1572684205] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (800436 kB) free!
[1572684285] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239564 kB) free!
[1572684389] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 282.52, 91.65, 50.59
[1572684506] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20965684 kB) free!
[1572684509] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 62.64, 70.42, 47.77
[1572684626] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34566632 kB) free.
[1572684629] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.97, 58.12, 46.01
[1572685019] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (246512 kB) free!
[1572685113] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (221612 kB) free!
[1572685139] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214416 kB) free!
[1572685180] SERVICE ALERT: sig3-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 22538 MB (18% inode=96%):
[1572685226] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8541972 kB) free!
[1572685233] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225860 kB) free!
[1572685259] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224860 kB) free!
[1572685300] SERVICE ALERT: sig3-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 21806 MB (18% inode=96%):
[1572685346] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4532408 kB) free!
[1572685353] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241240 kB) free!
[1572685420] SERVICE ALERT: sig3-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 21067 MB (17% inode=96%):
[1572685467] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27839996 kB) free.
[1572685568] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437252 kB) free!
[1572685859] SERVICE FLAPPING ALERT: qservercluster2;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1572685859] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (8885916 kB) free!
[1572685935] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 13.3% (4642420 kB) free!
[1572685954] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 45.2% (22267544 kB) free.
[1572686038] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.1% (21496844 kB) free.
[1572686055] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 11.8% (4117020 kB) free!
[1572686066] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1495304 kB) free!
[1572686068] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 63.7% (31410200 kB) free.
[1572686105] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.2% (19667288 kB) free.
[1572686148] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 84.7% (19013376 kB) free.
[1572686175] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (3523212 kB) free!
[1572686182] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.6% (23257880 kB) free!
[1572686185] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42202576 kB) free.
[1572686357] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4315692 kB) free!
[1572686459] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 37.3% (18369608 kB) free.
[1572686477] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3481172 kB) free!
[1572686503] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (5141916 kB) free!
[1572686597] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.3% (2670956 kB) free!
[1572686623] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 13.1% (4574500 kB) free!
[1572686743] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 11.5% (4030440 kB) free!
[1572686775] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (382128 kB) free!
[1572687032] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 93.5% (32574060 kB) free.
[1572687197] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.5% (22751344 kB) free.
[1572687220] SERVICE ALERT: sig3-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 90231 MB (75% inode=96%):
[1572687244] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189076 kB) free!
[1572687343] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.1% (31896540 kB) free.
[1572687365] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183836 kB) free!
[1572687485] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166064 kB) free!
[1572688087] SERVICE ALERT: sig3-vm2;CPU Load;OK;HARD;1;OK - load average: 2.12, 3.18, 4.30
[1572688210] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2019/11/02 15:14:38 [error] 24134#0: *1982322 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572688259] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2898312 kB) free!
[1572688328] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 63.1% (15567756 kB) free.
[1572688330] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572688374] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 88.2% (43443556 kB) free.
[1572688585] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18364040 kB) free!
[1572688687] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.56, 11.72, 8.26
[1572688706] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (14102348 kB) free!
[1572688807] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.83, 10.01, 7.98
[1572688826] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (9858740 kB) free!
[1572688927] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.41, 9.46, 8.09
[1572688934] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173936 kB) free!
[1572689054] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169536 kB) free!
[1572689175] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174028 kB) free!
[1572689426] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16298120 kB) free!
[1572689755] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1572690130] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(59) < 2019/11/02 15:50:47 [error] 24130#0: *1983708 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572690185] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8288428 kB) free!
[1572690223] SERVICE ALERT: secondary-server;Current Users;CRITICAL;SOFT;1;USERS CRITICAL - 13 users currently logged in
[1572690250] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572690305] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7962140 kB) free!
[1572690344] SERVICE ALERT: secondary-server;Current Users;CRITICAL;SOFT;2;USERS CRITICAL - 13 users currently logged in
[1572690425] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (7674704 kB) free!
[1572690464] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1572690626] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38657336 kB) free.
[1572690648] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1616 processes
[1572690768] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1615 processes
[1572690888] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1611 processes
[1572691226] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10018124 kB) free!
[1572691346] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6239696 kB) free!
[1572691435] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.51, 5.63, 5.29
[1572691466] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36685352 kB) free.
[1572691485] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 78.7% (19411460 kB) free.
[1572691555] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.26, 5.23, 5.21
[1572691625] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4744988 kB) free!
[1572691675] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.92, 5.29, 5.24
[1572692066] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9039476 kB) free!
[1572692095] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195884 kB) free!
[1572692186] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4590976 kB) free!
[1572692215] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193188 kB) free!
[1572692306] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400532 kB) free!
[1572692335] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190852 kB) free!
[1572692554] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 13.5% (6655968 kB) free!
[1572692675] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226636 kB) free!
[1572692795] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 10.1% (4969616 kB) free!
[1572692804] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21790508 kB) free.
[1572692839] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 38.9% (19156440 kB) free.
[1572692869] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22931872 kB) free!
[1572692874] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43436384 kB) free.
[1572693236] SERVICE ALERT: sig3-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 575 processes
[1572693357] SERVICE ALERT: sig3-vm2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 575 processes
[1572693395] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1572693395] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3612004 kB) free!
[1572693477] SERVICE ALERT: sig3-vm2;Total Processes;WARNING;HARD;3;PROCS WARNING: 575 processes
[1572693629] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 26.24, 26.42, 35.05
[1572694016] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175928 kB) free!
[1572694074] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16353172 kB) free!
[1572694136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185216 kB) free!
[1572694194] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36067164 kB) free.
[1572694256] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184748 kB) free!
[1572694397] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1608688 kB) free!
[1572694517] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (944848 kB) free!
[1572694554] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.1% (3834744 kB) free!
[1572694637] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (475060 kB) free!
[1572694674] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3364528 kB) free!
[1572694677] SERVICE ALERT: sig3-vm2;Total Processes;OK;HARD;1;PROCS OK: 520 processes
[1572694794] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (2880252 kB) free!
[1572694794] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6986504 kB) free!
[1572694830] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 26.93, 27.04, 29.22
[1572694914] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2572604 kB) free!
[1572695034] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20844936 kB) free!
[1572695035] SERVICE ALERT: sigqscluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1572695155] SERVICE ALERT: sigqscluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 299 processes
[1572695275] SERVICE ALERT: sigqscluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1572695394] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (564540 kB) free!
[1572695491] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.5% (23319384 kB) free.
[1572695519] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 31.4% (15488636 kB) free.
[1572695557] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1572695557] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.0% (22631592 kB) free.
[1572695594] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1572695594] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 83.7% (18791484 kB) free.
[1572695634] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396392 kB) free!
[1572696636] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 17.9% (6249372 kB) free!
[1572696756] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (5709784 kB) free!
[1572696834] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28351284 kB) free.
[1572696850] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2019/11/02 17:41:46 [error] 24131#0: *1966233 recv() failed (104: Connection reset by peer) while proxying connection, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.143:13001", bytes from/to client:40847/588606, bytes from/to upstream:588606/40847
[1572696876] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (5160116 kB) free!
[1572696970] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572697434] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (375980 kB) free!
[1572697477] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1572697477] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (2415004 kB) free!
[1572697543] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 13.9% (4869748 kB) free!
[1572697554] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389092 kB) free!
[1572697663] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 12.6% (4399136 kB) free!
[1572697675] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34314864 kB) free.
[1572697783] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 11.2% (3919660 kB) free!
[1572697976] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1572698077] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 93.0% (32383224 kB) free.
[1572698097] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.3% (22691920 kB) free.
[1572698274] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2906052 kB) free!
[1572698384] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1572698384] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.6% (31727908 kB) free.
[1572698394] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (387912 kB) free!
[1572698430] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 38.96, 34.40, 31.78
[1572698514] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37001144 kB) free.
[1572698551] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 31.47, 33.18, 31.66
[1572698671] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 30.03, 32.46, 31.61
[1572698706] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1572698829] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.8% (22814788 kB) free.
[1572699114] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5906812 kB) free!
[1572699234] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1439708 kB) free!
[1572699354] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398352 kB) free!
[1572700554] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24126308 kB) free!
[1572701154] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (421872 kB) free!
[1572701594] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41040868 kB) free.
[1572701671] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 28.10, 29.20, 29.75
[1572701705] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.5% (23046008 kB) free!
[1572701826] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 36.6% (48214724 kB) free.
[1572701845] SERVICE FLAPPING ALERT: qservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1572702804] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17577084 kB) free!
[1572702925] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13050172 kB) free!
[1572703045] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8808568 kB) free!
[1572703626] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1572703645] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13986980 kB) free!
[1572703894] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163384 kB) free!
[1572703925] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1572704014] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171220 kB) free!
[1572704134] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183848 kB) free!
[1572704245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27288332 kB) free.
[1572704755] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4362184 kB) free!
[1572704845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (778996 kB) free!
[1572704875] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.4% (4272680 kB) free!
[1572704965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38776556 kB) free.
[1572704995] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (4096404 kB) free!
[1572705202] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.0% (3134432 kB) free!
[1572705271] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 30.43, 37.59, 33.78
[1572705322] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.8% (2647256 kB) free!
[1572705391] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 28.01, 34.13, 32.97
[1572705442] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2179924 kB) free!
[1572705511] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 31.61, 34.24, 33.21
[1572705565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7530428 kB) free!
[1572705685] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2627712 kB) free!
[1572705805] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390244 kB) free!
[1572706026] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.3% (2436988 kB) free!
[1572706146] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (1877856 kB) free!
[1572706267] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1317988 kB) free!
[1572706615] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.1% (18536580 kB) free.
[1572706625] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (26000664 kB) free!
[1572706653] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21342808 kB) free!
[1572706656] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35419884 kB) free.
[1572706691] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.4% (22726216 kB) free.
[1572706720] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.0% (15488796 kB) free.
[1572706745] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.1% (19891984 kB) free!
[1572706865] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.5% (24432024 kB) free!
[1572707077] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (6881724 kB) free!
[1572707170] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2019/11/02 20:32:15 [error] 24134#0: *2004679 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:13001, upstream: "10.147.212.155:13001", bytes from/to client:6343/18733, bytes from/to upstream:18733/6343
[1572707197] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 18.3% (6360644 kB) free!
[1572707226] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183264 kB) free!
[1572707290] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572707317] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 16.7% (5832864 kB) free!
[1572707346] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177600 kB) free!
[1572707384] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1572707395] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2343948 kB) free!
[1572707465] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 34.9% (46025672 kB) free.
[1572707466] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197124 kB) free!
[1572707504] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1572707624] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1572707856] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24295820 kB) free!
[1572707918] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (3246996 kB) free!
[1572707976] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.9% (18283848 kB) free!
[1572707984] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (5200876 kB) free!
[1572708096] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12957984 kB) free!
[1572708104] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 13.6% (4746712 kB) free!
[1572708225] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 12.2% (4281536 kB) free!
[1572708696] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28989176 kB) free.
[1572708825] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.3% (31615060 kB) free.
[1572709118] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.9% (32358704 kB) free.
[1572709204] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4472532 kB) free!
[1572709296] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3011448 kB) free!
[1572709324] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.3% (4238768 kB) free!
[1572709416] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33182260 kB) free.
[1572709444] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (4017996 kB) free!
[1572709476] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9702884 kB) free!
[1572709596] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9705016 kB) free!
[1572709716] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (9702848 kB) free!
[1572709865] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.3% (25413928 kB) free!
[1572709985] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.7% (19397676 kB) free!
[1572710016] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8906876 kB) free!
[1572710105] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.8% (47155768 kB) free.
[1572710136] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4058536 kB) free!
[1572710256] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391348 kB) free!
[1572710644] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1852500 kB) free!
[1572711125] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9082588 kB) free!
[1572711245] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.7% (8713752 kB) free!
[1572711365] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.0% (8371080 kB) free!
[1572711457] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28562840 kB) free.
[1572712057] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1546968 kB) free!
[1572712177] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31559108 kB) free.
[1572712565] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4902948 kB) free!
[1572712777] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8532204 kB) free!
[1572712897] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3724716 kB) free!
[1572713017] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432912 kB) free!
[1572713656] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9852580 kB) free!
[1572713776] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9421208 kB) free!
[1572713896] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.2% (8981012 kB) free!
[1572713911] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1572714217] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28048852 kB) free.
[1572714496] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4668960 kB) free!
[1572714511] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 25.59, 27.54, 29.08
[1572714817] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2214076 kB) free!
[1572714905] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.0% (23700580 kB) free!
[1572714937] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (430700 kB) free!
[1572715025] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.5% (17826216 kB) free!
[1572715046] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220416 kB) free!
[1572715057] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33050524 kB) free.
[1572715145] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.8% (47248488 kB) free.
[1572715166] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241912 kB) free!
[1572715286] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236288 kB) free!
[1572715657] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4365664 kB) free!
[1572715777] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (392236 kB) free!
[1572715897] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36159828 kB) free.
[1572716088] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1436 processes
[1572716275] SERVICE ALERT: sig2-vm2;CPU Load;OK;HARD;1;OK - load average: 1.67, 3.88, 4.89
[1572716277] SERVICE ALERT: sig3-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 569 processes
[1572716311] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 31.46, 32.13, 30.26
[1572716323] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (3764636 kB) free!
[1572716345] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.3% (16179468 kB) free!
[1572716397] SERVICE ALERT: sig3-vm2;Total Processes;OK;HARD;1;PROCS OK: 457 processes
[1572716432] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 30.29, 30.79, 29.96
[1572716443] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3195544 kB) free!
[1572716465] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400196 kB) free!
[1572716498] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7995232 kB) free!
[1572716563] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.8% (2426216 kB) free!
[1572716585] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392404 kB) free!
[1572716618] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3397864 kB) free!
[1572716689] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 2248 processes
[1572716738] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (428268 kB) free!
[1572716809] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1476 processes
[1572716893] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4341428 kB) free!
[1572716896] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.3% (5075884 kB) free!
[1572717013] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.4% (4103712 kB) free!
[1572717031] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 27.92, 30.68, 30.69
[1572717066] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1572717127] SERVICE ALERT: sig3-vm2;CPU Load;OK;HARD;1;OK - load average: 0.76, 2.90, 4.66
[1572717133] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 16.2% (3821364 kB) free!
[1572717151] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 22.76, 28.13, 29.78
[1572717163] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (632400 kB) free!
[1572717496] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1572717496] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 27.3% (13463288 kB) free.
[1572717727] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.27, 7.93, 6.88
[1572717733] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1204584 kB) free!
[1572717847] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.42, 5.62, 6.15
[1572717938] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (25396156 kB) free!
[1572717969] SERVICE ALERT: sig3-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.69, 4.93, 5.85
[1572718228] SERVICE ALERT: sig3-vm2;CPU Load;UNKNOWN;HARD;3;Parameter inconsistency: 1-minute "warning load" is greater than "critical load"
[1572718349] SERVICE FLAPPING ALERT: sig3-vm2;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1572718349] SERVICE ALERT: sig3-vm2;CPU Load;OK;HARD;1;OK - load average: 10.15, 8.80, 7.38
[1572718385] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1572718385] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (28458660 kB) free.
[1572718538] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (470736 kB) free!
[1572718696] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8051716 kB) free!
[1572718816] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7677360 kB) free!
[1572718936] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.8% (7298152 kB) free!
[1572718985] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.9% (21013968 kB) free!
[1572719105] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13047140 kB) free!
[1572719225] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (6919008 kB) free!
[1572719346] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 362 processes