[1544384145] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1544385903] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1544415303] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544415423] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544415543] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544416143] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544416365] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 44.5% (21917368 kB) free.
[1544416599] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1544416660] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.78, 11.38, 10.07
[1544416720] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1544416780] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.95, 11.85, 10.39
[1544416840] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1544416900] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.07, 12.26, 10.71
[1544420744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544420864] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544422290] SERVICE ALERT: sig5;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig5' was out of bounds)
[1544422342] SERVICE ALERT: sig5;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544422345] HOST ALERT: sig5;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.57)
[1544422349] SERVICE ALERT: sig5;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig5' was out of bounds)
[1544422349] SERVICE ALERT: sig5;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig5' was out of bounds)
[1544422409] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1544422410] SERVICE ALERT: sig5;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.13, 0.04
[1544422463] SERVICE ALERT: sig5;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1544422469] SERVICE ALERT: sig5;Sensors Status;OK;HARD;1;SENSORS OK
[1544422470] SERVICE ALERT: sig5;Memory;OK;HARD;1;OK - 94.4% (15346424 kB) free.
[1544424007] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544424010] HOST ALERT: sig5;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.57)
[1544424074] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1544424127] SERVICE ALERT: sig5;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1544427676] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 45.5% (7389280 kB) free.
[1544428064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1544428184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544428304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544428802] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 301 processes
[1544428904] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544428923] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 300 processes
[1544429039] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 2 warn out of 344 processes
[1544429043] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 300 processes
[1544429344] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544429464] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544429504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544429624] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544429744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544430064] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544430184] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544430238] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1545936 kB) free!
[1544430344] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1544430344] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544431230] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 14.05, 24.20, 11.92
[1544431350] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.02, 18.12, 11.20
[1544431419] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 465 processes
[1544431443] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1544431470] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.43, 12.41, 9.95
[1544431479] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 468 processes
[1544431539] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 1 warn out of 466 processes
[1544431544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544431599] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 463 processes
[1544431664] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544431984] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544432043] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 300 processes
[1544432104] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544432476] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3232376 kB) free!
[1544432596] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3099264 kB) free!
[1544432670] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.24, 1.17, 3.40
[1544432716] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.1% (2944072 kB) free!
[1544432864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544432984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544433104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544433704] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544433904] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544433916] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1098124 kB) free!
[1544434024] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544434305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544434425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544434545] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544434624] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544434744] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544434864] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544434987] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (2898224 kB) free!
[1544435107] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.0% (2759504 kB) free!
[1544435227] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.1% (2619744 kB) free!
[1544435716] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1544435716] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 29.3% (4761940 kB) free.
[1544435745] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544435865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544435985] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544436064] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1544436239] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 341 processes
[1544436585] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544436839] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 341 processes
[1544436843] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1544436843] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1544437186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544437306] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544437426] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544437864] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544437984] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1544438039] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1544438039] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 345 processes
[1544438584] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544438704] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544440426] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544440470] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 522.17, 321.77, 126.72
[1544440547] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 287 processes
[1544440590] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 70.91, 215.46, 111.41
[1544440710] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 9.94, 144.31, 97.95
[1544441704] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544441825] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544442424] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544442510] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1544442510] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.13, 1.54, 15.10
[1544442544] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544442562] SERVICE ALERT: sig2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442570] SERVICE ALERT: sig2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442625] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442625] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442637] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442692] SERVICE ALERT: sig2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442700] SERVICE ALERT: sig2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442702] SERVICE ALERT: sig2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442755] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442755] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442767] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442798] SERVICE ALERT: sig2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442822] SERVICE ALERT: sig2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442830] SERVICE ALERT: sig2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442832] SERVICE ALERT: sig2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442867] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442885] SERVICE ALERT: sig2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442885] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442897] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442928] SERVICE ALERT: sig2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442962] SERVICE ALERT: sig2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1544442997] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1544443025] SERVICE ALERT: sig2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544443058] SERVICE ALERT: sig2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1544443127] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1544443144] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544443155] SERVICE ALERT: sig2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1544443264] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544443285] SERVICE ALERT: sig2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1544443658] SERVICE ALERT: sig2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1544443710] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 0.85, 4.80
[1544443727] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1544443840] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 316 processes
[1544443885] SERVICE ALERT: sig2;Sensors Status;OK;HARD;1;SENSORS OK
[1544444022] SERVICE ALERT: sig2;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.11, 0.10
[1544444030] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 94.4% (15349520 kB) free.
[1544444043] SERVICE FLAPPING ALERT: sservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1544444085] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38579 MB (75% inode=98%):
[1544444085] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 299 processes
[1544444097] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1544444162] SERVICE ALERT: sig2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 219768 MB (97% inode=99%):
[1544444440] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1544445148] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.4% (23383792 kB) free.
[1544445916] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1544448346] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1544448840] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1544451510] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 245.39, 309.45, 137.79
[1544451630] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 34.47, 207.51, 121.24
[1544451750] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.93, 139.35, 106.72
[1544452865] SERVICE ALERT: db1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1544452985] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1544453550] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.99, 1.34, 16.27
[1544453640] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1544453916] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1544454036] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1544454156] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1544454756] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1544455350] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.85, 0.94, 3.15
[1544457640] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 317 processes
[1544457946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1544458066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544458186] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544458441] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1544458561] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 338 processes
[1544458681] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 338 processes
[1544460641] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1544461241] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1544465550] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)