[1547592191] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 20.31, 16.34, 7.62
[1547592311] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.92, 11.37, 6.86
[1547592432] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.82, 8.13, 6.22
[1547592487] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 7 warn out of 355 processes
[1547593032] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.50, 2.37, 3.95
[1547597887] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 6 warn out of 354 processes
[1547600988] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1547602087] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 7 warn out of 357 processes
[1547602252] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.16, 4.06, 4.73
[1547604488] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 6 warn out of 355 processes
[1547606580] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547606700] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547606820] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547606839] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3393776 kB) free!
[1547606959] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.8% (3390824 kB) free!
[1547607079] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.5% (3052184 kB) free!
[1547607420] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547607540] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547612340] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547612460] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547615460] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547615479] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2429376 kB) free!
[1547615580] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547615700] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547616300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547616421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547616541] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1547617741] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547618288] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 361 processes
[1547619080] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547619200] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547620141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547620261] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547620382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547620539] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47086 MB (20% inode=99%):
[1547620659] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47094 MB (20% inode=99%):
[1547620779] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49174 MB (21% inode=99%):
[1547621000] SERVICE ALERT: db1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1547621120] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547622320] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547622440] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547623688] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 7 warn out of 364 processes
[1547624379] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47086 MB (20% inode=99%):
[1547624499] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47084 MB (20% inode=99%):
[1547624619] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47068 MB (20% inode=99%):
[1547625181] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 92.4% (15021828 kB) free.
[1547625219] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 98366 MB (43% inode=99%):
[1547625782] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1547625782] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1547626879] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (4612536 kB) free!
[1547627288] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 8 warn out of 360 processes
[1547628782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547628902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547629022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547629041] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547629161] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547629755] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4671320 kB) free!
[1547630222] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547630288] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 6 warn out of 362 processes
[1547630355] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5327264 kB) free!
[1547630581] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3110360 kB) free!
[1547630701] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3041728 kB) free!
[1547630821] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.4% (2985760 kB) free!
[1547630822] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547630895] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 6 warn out of 340 processes
[1547630942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547631062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1547631079] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2052452 kB) free!
[1547631495] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 7 warn out of 338 processes
[1547632161] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547632281] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547634555] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4886032 kB) free!
[1547635262] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547636288] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 358 processes
[1547636462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1547636479] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 24.0% (5875304 kB) free.
[1547636582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547636661] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547636678] SERVICE ALERT: sig1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547636691] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1547636702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547636781] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1547636798] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.50, 0.16, 0.06
[1547636802] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1547636855] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.7% (15385424 kB) free.
[1547636889] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1547636889] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 352 processes
[1547637079] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (4049912 kB) free!
[1547637199] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.7% (3612728 kB) free!
[1547637319] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1547637319] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (3209396 kB) free!
[1547637555] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1547637555] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6387688 kB) free!
[1547637919] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1567748 kB) free!
[1547638155] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4639460 kB) free!
[1547638163] HOST ALERT: sig2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1547638185] SERVICE ALERT: sig2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig2' was out of bounds)
[1547638195] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1547638281] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547638305] SERVICE ALERT: sig2;Sensors Status;OK;HARD;1;SENSORS OK
[1547638401] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547638608] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (3190748 kB) free!
[1547638621] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1618388 kB) free!
[1547638755] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6164752 kB) free!
[1547638840] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46595 MB (20% inode=99%):
[1547638960] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46595 MB (20% inode=99%):
[1547639080] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46456 MB (20% inode=99%):
[1547639208] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 21.6% (3504108 kB) free.
[1547639295] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 336 processes
[1547639470] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52958 MB (23% inode=99%):
[1547639633] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 479 processes
[1547639680] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49423 MB (21% inode=99%):
[1547639693] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 479 processes
[1547639702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1547639719] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3337492 kB) free!
[1547639754] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 2 warn out of 477 processes
[1547639814] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 1 warn out of 475 processes
[1547639955] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4734776 kB) free!
[1547640052] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.14, 6.32, 5.63
[1547640114] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 477 processes
[1547640172] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.14, 6.30, 5.70
[1547640292] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.19, 6.25, 5.75
[1547640324] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.57, 9.35, 9.80
[1547640489] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 347 processes
[1547640589] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1547641189] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1547641309] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1547641429] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1547641492] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.32, 3.52, 4.31
[1547641755] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6907288 kB) free!
[1547642001] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 336 processes
[1547642121] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1547642241] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1547642289] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 347 processes
[1547642630] SERVICE ALERT: db1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 5939 MB (11% inode=98%):
[1547642719] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2118208 kB) free!
[1547643302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1547643422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1547644430] SERVICE ALERT: db1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 11906 MB (23% inode=99%):
[1547644490] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9381820 kB) free!
[1547644610] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9330832 kB) free!
[1547644622] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1547644641] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547644730] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.5% (9101240 kB) free!
[1547644742] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1547644755] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4167708 kB) free!
[1547644761] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1547644862] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1547645355] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5577224 kB) free!
[1547645462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1547645583] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547645630] SERVICE ALERT: db1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10465 MB (20% inode=99%):
[1547645703] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547645724] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.96, 11.83, 10.85
[1547645750] SERVICE ALERT: db1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 10465 MB (20% inode=99%):
[1547645844] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.03, 11.66, 10.91
[1547645870] SERVICE ALERT: db1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 10634 MB (20% inode=99%):
[1547645961] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547645964] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.00, 11.66, 10.99
[1547646081] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1547646292] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.19, 7.10, 5.90
[1547646412] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.16, 6.48, 5.83
[1547646532] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.04, 5.59, 5.58
[1547646670] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46880 MB (20% inode=99%):
[1547646790] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46880 MB (20% inode=99%):
[1547646910] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46880 MB (20% inode=99%):
[1547646962] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1547647014] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 479 processes
[1547647074] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 477 processes
[1547647082] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 304 processes
[1547647155] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4192676 kB) free!
[1547647202] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1547647282] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1547647402] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1547647480] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46112 MB (20% inode=99%):
[1547647522] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1547647600] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45748 MB (20% inode=99%):
[1547647720] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45748 MB (20% inode=99%):
[1547647755] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5667404 kB) free!
[1547647802] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 287 processes
[1547648005] SERVICE ALERT: sig4;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547648008] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4' was out of bounds)
[1547648008] HOST ALERT: sig4;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.56)
[1547648008] SERVICE ALERT: sig4;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig4' was out of bounds)
[1547648033] SERVICE ALERT: sig4;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig4' was out of bounds)
[1547648040] SERVICE ALERT: sig4;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig4' was out of bounds)
[1547648046] SERVICE ALERT: sig4;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4' was out of bounds)
[1547648105] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1547648122] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1547648125] SERVICE ALERT: sig4;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 221089 MB (98% inode=99%):
[1547648128] SERVICE ALERT: sig4;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37513 MB (73% inode=98%):
[1547648128] SERVICE ALERT: sig4;CPU Load;OK;HARD;1;OK - load average: 0.45, 0.18, 0.07
[1547648153] SERVICE ALERT: sig4;Sensors Status;OK;HARD;1;SENSORS OK
[1547648160] SERVICE ALERT: sig4;Total Processes;OK;HARD;1;PROCS OK: 312 processes
[1547648166] SERVICE ALERT: sig4;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547648242] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1547648330] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3734456 kB) free!
[1547648355] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3461104 kB) free!
[1547648403] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547648523] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 313 processes
[1547648653] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1547648683] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1547648703] SERVICE ALERT: sig4;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4' was out of bounds)
[1547648706] SERVICE ALERT: sig4;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig4' was out of bounds)
[1547648728] SERVICE ALERT: sig4;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig4' was out of bounds)
[1547648728] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4' was out of bounds)
[1547648731] SERVICE ALERT: sig4;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig4' was out of bounds)
[1547648747] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1547648823] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 94.8% (15412544 kB) free.
[1547648826] SERVICE ALERT: sig4;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1547648842] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1547648848] SERVICE ALERT: sig4;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 221080 MB (98% inode=99%):
[1547648848] SERVICE ALERT: sig4;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37513 MB (73% inode=98%):
[1547648851] SERVICE FLAPPING ALERT: sig4;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1547648851] SERVICE ALERT: sig4;CPU Load;OK;HARD;1;OK - load average: 0.08, 0.08, 0.04
[1547648962] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547649082] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1547649474] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 486 processes
[1547649534] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 485 processes
[1547649555] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5612784 kB) free!
[1547649683] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1547649734] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1547649853] SERVICE FLAPPING ALERT: sig4;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1547649853] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1547649854] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1547649903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1547649974] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1547650155] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (3003704 kB) free!
[1547650883] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1547651003] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1547651123] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1547651446] HOST ALERT: sig4;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.56)
[1547651510] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1547651703] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547651719] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1547651774] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1547651808] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.9% (2586452 kB) free!
[1547651823] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547651929] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.5% (2350688 kB) free!
[1547652049] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.1% (2132200 kB) free!
[1547652323] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1547652323] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1547652423] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547652489] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1547652543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547652649] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1280488 kB) free!
[1547652663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547652923] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1547653021] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.3% (2168768 kB) free!
[1547653043] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547653164] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1547653263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547653755] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5618808 kB) free!
[1547653763] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547653863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547653883] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1547653984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547654104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547654483] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547654603] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547654821] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1546616 kB) free!
[1547654895] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 8 warn out of 337 processes
[1547655203] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547655323] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547656155] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3087520 kB) free!
[1547656523] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547656643] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547656764] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547657104] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547658563] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547658683] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547660110] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47992 MB (21% inode=99%):
[1547660119] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (3050436 kB) free!
[1547660236] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 2229 MB (4% inode=94%):
[1547660251] SERVICE FLAPPING ALERT: sig4;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547660449] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.2% (2627696 kB) free!
[1547660653] SERVICE FLAPPING ALERT: sig4;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1547660719] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2442740 kB) free!
[1547661049] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1547661049] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.9% (5022940 kB) free.
[1547661310] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47160 MB (20% inode=99%):
[1547661430] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46905 MB (20% inode=99%):
[1547661551] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46905 MB (20% inode=99%):
[1547662155] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5044528 kB) free!
[1547662504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547662624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547663201] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547663201] SERVICE ALERT: sig1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig1' was out of bounds)
[1547663260] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1547663321] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.15, 0.06