[1584988453] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.4% (2007900 kB) free!
[1584988573] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (2004204 kB) free!
[1584988693] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.3% (2005508 kB) free!
[1584988765] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24529536 kB) free!
[1584989104] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1584989224] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 331 processes
[1584989345] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1584989365] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386696 kB) free!
[1584989965] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23329920 kB) free!
[1584990241] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1584990565] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399524 kB) free!
[1584990841] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1584990890] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 65.1% (15955764 kB) free.
[1584991166] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18817492 kB) free!
[1584991766] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394376 kB) free!
[1584992233] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (4132328 kB) free!
[1584992366] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15685116 kB) free!
[1584992833] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1425992 kB) free!
[1584992966] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407668 kB) free!
[1584993154] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 40 users currently logged in
[1584993541] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236704 kB) free!
[1584993661] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (254328 kB) free!
[1584993781] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (221828 kB) free!
[1584994093] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (449896 kB) free!
[1584994213] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (449460 kB) free!
[1584994333] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (445448 kB) free!
[1584994354] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 32 users currently logged in
[1584994954] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 21 users currently logged in
[1584995366] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24437376 kB) free!
[1584995452] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 752.32, 163.67, 55.68
[1584995516] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3226.04, 954.31, 331.22
[1584995556] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [24/Mar/2020:02:01:57 +0530] TCP 502 0 0 60.002 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.001, 0.000"
[1584995572] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4319.21, 1722.68, 640.08
[1584995581] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 11.5% (5676548 kB) free!
[1584995636] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4649.53, 2250.38, 881.72
[1584995676] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584995692] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1442.50, 1747.86, 799.76
[1584995696] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1894.48, 1968.50, 883.28
[1584995756] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 628.28, 1506.22, 774.97
[1584995816] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 258.30, 1318.25, 776.58
[1584995936] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 37.78, 883.29, 682.87
[1584995966] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1581476 kB) free!
[1584996181] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4477696 kB) free!
[1584996276] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(57) < 182.18.144.246 [24/Mar/2020:02:06:10 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1584996356] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.09, 203.54, 406.90
[1584996396] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584996536] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.37, 121.36, 359.44
[1584996566] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16549736 kB) free!
[1584996892] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.94, 39.18, 224.50
[1584997166] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406972 kB) free!
[1584997868] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 35.3% (8660036 kB) free.
[1584998023] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 79.0% (38918320 kB) free.
[1584998547] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.5% (45109260 kB) free.
[1584999182] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1584999182] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 72.2% (35579648 kB) free.
[1585000167] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13791396 kB) free!
[1585000767] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395976 kB) free!
[1585001967] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15283304 kB) free!
[1585002567] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402016 kB) free!
[1585003767] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13838704 kB) free!
[1585004367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26497204 kB) free.
[1585004967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4513976 kB) free!
[1585005087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (823664 kB) free!
[1585005207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26813400 kB) free.
[1585005807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5100348 kB) free!
[1585005927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1588540 kB) free!
[1585006047] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26298360 kB) free!
[1585006268] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176996 kB) free!
[1585006354] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585006388] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166708 kB) free!
[1585006508] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (256696 kB) free!
[1585006647] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5457520 kB) free!
[1585007745] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.0% (18863236 kB) free.
[1585007833] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.2% (11648228 kB) free.
[1585007847] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21649376 kB) free!
[1585008182] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1585008218] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 71.5% (17523644 kB) free.
[1585008447] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2102088 kB) free!
[1585009647] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13264724 kB) free!
[1585010247] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33913928 kB) free.
[1585010399] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1585010519] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1585010639] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1585010641] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1585010745] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167896 kB) free!
[1585010761] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1585010847] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6661784 kB) free!
[1585010865] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167132 kB) free!
[1585010881] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1585010967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2331816 kB) free!
[1585010985] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195056 kB) free!
[1585011087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29718836 kB) free.
[1585011134] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 62.8% (82829728 kB) free.
[1585011239] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1585011481] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1585011687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2923204 kB) free!
[1585011807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393596 kB) free!
[1585011927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29364900 kB) free.
[1585012144] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1585012264] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1585012527] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3929412 kB) free!
[1585012647] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30966404 kB) free.
[1585013247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11054496 kB) free!
[1585013281] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1585013367] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7428364 kB) free!
[1585013401] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1585013488] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31845524 kB) free.
[1585013521] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1585013796] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(111) < 92.42.104.166 [24/Mar/2020:07:00:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585013916] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585014087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10246428 kB) free!
[1585014121] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1585014121] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1585014207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28418496 kB) free.
[1585014807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12330768 kB) free!
[1585014927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8443468 kB) free!
[1585015047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31389936 kB) free.
[1585015647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9555408 kB) free!
[1585015767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31884780 kB) free.
[1585015864] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1585015984] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1585016367] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11896916 kB) free!
[1585016487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31899900 kB) free.
[1585016687] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 79.6% (19520924 kB) free.
[1585017087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11976292 kB) free!
[1585017121] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1585017207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31847460 kB) free.
[1585017241] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1585017290] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3958152 kB) free!
[1585017350] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3953644 kB) free!
[1585017361] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1585017410] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.1% (3958140 kB) free!
[1585017470] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.1% (3944536 kB) free!
[1585017807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11952628 kB) free!
[1585017927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8155364 kB) free!
[1585018047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29741848 kB) free.
[1585018647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9239560 kB) free!
[1585018768] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30797676 kB) free.
[1585018984] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1585019105] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1585019367] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9777232 kB) free!
[1585019487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31444236 kB) free.
[1585020087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10637300 kB) free!
[1585020207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31357360 kB) free.
[1585020807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9864620 kB) free!
[1585020818] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4647936 kB) free!
[1585020927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31002896 kB) free.
[1585020938] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4547136 kB) free!
[1585021058] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4444916 kB) free!
[1585021240] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/03/24 09:06:23 [error] 22121#0: *76284446 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585021360] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585021527] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8874064 kB) free!
[1585021647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4579468 kB) free!
[1585021767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28472292 kB) free.
[1585022161] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1585022367] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5073604 kB) free!
[1585022487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28938176 kB) free.
[1585022893] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (275252 kB) free!
[1585023087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (4951764 kB) free!
[1585023207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29184688 kB) free.
[1585023458] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1757764 kB) free!
[1585023807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2228084 kB) free!
[1585023928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (383328 kB) free!
[1585024048] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28101644 kB) free.
[1585024639] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4769680 kB) free!
[1585024648] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (409448 kB) free!
[1585024759] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (4610136 kB) free!
[1585024768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385684 kB) free!
[1585024879] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2883740 kB) free!
[1585024888] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585024888] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26487784 kB) free.
[1585025479] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2061084 kB) free!
[1585025488] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (404940 kB) free!
[1585025608] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28679648 kB) free.
[1585026208] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (378432 kB) free!
[1585026328] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (373996 kB) free!
[1585026448] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385232 kB) free!
[1585026645] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.55, 6.35, 5.71
[1585026765] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.33, 5.81, 5.61
[1585026885] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.10, 4.95, 5.32
[1585026887] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.4% (4765608 kB) free!
[1585027007] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.3% (4490536 kB) free!
[1585027048] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15100360 kB) free!
[1585027127] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.3% (4229584 kB) free!
[1585027485] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.04, 4.08, 4.50
[1585028248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28752172 kB) free.
[1585028327] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1388060 kB) free!
[1585028685] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.12, 5.99, 5.18
[1585028805] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.58, 5.57, 5.14
[1585028848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (932080 kB) free!
[1585028925] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.76, 5.24, 5.07
[1585028969] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (427936 kB) free!
[1585029089] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1585029089] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (25453200 kB) free!
[1585029192] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.0% (3137288 kB) free!
[1585029312] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.3% (2766948 kB) free!
[1585029361] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1585029432] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.0% (2459552 kB) free!
[1585029526] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1585029526] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.33, 5.00, 4.91
[1585029689] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376548 kB) free!
[1585030032] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (936700 kB) free!
[1585030126] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.61, 5.58, 5.25
[1585030246] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.71, 4.99, 5.08
[1585030366] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.76, 4.61, 4.93
[1585030716] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(27) < 103.255.146.164 [24/Mar/2020:11:42:27 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585030836] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585030889] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17948196 kB) free!
[1585031064] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4644272 kB) free!
[1585031184] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4433036 kB) free!
[1585031304] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 17.2% (4236056 kB) free!
[1585031489] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381888 kB) free!
[1585031832] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.9% (15478304 kB) free.
[1585032304] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1585032424] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1585032504] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22867072 kB) free.
[1585032637] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(46) < 103.255.146.164 [24/Mar/2020:12:15:47 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585032689] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16293424 kB) free!
[1585032757] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585033289] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27076500 kB) free.
[1585033360] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/24 12:26:23 [error] 22133#0: *76342471 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: rq-1-1gbps.server, request: "GET /favicon.ico HTTP/1.0", host: "rq-1-1gbps.server", referrer: "http://rq-1-1gbps.server/"
[1585033481] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585033889] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1185668 kB) free!
[1585034009] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (392828 kB) free!
[1585034129] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402916 kB) free!
[1585034557] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(34) < 103.255.146.164 [24/Mar/2020:12:50:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585034677] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585035329] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20133084 kB) free!
[1585035766] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.95, 6.04, 5.13
[1585035886] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.58, 5.85, 5.17
[1585035929] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418820 kB) free!
[1585036006] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.07, 5.67, 5.17
[1585036477] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(19) < 103.255.146.164 [24/Mar/2020:13:19:06 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585036597] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585037129] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27106304 kB) free.
[1585037400] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 17.82, 51.76, 37.38
[1585037521] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.20, 37.92, 34.08
[1585037547] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9329772 kB) free!
[1585037641] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.93, 29.39, 31.45
[1585037667] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8915004 kB) free!
[1585037729] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (492716 kB) free!
[1585037787] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.2% (8481580 kB) free!
[1585037797] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(39) < 103.255.146.164 [24/Mar/2020:13:41:54 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585037849] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (436388 kB) free!
[1585037918] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585037969] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (428780 kB) free!
[1585038241] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 23.45, 27.43, 28.19
[1585038432] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4484372 kB) free!
[1585038552] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4050364 kB) free!
[1585038672] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.7% (3746440 kB) free!
[1585039169] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24788120 kB) free!
[1585039272] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1585039272] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2054144 kB) free!
[1585039370] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.0% (2208272 kB) free!
[1585039587] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3310356 kB) free!
[1585039717] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(37) < 103.255.146.164 [24/Mar/2020:14:14:14 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585039769] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399068 kB) free!
[1585039837] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585040369] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14156028 kB) free!
[1585040623] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9285028 kB) free!
[1585040743] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (9008816 kB) free!
[1585040863] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (8742316 kB) free!
[1585040904] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 10.6% (2618140 kB) free!
[1585041024] SERVICE ALERT: db1-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (2164240 kB) free!
[1585041144] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1710296 kB) free!
[1585041569] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (7026020 kB) free!
[1585041637] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(56) < 103.255.146.164 [24/Mar/2020:14:47:50 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585041757] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585042006] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1585042103] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (5424016 kB) free!
[1585042169] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34847668 kB) free.
[1585042224] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.2% (32130252 kB) free.
[1585042272] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.6% (15636468 kB) free.
[1585042355] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 31 users currently logged in
[1585042475] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 33 users currently logged in
[1585042595] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 33 users currently logged in
[1585042603] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1585042723] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1585042769] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (395184 kB) free!
[1585042843] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1585042890] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (414484 kB) free!
[1585043010] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386600 kB) free!
[1585043263] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4688340 kB) free!
[1585043354] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1502 processes
[1585043475] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1502 processes
[1585043544] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.9% (22922504 kB) free.
[1585043557] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(31) < 103.255.146.164 [24/Mar/2020:15:18:30 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585043595] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1498 processes
[1585043677] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585044259] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(350) < 103.16.101.72 [24/Mar/2020:15:32:10 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1585044277] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1284) < 92.42.104.166 [24/Mar/2020:15:34:37 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585044280] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(472) < 2020/03/24 15:32:10 [error] 22128#0: *76400649 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585044361] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1585044379] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585044397] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 92.42.104.166 [24/Mar/2020:15:34:38 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1585044400] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585044481] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1585044517] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(61) < 103.16.101.72 [24/Mar/2020:15:37:08 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1585044601] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1585044792] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1585044921] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 72.1% (35550720 kB) free.
[1585045717] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585045798] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1585045918] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1585046041] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 23.32, 45.53, 39.75
[1585046161] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 118.21, 58.82, 44.44
[1585046224] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1585046281] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 93.54, 92.95, 60.48
[1585046344] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1585046610] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25788508 kB) free!
[1585046763] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.7% (2223684 kB) free!
[1585046795] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 26 users currently logged in
[1585046884] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2880296 kB) free!
[1585046917] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(25) < 103.255.146.164 [24/Mar/2020:16:16:06 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585047004] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3640888 kB) free.
[1585047037] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585047195] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1506 processes
[1585047210] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399808 kB) free!
[1585047316] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1513 processes
[1585047400] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2020/03/24 16:21:28 [error] 22127#0: *76416403 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1585047436] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1511 processes
[1585047520] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585047604] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.7% (2388044 kB) free!
[1585047724] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3266536 kB) free.
[1585047810] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18311528 kB) free!
[1585048324] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2721508 kB) free!
[1585048410] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410124 kB) free!
[1585048444] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.3% (2646132 kB) free!
[1585048564] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3358228 kB) free.
[1585048837] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(31) < 103.255.146.164 [24/Mar/2020:16:41:18 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585048957] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585049010] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17475244 kB) free!
[1585049164] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3007344 kB) free!
[1585049281] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.68, 10.85, 35.46
[1585049284] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4061780 kB) free.
[1585049344] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 336 processes
[1585049464] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1585049472] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4232964 kB) free!
[1585049592] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (3934848 kB) free!
[1585049610] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406840 kB) free!
[1585049712] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (3657096 kB) free!
[1585049881] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1585049881] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 10.56, 11.48, 24.28
[1585050007] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (5012740 kB) free!
[1585050047] SERVICE ALERT: redis-2;CPU Load;WARNING;SOFT;1;WARNING - load average: 17.16, 12.39, 7.47
[1585050167] SERVICE ALERT: redis-2;CPU Load;WARNING;SOFT;2;WARNING - load average: 17.67, 14.18, 8.73
[1585050231] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 712 processes
[1585050287] SERVICE ALERT: redis-2;CPU Load;WARNING;HARD;3;WARNING - load average: 21.16, 16.31, 10.14
[1585050312] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2099588 kB) free!
[1585050352] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 745 processes
[1585050472] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 740 processes
[1585050607] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (1137336 kB) free!
[1585050620] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1585050740] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1585050757] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(30) < 103.255.146.164 [24/Mar/2020:17:21:47 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585050860] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1585050877] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585051063] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (6489648 kB) free!
[1585051411] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14161336 kB) free!
[1585051684] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (678756 kB) free!
[1585051804] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.5% (1873564 kB) free!
[1585051843] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1585051924] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (3223012 kB) free!
[1585051944] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.6% (4347868 kB) free!
[1585052011] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374948 kB) free!
[1585052064] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.8% (3895020 kB) free!
[1585052077] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(102) < 103.255.146.164 [24/Mar/2020:17:42:33 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585052184] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3478780 kB) free!
[1585052198] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585052236] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 2153 processes
[1585052263] SERVICE FLAPPING ALERT: qservercluster2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1585052263] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4499416 kB) free!
[1585052611] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23111392 kB) free!
[1585052784] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1585052784] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1842648 kB) free!
[1585053211] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396464 kB) free!
[1585053312] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.8% (15679748 kB) free.
[1585053322] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1585053446] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 70.9% (34930588 kB) free.
[1585053689] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1585053724] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4069544 kB) free.
[1585053811] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17210920 kB) free!
[1585053819] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1585053940] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 167132 MB (74% inode=99%):
[1585053984] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 32.6% (8055108 kB) free.
[1585054003] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1585054050] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1585054123] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38773 MB (75% inode=98%):
[1585054171] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1585054324] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (211000 kB) free!
[1585054411] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (430744 kB) free!
[1585054444] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.9% (1763676 kB) free!
[1585054565] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3675464 kB) free.
[1585054597] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(126) < 103.255.146.164 [24/Mar/2020:18:24:02 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585054718] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585055011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33712156 kB) free.
[1585055611] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10220680 kB) free!
[1585055731] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5840908 kB) free!
[1585055764] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.0% (1953116 kB) free!
[1585055851] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1871156 kB) free!
[1585055884] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.4% (4294812 kB) free.
[1585055917] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(82) < 103.255.146.164 [24/Mar/2020:18:43:42 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1585056037] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585056451] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26347880 kB) free!
[1585056452] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1585056582] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1585056706] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 70.0% (34509352 kB) free.
[1585056860] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1585057006] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 148.33, 87.87, 42.10
[1585057051] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1210596 kB) free!
[1585057238] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(64) < 103.255.146.164 [24/Mar/2020:19:01:00 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1585057358] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585057651] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30348692 kB) free.
[1585057684] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (241240 kB) free!
[1585057804] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.0% (1952828 kB) free!
[1585057924] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.3% (3957604 kB) free.
[1585058251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2924224 kB) free!
[1585058272] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1585058371] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402412 kB) free!
[1585058492] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384704 kB) free!
[1585058687] SERVICE ALERT: redis-2;CPU Load;OK;HARD;1;OK - load average: 1.96, 1.91, 4.14
[1585059092] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19131360 kB) free!
[1585059124] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.5% (1875320 kB) free!
[1585059244] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3984872 kB) free.
[1585059437] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1523 processes
[1585059458] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 24.2% (5937704 kB) free.
[1585059601] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1585059692] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411640 kB) free!
[1585059759] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 182.18.144.246 [24/Mar/2020:19:45:42 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1585059879] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585060037] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1907 processes
[1585060292] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21527300 kB) free!
[1585060444] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (1724668 kB) free!
[1585060564] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3586684 kB) free.
[1585060658] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3558816 kB) free!
[1585060778] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.4% (3289920 kB) free!
[1585060892] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419060 kB) free!
[1585060898] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (3109808 kB) free!
[1585061164] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (1241276 kB) free!
[1585061282] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585061285] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3044212 kB) free!
[1585061405] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.2% (4581284 kB) free.
[1585061492] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21528260 kB) free!
[1585061498] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2263852 kB) free!
[1585061712] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4265212 kB) free!
[1585061806] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.02, 10.37, 17.17
[1585061832] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4200112 kB) free!
[1585061952] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.3% (4100944 kB) free!
[1585062092] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (436868 kB) free!
[1585062605] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2489952 kB) free!
[1585062692] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21338100 kB) free!
[1585062725] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3798904 kB) free.
[1585063064] SERVICE FLAPPING ALERT: qservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585063292] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382096 kB) free!
[1585063492] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 887.72, 551.35, 218.77
[1585063612] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 121.65, 369.70, 192.53
[1585063637] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1585063637] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1574 processes
[1585063892] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21533456 kB) free!
[1585064079] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(18) < 203.212.70.170 [24/Mar/2020:20:56:17 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1585064199] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1585064237] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1210 processes
[1585064352] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.8% (15663812 kB) free.
[1585064492] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391812 kB) free!
[1585064525] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.4% (1690128 kB) free!
[1585064645] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2860312 kB) free!
[1585064765] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3975924 kB) free.
[1585064784] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585064806] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.52, 3.01, 4.47
[1585065092] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20833192 kB) free!
[1585065692] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390952 kB) free!
[1585065965] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (786184 kB) free!
[1585066085] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1905668 kB) free!
[1585066205] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3022000 kB) free!
[1585066265] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1585066292] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20426912 kB) free!
[1585066385] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1585066805] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3579508 kB) free.
[1585066892] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406760 kB) free!
[1585067237] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1530 processes
[1585067358] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1521 processes
[1585067478] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1517 processes
[1585067493] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18856676 kB) free!
[1585068093] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406500 kB) free!
[1585068605] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (901684 kB) free!
[1585068693] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16428680 kB) free!
[1585068725] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1894680 kB) free!
[1585068845] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (2880900 kB) free!
[1585069293] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (441212 kB) free!
[1585069893] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21715520 kB) free!
[1585070493] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397860 kB) free!
[1585071093] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20450144 kB) free!
[1585071693] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (445732 kB) free!
[1585071845] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1585071845] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (365808 kB) free!
[1585072293] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21091576 kB) free!
[1585072893] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410296 kB) free!
[1585073045] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.0% (4231164 kB) free.
[1585073493] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20687708 kB) free!
[1585073952] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1585074093] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388240 kB) free!