[1589999591] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17313840 kB) free!
[1590000105] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.37, 5.78, 5.34
[1590000191] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (441460 kB) free!
[1590000225] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.08, 4.83, 5.05
[1590000345] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.20, 4.98, 5.07
[1590000791] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18708148 kB) free!
[1590000946] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.78, 4.26, 4.70
[1590001391] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403896 kB) free!
[1590001546] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.94, 6.47, 5.49
[1590001666] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.56, 6.92, 5.76
[1590001786] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.93, 6.83, 5.89
[1590001991] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16585572 kB) free!
[1590002292] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (693488 kB) free!
[1590002412] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (423580 kB) free!
[1590002532] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.8% (1760240 kB) free!
[1590002592] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416172 kB) free!
[1590003192] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15134836 kB) free!
[1590003217] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.9% (220268 kB) free!
[1590003680] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 33350 MB (14.83% inode=100%):
[1590003732] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (556264 kB) free!
[1590003740] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 31302 MB (13.92% inode=100%):
[1590003792] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423888 kB) free!
[1590003800] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;3;DISK WARNING - free space: /home 27085 MB (12.05% inode=100%):
[1590003861] SERVICE ALERT: webserver2;HDD_Home status;WARNING;HARD;4;DISK WARNING - free space: /home 23488 MB (10.45% inode=100%):
[1590004161] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 78157 MB (34.77% inode=100%):
[1590004333] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1672832 kB) free!
[1590004392] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14400592 kB) free!
[1590004418] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 82.8% (20297448 kB) free.
[1590004933] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1632616 kB) free!
[1590004992] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (454220 kB) free!
[1590005070] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (247084 kB) free!
[1590005190] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216044 kB) free!
[1590005264] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (459780 kB) free!
[1590005310] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234512 kB) free!
[1590005384] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389568 kB) free!
[1590005504] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1519624 kB) free!
[1590005986] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.72, 4.20, 4.63
[1590006192] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32596768 kB) free.
[1590006792] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10134804 kB) free!
[1590006912] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6239644 kB) free!
[1590007032] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2556168 kB) free!
[1590007110] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.4% (8080728 kB) free!
[1590007632] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19080060 kB) free!
[1590007912] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (248212 kB) free!
[1590007933] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2845336 kB) free!
[1590008032] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232268 kB) free!
[1590008061] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (805768 kB) free!
[1590008152] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236812 kB) free!
[1590008181] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (249340 kB) free!
[1590008232] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407360 kB) free!
[1590008301] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250148 kB) free!
[1590008832] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15702696 kB) free!
[1590009432] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33375876 kB) free.
[1590009501] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 86.5% (42628424 kB) free.
[1590009952] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.4% (45019376 kB) free.
[1590010032] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11819896 kB) free!
[1590010153] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8010040 kB) free!
[1590010273] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4486304 kB) free!
[1590010304] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(353) < 172.105.57.57 [21/May/2020:03:01:32 +0530] TCP 502 0 0 60.000 "10.147.212.49:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1590010424] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590010632] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2537) < 203.212.70.170 [21/May/2020:03:03:32 +0530] TCP 502 0 0 180.000 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1590010710] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.8% (44279008 kB) free.
[1590010752] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590010792] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(451) < 2020/05/21 03:01:32 [error] 22131#0: *89687728 no live upstreams while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590010873] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19389272 kB) free!
[1590010913] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590011473] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393072 kB) free!
[1590012073] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14362108 kB) free!
[1590012673] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31652612 kB) free.
[1590013273] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9894892 kB) free!
[1590013393] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6166012 kB) free!
[1590013513] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2496468 kB) free!
[1590013787] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1590014113] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15634808 kB) free!
[1590014713] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28614076 kB) free.
[1590015133] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1340992 kB) free!
[1590015313] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1749544 kB) free!
[1590015433] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400372 kB) free!
[1590015553] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389860 kB) free!
[1590015733] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (2201540 kB) free!
[1590016153] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13745412 kB) free!
[1590016333] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4090216 kB) free.
[1590016754] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32510196 kB) free.
[1590017310] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.8% (9755680 kB) free.
[1590017354] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6129204 kB) free!
[1590017474] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2071784 kB) free!
[1590017594] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394488 kB) free!
[1590017912] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (218328 kB) free!
[1590018032] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (219456 kB) free!
[1590018152] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185884 kB) free!
[1590018194] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15331652 kB) free!
[1590018794] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33376120 kB) free.
[1590019003] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.6% (18780868 kB) free.
[1590019201] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.4% (10453432 kB) free.
[1590019394] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9468640 kB) free!
[1590019514] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5709804 kB) free!
[1590019595] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 49.6% (12165524 kB) free.
[1590019634] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2107612 kB) free!
[1590020234] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19876012 kB) free!
[1590020834] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35359428 kB) free.
[1590021204] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1590021324] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1590021434] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14152528 kB) free!
[1590021444] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1590021554] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10121552 kB) free!
[1590021674] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38366356 kB) free.
[1590022186] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.65, 6.06, 5.55
[1590022275] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17648500 kB) free!
[1590022307] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.93, 5.25, 5.32
[1590022395] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13496216 kB) free!
[1590022427] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.98, 5.03, 5.22
[1590022515] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9475936 kB) free!
[1590023115] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20753752 kB) free!
[1590023519] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4894812 kB) free!
[1590023589] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1590023640] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4898948 kB) free!
[1590023709] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1590023715] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34796772 kB) free.
[1590023760] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4898880 kB) free!
[1590023795] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.2% (2980424 kB) free!
[1590023803] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192640 kB) free!
[1590023829] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1590023845] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1590023915] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (2858700 kB) free!
[1590023923] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181908 kB) free!
[1590024035] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (2739360 kB) free!
[1590024043] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186236 kB) free!
[1590024218] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3974644 kB) free!
[1590024284] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (2203572 kB) free!
[1590024315] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12727364 kB) free!
[1590024344] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 31.1% (7616000 kB) free.
[1590024429] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1590024435] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8599796 kB) free!
[1590024445] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1590024555] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37111840 kB) free.
[1590024601] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197600 kB) free!
[1590024644] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.3% (4004128 kB) free!
[1590024704] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 71.1% (93706348 kB) free.
[1590024704] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.2% (3975836 kB) free!
[1590024721] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185528 kB) free!
[1590024764] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.3% (3988072 kB) free!
[1590024824] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.3% (3996724 kB) free!
[1590024841] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195080 kB) free!
[1590024842] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 77.2% (18926392 kB) free.
[1590025155] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15492720 kB) free!
[1590025236] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1976624 kB) free!
[1590025275] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (11452328 kB) free!
[1590025395] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7497456 kB) free!
[1590025995] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19706192 kB) free!
[1590026027] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.24, 4.42, 4.84
[1590026595] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33516196 kB) free.
[1590026627] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.07, 5.57, 5.17
[1590026747] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.93, 5.54, 5.21
[1590026790] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.51, 5.41, 5.07
[1590026867] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.12, 4.85, 4.99
[1590026910] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.14, 5.34, 5.09
[1590027030] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.13, 4.94, 4.98
[1590027175] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (505424 kB) free!
[1590027195] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12062776 kB) free!
[1590027315] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8065468 kB) free!
[1590027435] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26104108 kB) free!
[1590027467] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.92, 5.37, 5.24
[1590027587] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.80, 5.23, 5.21
[1590027707] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.64, 5.32, 5.24
[1590027733] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590028307] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1590028307] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.78, 4.72, 4.91
[1590028635] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27163212 kB) free.
[1590028907] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.11, 4.99, 5.08
[1590029027] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.44, 5.22, 5.15
[1590029147] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.81, 5.33, 5.22
[1590029235] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4649368 kB) free!
[1590029355] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27114460 kB) free.
[1590029955] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11848248 kB) free!
[1590030075] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7300364 kB) free!
[1590030195] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3028708 kB) free!
[1590030795] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17020296 kB) free!
[1590031395] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32357640 kB) free.
[1590031830] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.07, 6.58, 5.39
[1590031950] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.26, 8.29, 6.14
[1590031995] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7188024 kB) free!
[1590032070] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.11, 10.12, 7.07
[1590032115] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2493620 kB) free!
[1590032147] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 5.79, 5.07, 4.96
[1590032235] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425332 kB) free!
[1590032747] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.95, 6.84, 5.91
[1590032835] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14625824 kB) free!
[1590032867] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.41, 6.90, 6.04
[1590032987] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.51, 6.71, 6.05
[1590033435] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33831640 kB) free.
[1590033535] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.28352.664> (10.147.212.158:60748 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590033655] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590034035] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8483244 kB) free!
[1590034045] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1590034155] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3876480 kB) free!
[1590034276] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407040 kB) free!
[1590034333] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (823676 kB) free!
[1590034375] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (1728976 kB) free!
[1590034442] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.5% (4531992 kB) free!
[1590034453] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.9% (3717520 kB) free.
[1590034563] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.0% (4178736 kB) free!
[1590034683] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.6% (3832812 kB) free!
[1590034876] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19627704 kB) free!
[1590035053] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3042612 kB) free!
[1590035173] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3433040 kB) free.
[1590035283] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1925576 kB) free!
[1590035476] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410184 kB) free!
[1590035773] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (2906236 kB) free!
[1590035893] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3354396 kB) free.
[1590036270] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 39.78, 24.53, 17.61
[1590036493] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3227264 kB) free!
[1590036614] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.0% (2430264 kB) free!
[1590036712] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/21 10:21:47 [error] 22132#0: *89739015 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590036734] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3005776 kB) free!
[1590036832] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590037674] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 333 processes
[1590037698] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 758 processes
[1590037794] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1590037818] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 757 processes
[1590037938] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 756 processes
[1590038045] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.7% (2624164 kB) free!
[1590038165] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (2021524 kB) free!
[1590038285] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1446196 kB) free!
[1590039036] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (2599036 kB) free!
[1590039270] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.04, 11.41, 18.47
[1590039553] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(445) < 103.255.146.164 [21/May/2020:11:00:42 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1590039636] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (710064 kB) free!
[1590039673] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590040334] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3767136 kB) free.
[1590040432] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 11:14:41 [error] 22135#0: *89748246 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590040552] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590041035] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.4% (6811040 kB) free!
[1590041152] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 11:30:22 [error] 22135#0: *89750811 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590041155] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.3% (6412912 kB) free!
[1590041272] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590041275] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 17.0% (5964848 kB) free!
[1590041301] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9242784 kB) free!
[1590041421] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (8941580 kB) free!
[1590041541] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (8639144 kB) free!
[1590041575] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1606968 kB) free!
[1590042134] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3007940 kB) free!
[1590042175] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (1736696 kB) free!
[1590042184] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 12.3% (3045004 kB) free!
[1590042254] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (2786048 kB) free!
[1590042304] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 10.9% (2691832 kB) free!
[1590042374] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2167684 kB) free!
[1590042424] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2385352 kB) free!
[1590042472] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 11:56:46 [error] 22134#0: *89754738 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590042475] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (1222004 kB) free!
[1590042592] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590042974] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1590042974] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3291128 kB) free.
[1590043075] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.5% (31705352 kB) free.
[1590043086] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.5% (16060760 kB) free.
[1590043273] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(342) < 103.13.97.241 [21/May/2020:12:09:22 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590043276] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18112000 kB) free!
[1590043341] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3911568 kB) free!
[1590043393] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 203.212.70.170 [21/May/2020:12:11:37 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1590043513] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590043552] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8389204 kB) free!
[1590043560] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.6% (2359604 kB) free!
[1590043574] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2872900 kB) free!
[1590043672] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.2% (8004880 kB) free!
[1590043694] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2797972 kB) free!
[1590043792] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.6% (7696584 kB) free!
[1590043815] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2703276 kB) free!
[1590043876] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415484 kB) free!
[1590043975] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1590043975] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1612004 kB) free!
[1590044224] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.1% (22986076 kB) free.
[1590044387] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590044415] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (265468 kB) free!
[1590044476] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20703944 kB) free!
[1590044575] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (1723956 kB) free!
[1590044713] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(523) < 103.255.146.164 [21/May/2020:12:29:36 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590044834] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590044992] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4230484 kB) free!
[1590044992] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 12:32:23 [error] 22124#0: *89759719 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590045015] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3069488 kB) free!
[1590045076] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416256 kB) free!
[1590045112] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590045676] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24630036 kB) free!
[1590046215] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1324912 kB) free!
[1590046276] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (454400 kB) free!
[1590046815] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (2369776 kB) free!
[1590046876] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30524276 kB) free.
[1590047415] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3551644 kB) free.
[1590047476] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (404492 kB) free!
[1590047596] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (375272 kB) free!
[1590047670] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.43, 24.31, 17.75
[1590047716] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389136 kB) free!
[1590048015] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (1521924 kB) free!
[1590048135] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2094948 kB) free!
[1590048255] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2616492 kB) free!
[1590048316] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21115236 kB) free!
[1590048855] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1049060 kB) free!
[1590048916] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448856 kB) free!
[1590049086] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4198976 kB) free!
[1590049206] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3620552 kB) free!
[1590049326] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (3186716 kB) free!
[1590049424] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 3.6% (888404 kB) free!
[1590049516] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33639708 kB) free.
[1590049926] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1590049926] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (647704 kB) free!
[1590050055] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (2796068 kB) free!
[1590050116] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3510288 kB) free!
[1590050237] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386984 kB) free!
[1590050357] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425984 kB) free!
[1590050655] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1417908 kB) free!
[1590050670] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.77, 10.35, 17.32
[1590050833] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(124) < 103.255.146.164 [21/May/2020:14:15:34 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1590050953] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590050957] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22304400 kB) free!
[1590051112] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 14:13:48 [error] 22135#0: *89774908 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590051232] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590051557] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433144 kB) free!
[1590051855] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3169804 kB) free!
[1590052124] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.9% (4387020 kB) free!
[1590052157] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34900984 kB) free.
[1590052456] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167176 kB) free!
[1590052757] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4643680 kB) free!
[1590052877] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404592 kB) free!
[1590052997] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423996 kB) free!
[1590053275] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 18.0% (6293380 kB) free!
[1590053395] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.8% (5900096 kB) free!
[1590053515] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.6% (31718200 kB) free.
[1590053527] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.9% (16141120 kB) free.
[1590053597] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28842048 kB) free.
[1590053632] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/05/21 15:03:17 [error] 22129#0: *89782429 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590053753] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590053824] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.2% (3503856 kB) free!
[1590053945] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.1% (3238220 kB) free!
[1590053955] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(153) < 185.255.8.58 [21/May/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590054065] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.0% (2959092 kB) free!
[1590054076] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590054197] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (406620 kB) free!
[1590054256] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (2107352 kB) free!
[1590054317] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (414548 kB) free!
[1590054437] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432092 kB) free!
[1590054665] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.1% (22984824 kB) free.
[1590054856] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1431504 kB) free!
[1590055037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26692112 kB) free.
[1590055470] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 28.40, 28.60, 21.16
[1590055552] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 15:30:04 [error] 22135#0: *89787058 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590055637] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (420300 kB) free!
[1590055672] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590055758] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406080 kB) free!
[1590055794] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 333 processes
[1590055878] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (421532 kB) free!
[1590055914] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 336 processes
[1590055975] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590056034] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1590056056] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3057268 kB) free!
[1590056070] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1590056070] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.14, 13.12, 17.24
[1590056478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26904944 kB) free.
[1590056656] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161588 kB) free!
[1590057075] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [21/May/2020:15:57:27 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.000, 0.000"
[1590057078] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (423280 kB) free!
[1590057196] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590057198] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (440188 kB) free!
[1590057256] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.4% (2337508 kB) free!
[1590057318] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (449092 kB) free!
[1590057918] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23122644 kB) free!
[1590058518] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409712 kB) free!
[1590058975] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178416 kB) free!
[1590059034] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1590059056] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3382800 kB) free.
[1590059118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33730592 kB) free.
[1590059154] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1590059229] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1590059272] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 16:32:09 [error] 22135#0: *89798193 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590059349] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1590059392] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590059656] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2896696 kB) free!
[1590059718] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2713588 kB) free!
[1590059777] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2862740 kB) free!
[1590059838] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (413412 kB) free!
[1590059845] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1590059897] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (2191244 kB) free!
[1590059958] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406312 kB) free!
[1590060195] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 172.105.59.173 [21/May/2020:16:47:05 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1590060316] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590060445] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1590060497] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (953704 kB) free!
[1590060558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34039672 kB) free.
[1590060566] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1590060686] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1590060915] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(794) < 103.255.146.164 [21/May/2020:17:00:01 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590060960] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 42.7% (10462696 kB) free.
[1590061035] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590061097] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3739740 kB) free.
[1590061158] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2874000 kB) free!
[1590061193] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 17:09:32 [error] 22135#0: *89805168 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590061278] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (436424 kB) free!
[1590061313] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/21 17:10:57 [error] 22135#0: *89805443 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590061327] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2797704 kB) free!
[1590061399] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402740 kB) free!
[1590061416] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 24961 MB (19% inode=99%):
[1590061433] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590061447] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.6% (2369320 kB) free!
[1590061536] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 24957 MB (19% inode=99%):
[1590061567] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1941164 kB) free!
[1590061656] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 24910 MB (19% inode=99%):
[1590061683] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 42.6% (10448520 kB) free.
[1590061697] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2758316 kB) free!
[1590061817] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.0% (3903040 kB) free.
[1590061999] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34119176 kB) free.
[1590062417] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3175316 kB) free!
[1590062537] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4064044 kB) free.
[1590062599] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4481948 kB) free!
[1590062633] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 17:28:15 [error] 22135#0: *89808736 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590062719] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (421960 kB) free!
[1590062753] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590062839] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (488536 kB) free!
[1590063483] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.8% (4618032 kB) free!
[1590063603] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.2% (4206244 kB) free!
[1590063723] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.4% (3783880 kB) free!
[1590063738] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2634064 kB) free!
[1590063858] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.5% (4139304 kB) free.
[1590063870] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.36, 20.19, 16.84
[1590064323] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1683532 kB) free!
[1590064470] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.85, 14.29, 15.87
[1590064567] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16196920 kB) free.
[1590064865] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3910320 kB) free!
[1590064975] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (2766168 kB) free!
[1590064985] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.4% (3552884 kB) free!
[1590065105] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (3207156 kB) free!
[1590065705] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1590065705] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.1% (22989412 kB) free.
[1590066546] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 362 processes
[1590066666] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 362 processes
[1590066786] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 362 processes
[1590067039] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13777152 kB) free!
[1590067386] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1590067639] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399400 kB) free!
[1590068154] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1590068274] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1590068840] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15282268 kB) free!
[1590069258] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2998668 kB) free!
[1590069378] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.6% (5144236 kB) free.
[1590069440] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399672 kB) free!
[1590070338] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1590070578] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (150912 kB) free!
[1590070640] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14547328 kB) free!
[1590070698] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.4% (4125636 kB) free.
[1590070938] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1590071058] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1590071178] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1590071240] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384208 kB) free!
[1590071767] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (3782604 kB) free!
[1590071887] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.2% (3410172 kB) free!
[1590072007] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.6% (3061764 kB) free!
[1590072353] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/21 20:14:36 [error] 22135#0: *89833828 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590072440] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20139044 kB) free!
[1590072473] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590072607] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1257424 kB) free!
[1590072775] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (974040 kB) free!
[1590073040] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374628 kB) free!
[1590073098] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590073674] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1590073794] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1590074256] SERVICE ALERT: sig4-vm2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 14002 MB (10% inode=99%):
[1590075007] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16186592 kB) free.
[1590075270] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.79, 2.92, 4.90
[1590075440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31068304 kB) free.
[1590075456] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 14172 MB (11% inode=99%):
[1590076040] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1458960 kB) free!
[1590076056] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 113653 MB (88% inode=99%):
[1590076160] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384648 kB) free!
[1590076281] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399104 kB) free!
[1590076881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27424232 kB) free.
[1590077070] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.56, 5.69, 5.02
[1590077105] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590077190] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.57, 5.60, 5.06
[1590077298] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.1% (2461892 kB) free!
[1590077310] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.43, 5.50, 5.09
[1590077418] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.6% (4488040 kB) free.
[1590077481] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (386232 kB) free!
[1590077601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382916 kB) free!
[1590077721] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397416 kB) free!
[1590077911] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.27, 4.77, 4.91
[1590077987] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.48, 4.14, 4.77
[1590078321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32031704 kB) free.
[1590078587] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.58, 5.36, 5.12
[1590078608] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1695 processes
[1590078618] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (2609248 kB) free!
[1590078708] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.37, 6.03, 5.39
[1590078738] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176648 kB) free!
[1590078828] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.78, 5.85, 5.40
[1590078858] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (2684788 kB) free!
[1590078921] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (389940 kB) free!
[1590079041] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386480 kB) free!
[1590079161] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378924 kB) free!
[1590079458] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1072160 kB) free!
[1590079761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32144072 kB) free.
[1590079809] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1702 processes
[1590080058] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.2% (4259336 kB) free.
[1590080361] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2177556 kB) free!
[1590080481] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384412 kB) free!
[1590080601] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379032 kB) free!
[1590080628] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.50, 4.66, 5.00
[1590081009] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1699 processes
[1590081258] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (1282624 kB) free!
[1590081378] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3243928 kB) free!
[1590081498] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.7% (4183684 kB) free.
[1590082098] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2814608 kB) free!
[1590082218] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4060628 kB) free.
[1590082818] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.1% (1963012 kB) free!
[1590082938] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3713848 kB) free.
[1590083001] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14490536 kB) free!
[1590083409] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1590083409] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1702 processes
[1590083538] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (556548 kB) free!
[1590083601] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394992 kB) free!
[1590083658] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.8% (2083988 kB) free!
[1590083778] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3624092 kB) free.
[1590084010] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1699 processes
[1590084201] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22583676 kB) free!
[1590084378] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (749268 kB) free!
[1590084498] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.6% (2213796 kB) free!
[1590084607] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1590084619] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3673604 kB) free.
[1590084801] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395604 kB) free!
[1590085401] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30715864 kB) free.