[1590604256] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1345108 kB) free!
[1590604943] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31833412 kB) free.
[1590605456] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3323684 kB) free.
[1590605544] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8391888 kB) free!
[1590605664] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4466208 kB) free!
[1590605784] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (811232 kB) free!
[1590606384] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14595664 kB) free!
[1590606656] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171116 kB) free!
[1590606776] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (160852 kB) free!
[1590606896] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (433788 kB) free!
[1590606984] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30660744 kB) free.
[1590607584] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7805732 kB) free!
[1590607704] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3718564 kB) free!
[1590607824] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402108 kB) free!
[1590607857] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (1489884 kB) free!
[1590607917] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167244 kB) free!
[1590607977] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (171664 kB) free!
[1590608026] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590608037] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (192864 kB) free!
[1590608098] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1678160 kB) free!
[1590608146] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1590609024] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27400076 kB) free.
[1590609238] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 84.2% (20625384 kB) free.
[1590609624] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6390108 kB) free!
[1590609744] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2839556 kB) free!
[1590609864] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30545980 kB) free.
[1590610464] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8739332 kB) free!
[1590610584] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4963080 kB) free!
[1590610705] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1618848 kB) free!
[1590612505] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19293852 kB) free!
[1590613105] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1289192 kB) free!
[1590613498] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176380 kB) free!
[1590613705] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14340160 kB) free!
[1590614098] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3057664 kB) free!
[1590614274] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (44840208 kB) free.
[1590614305] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403860 kB) free!
[1590614423] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.0% (44833580 kB) free.
[1590614698] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3582564 kB) free.
[1590615090] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2544) < 49.50.86.152 [28/May/2020:03:00:37 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590615156] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 27.7% (6778092 kB) free.
[1590615210] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(208) < 203.212.70.170 [28/May/2020:03:01:52 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590615330] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590615368] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(363) < 2020/05/28 03:01:32 [error] 22134#0: *91424419 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590615480] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(295) < 182.18.144.246 [28/May/2020:03:00:25 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590615488] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590615505] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16982636 kB) free!
[1590615600] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590615756] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (241188 kB) free!
[1590615877] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (243180 kB) free!
[1590615997] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (245192 kB) free!
[1590616105] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26901624 kB) free.
[1590616705] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6972560 kB) free!
[1590616825] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3407712 kB) free!
[1590616945] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439824 kB) free!
[1590618145] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17054468 kB) free!
[1590618397] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.3% (10357448 kB) free.
[1590618745] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26832928 kB) free.
[1590618997] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (232328 kB) free!
[1590619117] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (240368 kB) free!
[1590619237] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1590619237] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (274968 kB) free!
[1590619346] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5923576 kB) free!
[1590619466] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2277716 kB) free!
[1590619586] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25549580 kB) free!
[1590620186] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5989904 kB) free!
[1590621386] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23168956 kB) free!
[1590621637] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.1% (9837104 kB) free.
[1590621986] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1843500 kB) free!
[1590622237] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194684 kB) free!
[1590622358] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184432 kB) free!
[1590622478] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174032 kB) free!
[1590623186] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17106228 kB) free!
[1590623643] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 55.6% (13728588 kB) free.
[1590623786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29223620 kB) free.
[1590623924] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 72.7% (17822824 kB) free.
[1590624284] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 56.7% (13894484 kB) free.
[1590624386] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11280628 kB) free!
[1590624506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7800852 kB) free!
[1590624626] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28745000 kB) free.
[1590625226] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10354600 kB) free!
[1590625346] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30145164 kB) free.
[1590625946] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11061556 kB) free!
[1590626067] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7152024 kB) free!
[1590626098] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590626187] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29893568 kB) free.
[1590626787] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10427192 kB) free!
[1590626907] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6603252 kB) free!
[1590627027] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29463560 kB) free.
[1590627284] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.5% (2578628 kB) free!
[1590627404] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (1934476 kB) free!
[1590627524] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1878660 kB) free!
[1590627627] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8623404 kB) free!
[1590627747] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4794560 kB) free!
[1590627867] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27882664 kB) free.
[1590628443] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4477388 kB) free!
[1590628467] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8159044 kB) free!
[1590628563] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (736548 kB) free!
[1590628587] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29107496 kB) free.
[1590628683] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (213364 kB) free!
[1590629187] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10506012 kB) free!
[1590629307] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6655240 kB) free!
[1590629427] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28836360 kB) free.
[1590629638] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 11.9% (2917736 kB) free!
[1590629699] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 32.7% (8015656 kB) free.
[1590630027] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8635704 kB) free!
[1590630147] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4760020 kB) free!
[1590630268] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27349380 kB) free.
[1590630598] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4897504 kB) free!
[1590630658] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4935552 kB) free.
[1590630867] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6981392 kB) free!
[1590630946] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590630987] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28261752 kB) free.
[1590631066] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590631124] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191000 kB) free!
[1590631170] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1590631187] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590631244] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182208 kB) free!
[1590631290] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1590631364] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194988 kB) free!
[1590631410] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1590631587] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8679700 kB) free!
[1590631707] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4644180 kB) free!
[1590631787] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1590631827] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27893360 kB) free.
[1590632010] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1590632427] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7498932 kB) free!
[1590632547] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3475756 kB) free!
[1590632667] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26659940 kB) free.
[1590633267] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4834180 kB) free!
[1590633387] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (381160 kB) free!
[1590633507] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25145792 kB) free!
[1590633878] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590634098] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 68.4% (16755540 kB) free.
[1590634107] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4167528 kB) free!
[1590634258] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4412896 kB) free!
[1590634318] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4407816 kB) free!
[1590634378] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 18.0% (4413156 kB) free!
[1590634438] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 18.0% (4418800 kB) free!
[1590635307] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18141340 kB) free!
[1590636507] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3622212 kB) free!
[1590637240] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.52, 6.68, 5.57
[1590637360] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.73, 7.19, 5.89
[1590637480] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.31, 7.33, 6.11
[1590637707] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19741152 kB) free!
[1590638307] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388676 kB) free!
[1590639507] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25853924 kB) free!
[1590639880] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.12, 2.57, 4.03
[1590640107] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4376740 kB) free!
[1590641307] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20292608 kB) free!
[1590641680] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.44, 5.65, 5.35
[1590641800] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.94, 5.68, 5.40
[1590641898] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.0% (3928532 kB) free!
[1590641907] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380648 kB) free!
[1590641920] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.72, 5.97, 5.53
[1590642018] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.9% (3649068 kB) free!
[1590642139] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.8% (3383448 kB) free!
[1590642587] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1590642707] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1590642739] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1787892 kB) free!
[1590642827] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1590642898] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3183388 kB) free!
[1590643018] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3117032 kB) free!
[1590643107] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23888964 kB) free!
[1590643138] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (2645668 kB) free!
[1590643707] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391656 kB) free!
[1590643738] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3263248 kB) free.
[1590644130] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(774) < 103.255.146.164 [28/May/2020:10:56:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590644185] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.5% (2576588 kB) free!
[1590644251] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590644305] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2162164 kB) free!
[1590644338] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.7% (2230152 kB) free!
[1590644425] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1853140 kB) free!
[1590644458] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2911288 kB) free!
[1590644578] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (2870852 kB) free!
[1590644907] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16066780 kB) free!
[1590645507] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377344 kB) free!
[1590646650] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(306) < 103.255.146.164 [28/May/2020:11:42:54 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1590646707] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28684628 kB) free.
[1590646770] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590647307] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1692496 kB) free!
[1590647426] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.0% (15935524 kB) free.
[1590647427] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (418388 kB) free!
[1590647547] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384776 kB) free!
[1590647578] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1590647578] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1203968 kB) free!
[1590648178] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2590592 kB) free!
[1590648778] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (710108 kB) free!
[1590649947] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29186164 kB) free.
[1590649978] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3416768 kB) free.
[1590650260] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 25239 MB (19% inode=99%):
[1590650275] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9113272 kB) free!
[1590650381] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 24983 MB (19% inode=99%):
[1590650395] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.7% (8744396 kB) free!
[1590650501] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 24271 MB (18% inode=99%):
[1590650515] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.1% (8405380 kB) free!
[1590650547] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2146804 kB) free!
[1590650578] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2955732 kB) free!
[1590650627] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1590650667] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (445216 kB) free!
[1590650698] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2874632 kB) free!
[1590650787] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420836 kB) free!
[1590650818] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2777056 kB) free!
[1590650920] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.55, 26.53, 20.58
[1590651102] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 93566 MB (72% inode=99%):
[1590651418] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3353876 kB) free.
[1590651520] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1590651520] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.16, 11.78, 16.29
[1590651601] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 332 processes
[1590651623] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8747816 kB) free!
[1590651721] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1590651743] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8467432 kB) free!
[1590651863] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (8166664 kB) free!
[1590652018] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (2927468 kB) free!
[1590652138] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2859664 kB) free!
[1590652258] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (2793548 kB) free!
[1590652315] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3732076 kB) free!
[1590653188] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30206248 kB) free.
[1590653458] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (344076 kB) free!
[1590653521] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 339 processes
[1590653527] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.50, 22.10, 17.07
[1590653641] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1590653663] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3670236 kB) free!
[1590653788] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (471788 kB) free!
[1590653908] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (388616 kB) free!
[1590654026] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.0% (3362728 kB) free!
[1590654028] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419860 kB) free!
[1590654058] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2631764 kB) free!
[1590654128] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.10, 15.13, 15.84
[1590654146] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (3001580 kB) free!
[1590654266] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.0% (2460304 kB) free!
[1590654841] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1590654866] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1590654866] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (179132 kB) free!
[1590654961] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1590655258] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1122532 kB) free!
[1590655561] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1590655682] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1590655689] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/28 14:17:54 [error] 22135#0: *91534336 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590655802] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590655809] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/28 14:19:06 [error] 22135#0: *91534663 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590655858] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.8% (2242196 kB) free!
[1590655929] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2020/05/28 14:21:47 [error] 22135#0: *91535453 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590656402] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1590656459] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1062820 kB) free!
[1590656648] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.9615.6110> (10.147.212.53:45418 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1590656769] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590657059] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.2% (2308600 kB) free!
[1590657129] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590657610] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (6087056 kB) free!
[1590657730] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (5706920 kB) free!
[1590657850] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 15.3% (5342048 kB) free!
[1590657968] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.27791.6382> (10.147.212.181:35558 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1590658088] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590658228] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15114992 kB) free!
[1590658451] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.4% (32018872 kB) free.
[1590658467] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.5% (16064968 kB) free.
[1590658770] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(162) < 185.255.8.58 [28/May/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590658828] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388304 kB) free!
[1590658859] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (459340 kB) free!
[1590658891] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590659459] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.1% (2446800 kB) free!
[1590661859] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3402448 kB) free.
[1590662402] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 337 processes
[1590662428] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30989312 kB) free.
[1590662459] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2998548 kB) free!
[1590662522] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 336 processes
[1590662579] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (685004 kB) free!
[1590662642] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1590662699] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1494804 kB) free!
[1590662921] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590663028] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (777700 kB) free!
[1590663148] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (383752 kB) free!
[1590663268] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392024 kB) free!
[1590663299] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.2% (4263760 kB) free.
[1590663842] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1590663868] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (22036992 kB) free!
[1590663900] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.2% (1986760 kB) free!
[1590663962] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1590664020] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.3% (4437812 kB) free.
[1590664468] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (438524 kB) free!
[1590664620] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3073468 kB) free!
[1590664740] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.5% (4954256 kB) free.
[1590665067] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 10.6% (2371760 kB) free!
[1590665068] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25738712 kB) free!
[1590665187] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (2049116 kB) free!
[1590665307] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1580692 kB) free!
[1590665668] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397104 kB) free!
[1590666268] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30023688 kB) free.
[1590666541] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (2573620 kB) free!
[1590666661] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.9% (3889164 kB) free.
[1590666868] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (375056 kB) free!
[1590666962] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1590666988] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (376640 kB) free!
[1590667082] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1590667109] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396396 kB) free!
[1590667709] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17432064 kB) free!
[1590667861] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.0% (2270276 kB) free!
[1590667981] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3451488 kB) free.
[1590668051] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (6907764 kB) free!
[1590668171] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.8% (6598412 kB) free!
[1590668292] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 17.9% (6281824 kB) free!
[1590668309] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420320 kB) free!
[1590668530] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/28 17:51:19 [error] 22135#0: *91596572 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590668581] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (955324 kB) free!
[1590668650] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590668701] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.4% (2500988 kB) free!
[1590668821] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3485340 kB) free.
[1590668909] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19056992 kB) free!
[1590669250] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/28 18:01:09 [error] 22135#0: *91598608 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590669370] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590669421] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.5% (1714676 kB) free!
[1590669482] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1590669492] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.5% (32053208 kB) free.
[1590669507] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.7% (16102532 kB) free.
[1590669509] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420848 kB) free!
[1590669541] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3540912 kB) free.
[1590669602] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1590670109] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19651700 kB) free!
[1590670141] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (816884 kB) free!
[1590670261] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.7% (2387172 kB) free!
[1590670382] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3544372 kB) free.
[1590670709] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406664 kB) free!
[1590671581] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1107696 kB) free!
[1590671701] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2813148 kB) free!
[1590671821] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3771228 kB) free.
[1590672421] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2214736 kB) free!
[1590672509] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14399084 kB) free!
[1590672541] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.2% (3939116 kB) free.
[1590673109] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (429928 kB) free!
[1590673141] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2202552 kB) free!
[1590673261] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3958264 kB) free.
[1590673709] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27011948 kB) free.
[1590673862] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2086808 kB) free!
[1590673982] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3973168 kB) free.
[1590674309] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (470496 kB) free!
[1590674429] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405380 kB) free!
[1590674549] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397688 kB) free!
[1590674582] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2032552 kB) free!
[1590674702] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4189748 kB) free.
[1590675302] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (2641352 kB) free!
[1590675422] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.9% (4374800 kB) free.
[1590676022] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3027772 kB) free!
[1590676142] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.3% (4269424 kB) free.
[1590676349] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17826712 kB) free!
[1590676720] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.14, 4.08, 4.85
[1590676742] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1035064 kB) free!
[1590676863] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (150252 kB) free!
[1590676949] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396200 kB) free!
[1590676983] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1423488 kB) free!
[1590677307] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4179956 kB) free!
[1590677320] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.10, 6.90, 5.93
[1590677427] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (3930960 kB) free!
[1590677440] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.69, 6.10, 5.76
[1590677547] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.9% (3798724 kB) free!
[1590677549] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21726652 kB) free!
[1590677560] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.95, 5.04, 5.42
[1590678149] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419400 kB) free!
[1590678160] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 6.02, 4.50, 4.78
[1590678183] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.7% (4338608 kB) free.
[1590678749] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23095020 kB) free!
[1590678783] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.0% (2280032 kB) free!
[1590678903] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4186216 kB) free.
[1590679347] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2064816 kB) free!
[1590679349] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423268 kB) free!
[1590679503] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2073844 kB) free!
[1590679623] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3849812 kB) free.
[1590679947] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16066732 kB) free.
[1590679949] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24080848 kB) free!
[1590680223] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (1418408 kB) free!
[1590680343] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3373168 kB) free.
[1590680549] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402992 kB) free!
[1590681149] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31848292 kB) free.
[1590681543] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2982608 kB) free!
[1590681664] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4190776 kB) free.
[1590681749] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4130672 kB) free!
[1590681869] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402988 kB) free!
[1590681989] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410568 kB) free!
[1590682263] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (1198836 kB) free!
[1590682383] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (2772876 kB) free!
[1590682503] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4061944 kB) free.
[1590682589] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18247292 kB) free!
[1590683103] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (887844 kB) free!
[1590683190] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (421444 kB) free!
[1590683223] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.0% (2275428 kB) free!
[1590683343] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3668276 kB) free.
[1590683790] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19942972 kB) free!
[1590684390] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390712 kB) free!
[1590684543] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (1333528 kB) free!
[1590684663] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.6% (2699860 kB) free!
[1590684783] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3789124 kB) free.
[1590684990] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19894252 kB) free!
[1590685590] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403160 kB) free!
[1590685983] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (622204 kB) free!
[1590686103] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.4% (1694200 kB) free!
[1590686190] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20000080 kB) free!
[1590686224] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2840620 kB) free!
[1590686790] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403868 kB) free!
[1590686824] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3359448 kB) free.
[1590687390] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18595500 kB) free!
[1590687990] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391748 kB) free!
[1590688590] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17629772 kB) free!
[1590689190] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405000 kB) free!
[1590689224] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2951544 kB) free!
[1590689344] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3395868 kB) free.
[1590689790] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14929268 kB) free!
[1590689944] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.7% (2393776 kB) free!
[1590690064] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3067492 kB) free!
[1590690184] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1489888 kB) free!
[1590690390] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (441592 kB) free!