[1587062012] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8895264 kB) free!
[1587062133] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31770464 kB) free.
[1587062732] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10877436 kB) free!
[1587062852] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7208480 kB) free!
[1587062972] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30690516 kB) free.
[1587063572] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9808976 kB) free!
[1587063693] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31641416 kB) free.
[1587063826] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187612 kB) free!
[1587063886] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182628 kB) free!
[1587063947] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (179356 kB) free!
[1587064010] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (176432 kB) free!
[1587064277] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.8% (2244960 kB) free!
[1587064294] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11436552 kB) free!
[1587064397] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (736592 kB) free!
[1587064414] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7864632 kB) free!
[1587064517] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1134484 kB) free!
[1587064534] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30504000 kB) free.
[1587064610] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 78.5% (19241776 kB) free.
[1587065134] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9383968 kB) free!
[1587065254] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31481512 kB) free.
[1587065854] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11305528 kB) free!
[1587065974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7756864 kB) free!
[1587066094] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30344784 kB) free.
[1587066694] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9662840 kB) free!
[1587066814] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31380080 kB) free.
[1587067240] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (460312 kB) free!
[1587067360] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (559724 kB) free!
[1587067414] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10984988 kB) free!
[1587067440] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239164 kB) free!
[1587067480] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2387580 kB) free!
[1587067534] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7493624 kB) free!
[1587067560] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233776 kB) free!
[1587067654] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29343472 kB) free.
[1587067681] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231996 kB) free!
[1587068125] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587068254] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8876292 kB) free!
[1587068374] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31040700 kB) free.
[1587068974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9990792 kB) free!
[1587069095] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31140480 kB) free.
[1587069324] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(531) < 103.13.97.241 [17/Apr/2020:02:04:07 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1587069444] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(11) < 203.212.70.170 [17/Apr/2020:02:06:06 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1587069481] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.8% (9767004 kB) free!
[1587069564] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587069694] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9666748 kB) free!
[1587069814] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30993384 kB) free.
[1587070414] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10557800 kB) free!
[1587070534] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30965344 kB) free.
[1587071034] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 68.5% (33740584 kB) free.
[1587071134] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11457536 kB) free!
[1587071254] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30951016 kB) free.
[1587071575] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.2% (42960408 kB) free.
[1587071854] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11919656 kB) free!
[1587071974] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587071974] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30947244 kB) free.
[1587072574] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11942444 kB) free!
[1587072674] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(357) < 2020/04/17 03:00:21 [error] 22125#0: *84507053 no live upstreams while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1587072694] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30907112 kB) free.
[1587072794] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(22) < 2020/04/17 03:01:32 [error] 22129#0: *84507403 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1587072904] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(310) < 172.105.57.57 [17/Apr/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1587072914] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587072917] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (2428168 kB) free!
[1587073024] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587073082] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.4% (44527852 kB) free.
[1587073164] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2245) < 203.212.70.170 [17/Apr/2020:03:03:32 +0530] TCP 502 0 0 180.000 "10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1587073285] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587073294] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11322176 kB) free!
[1587073414] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30824872 kB) free.
[1587074014] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11455412 kB) free!
[1587074134] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30764484 kB) free.
[1587074735] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10913608 kB) free!
[1587074855] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30728124 kB) free.
[1587075455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10981376 kB) free!
[1587075575] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30674204 kB) free.
[1587075595] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.7% (9985576 kB) free.
[1587076175] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10737236 kB) free!
[1587076295] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29129348 kB) free.
[1587076517] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (561732 kB) free!
[1587076795] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (249236 kB) free!
[1587076915] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (227756 kB) free!
[1587077035] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (603344 kB) free!
[1587077117] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2784276 kB) free!
[1587077495] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11391836 kB) free!
[1587077615] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29354532 kB) free.
[1587077635] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 37.9% (9283660 kB) free.
[1587078215] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10148720 kB) free!
[1587078335] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28933172 kB) free.
[1587079535] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6064032 kB) free!
[1587079655] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30088908 kB) free.
[1587080035] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182196 kB) free!
[1587080155] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177828 kB) free!
[1587080255] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6036776 kB) free!
[1587080276] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1587080276] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187336 kB) free!
[1587080376] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29594324 kB) free.
[1587080896] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.8% (18831084 kB) free.
[1587080975] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5104380 kB) free!
[1587081095] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28219252 kB) free.
[1587081231] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 74.9% (18350828 kB) free.
[1587081342] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.6% (11749020 kB) free.
[1587082295] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9098060 kB) free!
[1587082415] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29462540 kB) free.
[1587083615] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14179476 kB) free!
[1587083717] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3280024 kB) free.
[1587083735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31137456 kB) free.
[1587084317] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3237104 kB) free!
[1587084437] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3233312 kB) free!
[1587084557] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (25.3% change >= 20.0% threshold)
[1587084557] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (3231192 kB) free!
[1587084935] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13816436 kB) free!
[1587085055] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30754772 kB) free.
[1587085480] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 41.7% (55006184 kB) free.
[1587085542] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (1817504 kB) free!
[1587085662] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193224 kB) free!
[1587085758] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3293248 kB) free.
[1587085782] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192796 kB) free!
[1587086255] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14457772 kB) free!
[1587086358] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3139868 kB) free!
[1587086375] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30587084 kB) free.
[1587086478] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (3161868 kB) free!
[1587086598] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (3157888 kB) free!
[1587087575] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11022616 kB) free!
[1587087695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29340756 kB) free.
[1587089876] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1587090095] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12956448 kB) free!
[1587090215] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29130408 kB) free.
[1587092015] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10916116 kB) free!
[1587092135] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28048672 kB) free.
[1587092557] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1587092631] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1842236 kB) free!
[1587092677] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1587092751] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (1768988 kB) free!
[1587092871] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1700864 kB) free!
[1587093088] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.99, 5.55, 5.15
[1587093208] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.91, 5.17, 5.05
[1587093328] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.13, 4.46, 4.81
[1587093335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9334156 kB) free!
[1587093455] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.5% (24423636 kB) free!
[1587093576] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20110132 kB) free!
[1587094176] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30352684 kB) free.
[1587094529] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.23, 6.32, 5.18
[1587094576] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 366 processes
[1587094649] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.67, 6.29, 5.31
[1587094696] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1587094769] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.18, 6.37, 5.47
[1587097176] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15133184 kB) free!
[1587097296] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34656876 kB) free.
[1587097896] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4173396 kB) free!
[1587097998] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587098016] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33160216 kB) free.
[1587098616] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5131720 kB) free!
[1587098736] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (634280 kB) free!
[1587098856] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32072520 kB) free.
[1587099396] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1587099456] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (9003912 kB) free!
[1587099496] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168816 kB) free!
[1587099504] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1587099517] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1587099576] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31651488 kB) free.
[1587099616] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (221640 kB) free!
[1587099624] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1587099637] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1587099684] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(194) < 103.255.146.164 [17/Apr/2020:10:30:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587099736] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194588 kB) free!
[1587099804] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587100177] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7733056 kB) free!
[1587100224] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 353 processes
[1587100237] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1587100297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30419700 kB) free.
[1587100344] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1587100357] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1587100464] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1587100477] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1587100897] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7060644 kB) free!
[1587100999] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (835212 kB) free!
[1587101005] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(207) < 103.255.146.164 [17/Apr/2020:10:51:36 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587101017] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2883064 kB) free!
[1587101064] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1587101077] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1587101125] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587101137] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25696760 kB) free!
[1587101413] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3415776 kB) free!
[1587101480] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 63.7% (15603104 kB) free.
[1587101533] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2894240 kB) free!
[1587101599] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2624360 kB) free!
[1587101653] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.7% (2395592 kB) free!
[1587101737] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1587101737] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (1955084 kB) free!
[1587102080] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183372 kB) free!
[1587102200] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167088 kB) free!
[1587102253] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (191500 kB) free!
[1587102320] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193908 kB) free!
[1587102920] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 56.0% (13713768 kB) free.
[1587102925] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(93) < 103.255.146.164 [17/Apr/2020:11:15:33 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587102937] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15227220 kB) free!
[1587103045] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587104137] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376644 kB) free!
[1587104599] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (1471588 kB) free!
[1587104664] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1587104784] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1587104845] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(126) < 103.255.146.164 [17/Apr/2020:11:57:00 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587104904] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1587104965] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587105253] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.3% (16232368 kB) free.
[1587105278] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1587105337] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16470192 kB) free!
[1587105398] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1587105458] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 36.3% (8906252 kB) free.
[1587105799] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1587105799] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2169424 kB) free!
[1587106538] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29374516 kB) free.
[1587106765] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(164) < 103.255.146.164 [17/Apr/2020:12:24:35 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587106886] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587107138] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7804112 kB) free!
[1587107258] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3830988 kB) free!
[1587107378] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27718680 kB) free.
[1587107599] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1516452 kB) free!
[1587107634] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9182116 kB) free!
[1587107754] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (9056792 kB) free!
[1587107874] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8944144 kB) free!
[1587107978] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4677428 kB) free!
[1587108085] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(92) < 103.255.146.164 [17/Apr/2020:12:42:17 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587108098] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (752748 kB) free!
[1587108199] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (1732128 kB) free!
[1587108205] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587108218] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25349656 kB) free!
[1587108818] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2775352 kB) free!
[1587109520] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.8% (4602816 kB) free!
[1587109640] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.3% (4479940 kB) free!
[1587109760] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1587109760] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.6% (4308804 kB) free!
[1587110000] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (500160 kB) free!
[1587110018] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19315808 kB) free!
[1587110576] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9314264 kB) free!
[1587110600] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3297968 kB) free.
[1587110605] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(118) < 103.255.146.164 [17/Apr/2020:13:29:19 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587110618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27572504 kB) free.
[1587110696] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9169840 kB) free!
[1587110725] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587110816] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (8995796 kB) free!
[1587111200] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.7% (2065104 kB) free!
[1587111218] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5092784 kB) free!
[1587111314] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/17 13:42:00 [error] 22135#0: *84573534 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/uuid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/uuid"
[1587111320] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.4% (1855772 kB) free!
[1587111325] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(102) < 103.255.146.164 [17/Apr/2020:13:41:54 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587111338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1156372 kB) free!
[1587111434] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587111440] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (1875080 kB) free!
[1587111445] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587111458] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27779932 kB) free.
[1587111474] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4924112 kB) free!
[1587111560] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1844684 kB) free!
[1587112040] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1269020 kB) free!
[1587112058] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4540816 kB) free!
[1587112179] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (575716 kB) free!
[1587112299] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28485532 kB) free.
[1587112640] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3163488 kB) free!
[1587112899] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5248380 kB) free!
[1587113019] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1147920 kB) free!
[1587113139] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28630748 kB) free.
[1587113739] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5906376 kB) free!
[1587113845] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(66) < 103.255.146.164 [17/Apr/2020:14:19:39 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587113859] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1701496 kB) free!
[1587113965] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587113979] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28415728 kB) free.
[1587114253] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (3705272 kB) free!
[1587114373] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3436320 kB) free!
[1587114416] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4267328 kB) free!
[1587114493] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (3188788 kB) free!
[1587114579] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5031176 kB) free!
[1587114699] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (915296 kB) free!
[1587114819] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29053324 kB) free.
[1587115093] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1581004 kB) free!
[1587115165] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(60) < 103.255.146.164 [17/Apr/2020:14:42:29 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587115285] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587115419] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4678852 kB) free!
[1587115539] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (454664 kB) free!
[1587115659] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28911960 kB) free.
[1587115886] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(207) < 185.255.8.58 [17/Apr/2020:15:00:06 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1587116006] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587116259] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3478004 kB) free!
[1587116293] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1587116293] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.8% (16347524 kB) free.
[1587116380] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403420 kB) free!
[1587116500] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29682816 kB) free.
[1587116606] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(60) < 103.255.146.164 [17/Apr/2020:15:11:55 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587116726] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587117100] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5588916 kB) free!
[1587117220] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1603196 kB) free!
[1587117326] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(61) < 103.255.146.164 [17/Apr/2020:15:22:03 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587117340] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29168308 kB) free.
[1587117446] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587117940] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5381136 kB) free!
[1587118060] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (914100 kB) free!
[1587118180] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28982108 kB) free.
[1587118780] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6151456 kB) free!
[1587118900] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1888556 kB) free!
[1587119020] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391208 kB) free!
[1587119680] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.1% (22518820 kB) free!
[1587119800] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.5% (33559520 kB) free.
[1587119840] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1295908 kB) free!
[1587119846] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(74) < 103.255.146.164 [17/Apr/2020:15:58:14 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587119966] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587120196] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1587120316] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1587120440] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.8% (1755524 kB) free!
[1587120566] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(71) < 103.255.146.164 [17/Apr/2020:16:11:43 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587120686] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587120820] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33851356 kB) free.
[1587121420] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (2036532 kB) free!
[1587121541] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400968 kB) free!
[1587121661] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408060 kB) free!
[1587121760] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1587122240] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (887900 kB) free!
[1587122261] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15517536 kB) free!
[1587122486] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(98) < 203.212.70.170 [17/Apr/2020:16:47:10 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1587122606] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587122840] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.4% (2173740 kB) free!
[1587122861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33554888 kB) free.
[1587123440] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (259124 kB) free!
[1587123461] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4332420 kB) free!
[1587123581] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399160 kB) free!
[1587123701] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6444088 kB) free!
[1587124640] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (2066504 kB) free!
[1587124901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28181532 kB) free.
[1587125006] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(185) < 103.255.146.164 [17/Apr/2020:17:24:36 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587125116] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1587125126] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587125236] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1587125293] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (3706696 kB) free!
[1587125413] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3499608 kB) free!
[1587125501] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (2017868 kB) free!
[1587125533] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3267848 kB) free!
[1587125621] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39741568 kB) free.
[1587126133] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1583216 kB) free!
[1587126221] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13026260 kB) free!
[1587126326] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(234) < 103.255.146.164 [17/Apr/2020:17:50:14 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587126341] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8258992 kB) free!
[1587126446] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587126461] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3910824 kB) free!
[1587126733] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.9% (16359264 kB) free.
[1587127061] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16209468 kB) free!
[1587127104] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1587127640] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (640444 kB) free!
[1587127647] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(115) < 103.255.146.164 [17/Apr/2020:18:12:45 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587127662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37319412 kB) free.
[1587127767] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587128240] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3695220 kB) free.
[1587128262] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6071132 kB) free!
[1587128382] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1528884 kB) free!
[1587128502] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401276 kB) free!
[1587129440] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2640632 kB) free!
[1587129560] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3148088 kB) free!
[1587129680] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3170684 kB) free!
[1587129702] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35176152 kB) free.
[1587130000] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25887700 kB) free!
[1587130120] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.5% (19068352 kB) free!
[1587130240] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.3% (33330136 kB) free.
[1587130302] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5984612 kB) free!
[1587130422] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1288640 kB) free!
[1587130542] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407392 kB) free!
[1587130768] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(85) < 103.255.146.164 [17/Apr/2020:19:01:12 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587130888] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587131742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33964848 kB) free.
[1587132080] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (826016 kB) free!
[1587132342] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5180248 kB) free!
[1587132462] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (521544 kB) free!
[1587132582] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375744 kB) free!
[1587132640] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.7% (23310368 kB) free!
[1587132680] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.7% (4016888 kB) free.
[1587132760] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.7% (32536956 kB) free.
[1587133783] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32621548 kB) free.
[1587133880] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2627892 kB) free!
[1587133960] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.8% (23401196 kB) free!
[1587134000] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3452028 kB) free.
[1587134080] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 29.4% (38790516 kB) free.
[1587134383] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5661020 kB) free!
[1587134503] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (964260 kB) free!
[1587134600] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2585484 kB) free!
[1587134623] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37909708 kB) free.
[1587134720] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3364804 kB) free.
[1587135223] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10394708 kB) free!
[1587135320] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.0% (2431200 kB) free!
[1587135343] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5638344 kB) free!
[1587135436] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1587135440] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3159580 kB) free!
[1587135463] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (847628 kB) free!
[1587135557] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1587135560] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3376872 kB) free.
[1587136160] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (2749384 kB) free!
[1587136280] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3192940 kB) free!
[1587136400] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3172800 kB) free!
[1587136663] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17600124 kB) free!
[1587137263] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417444 kB) free!
[1587137969] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.79, 2.69, 3.97
[1587138134] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587138200] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (944728 kB) free!
[1587138463] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33957208 kB) free.
[1587138800] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2934684 kB) free!
[1587139063] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8500152 kB) free!
[1587139183] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4223240 kB) free!
[1587139303] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398064 kB) free!
[1587139903] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13983628 kB) free!
[1587140000] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1382192 kB) free!
[1587140504] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30186200 kB) free.
[1587140600] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2765760 kB) free!
[1587140726] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1668 processes
[1587141104] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5390764 kB) free!
[1587141224] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1507660 kB) free!
[1587141326] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1841 processes
[1587141344] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34916124 kB) free.
[1587141800] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (788400 kB) free!
[1587141944] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8702308 kB) free!
[1587142064] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4485096 kB) free!
[1587142184] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (851680 kB) free!
[1587142784] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16961532 kB) free!
[1587143000] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3443480 kB) free.
[1587143384] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32386388 kB) free.
[1587143984] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8378544 kB) free!
[1587144104] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4591712 kB) free!
[1587144200] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2757852 kB) free!
[1587144224] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21146868 kB) free!
[1587144320] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3222456 kB) free!
[1587144441] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3539012 kB) free.
[1587144824] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (10677544 kB) free!
[1587145424] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (25445300 kB) free!
[1587146024] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4779108 kB) free!
[1587146240] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (2918016 kB) free!
[1587146360] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3053148 kB) free!
[1587146480] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3165976 kB) free!
[1587146625] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16308852 kB) free!
[1587147080] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3747952 kB) free.
[1587147225] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29883700 kB) free.
[1587147825] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7915172 kB) free!
[1587147945] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4361912 kB) free!
[1587148065] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34566276 kB) free.