[1561487489] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1586812 kB) free!
[1561487630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561487721] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.58, 2.30, 4.92
[1561487750] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561488026] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (1251340 kB) free!
[1561488145] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.46, 2.93
[1561488146] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (1243776 kB) free!
[1561488266] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1239524 kB) free!
[1561488350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561488470] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561489070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561489190] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561489790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561489910] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561489959] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.1% (18788372 kB) free.
[1561490030] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.1% (22327584 kB) free.
[1561490097] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.9% (22285528 kB) free.
[1561490511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561490553] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8702120 kB) free!
[1561490631] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561490673] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13053744 kB) free.
[1561491080] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1888640 kB) free!
[1561491141] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 11.4% (2788852 kB) free!
[1561491170] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (807400 kB) free!
[1561491201] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 11.4% (2789012 kB) free!
[1561491231] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561491261] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 11.4% (2788424 kB) free!
[1561491273] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8424996 kB) free!
[1561491290] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (738004 kB) free!
[1561491351] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561491393] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13043488 kB) free.
[1561491410] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (728264 kB) free!
[1561491715] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.6% (18896484 kB) free.
[1561491951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561492071] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561492593] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9873332 kB) free!
[1561492671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561492713] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13055116 kB) free.
[1561492791] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561492791] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561492920] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.1% (3238756 kB) free!
[1561493035] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (1652408 kB) free!
[1561493040] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (1773632 kB) free!
[1561493155] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (785000 kB) free!
[1561493160] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (1003388 kB) free!
[1561493277] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173864 kB) free!
[1561493313] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9729692 kB) free!
[1561493392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561493433] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13046976 kB) free.
[1561493512] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561493575] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561493695] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561493754] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1561494033] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9586912 kB) free!
[1561494112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561494153] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13053076 kB) free.
[1561494232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1561494753] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9438572 kB) free!
[1561494818] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169708 kB) free!
[1561494874] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13052016 kB) free.
[1561494938] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (182492 kB) free!
[1561495058] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (190860 kB) free!
[1561495432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561495473] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9275836 kB) free!
[1561495552] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561495593] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13041548 kB) free.
[1561496152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561496193] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9025220 kB) free!
[1561496273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561496313] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13050752 kB) free.
[1561496872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561496913] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8773596 kB) free!
[1561496992] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561497033] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13040952 kB) free.
[1561497102] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.8% (18209076 kB) free.
[1561497167] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.0% (21812160 kB) free.
[1561497239] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 72.5% (11782036 kB) free.
[1561497592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561497633] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8523500 kB) free!
[1561497706] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (1369312 kB) free!
[1561497712] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561497753] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13051152 kB) free.
[1561497826] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (1076660 kB) free!
[1561497946] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197220 kB) free!
[1561498312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561498377] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4843156 kB) free!
[1561498432] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561498497] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.6% (3096768 kB) free!
[1561498617] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1421748 kB) free!
[1561498953] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9785284 kB) free!
[1561499032] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561499073] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13052940 kB) free.
[1561499153] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561499674] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9715640 kB) free!
[1561499752] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561499794] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13034752 kB) free.
[1561499872] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561500354] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1561500394] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9720380 kB) free!
[1561500472] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561500514] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13043596 kB) free.
[1561500592] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561500988] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 71.9% (17753844 kB) free.
[1561501047] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.6% (21473148 kB) free.
[1561501114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9586524 kB) free!
[1561501192] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561501234] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13030432 kB) free.
[1561501312] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561501598] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4771056 kB) free!
[1561501718] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4752292 kB) free!
[1561501834] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9460572 kB) free!
[1561501838] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (4734544 kB) free!
[1561501913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561501954] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13039992 kB) free.
[1561502033] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561502554] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9700468 kB) free!
[1561502633] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561502674] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13042244 kB) free.
[1561502753] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561503274] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9544640 kB) free!
[1561503353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561503395] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13034340 kB) free.
[1561503473] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561503994] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9404552 kB) free!
[1561504114] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13042768 kB) free.
[1561504673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561504714] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9192720 kB) free!
[1561504794] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561504834] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13030268 kB) free.
[1561504945] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 34.96, 58.16, 33.40
[1561505065] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.88, 38.97, 29.37
[1561505121] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 94.77, 68.44, 32.91
[1561505185] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.97, 26.20, 25.85
[1561505241] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 94.76, 77.14, 40.40
[1561505361] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 95.07, 83.11, 47.05
[1561505393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561505434] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8948940 kB) free!
[1561505513] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561505554] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13033520 kB) free.
[1561505785] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.24, 3.73, 13.66
[1561506113] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561506154] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8419396 kB) free!
[1561506233] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561506274] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12477348 kB) free.
[1561506561] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.01, 3.55, 14.89
[1561506833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561506953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561506985] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.13, 0.36, 4.00
[1561507073] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561507674] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561508074] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8056424 kB) free!
[1561508194] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12166948 kB) free.
[1561508274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561508362] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.48, 2.51, 4.24
[1561508394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561508514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561508858] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1561509114] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1561509114] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561509251] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 289.35, 117.82, 44.16
[1561509371] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 318.95, 185.86, 78.02
[1561509394] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9861640 kB) free!
[1561509491] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 322.68, 231.33, 107.76
[1561509514] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8263196 kB) free!
[1561509562] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 72.57, 17.91, 8.02
[1561509634] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12265440 kB) free.
[1561509682] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 285.02, 117.27, 45.63
[1561509714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561509802] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1561509802] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 314.21, 183.93, 78.74
[1561509834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561509954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561510554] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561511155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561511185] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 292.51, 188.48, 82.70
[1561511275] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561511305] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 193.82, 199.90, 100.55
[1561511395] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561511425] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1561511425] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.42, 133.87, 88.39
[1561511517] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (2964584 kB) free!
[1561511637] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (2942732 kB) free!
[1561511757] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (2920720 kB) free!
[1561512034] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9344536 kB) free!
[1561512038] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561512154] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12815064 kB) free.
[1561512202] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.88, 2.22, 19.83
[1561512491] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 2.36, 10.82
[1561512754] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9289128 kB) free!
[1561512874] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12815412 kB) free.
[1561513225] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 0.89, 14.45
[1561513474] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9164388 kB) free!
[1561513595] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12804780 kB) free.
[1561513691] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.48, 0.55, 3.37
[1561513795] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561514002] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.78, 1.90, 4.50
[1561514194] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9046432 kB) free!
[1561514314] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12810264 kB) free.
[1561514395] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561514425] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.35, 4.22
[1561514516] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561514636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561517157] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1565260 kB) free!
[1561517914] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9229100 kB) free!
[1561518034] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12800612 kB) free.
[1561518236] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561518561] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 34.2% (8377204 kB) free.
[1561518634] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9116716 kB) free!
[1561518754] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12808744 kB) free.
[1561519003] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 561 processes
[1561519063] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 559 processes
[1561519355] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8944424 kB) free!
[1561519475] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12809224 kB) free.
[1561519838] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2469572 kB) free!
[1561520036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561520075] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8677164 kB) free!
[1561520156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561520195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12800904 kB) free.
[1561520756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561520795] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8328212 kB) free!
[1561520876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561520915] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12809468 kB) free.
[1561521476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561521596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561521954] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1561522074] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1561522194] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1561522197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561522317] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561522917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561523037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561523394] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1561523637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561523757] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561524182] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 356 processes
[1561524202] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561524302] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 2 warn out of 358 processes
[1561524422] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1561524487] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (314640 kB) free!
[1561524607] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178608 kB) free!
[1561524727] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155632 kB) free!
[1561525225] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561525558] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561525678] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561526257] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1561526878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561526915] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8604652 kB) free!
[1561526998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561527035] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12789484 kB) free.
[1561527108] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.99, 6.41, 5.52
[1561527228] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.87, 6.10, 5.50
[1561527348] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.64, 5.99, 5.54
[1561527598] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1561527848] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168708 kB) free!
[1561527895] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1561527968] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173112 kB) free!
[1561528015] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561528088] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168100 kB) free!
[1561528198] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561528235] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8788988 kB) free!
[1561528318] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561528355] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12795152 kB) free.
[1561528558] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 31.7% (5154668 kB) free.
[1561528615] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1561528735] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1561528919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561528955] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8656224 kB) free!
[1561529039] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561529075] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12796996 kB) free.
[1561529394] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 10 users currently logged in
[1561529514] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 10 users currently logged in
[1561529634] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1561529639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561529675] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8591452 kB) free!
[1561529758] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2823244 kB) free!
[1561529759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561529796] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12786328 kB) free.
[1561529858] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4622868 kB) free!
[1561529878] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.9% (2581680 kB) free!
[1561529978] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4530020 kB) free!
[1561529998] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.4% (2335760 kB) free!
[1561530098] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4465908 kB) free!
[1561530359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561530395] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8398564 kB) free!
[1561530479] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561530515] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12794344 kB) free.
[1561530598] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1068892 kB) free!
[1561531079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561531115] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8434096 kB) free!
[1561531199] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561531235] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12784932 kB) free.
[1561531735] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1561531800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561531835] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8991324 kB) free!
[1561531855] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1561531898] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2397312 kB) free!
[1561531920] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561531955] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12790072 kB) free.
[1561532520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561532555] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9009072 kB) free!
[1561532602] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 35.21, 27.59, 14.69
[1561532640] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561532675] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12791860 kB) free.
[1561532689] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5226 MB (10% inode=97%):
[1561532722] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 34.05, 29.72, 17.03
[1561532842] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.87, 21.99, 15.81
[1561533098] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (2522392 kB) free!
[1561533240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561533275] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8451812 kB) free!
[1561533289] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6941 MB (13% inode=98%):
[1561533360] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561533395] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12784236 kB) free.
[1561533698] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1561533698] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2073160 kB) free!
[1561533719] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.6% (17914716 kB) free.
[1561533819] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.2% (21620204 kB) free.
[1561533927] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1561533927] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.1% (14652224 kB) free.
[1561533960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561533995] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8219760 kB) free!
[1561534025] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.4% (14859380 kB) free.
[1561534080] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561534116] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12789324 kB) free.
[1561534126] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 71.8% (11673492 kB) free.
[1561534329] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.0% (2723068 kB) free!
[1561534449] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (1476464 kB) free!
[1561534569] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1273612 kB) free!
[1561534642] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.81, 1.85, 3.89
[1561534681] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561534801] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561535089] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5553 MB (10% inode=97%):
[1561535225] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (1564256 kB) free!
[1561535242] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 307.99, 264.98, 146.54
[1561535315] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9580464 kB) free!
[1561535345] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (526192 kB) free!
[1561535362] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 46.88, 181.43, 130.47
[1561535370] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 678 processes
[1561535401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561535435] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12781204 kB) free.
[1561535465] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174076 kB) free!
[1561535482] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.12, 122.03, 114.88
[1561535490] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 679 processes
[1561535521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561535610] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 370 processes
[1561535689] SERVICE FLAPPING ALERT: pservercluster2;HDD_Root status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1561535689] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6940 MB (13% inode=98%):
[1561536035] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9687532 kB) free!
[1561536122] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561536155] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12677684 kB) free.
[1561536242] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561536755] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9401360 kB) free!
[1561536842] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561536875] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12782444 kB) free.
[1561536962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561537082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561537283] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1561537283] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.16, 2.38, 18.36
[1561537548] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 1.51, 2.16, 4.19
[1561538282] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561538675] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9727420 kB) free!
[1561538795] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8281160 kB) free!
[1561538891] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 324.75, 236.02, 111.90
[1561538916] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12776104 kB) free.
[1561539011] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 325.17, 265.48, 137.74
[1561539083] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.90, 1.79, 4.28
[1561539132] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 68.57, 197.28, 129.07
[1561539392] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.3% (14832864 kB) free.
[1561539428] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.3% (18331196 kB) free.
[1561539482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561539515] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8218816 kB) free!
[1561539602] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561539635] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12764356 kB) free.
[1561540202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561540235] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8498068 kB) free!
[1561540322] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561540355] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12770500 kB) free.
[1561540638] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.0% (2950256 kB) free!
[1561540730] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (1470968 kB) free!
[1561540758] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (826620 kB) free!
[1561540825] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14.76, 68.54, 39.44
[1561540850] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (987392 kB) free!
[1561540878] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1561540878] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190888 kB) free!
[1561540922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561540945] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.04, 46.24, 34.80
[1561540955] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8554960 kB) free!
[1561540970] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (405972 kB) free!
[1561541042] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561541065] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.84, 31.11, 30.64
[1561541075] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12775760 kB) free.
[1561541138] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1561541148] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.76, 9.50, 7.11
[1561541197] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (640732 kB) free!
[1561541268] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.12, 9.19, 7.27
[1561541317] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164876 kB) free!
[1561541388] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.08, 8.49, 7.25
[1561541437] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1561541437] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160704 kB) free!
[1561541532] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.79, 3.87, 15.64
[1561541571] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.1% (6697212 kB) free.
[1561541642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561541665] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 4.49, 16.23
[1561541675] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8378012 kB) free!
[1561541738] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (3035076 kB) free!
[1561541763] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561541795] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12763860 kB) free.
[1561541858] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.9% (2744480 kB) free!
[1561541978] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.0% (2442336 kB) free!
[1561542229] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3914128 kB) free!
[1561542349] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.5% (3806076 kB) free!
[1561542362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561542395] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8214284 kB) free!
[1561542469] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (3740440 kB) free!
[1561542482] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561542515] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12772768 kB) free.
[1561542578] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1304548 kB) free!
[1561542865] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.49, 4.78
[1561543057] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.5% (14865100 kB) free.
[1561543082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561543083] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 78.1% (19281544 kB) free.
[1561543115] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8374680 kB) free!
[1561543116] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.3% (21395992 kB) free.
[1561543159] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1561543159] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.3% (14998860 kB) free.
[1561543188] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 1.99, 2.35, 4.24
[1561543202] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1561543210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 35.9% (17683024 kB) free.
[1561543322] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561543333] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.63, 3.02, 4.92
[1561543688] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (233780 kB) free!
[1561543808] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174428 kB) free!
[1561543928] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182236 kB) free!
[1561544522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561544634] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1561544643] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561544822] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1561544882] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1561544942] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1561545002] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1561545525] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4422672 kB) free!
[1561545588] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.63, 10.28, 7.23
[1561545619] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9852128 kB) free!
[1561545645] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (4218844 kB) free!
[1561545708] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.22, 9.77, 7.43
[1561545739] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.5% (13064800 kB) free.
[1561545765] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (4011708 kB) free!
[1561545828] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1561545828] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.79, 8.57, 7.26
[1561546068] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (184428 kB) free!
[1561546189] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174860 kB) free!
[1561546309] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181264 kB) free!
[1561546442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561546442] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2132240 kB) free!
[1561546562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561546562] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 10.5% (1702012 kB) free!
[1561546682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561546682] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1380392 kB) free!
[1561546855] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1561546928] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 378 processes
[1561546933] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.98, 7.92, 7.05
[1561546975] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1561547048] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 373 processes
[1561547053] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.98, 6.74, 6.74
[1561547089] SERVICE FLAPPING ALERT: pservercluster2;HDD_Root status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561547168] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 371 processes
[1561547175] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1561547175] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.01, 5.70, 6.34
[1561547565] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2068184 kB) free!
[1561547571] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (1364080 kB) free!
[1561547691] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.0% (1104084 kB) free!
[1561547702] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 59 users currently logged in
[1561547769] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1561547811] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.6% (852588 kB) free!
[1561548211] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.0% (14952364 kB) free.
[1561548249] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.0% (14470224 kB) free.
[1561548288] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1561548288] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6837756 kB) free.
[1561548291] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 377 processes
[1561548361] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.7% (21977152 kB) free.
[1561548375] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.51, 2.94, 4.01
[1561548377] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.1% (18280376 kB) free.
[1561548412] SERVICE ALERT: sig3-vm2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1561548573] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 13.6% (2204208 kB) free!
[1561548693] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 21.8% (3539984 kB) free.
[1561548739] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8331756 kB) free!
[1561548816] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (222228 kB) free!
[1561548860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12275652 kB) free.
[1561548936] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (153344 kB) free!
[1561548988] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196268 kB) free!
[1561549056] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (177696 kB) free!
[1561549108] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195296 kB) free!
[1561549228] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171256 kB) free!
[1561549283] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561549293] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 13.4% (2182964 kB) free!
[1561549413] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 11.0% (1787256 kB) free!
[1561549533] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1410256 kB) free!
[1561549566] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (1000916 kB) free!
[1561549575] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.55, 8.46, 5.64
[1561549686] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (821556 kB) free!
[1561549695] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.64, 6.59, 5.30
[1561549806] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (595756 kB) free!
[1561549815] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.72, 5.24, 4.96
[1561550415] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.01, 6.04, 5.31
[1561550535] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.90, 4.90, 4.98
[1561550610] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 817 processes
[1561550659] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8206652 kB) free!
[1561550697] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (255148 kB) free!
[1561550730] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 380 processes
[1561550779] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12283908 kB) free.
[1561550817] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (126568 kB) free!
[1561550937] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (117152 kB) free!
[1561551135] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.52, 8.98, 6.32
[1561551255] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.72, 6.88, 5.87
[1561551375] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.10, 5.57, 5.51
[1561551860] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 10.7% (1747004 kB) free!
[1561551975] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.72, 3.32, 4.33
[1561551979] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9859836 kB) free!
[1561551980] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (1261764 kB) free!
[1561552099] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8269588 kB) free!
[1561552100] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (828672 kB) free!
[1561552219] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12179752 kB) free.
[1561552700] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1561552700] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.6% (14242264 kB) free.
[1561552737] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.3% (6711192 kB) free.
[1561553107] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4630152 kB) free!
[1561553227] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4492956 kB) free!
[1561553347] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (4358516 kB) free!
[1561553419] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9166484 kB) free!
[1561553539] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12954584 kB) free.
[1561554133] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.8% (14600188 kB) free.
[1561554150] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 61.2% (15090520 kB) free.
[1561554171] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.8% (22496792 kB) free.
[1561554230] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.4% (14524920 kB) free.
[1561554248] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.9% (22045840 kB) free.
[1561554482] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561554739] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.9% (2745672 kB) free!
[1561554754] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193912 kB) free!
[1561554859] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (467800 kB) free!
[1561554875] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175132 kB) free!
[1561554979] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166068 kB) free!
[1561554995] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163592 kB) free!
[1561555082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561555202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561555322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561555428] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1561555939] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8585056 kB) free!
[1561556059] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12479076 kB) free.
[1561556522] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1561556522] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1561556969] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (11298216 kB) free!
[1561557089] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (18776924 kB) free.
[1561557123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561557243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561557363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561557859] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9099380 kB) free!
[1561557963] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561557979] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12899280 kB) free.
[1561558137] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (349956 kB) free!
[1561558257] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (129104 kB) free!
[1561558377] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (131592 kB) free!
[1561558563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561558683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561558803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561559034] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.1% (2784068 kB) free!
[1561559154] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.8% (2405864 kB) free!
[1561559274] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.6% (2050032 kB) free!
[1561559403] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561559577] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 90.8% (6674908 kB) free.
[1561559779] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9017740 kB) free!
[1561559874] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (574928 kB) free!
[1561559900] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12869704 kB) free.
[1561559900] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2582044 kB) free!
[1561560004] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561560021] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 13.9% (2259556 kB) free!
[1561560122] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.90, 5.44, 5.03
[1561560124] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561560141] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.3% (2003424 kB) free!
[1561560242] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.19, 5.43, 5.08
[1561560244] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561560362] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.09, 5.29, 5.07
[1561560741] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (719544 kB) free!
[1561560777] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1561560962] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 4.22, 4.34, 4.69
[1561561022] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1561561657] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1561562257] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1561563228] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.53, 3.03, 4.14
[1561563376] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561563499] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9094220 kB) free!
[1561563619] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12855276 kB) free.
[1561563844] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561564219] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9006056 kB) free!
[1561564339] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12862216 kB) free.
[1561564939] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9110800 kB) free!
[1561565060] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12864108 kB) free.
[1561565644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561565659] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9231924 kB) free!
[1561565764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561565779] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12856052 kB) free.
[1561565779] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561565795] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561566364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561566379] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9282224 kB) free!
[1561566484] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561566499] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12862028 kB) free.
[1561567084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561567099] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9293040 kB) free!
[1561567204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561567219] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12853760 kB) free.
[1561567805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561567819] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9267808 kB) free!
[1561567925] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561567939] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12862648 kB) free.
[1561568525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561568539] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9252776 kB) free!
[1561568645] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1561568659] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12863648 kB) free.
[1561568777] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.6% (22348520 kB) free.
[1561568864] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.9% (14280624 kB) free.
[1561569012] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 85.1% (13827888 kB) free.
[1561569133] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.6% (14888032 kB) free.
[1561569245] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561569259] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9230744 kB) free!
[1561569365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561569379] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12846456 kB) free.
[1561569965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561569980] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9245336 kB) free!
[1561570085] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561570100] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12849892 kB) free.
[1561570377] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561570686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561570700] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9224592 kB) free!
[1561570704] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1561570806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561570820] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12842180 kB) free.
[1561570824] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1561571406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561571420] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9129772 kB) free!
[1561571526] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561571540] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12853868 kB) free.
[1561572126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561572140] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9026940 kB) free!
[1561572246] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561572260] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12852604 kB) free.
[1561572846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561572860] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9035148 kB) free!
[1561572966] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561572980] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12844184 kB) free.
[1561573567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561573580] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9477524 kB) free!
[1561573687] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1561573687] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561573700] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12849040 kB) free.