[1551810623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551811552] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.0% (9386884 kB) free!
[1551811823] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551812370] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1786500 kB) free!
[1551812752] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2759372 kB) free!
[1551812886] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.44, 4.79, 4.96
[1551812970] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6365236 kB) free!
[1551813023] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551813143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551813263] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1551813263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551813352] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.8% (9281056 kB) free!
[1551813570] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1551813570] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9931304 kB) free.
[1551814552] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2915992 kB) free!
[1551814686] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.42, 5.69, 5.14
[1551814770] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5103444 kB) free!
[1551814806] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.10, 5.84, 5.26
[1551814890] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (3660620 kB) free!
[1551814926] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.73, 5.79, 5.31
[1551815010] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9929952 kB) free.
[1551815063] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551815152] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.8% (9764744 kB) free!
[1551815611] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4935616 kB) free!
[1551815663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551815731] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (3370492 kB) free!
[1551815783] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551815851] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9942764 kB) free.
[1551816352] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3806164 kB) free!
[1551816383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551816451] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4689816 kB) free!
[1551816503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1551816571] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2977328 kB) free!
[1551816623] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551816691] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9924016 kB) free.
[1551816726] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.02, 4.31, 4.84
[1551816952] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.9% (10788372 kB) free.
[1551817223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551817291] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4478764 kB) free!
[1551817344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551817411] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9943436 kB) free.
[1551817464] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1551818011] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5574048 kB) free!
[1551818064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551818131] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4272772 kB) free!
[1551818152] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9193632 kB) free!
[1551818184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551818251] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9927768 kB) free.
[1551818272] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6815008 kB) free!
[1551818304] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551818392] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4122008 kB) free!
[1551818851] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5359704 kB) free!
[1551818904] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551818972] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (4067388 kB) free!
[1551818992] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.9% (10793348 kB) free.
[1551819024] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551819092] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9935128 kB) free.
[1551819144] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551819592] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8194508 kB) free!
[1551819692] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5142916 kB) free!
[1551819712] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.1% (5455136 kB) free!
[1551819744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551819812] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3718060 kB) free!
[1551819832] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2753844 kB) free!
[1551819865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551819932] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9912464 kB) free.
[1551819985] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551820326] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.00, 5.93, 5.45
[1551820432] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.1% (8898272 kB) free!
[1551820446] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.93, 5.90, 5.49
[1551820532] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4889920 kB) free!
[1551820566] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.37, 6.10, 5.61
[1551820585] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551820652] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (3286708 kB) free!
[1551820705] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551820772] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9930860 kB) free.
[1551821305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551821372] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4689224 kB) free!
[1551821425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551821492] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (2994380 kB) free!
[1551821545] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551821612] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9915156 kB) free.
[1551821632] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10821976 kB) free.
[1551821766] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.46, 4.50, 4.88
[1551822145] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551822213] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4520472 kB) free!
[1551822232] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8370800 kB) free!
[1551822266] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551822333] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 20.0% (9867176 kB) free!
[1551822352] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.0% (5919700 kB) free!
[1551822386] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551822453] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9928872 kB) free.
[1551822473] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3494528 kB) free!
[1551822986] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551823053] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4369764 kB) free!
[1551823073] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10040276 kB) free.
[1551823106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551823173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9916156 kB) free.
[1551823226] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551823673] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.8% (6790956 kB) free!
[1551823773] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5456800 kB) free!
[1551823793] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (4104876 kB) free!
[1551823826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551823893] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4172300 kB) free!
[1551823913] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10829744 kB) free.
[1551823946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551824013] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9931488 kB) free.
[1551824066] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551824513] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8417372 kB) free!
[1551824613] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.7% (5289196 kB) free!
[1551824633] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.7% (5783204 kB) free!
[1551824667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551824733] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (4004264 kB) free!
[1551824753] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3186384 kB) free!
[1551824787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551824853] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9906972 kB) free.
[1551824907] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551825366] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.19, 5.93, 5.45
[1551825453] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5144648 kB) free!
[1551825486] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.76, 5.89, 5.50
[1551825507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551825574] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (3769544 kB) free!
[1551825606] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.04, 5.92, 5.55
[1551825627] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551825694] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9909528 kB) free.
[1551825747] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551826294] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4977188 kB) free!
[1551826347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551826414] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3474444 kB) free!
[1551826467] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551826534] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9899856 kB) free.
[1551826553] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.3% (6074196 kB) free!
[1551826806] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.42, 4.30, 4.71
[1551827067] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551827134] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4824096 kB) free!
[1551827188] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551827254] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (3190564 kB) free!
[1551827308] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551827374] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9915272 kB) free.
[1551827908] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551827974] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4572348 kB) free!
[1551828028] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551828094] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (2737512 kB) free!
[1551828148] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551828214] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9905724 kB) free.
[1551828354] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10848088 kB) free.
[1551828748] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551828815] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4396848 kB) free!
[1551828868] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551828935] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551828935] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9914956 kB) free.
[1551828988] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1551828988] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551829535] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5572384 kB) free!
[1551829588] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551829655] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4267936 kB) free!
[1551829709] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551829775] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9906068 kB) free.
[1551829829] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551830154] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (3321008 kB) free!
[1551830274] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.1% (10880036 kB) free.
[1551830375] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5415240 kB) free!
[1551830406] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.11, 5.88, 5.41
[1551830429] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551830495] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (4137168 kB) free!
[1551830526] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.42, 5.98, 5.50
[1551830549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551830615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9913684 kB) free.
[1551830646] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.31, 6.05, 5.58
[1551830669] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551831215] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5221204 kB) free!
[1551831269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551831335] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (3911608 kB) free!
[1551831389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551831455] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9900792 kB) free.
[1551831474] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3760484 kB) free!
[1551831509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551831594] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.1% (10904668 kB) free.
[1551831846] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.09, 4.23, 4.73
[1551832055] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5082480 kB) free!
[1551832110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551832176] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3686172 kB) free!
[1551832230] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551832296] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9914276 kB) free.
[1551832794] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3689984 kB) free!
[1551832830] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551832896] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4874608 kB) free!
[1551832914] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10239104 kB) free.
[1551832950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551833016] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (3320540 kB) free!
[1551833070] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551833136] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9895776 kB) free.
[1551833514] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.5% (6147476 kB) free!
[1551833634] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (3432460 kB) free!
[1551833670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551833736] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4598400 kB) free!
[1551833755] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10264808 kB) free.
[1551833790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551833856] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2860848 kB) free!
[1551833910] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551833976] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9915660 kB) free.
[1551834354] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.9% (5878364 kB) free!
[1551834474] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3411024 kB) free!
[1551834511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551834534] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 556 processes
[1551834576] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4381520 kB) free!
[1551834594] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.2% (10947276 kB) free.
[1551834595] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 559 processes
[1551834631] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551834696] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9899100 kB) free.
[1551834751] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551835194] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.8% (6809340 kB) free!
[1551835296] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5500608 kB) free!
[1551835314] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4318204 kB) free!
[1551835351] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551835416] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (4226700 kB) free!
[1551835434] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10847688 kB) free.
[1551835446] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.79, 5.85, 5.31
[1551835471] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1551835537] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9910792 kB) free.
[1551835566] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.69, 5.81, 5.36
[1551835591] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551835686] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.92, 5.89, 5.44
[1551836034] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7652072 kB) free!
[1551836136] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5338596 kB) free!
[1551836154] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.3% (5072816 kB) free!
[1551836191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551836256] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (4046988 kB) free!
[1551836274] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2492092 kB) free!
[1551836311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551836376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9900156 kB) free.
[1551836431] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551836874] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.8% (8762796 kB) free!
[1551836886] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.63, 4.34, 4.64
[1551836976] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5166928 kB) free!
[1551837032] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551837096] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (3857384 kB) free!
[1551837152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551837216] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9912424 kB) free.
[1551837272] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1551837816] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3724920 kB) free!
[1551837872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551837936] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.9% (9805156 kB) free!
[1551837992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551838056] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.0% (7877664 kB) free!
[1551838074] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.7% (10685656 kB) free.
[1551838112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551838674] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7519808 kB) free!
[1551838795] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.2% (5021556 kB) free!
[1551838915] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.6% (11138576 kB) free.
[1551839515] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.2% (7986964 kB) free!
[1551839579] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 84.18, 21.58, 7.76
[1551839635] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.9% (5371212 kB) free!
[1551839699] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 32.10, 30.29, 13.24
[1551839755] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11177600 kB) free.
[1551839819] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.54, 20.35, 11.66
[1551839856] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9907628 kB) free.
[1551839886] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.33, 5.81, 5.22
[1551839913] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551840006] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.98, 5.83, 5.30
[1551840126] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.95, 5.85, 5.37
[1551840355] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7404564 kB) free!
[1551840475] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4778312 kB) free!
[1551840595] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.9% (10802248 kB) free.
[1551841019] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.62, 0.77, 3.51
[1551841056] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3069544 kB) free!
[1551841176] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9905564 kB) free.
[1551841195] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.5% (6641636 kB) free!
[1551841315] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.6% (11127128 kB) free.
[1551841713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1551841776] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4209556 kB) free!
[1551841833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551841896] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9902840 kB) free.
[1551841915] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8519660 kB) free!
[1551841927] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.40, 4.42, 4.75
[1551841953] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1551842035] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5952932 kB) free!
[1551842155] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11183288 kB) free.
[1551842496] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3481992 kB) free!
[1551842616] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9646512 kB) free!
[1551842736] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.5% (7653360 kB) free!
[1551842756] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7658092 kB) free!
[1551842876] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4951600 kB) free!
[1551842996] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10846848 kB) free.
[1551843596] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.8% (6809048 kB) free!
[1551843716] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (4052448 kB) free!
[1551843754] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551843836] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10267084 kB) free.
[1551843874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1551843994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551844327] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.39, 5.64, 5.39
[1551844436] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.8% (5833172 kB) free!
[1551844447] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.17, 5.52, 5.38
[1551844556] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.4% (11030104 kB) free.
[1551844567] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.81, 5.58, 5.42
[1551845156] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.1% (6925316 kB) free!
[1551845276] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (4115772 kB) free!
[1551845396] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10382484 kB) free.
[1551845997] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229108 kB) free!
[1551846117] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.9% (5867988 kB) free!
[1551846237] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (1630196 kB) free!
[1551846473] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 19.4% (1429296 kB) free!
[1551846593] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 19.4% (1423336 kB) free!
[1551846713] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 19.2% (1413176 kB) free!
[1551847438] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.5% (7631268 kB) free!
[1551848038] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2198388 kB) free!
[1551849295] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 588 processes
[1551849355] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 599 processes
[1551849977] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 76.5% (5627964 kB) free.
[1551851513] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (699036 kB) free!
[1551851599] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (1418256 kB) free!
[1551851720] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.1% (1405112 kB) free!
[1551851840] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.9% (1389620 kB) free!
[1551852525] SERVICE ALERT: sig5-vm1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1551852555] HOST ALERT: sig5-vm1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1551852575] HOST ALERT: sig5-vm1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1551852646] SERVICE ALERT: sig5-vm1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 103292 MB (86% inode=96%):
[1551854167] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1551854277] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 72.9% (5359948 kB) free.
[1551856040] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 94.8% (6975348 kB) free.
[1551856549] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46736 MB (20% inode=99%):
[1551856669] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47250 MB (21% inode=99%):
[1551857269] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47076 MB (20% inode=99%):
[1551857389] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46943 MB (20% inode=99%):
[1551857509] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48882 MB (21% inode=99%):
[1551857571] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.8% (1748660 kB) free!
[1551857991] SERVICE ALERT: sig2-vm1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2-vm1' was out of bounds)
[1551858111] SERVICE ALERT: sig2-vm1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1551858172] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1317092 kB) free!
[1551858315] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 93.9% (6904864 kB) free.
[1551858917] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 38.6% (6270820 kB) free.
[1551858936] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4902264 kB) free!
[1551858967] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.68, 4.53, 4.82
[1551859438] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551860167] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.57, 6.30, 5.63
[1551860287] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.15, 6.63, 5.83
[1551860407] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.90, 6.73, 5.97
[1551860509] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45805 MB (20% inode=99%):
[1551860629] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45768 MB (20% inode=99%):
[1551860750] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45420 MB (20% inode=99%):
[1551861172] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.2% (2466816 kB) free!
[1551861772] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1551861772] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1255568 kB) free!
[1551863622] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1551863682] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1551863742] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1551863802] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1551864238] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.0% (13318400 kB) free.
[1551864838] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8896804 kB) free!
[1551864873] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 48.7% (24009632 kB) free.
[1551864922] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 49.6% (24428192 kB) free.
[1551864967] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21986712 kB) free.
[1551865000] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.3% (22857340 kB) free.
[1551865036] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.2% (7004948 kB) free.
[1551865070] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1551865080] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.8% (22985488 kB) free.
[1551865130] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.0% (14785424 kB) free.
[1551865170] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.6% (22198056 kB) free.
[1551865190] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1551865302] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1551866150] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48093 MB (21% inode=99%):
[1551866502] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1551866562] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1551866622] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1551866682] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1551866733] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (1297592 kB) free!
[1551866750] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45943 MB (20% inode=99%):
[1551866853] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228104 kB) free!
[1551866870] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45911 MB (20% inode=99%):
[1551866973] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.7% (9216856 kB) free!
[1551866990] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52365 MB (23% inode=99%):
[1551867049] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 50.1% (24704956 kB) free.
[1551867282] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1551867282] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1551867377] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (1354492 kB) free!
[1551867497] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (1320268 kB) free!
[1551867617] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 17.5% (1287656 kB) free!
[1551867917] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.5% (2525096 kB) free!
[1551868005] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177612 kB) free!
[1551868037] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.6% (2377532 kB) free!
[1551868125] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173432 kB) free!
[1551868133] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (177872 kB) free!
[1551868157] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.7% (2223712 kB) free!
[1551868245] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175156 kB) free!
[1551868253] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8351192 kB) free!
[1551868253] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 11.2% (1826848 kB) free!
[1551868373] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 34.2% (16849556 kB) free.
[1551868373] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.6% (1718852 kB) free!
[1551868757] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1422732 kB) free!
[1551868973] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1032528 kB) free!
[1551869382] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 63 users currently logged in
[1551869442] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 63 users currently logged in
[1551869502] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 63 users currently logged in
[1551869562] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1551869573] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1778888 kB) free!
[1551869693] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227300 kB) free!
[1551869813] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1551869813] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.4% (8080356 kB) free!
[1551869957] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1551869957] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 44.0% (7144316 kB) free.
[1551870017] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (646776 kB) free!
[1551870413] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4092568 kB) free!
[1551872155] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 656 processes
[1551872215] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 658 processes
[1551872276] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 2 warn out of 653 processes
[1551872336] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 651 processes
[1551872813] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.9% (9826600 kB) free!
[1551872877] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 19.1% (1402772 kB) free!
[1551872997] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 18.8% (1380700 kB) free!
[1551873117] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 18.6% (1368044 kB) free!
[1551873462] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1551874613] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3672124 kB) free!
[1551874817] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 15.7% (1155200 kB) free!
[1551875813] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9847312 kB) free!
[1551877062] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 59 users currently logged in
[1551877613] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3644848 kB) free!
[1551878709] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22285 MB (10% inode=99%):
[1551879309] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23731 MB (11% inode=99%):
[1551879413] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10067884 kB) free.
[1551880013] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8873084 kB) free!
[1551880133] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.9% (5859640 kB) free!
[1551880253] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2941976 kB) free!
[1551880373] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551880757] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1551880817] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (644436 kB) free!
[1551880917] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (658400 kB) free!
[1551881957] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (3086788 kB) free!
[1551882007] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.41, 4.15, 4.96
[1551882053] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10446640 kB) free.
[1551882077] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (3002052 kB) free!
[1551882197] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (2919304 kB) free!
[1551882653] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8979240 kB) free!
[1551882773] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.3% (5580964 kB) free!
[1551882893] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2545720 kB) free!
[1551883509] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19171 MB (9% inode=99%):
[1551883997] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1090928 kB) free!
[1551884693] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.4% (7598664 kB) free!
[1551884709] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551884709] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 37778 MB (18% inode=99%):
[1551885007] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.94, 5.66, 5.32
[1551885127] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.30, 5.51, 5.31
[1551885247] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.77, 5.25, 5.24
[1551886176] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 19.4% (1423332 kB) free!
[1551886296] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 19.2% (1410916 kB) free!
[1551886416] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 19.0% (1399384 kB) free!
[1551886448] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.93, 4.14, 4.71
[1551888394] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551888893] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4779968 kB) free!
[1551889594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551889714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551889834] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551891034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551891155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551891275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551891875] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551891893] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.4% (11037020 kB) free.
[1551891995] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551892115] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551892715] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551892835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551892955] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551893093] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8832168 kB) free!
[1551893213] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.5% (6170232 kB) free!
[1551893333] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3490140 kB) free!
[1551893556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551893676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551893683] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8922532 kB) free!
[1551893796] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551893803] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.5% (16017800 kB) free.
[1551894224] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 140.91, 33.33, 12.08
[1551894344] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 444.87, 185.20, 70.31
[1551894396] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551894464] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 486.27, 286.97, 121.54
[1551894516] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551894636] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1551895133] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.3% (11005340 kB) free.
[1551895236] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1551895356] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551895476] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551895509] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1551896077] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551896197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551896203] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8903276 kB) free!
[1551896317] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551896323] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.5% (16025052 kB) free.
[1551896334] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9116608 kB) free!
[1551896454] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.0% (6385828 kB) free!
[1551896574] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3648276 kB) free!
[1551896917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes