[1568486544] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176604 kB) free!
[1568486604] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193016 kB) free!
[1568486664] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169116 kB) free!
[1568486724] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199768 kB) free!
[1568486725] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184056 kB) free!
[1568486784] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176620 kB) free!
[1568486844] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198332 kB) free!
[1568486845] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (214780 kB) free!
[1568486964] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195624 kB) free!
[1568487813] SERVICE ALERT: rq-1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.82, 20.46, 20.02
[1568488413] SERVICE ALERT: rq-1;CPU Load;WARNING;HARD;3;WARNING - load average: 19.55, 19.41, 19.61
[1568489520] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 509 processes
[1568489580] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 506 processes
[1568490845] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9280848 kB) free!
[1568490965] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9122652 kB) free!
[1568491076] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.52, 6.99, 5.36
[1568491085] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (7973000 kB) free!
[1568491196] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.80, 7.18, 5.62
[1568491316] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.17, 7.12, 5.79
[1568491636] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.4% (4030864 kB) free!
[1568491756] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3576836 kB) free!
[1568491876] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (3135096 kB) free!
[1568492476] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (907324 kB) free!
[1568492549] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.35, 12.26, 7.26
[1568492669] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.43, 8.45, 6.47
[1568492789] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 5.78, 5.73
[1568492885] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3320460 kB) free!
[1568493389] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1568493389] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.46, 1.26, 3.27
[1568493984] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4234748 kB) free!
[1568493989] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 59.81, 19.63, 8.63
[1568494056] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1568494110] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.83, 13.39, 7.67
[1568494230] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.12, 9.28, 6.86
[1568494316] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1568494316] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.29, 4.10, 4.66
[1568494488] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8340664 kB) free!
[1568494584] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169312 kB) free!
[1568494608] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (780148 kB) free!
[1568494728] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233160 kB) free!
[1568494830] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.73, 1.73, 3.85
[1568495929] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 94.9% (46746228 kB) free.
[1568495996] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 456 processes
[1568496118] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 455 processes
[1568496240] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1568496476] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;DISK CRITICAL - free space: / 5064 MB (9% inode=97%):
[1568496597] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;DISK CRITICAL - free space: / 4242 MB (8% inode=96%):
[1568496719] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 3356 MB (6% inode=95%):
[1568496842] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1568497085] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.46, 6.26, 5.14
[1568497206] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.33, 6.29, 5.29
[1568497276] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 97.5% (23903044 kB) free.
[1568497319] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 35649 MB (69% inode=98%):
[1568497326] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.41, 4.68, 4.84
[1568497442] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1568497645] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 60.7% (14882364 kB) free.
[1568497764] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 48.2% (11811360 kB) free.
[1568499486] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 862 processes
[1568499606] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 863 processes
[1568499726] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 342 processes
[1568500301] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1568500301] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1568501430] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.16, 7.64, 5.85
[1568501550] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.58, 5.27, 5.20
[1568501670] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.94, 3.80, 4.67
[1568503316] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1568506471] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568509080] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 415 processes
[1568509141] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 413 processes
[1568510784] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.6% (10748028 kB) free.
[1568511541] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 390 processes
[1568511595] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 326 processes
[1568511601] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 404 processes
[1568511661] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 407 processes
[1568511715] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 327 processes
[1568511721] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568511835] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1568511985] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164084 kB) free!
[1568512051] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168748 kB) free!
[1568512105] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197892 kB) free!
[1568512165] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181596 kB) free!
[1568512175] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196208 kB) free!
[1568512227] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194744 kB) free!
[1568512285] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192804 kB) free!
[1568512305] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1568512405] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170096 kB) free!
[1568515217] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1568515337] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1568520000] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.3% (23508328 kB) free.
[1568520084] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.1% (21113276 kB) free.
[1568520156] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.4% (40629420 kB) free.
[1568520224] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.6% (21713216 kB) free.
[1568520356] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1568520357] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1568520357] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1568520476] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1568520477] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1568520477] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 157005 MB (69% inode=99%):
[1568520698] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175032 kB) free!
[1568520820] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177800 kB) free!
[1568520834] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194284 kB) free!
[1568520940] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183880 kB) free!
[1568520954] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194736 kB) free!
[1568521074] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199592 kB) free!
[1568522039] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1568522159] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1568522506] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.58, 6.69, 5.41
[1568522626] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.51, 6.84, 5.62
[1568522746] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.22, 5.90, 5.43
[1568524205] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (4205892 kB) free!
[1568524325] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (4169732 kB) free!
[1568524445] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4145884 kB) free!
[1568524546] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 0.98, 2.64, 4.35
[1568524670] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 103.74, 54.94, 24.32
[1568524790] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 107.24, 72.50, 34.47
[1568524910] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 106.97, 83.91, 43.26
[1568525070] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 84.91, 52.53, 23.30
[1568525190] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 86.97, 63.97, 31.03
[1568525310] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 87.47, 71.54, 37.79
[1568525740] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.8% (2656304 kB) free!
[1568525746] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.88, 5.68, 5.08
[1568525866] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.05, 6.67, 5.51
[1568525986] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.35, 6.37, 5.53
[1568526245] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1577184 kB) free!
[1568526340] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1568526340] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2062920 kB) free!
[1568526510] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.26, 3.25, 14.01
[1568526710] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.37, 2.45, 15.82
[1568528310] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.99, 1.60, 3.46
[1568528386] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1568528386] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.21, 3.95, 4.92
[1568528510] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.62, 1.89, 3.68
[1568533729] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9776316 kB) free!
[1568533849] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9539496 kB) free!
[1568533969] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9297948 kB) free!
[1568535701] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568535821] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1568535941] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1568536369] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4952468 kB) free!
[1568537740] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568537986] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1568540743] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1568540863] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1568540983] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1568540991] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.7% (23133952 kB) free.
[1568541028] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.7% (22722648 kB) free.
[1568541060] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 43.1% (21216156 kB) free.
[1568541095] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.2% (21630476 kB) free.
[1568545368] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9674856 kB) free!
[1568545488] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9676948 kB) free!
[1568545608] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (9676432 kB) free!
[1568549826] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1080760 kB) free!
[1568549947] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (1078892 kB) free!
[1568550067] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1076532 kB) free!
[1568550710] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 29.45, 17.88, 7.54
[1568550830] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.51, 12.17, 6.70
[1568550950] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.61, 8.50, 6.02
[1568551550] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 1.64, 3.47
[1568552395] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (1654896 kB) free!
[1568552433] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175752 kB) free!
[1568552505] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190136 kB) free!
[1568552515] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (1486796 kB) free!
[1568552553] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194476 kB) free!
[1568552625] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170496 kB) free!
[1568552635] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1475944 kB) free!
[1568552673] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167224 kB) free!
[1568552745] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183724 kB) free!
[1568555516] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.59, 6.75, 5.85
[1568555636] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.49, 4.65, 5.19
[1568555756] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 3.30, 4.63
[1568557110] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.73, 13.05, 6.90
[1568557230] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.39, 8.99, 6.16
[1568557350] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.59, 6.18, 5.47
[1568557950] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.34, 1.43, 3.25
[1568559956] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 267.43, 204.29, 99.70
[1568560076] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 149.31, 195.67, 110.17
[1568560196] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.44, 131.06, 96.86
[1568560435] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2726552 kB) free!
[1568560950] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.33, 11.68, 9.81
[1568561035] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2446928 kB) free!
[1568561070] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.01, 8.09, 8.72
[1568561190] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.68, 5.61, 7.73
[1568561790] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1568561790] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.72, 1.37, 4.40
[1568561996] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.21, 0.58, 14.20
[1568563197] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.53, 4.42
[1568564279] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.5% (23567640 kB) free.
[1568564312] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.0% (23276536 kB) free.
[1568564344] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.9% (23507636 kB) free.
[1568564380] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.8% (14432928 kB) free.
[1568564383] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1568564503] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes