[1589481223] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589481365] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24268016 kB) free!
[1589481965] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1597100 kB) free!
[1589482565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26509252 kB) free.
[1589482821] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.0% (2205292 kB) free!
[1589483165] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3907476 kB) free!
[1589483285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393164 kB) free!
[1589483405] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403616 kB) free!
[1589483972] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167216 kB) free!
[1589484005] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18457544 kB) free!
[1589484021] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 80.4% (19715976 kB) free.
[1589484572] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (1854352 kB) free!
[1589484605] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416660 kB) free!
[1589485172] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (1485676 kB) free!
[1589485205] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19125852 kB) free!
[1589485772] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.8% (1916940 kB) free!
[1589485805] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390980 kB) free!
[1589486372] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1501456 kB) free!
[1589486405] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19532748 kB) free!
[1589486427] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235788 kB) free!
[1589486547] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217568 kB) free!
[1589486618] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (429656 kB) free!
[1589486667] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231836 kB) free!
[1589486738] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378232 kB) free!
[1589486858] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2296284 kB) free!
[1589486972] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (1846968 kB) free!
[1589487005] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399760 kB) free!
[1589487572] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3658808 kB) free.
[1589487606] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18696812 kB) free!
[1589488206] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381384 kB) free!
[1589488467] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.8% (8259400 kB) free!
[1589488772] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (1013752 kB) free!
[1589488806] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18427960 kB) free!
[1589488892] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (539496 kB) free!
[1589489012] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 32.2% (5227892 kB) free.
[1589489406] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410260 kB) free!
[1589490006] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17856924 kB) free!
[1589490606] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410392 kB) free!
[1589490740] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 89.7% (44197104 kB) free.
[1589491174] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (44841524 kB) free.
[1589491206] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16702564 kB) free!
[1589491806] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27033404 kB) free.
[1589492067] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.0% (44375920 kB) free.
[1589492259] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(388) < 172.105.57.57 [15/May/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589492340] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3015) < 172.105.59.173 [15/May/2020:03:02:42 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1589492380] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589492382] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(529) < 2020/05/15 03:01:32 [error] 22133#0: *88519104 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589492406] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14712572 kB) free!
[1589492460] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589492503] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589492526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10780784 kB) free!
[1589492646] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (6932692 kB) free!
[1589493012] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.66, 5.00, 5.05
[1589493132] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.96, 4.66, 4.92
[1589493246] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (24019432 kB) free!
[1589493846] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4734304 kB) free!
[1589494446] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21044544 kB) free!
[1589495047] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2178368 kB) free!
[1589495647] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17334828 kB) free!
[1589496247] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35529288 kB) free.
[1589496847] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14006160 kB) free!
[1589496967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10194904 kB) free!
[1589497087] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6542584 kB) free!
[1589497687] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21636396 kB) free!
[1589498287] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3645536 kB) free!
[1589498300] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.1% (9573964 kB) free.
[1589498887] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17314748 kB) free!
[1589499212] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589499487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32165288 kB) free.
[1589499500] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182576 kB) free!
[1589499620] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (240396 kB) free!
[1589499740] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200140 kB) free!
[1589500087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12910912 kB) free!
[1589500207] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9375128 kB) free!
[1589500327] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5793996 kB) free!
[1589500720] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.0% (18862852 kB) free.
[1589500927] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18340472 kB) free!
[1589500939] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 53.2% (13048536 kB) free.
[1589501037] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.3% (4022788 kB) free!
[1589501527] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32251176 kB) free.
[1589501637] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193440 kB) free!
[1589502127] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4906596 kB) free!
[1589502248] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (486328 kB) free!
[1589502368] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388428 kB) free!
[1589502968] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14765268 kB) free!
[1589503568] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36751652 kB) free.
[1589504168] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9054896 kB) free!
[1589504288] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4847032 kB) free!
[1589504408] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (850472 kB) free!
[1589504421] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (588324 kB) free!
[1589504481] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 29.8% (7296836 kB) free.
[1589504781] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3557860 kB) free!
[1589504841] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3332512 kB) free!
[1589504858] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 53.9% (71104388 kB) free.
[1589504901] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 13.6% (3325080 kB) free!
[1589504961] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 13.6% (3329476 kB) free!
[1589505008] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22802144 kB) free!
[1589505309] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1589505429] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1589505549] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1589505608] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38931568 kB) free.
[1589506208] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19294628 kB) free!
[1589506328] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15212092 kB) free!
[1589506448] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (11180916 kB) free!
[1589507048] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23475860 kB) free!
[1589507648] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36489220 kB) free.
[1589508248] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (14996452 kB) free!
[1589508368] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11053544 kB) free!
[1589508488] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (7293580 kB) free!
[1589509089] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20460908 kB) free!
[1589509689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33703724 kB) free.
[1589510289] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10014664 kB) free!
[1589510409] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5610492 kB) free!
[1589510529] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1477084 kB) free!
[1589511129] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16824884 kB) free!
[1589511549] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1589511729] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30867372 kB) free.
[1589511732] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.87, 5.66, 5.03
[1589511852] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.69, 6.70, 5.48
[1589511972] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.37, 6.68, 5.61
[1589512260] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(544) < 103.255.146.164 [15/May/2020:08:31:40 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589512329] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8337612 kB) free!
[1589512380] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589512449] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3764872 kB) free!
[1589512569] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403216 kB) free!
[1589513539] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3967140 kB) free!
[1589513580] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(655) < 103.255.146.164 [15/May/2020:09:01:24 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589513659] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (3692372 kB) free!
[1589513701] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589513769] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38041516 kB) free.
[1589513779] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.9% (3417752 kB) free!
[1589514369] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13439496 kB) free!
[1589514379] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2354432 kB) free!
[1589514489] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8931760 kB) free!
[1589514609] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4553564 kB) free!
[1589515210] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18322044 kB) free!
[1589515810] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388112 kB) free!
[1589516410] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13260076 kB) free!
[1589517300] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(439) < 103.255.146.164 [15/May/2020:10:01:49 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589517421] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589517520] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182772 kB) free!
[1589517610] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385292 kB) free!
[1589517640] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184284 kB) free!
[1589517657] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1589517760] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170208 kB) free!
[1589517777] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1589517897] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1589518210] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14140676 kB) free!
[1589518498] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1589518810] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399192 kB) free!
[1589519410] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24325984 kB) free!
[1589519511] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 60.4% (14811396 kB) free.
[1589520010] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387028 kB) free!
[1589520223] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.6% (3057520 kB) free!
[1589520343] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.9% (2437940 kB) free!
[1589520463] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1881232 kB) free!
[1589520610] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21572832 kB) free!
[1589521021] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(409) < 103.255.146.164 [15/May/2020:11:01:35 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589521141] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589521210] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411904 kB) free!
[1589521741] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 182.18.144.246 [15/May/2020:11:19:00 +0530] TCP 502 0 0 120.001 "10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.154:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 60.000, 0.000, 60.000"
[1589521810] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24428944 kB) free!
[1589521861] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589522410] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388672 kB) free!
[1589522461] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(402) < 103.255.146.164 [15/May/2020:11:21:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589522581] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(650) < 103.255.146.164 [15/May/2020:11:31:14 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589522702] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589523010] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22641664 kB) free!
[1589523301] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [15/May/2020:11:41:10 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589523421] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589523610] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408632 kB) free!
[1589523711] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.3% (3755108 kB) free!
[1589523824] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (6836628 kB) free!
[1589523831] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 13.7% (3348280 kB) free!
[1589523944] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 17.9% (6282760 kB) free!
[1589523951] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 12.0% (2948452 kB) free!
[1589524019] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.3% (3769504 kB) free!
[1589524064] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (5704864 kB) free!
[1589524139] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.5% (3331300 kB) free!
[1589524174] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9713780 kB) free!
[1589524211] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26746376 kB) free.
[1589524259] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 11.6% (2852820 kB) free!
[1589524294] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9336368 kB) free!
[1589524414] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8908608 kB) free!
[1589524421] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 326 processes
[1589524541] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1589524551] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (873504 kB) free!
[1589524663] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.5% (16045212 kB) free.
[1589524664] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.8% (31815892 kB) free.
[1589524811] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (1075212 kB) free!
[1589524859] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (451904 kB) free!
[1589524931] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408460 kB) free!
[1589525051] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387840 kB) free!
[1589525459] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.8% (22902184 kB) free.
[1589525651] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17618232 kB) free!
[1589526214] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3207492 kB) free!
[1589526251] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (471276 kB) free!
[1589526421] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(267) < 103.255.146.164 [15/May/2020:12:31:36 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589526541] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589526812] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (2901728 kB) free!
[1589526851] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24088820 kB) free!
[1589526932] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2909988 kB) free!
[1589527052] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (2988800 kB) free!
[1589527451] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (426488 kB) free!
[1589527652] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3331576 kB) free.
[1589527741] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(382) < 103.255.146.164 [15/May/2020:12:49:28 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589527861] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589527940] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9210200 kB) free!
[1589528051] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20841056 kB) free!
[1589528060] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (8957296 kB) free!
[1589528180] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (8694200 kB) free!
[1589528560] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (3928428 kB) free!
[1589528651] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405700 kB) free!
[1589528852] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (2908032 kB) free!
[1589528973] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3187788 kB) free!
[1589529061] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(435) < 103.255.146.164 [15/May/2020:13:14:05 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589529093] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3403504 kB) free.
[1589529181] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589529760] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1587464 kB) free!
[1589530293] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3058372 kB) free!
[1589530413] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3302436 kB) free.
[1589530541] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1589530580] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4060172 kB) free!
[1589530661] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1589531013] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (1538660 kB) free!
[1589531051] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20288988 kB) free!
[1589531133] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.8% (2088580 kB) free!
[1589531253] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (2519132 kB) free!
[1589531263] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.7% (2850272 kB) free!
[1589531384] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.3% (2542900 kB) free!
[1589531504] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2100268 kB) free!
[1589531651] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (452004 kB) free!
[1589531853] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1589531853] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (891552 kB) free!
[1589532252] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29355972 kB) free.
[1589532453] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (1869828 kB) free!
[1589532781] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(348) < 103.255.146.164 [15/May/2020:14:20:50 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589532852] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (401684 kB) free!
[1589532901] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589532972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405508 kB) free!
[1589533053] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3511072 kB) free.
[1589533092] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393232 kB) free!
[1589533653] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3177616 kB) free!
[1589533692] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (21998964 kB) free!
[1589533773] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3037188 kB) free!
[1589533893] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1395916 kB) free!
[1589534101] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(391) < 103.255.146.164 [15/May/2020:14:43:48 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589534222] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589534292] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433440 kB) free!
[1589534861] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1589534892] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31760084 kB) free.
[1589534981] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1589535104] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.9% (16145592 kB) free.
[1589535421] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 185.255.8.58 [15/May/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1589535492] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (403748 kB) free!
[1589535541] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589535612] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404940 kB) free!
[1589535732] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410372 kB) free!
[1589536141] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(459) < 103.255.146.164 [15/May/2020:15:14:34 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589536181] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1589536261] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589536293] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3583532 kB) free.
[1589536301] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1589536332] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30580820 kB) free.
[1589536893] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (153116 kB) free!
[1589536932] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (385928 kB) free!
[1589537013] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (1133868 kB) free!
[1589537052] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407564 kB) free!
[1589537133] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.7% (4669924 kB) free.
[1589537172] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420736 kB) free!
[1589537733] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3238864 kB) free!
[1589537772] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20452068 kB) free!
[1589537854] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (160472 kB) free!
[1589537974] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1409916 kB) free!
[1589538373] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409708 kB) free!
[1589538661] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(372) < 103.255.146.164 [15/May/2020:15:58:10 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589538781] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589538973] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32303372 kB) free.
[1589539174] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.0% (5036968 kB) free.
[1589539573] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1640924 kB) free!
[1589539693] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (412420 kB) free!
[1589539733] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.12696.436> (10.147.212.147:37722 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1589539774] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.8% (1915280 kB) free!
[1589539813] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (436072 kB) free!
[1589539853] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589539894] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3125088 kB) free!
[1589540014] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.3% (4276436 kB) free.
[1589540413] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32647040 kB) free.
[1589540581] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(376) < 103.255.146.164 [15/May/2020:16:31:30 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589540614] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (956656 kB) free!
[1589540661] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 8.5% (2084220 kB) free!
[1589540701] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589540734] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.2% (2476924 kB) free!
[1589540854] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1686132 kB) free!
[1589541013] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (388432 kB) free!
[1589541133] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402592 kB) free!
[1589541253] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400060 kB) free!
[1589541853] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27915744 kB) free.
[1589542054] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3713812 kB) free.
[1589542304] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3587544 kB) free!
[1589542424] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.2% (2960536 kB) free!
[1589542453] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (405068 kB) free!
[1589542544] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.2% (2506748 kB) free!
[1589542573] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (413400 kB) free!
[1589542654] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.5% (2522792 kB) free!
[1589542693] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415228 kB) free!
[1589542774] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3670972 kB) free.
[1589543144] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1589543144] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (761076 kB) free!
[1589543293] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27701476 kB) free.
[1589543374] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (1622436 kB) free!
[1589543494] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3197292 kB) free!
[1589543502] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/15 17:20:23 [error] 22134#0: *88689536 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1589543614] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3618896 kB) free.
[1589543623] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589543858] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.8% (26083700 kB) free!
[1589543893] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (380164 kB) free!
[1589543978] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.8% (40642044 kB) free.
[1589544014] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (426828 kB) free!
[1589544134] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425040 kB) free!
[1589544214] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2953676 kB) free!
[1589544334] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3672820 kB) free.
[1589544578] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.6% (20578832 kB) free!
[1589544580] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(593) < 49.50.86.152 [15/May/2020:17:39:33 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589544698] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 33.1% (43629232 kB) free.
[1589544700] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589544734] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25650608 kB) free!
[1589544822] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(777) < 2020/05/15 17:40:43 [error] 22129#0: *88695238 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589544901] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(655) < 203.212.70.170 [15/May/2020:17:37:35 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1589544934] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (2727728 kB) free!
[1589544942] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589545021] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589545055] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3968744 kB) free.
[1589545334] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (422488 kB) free!
[1589545654] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2597412 kB) free!
[1589545774] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4099384 kB) free.
[1589545973] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (1634988 kB) free!
[1589546028] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 829 processes
[1589546145] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.3% (16226396 kB) free.
[1589546148] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 827 processes
[1589546221] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(921) < 203.212.70.170 [15/May/2020:18:03:42 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.000, 0.000"
[1589546268] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 337 processes
[1589546341] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589546349] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1589546374] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3122436 kB) free!
[1589546494] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3115056 kB) free!
[1589546614] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1342852 kB) free!
[1589546941] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(654) < 103.255.146.164 [15/May/2020:18:14:04 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589547061] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589547098] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26284836 kB) free!
[1589547173] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (326392 kB) free!
[1589547218] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.6% (40315276 kB) free.
[1589547734] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16339736 kB) free!
[1589547814] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.3% (4282104 kB) free.
[1589547818] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.1% (23810996 kB) free!
[1589547938] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.0% (36896212 kB) free.
[1589548334] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (458532 kB) free!
[1589548934] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28172104 kB) free.
[1589549534] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (382936 kB) free!
[1589549614] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (1715484 kB) free!
[1589549654] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (450128 kB) free!
[1589549735] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.5% (4469472 kB) free.
[1589549774] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417236 kB) free!
[1589550334] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2638972 kB) free!
[1589550372] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.96, 4.56, 4.98
[1589550374] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26741180 kB) free.
[1589550454] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.0% (4545708 kB) free.
[1589550972] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.02, 4.54, 5.09
[1589550974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (375384 kB) free!
[1589551054] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2771152 kB) free!
[1589551092] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.10, 4.08, 4.86
[1589551095] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395668 kB) free!
[1589551174] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.3% (4441320 kB) free.
[1589551215] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404424 kB) free!
[1589551692] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.16, 5.36, 5.02
[1589551774] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3096072 kB) free!
[1589551812] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.74, 5.99, 5.27
[1589551815] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19162288 kB) free!
[1589551894] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.5% (4628824 kB) free.
[1589551932] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.05, 6.37, 5.50
[1589552014] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 38.1% (18786500 kB) free.
[1589552180] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 31.5% (15499592 kB) free.
[1589552415] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393516 kB) free!
[1589552627] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 17.4% (4266768 kB) free!
[1589552747] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 17.4% (4267456 kB) free!
[1589552867] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 17.4% (4266652 kB) free!
[1589553094] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (1445324 kB) free!
[1589553214] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4190124 kB) free.
[1589553814] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (1504872 kB) free!
[1589553934] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3967604 kB) free.
[1589554534] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (1148160 kB) free!
[1589554654] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3649452 kB) free.
[1589554815] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (14037176 kB) free!
[1589555145] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1589555415] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383864 kB) free!
[1589555532] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.86, 4.58, 4.86
[1589555854] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2492204 kB) free!
[1589555974] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.2% (4748672 kB) free.
[1589556015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35611532 kB) free.
[1589556132] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.01, 5.18, 5.04
[1589556252] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.93, 4.82, 4.93
[1589556345] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3660600 kB) free!
[1589556465] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3444452 kB) free!
[1589556574] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1372072 kB) free!
[1589556585] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (3190036 kB) free!
[1589556615] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (721160 kB) free!
[1589556694] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3347492 kB) free.
[1589556735] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397272 kB) free!
[1589556740] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4439200 kB) free!
[1589556852] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.48, 5.40, 5.06
[1589556855] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399536 kB) free!
[1589556973] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.07, 4.97, 4.95
[1589557186] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16195512 kB) free.
[1589557455] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24449748 kB) free!
[1589557894] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2538664 kB) free!
[1589557940] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170388 kB) free!
[1589558014] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4192748 kB) free.
[1589558055] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6065716 kB) free!
[1589558614] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (289964 kB) free!
[1589558734] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.8% (2244368 kB) free!
[1589558854] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.3% (3951152 kB) free.
[1589559549] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1589559669] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589559789] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1589559951] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 55.7% (13643376 kB) free.
[1589559980] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9809900 kB) free!
[1589560054] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (1484020 kB) free!
[1589560100] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (9776204 kB) free!
[1589560174] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (2946180 kB) free!
[1589560220] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (9738216 kB) free!
[1589560295] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589560295] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3818388 kB) free.
[1589560389] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1589560455] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21612916 kB) free!
[1589561055] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (422932 kB) free!
[1589561494] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (969856 kB) free!
[1589561614] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.7% (2231152 kB) free!
[1589561655] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31562184 kB) free.
[1589561734] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3537292 kB) free.
[1589562255] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2580948 kB) free!
[1589562375] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391772 kB) free!
[1589562495] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423944 kB) free!
[1589562814] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9790708 kB) free!
[1589562934] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9743828 kB) free!
[1589563054] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.7% (9692600 kB) free!
[1589563095] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27382260 kB) free.
[1589563695] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1592756 kB) free!
[1589563815] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382796 kB) free!
[1589563935] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397272 kB) free!
[1589564134] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (758820 kB) free!
[1589564254] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.3% (3944656 kB) free.
[1589564535] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21863016 kB) free!
[1589565135] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425488 kB) free!
[1589565454] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2580180 kB) free!
[1589565575] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3358180 kB) free.
[1589565735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32746232 kB) free.
[1589566174] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1359080 kB) free!
[1589566294] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1893792 kB) free!
[1589566336] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3633280 kB) free!
[1589566414] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (2382632 kB) free!
[1589566456] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404952 kB) free!
[1589566576] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406640 kB) free!
[1589567014] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3568424 kB) free.
[1589567176] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24062340 kB) free!