[1551205851] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551205971] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551206571] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551206691] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551207291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551207411] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1551207411] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551208012] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551208132] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551208252] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551208852] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551208972] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551209572] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551209692] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551210292] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551210413] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551211012] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551211132] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551211732] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551211852] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551212452] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551212572] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551213172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551213292] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551213893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551214013] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551214133] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551214733] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551214853] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551215453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551215573] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551216173] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551216294] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551216893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551217013] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551217133] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551217733] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551217853] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551218453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551218573] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551218984] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (1423028 kB) free!
[1551219104] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 19.3% (1423052 kB) free!
[1551219174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551219224] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 19.3% (1422540 kB) free!
[1551219294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551219414] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551220014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551220134] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551220734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551220854] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551221454] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551221575] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551222174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551222294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551222414] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551223014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551223134] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551223734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551223854] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551224454] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551224575] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551225174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551225294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551225414] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551226014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551226134] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551226734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551226854] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551227455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551227575] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551228175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551228295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551228415] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551229015] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551229135] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551229735] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551229856] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551230455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551230575] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551230695] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551231295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551231415] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551231714] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 365 processes
[1551231835] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1551232015] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551232135] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551232434] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 356 processes
[1551232554] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1551232736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551232856] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551232976] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551233576] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551233696] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551233754] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1551233874] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1551234296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551234416] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551235016] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551235137] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551235736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551235856] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551236456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551236576] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551237176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551237296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551237416] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551238016] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551238137] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551239336] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551239456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551239576] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551241376] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551241977] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551242097] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551242217] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551243417] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1551243417] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551244017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551244137] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551244257] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551245429] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 256.85, 106.80, 40.45
[1551245549] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 280.23, 165.38, 69.96
[1551245669] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 283.36, 204.56, 95.88
[1551246057] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551246657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551246777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551246897] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551247136] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 220.10, 159.91, 75.25
[1551247200] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4735044 kB) free!
[1551247256] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 88.46, 143.10, 80.26
[1551247320] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4735900 kB) free!
[1551247376] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.29, 95.90, 70.59
[1551247440] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (4734136 kB) free!
[1551247469] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 1.43, 16.95
[1551248003] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 93.4% (6872204 kB) free.
[1551248087] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 91.0% (6688980 kB) free.
[1551248177] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.7% (22959544 kB) free.
[1551248269] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1551248389] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1551248509] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1551248576] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.75, 2.39, 19.87
[1551249052] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 19.8% (1455116 kB) free!
[1551249172] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.6% (1444104 kB) free!
[1551249269] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.35, 0.87, 3.07
[1551249284] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 94.9% (6982764 kB) free.
[1551249709] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1551250376] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.45, 3.30
[1551250961] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9651284 kB) free!
[1551251081] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.8% (10746080 kB) free.
[1551251773] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4876708 kB) free!
[1551251893] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4872272 kB) free!
[1551252013] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (4867436 kB) free!
[1551253497] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551254081] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8922880 kB) free!
[1551254098] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551254201] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10245356 kB) free.
[1551254218] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551254338] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551255363] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3240292 kB) free!
[1551255483] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.7% (3199028 kB) free!
[1551255603] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.4% (3155336 kB) free!
[1551255709] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1551255829] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1551255949] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1551256001] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7930552 kB) free!
[1551256121] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9496428 kB) free!
[1551256241] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.5% (10609264 kB) free.
[1551256549] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1551256549] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1551261013] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 22.2% (5429760 kB) free.
[1551261308] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.0% (14469636 kB) free.
[1551262065] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3234556 kB) free!
[1551262138] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551262185] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3229860 kB) free!
[1551262305] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.6% (3184168 kB) free!
[1551262738] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551262859] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551262979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551263413] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4771828 kB) free!
[1551263533] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4754660 kB) free!
[1551263628] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.2% (22850752 kB) free.
[1551263674] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.8% (14765476 kB) free.
[1551265379] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551265979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551266099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551266219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551267041] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9254596 kB) free!
[1551267161] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.1% (10881264 kB) free.
[1551267949] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551268019] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551268361] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9150532 kB) free!
[1551268481] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.4% (11019600 kB) free.
[1551268619] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551268739] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551268859] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551271859] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551272459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551272579] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551272699] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551273881] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8361620 kB) free!
[1551273949] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1551274001] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10380120 kB) free.
[1551274069] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1551274189] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1551275100] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551275700] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551275820] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551275940] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551276401] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.2% (7997968 kB) free!
[1551276521] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10170868 kB) free.
[1551276538] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46995 MB (20% inode=99%):
[1551276658] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46995 MB (20% inode=99%):
[1551276778] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54160 MB (24% inode=99%):
[1551277721] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.8% (7762236 kB) free!
[1551277841] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9936044 kB) free.
[1551279178] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46966 MB (20% inode=99%):
[1551279298] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46902 MB (20% inode=99%):
[1551279418] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54581 MB (24% inode=99%):
[1551280241] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.8% (7287684 kB) free!
[1551280361] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9527592 kB) free!
[1551280481] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10405724 kB) free.
[1551281081] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.7% (7247272 kB) free!
[1551281201] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9438940 kB) free!
[1551281321] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10369300 kB) free.
[1551281921] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7634356 kB) free!
[1551282041] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9779848 kB) free!
[1551282161] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.0% (6887268 kB) free!
[1551282418] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46691 MB (20% inode=99%):
[1551282538] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54882 MB (24% inode=99%):
[1551282540] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551283140] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551283260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551283381] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551284562] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10132664 kB) free.
[1551286962] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.3% (7042944 kB) free!
[1551287082] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9099732 kB) free!
[1551287202] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10231716 kB) free.
[1551287848] SERVICE ALERT: sig2-vm2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1551287968] SERVICE ALERT: sig2-vm2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1551288088] SERVICE ALERT: sig2-vm2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1551289002] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.9% (6861844 kB) free!
[1551289122] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8841312 kB) free!
[1551289242] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10230820 kB) free.
[1551290195] SERVICE ALERT: sigqscluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1551290742] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1551290802] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1551290862] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 63 users currently logged in
[1551290922] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1551291042] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.7% (7243324 kB) free!
[1551291162] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10146680 kB) free.
[1551291395] SERVICE ALERT: sigqscluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1551291515] SERVICE ALERT: sigqscluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1551291635] SERVICE ALERT: sigqscluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1551291762] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.3% (8051628 kB) free!
[1551291882] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10224096 kB) free.
[1551292122] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 58 users currently logged in