[1551292482] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.3% (8041196 kB) free!
[1551292602] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10073476 kB) free.
[1551293488] SERVICE ALERT: sig2-vm2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1551293802] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8691260 kB) free!
[1551293833] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1551293922] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10235516 kB) free.
[1551293938] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46672 MB (20% inode=99%):
[1551293953] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 336 processes
[1551294058] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53840 MB (23% inode=99%):
[1551294073] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1551294631] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551294751] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551294781] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1551294872] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551295722] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8938708 kB) free!
[1551295843] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10257996 kB) free.
[1551295981] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551296442] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9655304 kB) free!
[1551296562] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10263520 kB) free.
[1551297181] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551297301] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551297902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551298022] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1551298622] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1551298742] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1551298962] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9014796 kB) free!
[1551299082] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10288244 kB) free.
[1551299342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551299458] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46660 MB (20% inode=99%):
[1551299462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551299578] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54276 MB (24% inode=99%):
[1551299582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551299682] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9787620 kB) free!
[1551299802] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10294964 kB) free.
[1551300182] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551300402] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9368132 kB) free!
[1551300522] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10312588 kB) free.
[1551300778] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46660 MB (20% inode=99%):
[1551300782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551300898] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54851 MB (24% inode=99%):
[1551300903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551301498] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46660 MB (20% inode=99%):
[1551301502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551301618] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54851 MB (24% inode=99%):
[1551301622] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551301722] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7879732 kB) free!
[1551301843] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10336688 kB) free.
[1551302218] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46660 MB (20% inode=99%):
[1551302222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551302338] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54851 MB (24% inode=99%):
[1551302342] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551302942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551303042] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9298092 kB) free!
[1551303062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551303162] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10366716 kB) free.
[1551303182] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551303762] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9558616 kB) free!
[1551303782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551303882] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10381192 kB) free.
[1551303903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551304482] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8500864 kB) free!
[1551304502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551304602] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10397144 kB) free.
[1551304622] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551305222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551305338] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46656 MB (20% inode=99%):
[1551305342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551305458] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54848 MB (24% inode=99%):
[1551305462] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551306058] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46656 MB (20% inode=99%):
[1551306062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551306178] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54847 MB (24% inode=99%):
[1551306182] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551306778] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46655 MB (20% inode=99%):
[1551306783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551306898] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54847 MB (24% inode=99%):
[1551306903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551307498] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46655 MB (20% inode=99%):
[1551307503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551307618] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46655 MB (20% inode=99%):
[1551307623] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551307738] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54847 MB (24% inode=99%):
[1551308223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551308338] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46655 MB (20% inode=99%):
[1551308343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551308458] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54847 MB (24% inode=99%):
[1551308463] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551309058] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46655 MB (20% inode=99%):
[1551309063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551309178] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54847 MB (24% inode=99%):
[1551309184] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551309783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551309903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551310503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551310623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551310743] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551311343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551311463] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551312064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551312184] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551312778] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46428 MB (20% inode=99%):
[1551312784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551312898] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54844 MB (24% inode=99%):
[1551312904] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551313002] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8419572 kB) free!
[1551313122] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.9% (10776092 kB) free.
[1551313498] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46652 MB (20% inode=99%):
[1551313504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551313618] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46652 MB (20% inode=99%):
[1551313624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551313722] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9743552 kB) free!
[1551313738] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54844 MB (24% inode=99%):
[1551313842] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.9% (10809744 kB) free.
[1551314224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551314338] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46652 MB (20% inode=99%):
[1551314344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551314458] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54844 MB (24% inode=99%):
[1551314464] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551315058] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46652 MB (20% inode=99%):
[1551315065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551315178] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55354 MB (24% inode=99%):
[1551315185] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551315642] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9714032 kB) free!
[1551315762] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10858624 kB) free.
[1551315785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551315905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551316025] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1551316625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551316745] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551317345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551317465] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551318066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551318178] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46650 MB (20% inode=99%):
[1551318186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551318298] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54841 MB (24% inode=99%):
[1551318306] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551318898] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46648 MB (20% inode=99%):
[1551318906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551319018] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54840 MB (24% inode=99%):
[1551319026] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551319618] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46648 MB (20% inode=99%):
[1551319626] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551319738] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54840 MB (24% inode=99%):
[1551319746] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551320339] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46648 MB (20% inode=99%):
[1551320346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551320459] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46648 MB (20% inode=99%):
[1551320467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551320579] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54840 MB (24% inode=99%):
[1551320587] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1551321179] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46648 MB (20% inode=99%):
[1551321187] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551321299] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54776 MB (24% inode=99%):
[1551321307] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551321899] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46134 MB (20% inode=99%):
[1551321907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1551322019] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54774 MB (24% inode=99%):
[1551322027] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1551322619] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46582 MB (20% inode=99%):
[1551322627] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551322739] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54837 MB (24% inode=99%):
[1551322747] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551323339] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46645 MB (20% inode=99%):
[1551323348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551323459] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54837 MB (24% inode=99%):
[1551323468] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1551324059] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46581 MB (20% inode=99%):
[1551324068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551324162] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8974136 kB) free!
[1551324179] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53749 MB (23% inode=99%):
[1551324188] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551324282] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9925172 kB) free.
[1551324308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551325482] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9609204 kB) free!
[1551325602] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.5% (10586484 kB) free.
[1551326579] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46634 MB (20% inode=99%):
[1551326699] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54825 MB (24% inode=99%):
[1551326709] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551327309] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551327429] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551327549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551329202] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9803576 kB) free!
[1551329323] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10310096 kB) free.
[1551329922] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8831224 kB) free!
[1551329949] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1551329949] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1551330042] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9958576 kB) free.
[1551330550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1551330642] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8105364 kB) free!
[1551330670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551330762] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9437048 kB) free!
[1551330790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1551330882] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.8% (6284824 kB) free!
[1551330899] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46173 MB (20% inode=99%):
[1551331019] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54730 MB (24% inode=99%):
[1551331189] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1551331309] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1551331429] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1551332082] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.0% (12322924 kB) free.
[1551332590] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551333190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551333310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551333430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551333882] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8832532 kB) free!
[1551334002] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10428560 kB) free.
[1551334019] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46605 MB (20% inode=99%):
[1551334139] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54795 MB (24% inode=99%):
[1551334512] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3123800 kB) free!
[1551334602] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9705168 kB) free!
[1551334632] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (3067968 kB) free!
[1551334722] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11248444 kB) free.
[1551334752] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.4% (3158112 kB) free!
[1551335230] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551335353] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.2% (14658676 kB) free.
[1551335830] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551335950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551336070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551336522] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9152664 kB) free!
[1551336642] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10507628 kB) free.
[1551337474] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 19.0% (1395352 kB) free!
[1551337594] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 18.6% (1365600 kB) free!
[1551337714] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 18.1% (1327896 kB) free!
[1551337842] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.1% (5979656 kB) free!
[1551337870] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551337962] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.9% (7360752 kB) free!
[1551338082] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4220904 kB) free!
[1551338131] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 19.7% (1446408 kB) free!
[1551338252] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 19.1% (1403124 kB) free!
[1551338339] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46229 MB (20% inode=99%):
[1551338372] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 18.8% (1384284 kB) free!
[1551338459] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54358 MB (24% inode=99%):
[1551338682] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1551338682] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.4% (6114992 kB) free!
[1551338683] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.2% (2463712 kB) free!
[1551338803] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 13.2% (2145984 kB) free!
[1551338923] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 11.0% (1789524 kB) free!
[1551339059] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45870 MB (20% inode=99%):
[1551339070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1551339179] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53933 MB (23% inode=99%):
[1551339190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551339310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551340482] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3796584 kB) free!
[1551340615] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.3% (14669156 kB) free.
[1551340741] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 94.9% (6980148 kB) free.
[1551340778] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8080432 kB) free!
[1551340844] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 93.2% (6855412 kB) free.
[1551340898] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.3% (13469872 kB) free.
[1551341082] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.5% (5655172 kB) free!
[1551341682] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3961088 kB) free!
[1551342698] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9460296 kB) free!
[1551342819] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7736056 kB) free!
[1551342939] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6305956 kB) free!
[1551343483] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.9% (6347636 kB) free!
[1551344083] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2936964 kB) free!
[1551344139] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4914196 kB) free!
[1551344407] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 74.9% (36932360 kB) free.
[1551344458] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 56.2% (27697948 kB) free.
[1551344579] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44663 MB (19% inode=99%):
[1551344699] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52855 MB (23% inode=99%):
[1551344710] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1551345310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551345431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551345551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551345893] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 19.6% (1439828 kB) free!
[1551346013] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.1% (1407620 kB) free!
[1551346133] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 18.7% (1377224 kB) free!
[1551346753] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.5% (2681076 kB) free!
[1551346873] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.5% (2516152 kB) free!
[1551346993] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.2% (2305172 kB) free!
[1551347411] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8570072 kB) free!
[1551347531] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9415492 kB) free!
[1551347593] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1285556 kB) free!
[1551347651] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.8% (5829040 kB) free!
[1551347882] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4916196 kB) free!
[1551347951] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1551348002] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4875708 kB) free!
[1551348122] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4822108 kB) free!
[1551348251] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10453300 kB) free.
[1551348299] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44120 MB (19% inode=99%):
[1551348419] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44376 MB (19% inode=99%):
[1551348539] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51546 MB (22% inode=99%):
[1551348551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551348671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1551348791] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1551348851] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.2% (6499468 kB) free!
[1551348971] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9762944 kB) free!
[1551349091] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.2% (6510704 kB) free!
[1551349133] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (670920 kB) free!
[1551349537] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 70.7% (34854752 kB) free.
[1551349593] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.7% (22138468 kB) free.
[1551351191] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1551351540] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44106 MB (19% inode=99%):
[1551351660] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44362 MB (19% inode=99%):
[1551351776] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.43, 15.45, 9.87
[1551351777] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4756144 kB) free!
[1551351780] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52036 MB (23% inode=99%):
[1551351791] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551351896] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.63, 10.49, 8.73
[1551351898] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4742268 kB) free!
[1551351911] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1551352016] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.70, 7.25, 7.76
[1551352018] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (4722560 kB) free!
[1551352031] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551352527] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 268 processes
[1551352616] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.40, 1.40, 4.31
[1551352647] SERVICE ALERT: sig2-vm1;CPU_Procs;OK;HARD;1;CPU OK: 275 processes
[1551352733] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 94.7% (6962908 kB) free.
[1551352741] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 19.9% (1466420 kB) free!
[1551352765] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.7% (14897412 kB) free.
[1551352825] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.0% (6988572 kB) free.
[1551352879] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.7% (21984864 kB) free.
[1551353458] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42241 MB (20% inode=99%):
[1551353578] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46262 MB (22% inode=99%):
[1551353741] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8547600 kB) free!
[1551353861] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.3% (12456036 kB) free.
[1551354431] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1551354779] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42072 MB (20% inode=99%):
[1551354780] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46646 MB (20% inode=99%):
[1551354899] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46155 MB (22% inode=99%):
[1551354900] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46134 MB (20% inode=99%):
[1551355020] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51514 MB (22% inode=99%):
[1551355031] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551355061] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8404768 kB) free!
[1551355151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1551355181] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.9% (12279368 kB) free.
[1551355271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1551355305] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 617 processes
[1551355365] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 617 processes
[1551356069] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 42.77, 119.70, 62.26
[1551356099] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42006 MB (20% inode=99%):
[1551356189] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6.30, 80.29, 54.79
[1551356219] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46094 MB (22% inode=99%):
[1551356263] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4.68, 85.81, 66.00
[1551356309] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.13, 53.84, 48.19
[1551356382] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8535188 kB) free!
[1551356383] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.04, 57.93, 58.19
[1551356502] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11712472 kB) free.
[1551356503] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.47, 39.25, 51.31
[1551357102] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7884844 kB) free!
[1551357222] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4933364 kB) free!
[1551357342] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.4% (9057424 kB) free!
[1551357510] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.82, 1.58, 13.66
[1551357703] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.55, 2.26, 15.26
[1551358710] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.85, 0.66, 4.20
[1551359471] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1551359503] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.35, 1.37, 3.46
[1551359820] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46502 MB (20% inode=99%):
[1551359940] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44454 MB (19% inode=99%):
[1551360060] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50981 MB (22% inode=99%):
[1551360071] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1551360191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551360311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1551362111] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1551362711] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1551362831] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551362951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551363060] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43807 MB (19% inode=99%):
[1551363180] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51744 MB (23% inode=99%):
[1551363342] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1551365951] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1551366551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1551366671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1551366791] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551367380] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43804 MB (19% inode=99%):
[1551367500] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43804 MB (19% inode=99%):
[1551367542] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4903208 kB) free!
[1551367620] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51483 MB (22% inode=99%):
[1551369943] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.9% (7348908 kB) free!
[1551370620] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44097 MB (19% inode=99%):
[1551370740] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52288 MB (23% inode=99%):
[1551372343] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4239788 kB) free!
[1551373391] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1551373991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1551374112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1551374143] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1551374143] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.6% (7193164 kB) free!
[1551374232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1551374340] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46191 MB (20% inode=99%):
[1551374460] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52271 MB (23% inode=99%):
[1551375060] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44061 MB (19% inode=99%):
[1551375180] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52252 MB (23% inode=99%):
[1551376380] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43643 MB (19% inode=99%):
[1551376500] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52251 MB (23% inode=99%):
[1551376632] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551377100] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44059 MB (19% inode=99%):
[1551377144] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4629560 kB) free!
[1551377220] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52251 MB (23% inode=99%):
[1551377820] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46139 MB (20% inode=99%):
[1551377832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551377940] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46139 MB (20% inode=99%):
[1551377952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1551378060] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52251 MB (23% inode=99%):
[1551378073] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1551378344] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.8% (7275008 kB) free!