[1580927793] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23748872 kB) free!
[1580927844] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.0% (29592808 kB) free.
[1580927913] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43367704 kB) free.
[1580928513] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22861800 kB) free!
[1580928633] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42254792 kB) free.
[1580929284] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3382548 kB) free!
[1580929346] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188068 kB) free!
[1580929410] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (172712 kB) free!
[1580929470] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (187912 kB) free!
[1580929770] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 56.9% (13955920 kB) free.
[1580929833] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24339356 kB) free!
[1580929953] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43303952 kB) free.
[1580930553] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24083156 kB) free!
[1580930673] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43328700 kB) free.
[1580931100] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 35.7% (8798044 kB) free.
[1580931273] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23168352 kB) free!
[1580931393] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41542696 kB) free.
[1580931700] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187508 kB) free!
[1580931820] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200484 kB) free!
[1580931940] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168284 kB) free!
[1580932593] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23302284 kB) free!
[1580932713] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41027672 kB) free.
[1580935713] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24271572 kB) free!
[1580935833] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40757676 kB) free.
[1580936358] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235212 kB) free!
[1580936478] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (212744 kB) free!
[1580936549] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.5% (10179948 kB) free.
[1580936599] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234292 kB) free!
[1580937199] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.3% (41527748 kB) free.
[1580938459] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(245) < 103.16.101.72 [06/Feb/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1580938581] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580938659] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(318) < 2020/02/06 03:01:32 [error] 10211#0: *49466772 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580938779] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580938788] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2468) < 182.18.144.246 [06/Feb/2020:03:02:32 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1580938833] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25129388 kB) free!
[1580938908] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580938953] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39370056 kB) free.
[1580945553] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (157264 kB) free!
[1580945673] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166780 kB) free!
[1580945793] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190384 kB) free!
[1580946357] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.1% (21346376 kB) free.
[1580946940] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.8% (12043988 kB) free.
[1580949725] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 88.6% (116755728 kB) free.
[1580949766] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.9% (39853700 kB) free.
[1580949803] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.0% (22057060 kB) free.
[1580951140] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (3080484 kB) free!
[1580951260] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.5% (3079748 kB) free!
[1580951380] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.4% (3070024 kB) free!
[1580953780] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2450264 kB) free!
[1580956509] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1580956577] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1580956577] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22098040 kB) free.
[1580956629] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580956779] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(36) < 2020/02/06 08:09:35 [error] 10213#0: *49543632 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1580956899] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580957499] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(67) < 2020/02/06 08:21:26 [error] 10222#0: *49547961 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1580957619] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580958108] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 182.18.144.246 [06/Feb/2020:08:22:06 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1580958228] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580963830] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580963950] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580964087] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177352 kB) free!
[1580964135] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580964171] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1580964207] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196096 kB) free!
[1580964256] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580964291] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1580964327] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177580 kB) free!
[1580964357] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168792 kB) free!
[1580964376] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580964411] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1580964477] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189924 kB) free!
[1580964597] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189760 kB) free!
[1580964836] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.8% (21771332 kB) free.
[1580964875] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.1% (22077928 kB) free.
[1580964976] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1580964981] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1580965011] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1580969010] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163888 kB) free!
[1580969130] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180876 kB) free!
[1580969204] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.2% (21865296 kB) free.
[1580969207] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4170536 kB) free!
[1580969327] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.3% (3883852 kB) free!
[1580969447] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.6% (3505112 kB) free!
[1580970047] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1817412 kB) free!
[1580970409] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (1751944 kB) free!
[1580970507] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.0% (15712124 kB) free.
[1580970530] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 66.3% (16262236 kB) free.
[1580970767] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1580970819] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(137) < 2020/02/06 12:01:54 [error] 10221#0: *49640765 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580970887] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1580970939] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580970981] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.7% (4110388 kB) free!
[1580971101] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (3943508 kB) free!
[1580971221] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (3169764 kB) free!
[1580971735] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4755132 kB) free!
[1580971821] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (1137924 kB) free!
[1580971855] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4701588 kB) free!
[1580971975] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4618836 kB) free!
[1580972140] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/02/06 12:18:56 [error] 10218#0: *49652039 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.38412005109022873 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1580972260] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580972950] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580973070] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580973776] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2369768 kB) free!
[1580973827] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.2% (21527408 kB) free.
[1580973849] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.8% (21755316 kB) free.
[1580973878] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4826644 kB) free!
[1580973998] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4652584 kB) free!
[1580974119] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4485472 kB) free!
[1580975553] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12939028 kB) free!
[1580975568] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.2% (7968040 kB) free!
[1580975673] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40942656 kB) free.
[1580975688] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (7965996 kB) free!
[1580975808] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.3% (7033152 kB) free!
[1580975918] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2448900 kB) free!
[1580976078] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1580976198] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1580976273] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11713908 kB) free!
[1580976394] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.4% (49253040 kB) free.
[1580976993] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22766988 kB) free!
[1580977113] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.2% (50322376 kB) free.
[1580977375] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.8% (21777068 kB) free.
[1580977410] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39839260 kB) free.
[1580977427] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4757708 kB) free!
[1580977548] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.8% (4629920 kB) free!
[1580977668] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1580977668] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (4514460 kB) free!
[1580977713] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25850624 kB) free!
[1580977833] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.7% (48392428 kB) free.
[1580978433] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23502256 kB) free!
[1580978553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (47965276 kB) free.
[1580978599] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9028204 kB) free!
[1580978719] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8787364 kB) free!
[1580978776] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1580978839] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.2% (8500160 kB) free!
[1580978868] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (555020 kB) free!
[1580978896] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1580979153] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22578836 kB) free!
[1580979273] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.2% (47729640 kB) free.
[1580979496] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1580979615] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1580979616] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 4 crit, 0 warn out of 365 processes
[1580979736] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1580979737] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 0 warn out of 361 processes
[1580979778] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169100 kB) free!
[1580979856] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1580979873] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22030044 kB) free!
[1580979898] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190992 kB) free!
[1580979993] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.2% (46420128 kB) free.
[1580980019] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1580980019] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182800 kB) free!
[1580980110] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3419196 kB) free!
[1580980230] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3184972 kB) free!
[1580980350] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (2919268 kB) free!
[1580980440] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.6% (4568084 kB) free!
[1580980560] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.0% (4399992 kB) free!
[1580980593] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21782092 kB) free!
[1580980639] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4517956 kB) free!
[1580980680] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.3% (4238060 kB) free!
[1580980713] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19879096 kB) free!
[1580980833] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41576236 kB) free.
[1580980950] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1580980950] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1409196 kB) free!
[1580981470] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1580981550] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.2% (15752976 kB) free.
[1580981590] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1580981618] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (213032 kB) free!
[1580981628] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(160) < 92.42.104.166 [06/Feb/2020:14:57:54 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1580981738] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236744 kB) free!
[1580981749] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580981858] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240760 kB) free!
[1580982034] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20056904 kB) free!
[1580982137] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 359 processes
[1580982154] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.6% (21868600 kB) free!
[1580982274] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39625836 kB) free.
[1580982350] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (206964 kB) free!
[1580982470] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (206776 kB) free!
[1580982590] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (209176 kB) free!
[1580982737] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 355 processes
[1580982874] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15098132 kB) free!
[1580982994] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39836820 kB) free.
[1580983080] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2084788 kB) free!
[1580983246] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.8% (21760744 kB) free.
[1580983313] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.1% (20751300 kB) free.
[1580983370] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1580983370] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.5% (39646908 kB) free.
[1580983423] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.3% (43041944 kB) free.
[1580983451] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.2% (21629904 kB) free.
[1580983485] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.4% (14207108 kB) free.
[1580983594] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14950500 kB) free!
[1580983714] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40344012 kB) free.
[1580984058] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179180 kB) free!
[1580984178] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191744 kB) free!
[1580984298] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190044 kB) free!
[1580984314] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15168824 kB) free!
[1580984434] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39154672 kB) free.
[1580984519] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185976 kB) free!
[1580984639] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190704 kB) free!
[1580984759] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184196 kB) free!
[1580984894] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1580984894] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1580985034] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15418984 kB) free!
[1580985154] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11173548 kB) free!
[1580985274] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36843044 kB) free.
[1580985460] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2020/02/06 16:00:55 [error] 10214#0: *49777837 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580985494] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 356 processes
[1580985498] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 47.4% (11630336 kB) free.
[1580985580] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580985614] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1580985874] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11284740 kB) free!
[1580985994] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7446624 kB) free!
[1580986099] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (4322984 kB) free!
[1580986114] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33689984 kB) free.
[1580986214] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1580986219] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179104 kB) free!
[1580986334] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 4 crit, 0 warn out of 361 processes
[1580986339] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177632 kB) free!
[1580986454] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1580986494] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.4% (21556328 kB) free.
[1580986526] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.3% (21876784 kB) free.
[1580986569] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9166164 kB) free!
[1580986689] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (8997472 kB) free!
[1580986715] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8568560 kB) free!
[1580986809] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (8845300 kB) free!
[1580986835] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33315272 kB) free.
[1580987148] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(473) < 103.255.146.164 [06/Feb/2020:16:32:57 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1580987268] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580987435] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7592956 kB) free!
[1580987555] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3603476 kB) free!
[1580987675] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41874700 kB) free.
[1580988275] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18458920 kB) free!
[1580988395] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44605036 kB) free.
[1580988859] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4832812 kB) free!
[1580988887] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1580988979] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4706272 kB) free!
[1580988995] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19471956 kB) free!
[1580989007] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1580989099] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4626744 kB) free!
[1580989115] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28962960 kB) free.
[1580989715] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21614816 kB) free!
[1580989835] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17557968 kB) free!
[1580989955] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (44036100 kB) free.
[1580990054] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1580990094] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (1050020 kB) free!
[1580990130] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184920 kB) free!
[1580990174] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1580990215] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187520 kB) free!
[1580990250] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189564 kB) free!
[1580990335] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (245940 kB) free!
[1580990370] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182008 kB) free!
[1580990550] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.9% (3571740 kB) free!
[1580990555] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18976212 kB) free!
[1580990670] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3349992 kB) free!
[1580990675] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14827440 kB) free!
[1580990790] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3146316 kB) free!
[1580990795] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41576076 kB) free.
[1580991009] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4594784 kB) free!
[1580991371] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.8% (20676452 kB) free.
[1580991396] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17216052 kB) free!
[1580991424] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 88.3% (43531088 kB) free.
[1580991454] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.0% (21804396 kB) free.
[1580991516] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32873828 kB) free.
[1580991781] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1580991974] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1580991975] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191164 kB) free!
[1580991990] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1778792 kB) free!
[1580992095] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182872 kB) free!
[1580992116] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20033228 kB) free!
[1580992215] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177744 kB) free!
[1580992236] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15562840 kB) free!
[1580992356] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41337344 kB) free.
[1580992548] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.3% (21051816 kB) free.
[1580992580] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.1% (15750868 kB) free.
[1580992627] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.3% (22142152 kB) free.
[1580992956] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16696400 kB) free!
[1580993076] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12609128 kB) free!
[1580993157] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169804 kB) free!
[1580993174] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1580993196] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40103008 kB) free.
[1580993277] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175352 kB) free!
[1580993295] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1580993397] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173520 kB) free!
[1580993415] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1580993796] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15614088 kB) free!
[1580993861] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (2846832 kB) free!
[1580993916] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11724580 kB) free!
[1580993982] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (861124 kB) free!
[1580994036] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37763480 kB) free.
[1580994102] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172568 kB) free!
[1580994636] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13144216 kB) free!
[1580994756] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9367020 kB) free!
[1580994876] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36526148 kB) free.
[1580995476] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12334172 kB) free!
[1580995597] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38053880 kB) free.
[1580995798] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.3% (18087888 kB) free.
[1580996196] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14059016 kB) free!
[1580996316] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9798788 kB) free!
[1580996436] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37280532 kB) free.
[1580996581] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9615760 kB) free!
[1580996701] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 20.2% (9975596 kB) free.
[1580997015] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1580997036] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13555960 kB) free!
[1580997156] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9543972 kB) free!
[1580997276] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36697744 kB) free.
[1580997472] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2406776 kB) free!
[1580997592] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.6% (4479072 kB) free.
[1580997702] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 50.3% (12326168 kB) free.
[1580997856] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1580997876] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10775076 kB) free!
[1580997901] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8631968 kB) free!
[1580997996] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6229564 kB) free!
[1580998021] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8631780 kB) free!
[1580998116] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35810112 kB) free.
[1580998141] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (8627624 kB) free!
[1580998716] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10694672 kB) free!
[1580998836] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6762524 kB) free!
[1580998956] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34575700 kB) free.
[1580999556] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9694848 kB) free!
[1580999676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5656972 kB) free!
[1580999797] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33500280 kB) free.
[1580999868] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(15) < 203.212.70.170 [06/Feb/2020:20:02:25 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1580999980] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(52) < 2020/02/06 20:00:42 [error] 10218#0: *49908543 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580999988] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581000100] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581000396] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9209292 kB) free!
[1581000516] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5201204 kB) free!
[1581000636] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32752164 kB) free.
[1581001188] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(59) < 103.16.101.72 [06/Feb/2020:20:21:48 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1581001236] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7685308 kB) free!
[1581001309] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581001356] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3567048 kB) free!
[1581001476] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31474132 kB) free.
[1581002076] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5019428 kB) free!
[1581002188] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581002196] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14749204 kB) free!
[1581002316] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30871896 kB) free.
[1581002916] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7879988 kB) free!
[1581003036] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33127236 kB) free.
[1581003636] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10235820 kB) free!
[1581003756] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6450456 kB) free!
[1581003876] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32928836 kB) free.
[1581004477] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9831544 kB) free!
[1581004597] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6105032 kB) free!
[1581004717] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31982192 kB) free.
[1581005207] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1581005317] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9476708 kB) free!
[1581005327] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1581005398] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1581005437] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5828420 kB) free!
[1581005557] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30542748 kB) free.
[1581006048] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4867144 kB) free!
[1581006102] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1581006157] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7982048 kB) free!
[1581006168] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.8% (4854048 kB) free!
[1581006277] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32645384 kB) free.
[1581006288] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.7% (4838900 kB) free!
[1581006877] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9786720 kB) free!
[1581006997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5879636 kB) free!
[1581007117] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31506588 kB) free.
[1581007717] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7755412 kB) free!
[1581007837] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21362072 kB) free!
[1581007957] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28865320 kB) free.
[1581008341] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4926724 kB) free!
[1581008557] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5716016 kB) free!
[1581008677] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32179168 kB) free.
[1581009278] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8228136 kB) free!
[1581009398] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4331312 kB) free!
[1581009518] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29945216 kB) free.
[1581010118] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6498616 kB) free!
[1581010238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31900380 kB) free.
[1581010838] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9332832 kB) free!
[1581010958] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5476900 kB) free!
[1581011078] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30147588 kB) free.
[1581011678] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7274884 kB) free!
[1581011798] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31784804 kB) free.
[1581012398] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8913164 kB) free!
[1581012444] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2160888 kB) free!
[1581012518] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5250104 kB) free!
[1581012564] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240644 kB) free!
[1581012638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28364804 kB) free.
[1581012684] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (923184 kB) free!
[1581012790] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 410 processes
[1581012910] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 407 processes
[1581013238] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6937764 kB) free!
[1581013329] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.0% (18479252 kB) free!
[1581013358] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30987888 kB) free.
[1581013449] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.8% (16859968 kB) free!
[1581013569] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.4% (16363624 kB) free!