[1549478802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549478922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549479042] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549482042] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549482162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549482282] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549485282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549485402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549485522] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549488522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549488642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549488762] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549489762] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1549489882] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1549490002] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1549491762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549491882] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549492002] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549492402] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1549495002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549495122] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549495242] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549498243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549498363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549498483] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549501483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549501604] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549501724] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549505324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549505444] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549505565] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549508564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549508684] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549508804] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549510604] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549510724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549510844] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549511444] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549511564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549511684] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549511713] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.57, 14.63, 11.13
[1549511833] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.72, 10.30, 9.96
[1549511953] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.38, 7.31, 8.91
[1549512285] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549512885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549513005] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549513153] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.35, 1.53, 3.49
[1549514205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549514259] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.7% (22962016 kB) free.
[1549514325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549514925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549515046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549515166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549517859] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188084 kB) free!
[1549517979] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195092 kB) free!
[1549518099] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185176 kB) free!
[1549518896] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.1% (2780076 kB) free!
[1549519016] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.2% (2635976 kB) free!
[1549519136] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.5% (2524312 kB) free!
[1549519712] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.6% (22930548 kB) free.
[1549519736] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.7% (15059404 kB) free.
[1549521815] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 21 users currently logged in
[1549521935] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 21 users currently logged in
[1549522055] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 21 users currently logged in
[1549525713] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169924 kB) free!
[1549525833] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179628 kB) free!
[1549525953] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196684 kB) free!
[1549526415] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 409 processes
[1549526535] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1549526868] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 42 users currently logged in
[1549526928] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 42 users currently logged in
[1549526988] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 42 users currently logged in
[1549527048] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 42 users currently logged in
[1549527166] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549527766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549527886] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549528055] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 26 users currently logged in
[1549528486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549528607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549528727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549529927] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549530527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549530647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549530767] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1549530767] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549531968] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549532221] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.4% (22904720 kB) free.
[1549532336] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3113300 kB) free!
[1549532456] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (3046836 kB) free!
[1549532499] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.1% (15132488 kB) free.
[1549532568] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549532688] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549532808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549533408] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549533429] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175240 kB) free!
[1549533549] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178992 kB) free!
[1549533669] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193228 kB) free!
[1549533986] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1549533986] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 87.8% (21516760 kB) free.
[1549534008] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549534128] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549534249] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549539050] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1549539648] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 51 users currently logged in
[1549539650] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549539770] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549539890] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549539990] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (3168652 kB) free!
[1549540110] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.9% (3150028 kB) free!
[1549540230] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (3143728 kB) free!
[1549540490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549540610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549540730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549540830] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2393860 kB) free!
[1549541368] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.7% (22971764 kB) free.
[1549542531] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549543153] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 25.14, 33.01, 15.66
[1549543273] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.84, 22.61, 13.95
[1549543393] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.09, 15.69, 12.47
[1549543731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549543851] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549544594] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.97, 2.13, 4.75
[1549544748] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 50 users currently logged in
[1549545048] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 51 users currently logged in
[1549545648] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 50 users currently logged in
[1549545651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549545772] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549545948] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1549545948] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 51 users currently logged in
[1549546295] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1549546295] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1549546409] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 52 users currently logged in
[1549546744] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (27.8% change >= 20.0% threshold)
[1549546766] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 27 users currently logged in
[1549546971] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549547091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549547211] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549549113] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1549552113] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1549552611] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549552731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549552852] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549554194] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 54.05, 13.52, 5.57
[1549554314] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.33, 13.85, 6.87
[1549554434] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.97, 9.71, 6.20
[1549555035] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1549555035] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.36, 2.48, 3.90
[1549557652] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549557772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549557892] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549557962] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 537 processes
[1549558022] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 540 processes
[1549562092] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549562212] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549562332] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes