[1572719467] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1572719534] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.8% (22818016 kB) free.
[1572719537] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.8% (45231372 kB) free.
[1572719564] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 84.5% (18982432 kB) free.
[1572719738] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29352348 kB) free.
[1572719918] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1572720225] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572720319] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198048 kB) free!
[1572720338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3292548 kB) free!
[1572720440] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196496 kB) free!
[1572720458] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407808 kB) free!
[1572720560] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173116 kB) free!
[1572720578] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17612808 kB) free!
[1572721178] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396284 kB) free!
[1572721626] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.1% (33115228 kB) free.
[1572721678] HOST ALERT: sigqscluster2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.55)
[1572721723] SERVICE ALERT: sigqscluster2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sigqscluster2' was out of bounds)
[1572721925] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sigqscluster2' was out of bounds)
[1572721941] SERVICE ALERT: sigqscluster2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sigqscluster2' was out of bounds)
[1572722029] SERVICE ALERT: sigqscluster2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sigqscluster2' was out of bounds)
[1572722029] SERVICE ALERT: sigqscluster2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sigqscluster2' was out of bounds)
[1572722039] SERVICE ALERT: sigqscluster2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sigqscluster2' was out of bounds)
[1572722039] SERVICE ALERT: sigqscluster2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sigqscluster2' was out of bounds)
[1572722050] SERVICE ALERT: sigqscluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sigqscluster2' was out of bounds)
[1572722239] HOST ALERT: sigqscluster2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.55)
[1572722378] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27489708 kB) free.
[1572722826] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.0% (18503376 kB) free!
[1572722946] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10787768 kB) free!
[1572722978] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3597924 kB) free!
[1572723066] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4463756 kB) free!
[1572723098] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (398024 kB) free!
[1572723218] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (444644 kB) free!
[1572723516] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3497748 kB) free!
[1572723666] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.5% (23011412 kB) free!
[1572724266] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4891640 kB) free!
[1572724337] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229464 kB) free!
[1572724418] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21643404 kB) free!
[1572724457] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226700 kB) free!
[1572724577] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241424 kB) free!
[1572724795] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (2620780 kB) free!
[1572724866] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (3095200 kB) free!
[1572724943] HOST ALERT: sigqscluster2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1572725018] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388872 kB) free!
[1572725395] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2367348 kB) free!
[1572725466] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193688 kB) free!
[1572725569] HOST ALERT: sigqscluster2;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1572725633] HOST ALERT: sigqscluster2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1572726066] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26476316 kB) free.
[1572726149] SERVICE FLAPPING ALERT: sig3-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1572726218] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28449936 kB) free.
[1572726666] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.3% (18892020 kB) free!
[1572726786] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10621348 kB) free!
[1572726818] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3643404 kB) free!
[1572726844] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (4489308 kB) free!
[1572726906] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4107464 kB) free!
[1572726938] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403576 kB) free!
[1572726949] SERVICE ALERT: sigqscluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1572727058] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433468 kB) free!
[1572727090] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2019/11/03 02:08:08 [error] 24131#0: *2097225 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572727176] SERVICE ALERT: sigqscluster2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1572727210] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(19) < 2019/11/03 02:10:09 [error] 24130#0: *2097392 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572727243] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 0.17, 0.14, 0.14
[1572727327] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 325.89, 777.91, 518.23
[1572727331] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(17) < 2019/11/03 02:11:42 [error] 24130#0: *2097512 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572727335] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 11.9% (2807144 kB) free!
[1572727424] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 495 processes
[1572727436] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 96.4% (47499144 kB) free.
[1572727444] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 24.1% (5918728 kB) free.
[1572727447] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 48.37, 521.88, 455.87
[1572727455] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (2236024 kB) free!
[1572727462] SERVICE ALERT: sigqscluster2;Total Processes;OK;HARD;1;PROCS OK: 302 processes
[1572727463] SERVICE ALERT: sigqscluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 208589 MB (99% inode=99%):
[1572727474] SERVICE ALERT: sigqscluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45684 MB (89% inode=99%):
[1572727492] SERVICE ALERT: sigqscluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1572727575] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1622968 kB) free!
[1572727730] SERVICE ALERT: sigqscluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1572727800] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 16 users currently logged in
[1572727866] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.7% (18932048 kB) free.
[1572727901] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.6% (22441852 kB) free.
[1572727931] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572727954] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 92.6% (45645008 kB) free.
[1572727965] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.2% (2965900 kB) free!
[1572728021] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 58.1% (76576244 kB) free.
[1572728069] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (43973328 kB) free.
[1572728085] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.4% (2566452 kB) free!
[1572728106] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1572728106] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.5% (22735672 kB) free.
[1572728158] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 84.2% (18905252 kB) free.
[1572728185] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.4% (40594320 kB) free.
[1572728217] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.4% (45541748 kB) free.
[1572728247] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.8% (40827964 kB) free.
[1572728286] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.8% (22755480 kB) free.
[1572728473] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169904 kB) free!
[1572728523] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.40, 6.39, 5.07
[1572728593] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192256 kB) free!
[1572728643] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.75, 5.69, 4.97
[1572728713] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1572728713] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178872 kB) free!
[1572728997] SERVICE ALERT: sig3-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 564 processes
[1572729117] SERVICE ALERT: sig3-vm2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 564 processes
[1572729161] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236020 kB) free!
[1572729237] SERVICE ALERT: sig3-vm2;Total Processes;WARNING;HARD;3;PROCS WARNING: 564 processes
[1572729281] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233824 kB) free!
[1572729358] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1572729401] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224268 kB) free!
[1572729420] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231276 kB) free!
[1572729540] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.8% (44753572 kB) free.
[1572729843] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.20, 5.59, 5.08
[1572729869] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8205196 kB) free!
[1572729963] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.30, 5.61, 5.16
[1572729989] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4573912 kB) free!
[1572730001] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 86.8% (42793532 kB) free.
[1572730075] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.61, 8.02, 6.26
[1572730083] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.17, 5.14, 5.04
[1572730109] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34540536 kB) free.
[1572730119] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 17.0% (5925836 kB) free!
[1572730195] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.10, 7.78, 6.40
[1572730239] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.8% (32321024 kB) free.
[1572730315] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.91, 9.01, 7.04
[1572730351] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 32.53, 34.41, 30.29
[1572730471] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 26.13, 32.52, 30.18
[1572730591] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 32.03, 32.31, 30.36
[1572730610] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1669 processes
[1572730683] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.03, 4.26, 4.64
[1572730709] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9715516 kB) free!
[1572730730] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1668 processes
[1572730829] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6027780 kB) free!
[1572730850] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1668 processes
[1572730949] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2372768 kB) free!
[1572731191] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 28.43, 28.89, 29.80
[1572731713] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 66.6% (16438648 kB) free.
[1572732149] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28454144 kB) free.
[1572732244] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4888452 kB) free!
[1572732364] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4818856 kB) free!
[1572732391] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 28.93, 31.06, 30.27
[1572732484] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4770348 kB) free!
[1572732511] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 22.69, 28.31, 29.38
[1572732749] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4219228 kB) free!
[1572732869] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (558076 kB) free!
[1572732989] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397444 kB) free!
[1572734189] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21801432 kB) free!
[1572734789] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386260 kB) free!
[1572735451] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1572735989] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21390708 kB) free!
[1572736589] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399968 kB) free!
[1572736651] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8908900 kB) free!
[1572736771] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8285896 kB) free!
[1572736891] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.5% (7658228 kB) free!
[1572737109] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4478848 kB) free!
[1572737189] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15617732 kB) free!
[1572737229] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (3885020 kB) free!
[1572737285] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2315152 kB) free!
[1572737349] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.7% (3231176 kB) free!
[1572737789] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386860 kB) free!
[1572737952] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190580 kB) free!
[1572738691] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3063272 kB) free!
[1572738821] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1572738958] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2989268 kB) free!
[1572738989] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27834696 kB) free.
[1572739078] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.8% (2654016 kB) free!
[1572739198] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.3% (2305844 kB) free!
[1572739589] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4631136 kB) free!
[1572739709] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (962936 kB) free!
[1572739799] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (613032 kB) free!
[1572739829] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (442308 kB) free!
[1572740340] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9528220 kB) free!
[1572740439] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (5054816 kB) free!
[1572740461] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9172720 kB) free!
[1572740559] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 13.5% (4687588 kB) free!
[1572740581] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.9% (8835704 kB) free!
[1572740679] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 12.4% (4313764 kB) free!
[1572740713] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1572741029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26945212 kB) free.
[1572741092] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.8% (41315240 kB) free.
[1572741279] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 93.2% (32479436 kB) free.
[1572741313] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197816 kB) free!
[1572741433] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191000 kB) free!
[1572741489] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (2447500 kB) free!
[1572741553] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184248 kB) free!
[1572741555] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.4% (22724844 kB) free.
[1572741599] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 84.2% (18902336 kB) free.
[1572741609] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167692 kB) free!
[1572741629] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5983668 kB) free!
[1572741705] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188460 kB) free!
[1572741715] SERVICE ALERT: sig2-vm2;CPU Load;OK;HARD;1;OK - load average: 1.14, 2.76, 4.56
[1572741723] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.0% (19986536 kB) free.
[1572741729] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190464 kB) free!
[1572741749] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2682724 kB) free!
[1572741823] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.0% (22058532 kB) free.
[1572741866] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.4% (45350872 kB) free.
[1572741915] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.6% (20735292 kB) free.
[1572741978] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.1% (23060048 kB) free.
[1572742323] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191884 kB) free!
[1572742443] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163748 kB) free!
[1572742563] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1572742563] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186908 kB) free!
[1572743673] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15717632 kB) free!
[1572743794] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12534044 kB) free!
[1572743914] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9473832 kB) free!
[1572744514] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22338736 kB) free!
[1572745114] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9049448 kB) free!
[1572745205] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1572745325] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1572745557] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: /home 47110 MB (20% inode=99%):
[1572745677] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: /home 47110 MB (20% inode=99%):
[1572745714] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21292388 kB) free!
[1572745797] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: /home 47110 MB (20% inode=99%):
[1572745981] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4942076 kB) free!
[1572746314] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33517744 kB) free.
[1572746914] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19022744 kB) free!
[1572747034] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15995444 kB) free!
[1572747154] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (13066632 kB) free!
[1572747529] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.7% (22863900 kB) free.
[1572747570] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.1% (52912252 kB) free.
[1572747649] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.7% (47165856 kB) free.
[1572749374] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16283344 kB) free!
[1572749494] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (13241468 kB) free!
[1572749614] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33063672 kB) free.
[1572749955] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1572750214] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17760304 kB) free!
[1572750334] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14706624 kB) free!
[1572750455] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11671560 kB) free!
[1572750466] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 298 processes
[1572750586] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1572751055] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22921044 kB) free!
[1572751655] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33290456 kB) free.
[1572752255] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16337092 kB) free!
[1572752375] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12942448 kB) free!
[1572752495] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9440580 kB) free!
[1572753095] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20704280 kB) free!
[1572753533] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1572753695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31612676 kB) free.
[1572754295] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13654176 kB) free!
[1572754415] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10153388 kB) free!
[1572754535] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6629964 kB) free!
[1572755135] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16266132 kB) free!
[1572755735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27577436 kB) free.
[1572756335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7238456 kB) free!
[1572756455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3664840 kB) free!
[1572756575] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31449456 kB) free.
[1572757175] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14073172 kB) free!
[1572757295] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10167804 kB) free!
[1572757415] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5889556 kB) free!
[1572757583] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.1% (54177360 kB) free.
[1572758783] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25951060 kB) free!
[1572758904] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.8% (22084864 kB) free!
[1572758933] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (352404 kB) free!
[1572759024] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18191216 kB) free!
[1572759053] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163600 kB) free!
[1572759173] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191660 kB) free!
[1572759624] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30599068 kB) free.
[1572760224] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11067812 kB) free!
[1572760344] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7306980 kB) free!
[1572760464] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3657456 kB) free!
[1572760505] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.5% (17900972 kB) free.
[1572760553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.1% (54204508 kB) free.
[1572760872] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1572760992] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1572761115] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181256 kB) free!
[1572761235] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194456 kB) free!
[1572761355] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181684 kB) free!
[1572761756] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25344504 kB) free!
[1572761876] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (21007268 kB) free!
[1572761996] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16817272 kB) free!
[1572762596] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29862204 kB) free.
[1572762925] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.3% (3021572 kB) free!
[1572763045] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.8% (2892604 kB) free!
[1572763165] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2767004 kB) free!
[1572763196] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8126084 kB) free!
[1572763316] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4033400 kB) free!
[1572763436] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398940 kB) free!
[1572763765] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2248268 kB) free!
[1572764026] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4529508 kB) free!
[1572764146] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4438800 kB) free!
[1572764266] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4365048 kB) free!
[1572764306] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1572764306] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.7% (22392680 kB) free.
[1572764348] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43362544 kB) free.
[1572764393] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 80.8% (19807348 kB) free.
[1572764427] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.7% (21978924 kB) free.
[1572766155] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9086696 kB) free!
[1572766275] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4494508 kB) free!
[1572766325] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1572766395] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388400 kB) free!
[1572766445] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1572766672] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (2968956 kB) free!
[1572766793] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (2968348 kB) free!
[1572766913] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (2970600 kB) free!
[1572767419] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40109616 kB) free.
[1572768401] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9489296 kB) free!
[1572768521] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9430968 kB) free!
[1572768641] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9381780 kB) free!
[1572769224] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7151508 kB) free!
[1572769344] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3471740 kB) free!
[1572769464] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32510388 kB) free.
[1572769786] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 297 processes
[1572769906] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1572770065] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11464020 kB) free!
[1572770185] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7791952 kB) free!
[1572770305] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4299724 kB) free!
[1572770316] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181984 kB) free!
[1572770436] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187912 kB) free!
[1572770556] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176332 kB) free!
[1572770905] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17364172 kB) free!
[1572771031] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1572771152] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572771505] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31529784 kB) free.
[1572771641] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225172 kB) free!
[1572771657] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3033456 kB) free!
[1572771777] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237116 kB) free!
[1572771897] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228204 kB) free!
[1572772105] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11214140 kB) free!
[1572772225] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7560052 kB) free!
[1572772345] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3983820 kB) free!
[1572772892] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.0% (22449988 kB) free.
[1572772940] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 94.9% (46761520 kB) free.
[1572772945] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15933524 kB) free!
[1572772984] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.2% (47717396 kB) free.
[1572773040] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.6% (47113412 kB) free.
[1572774131] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(42) < 2019/11/03 15:12:09 [error] 24134#0: *2243288 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572774184] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18678736 kB) free!
[1572774192] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 369 processes
[1572774251] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2019/11/03 15:14:00 [error] 24130#0: *2243440 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572774304] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14634068 kB) free!
[1572774312] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 485 processes
[1572774371] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2019/11/03 15:14:16 [error] 24130#0: *2243424 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572774424] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33290704 kB) free.
[1572774971] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572775024] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13626332 kB) free!
[1572775144] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (9961060 kB) free!
[1572775264] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6296216 kB) free!
[1572775572] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(33) < 2019/11/03 15:31:41 [error] 24130#0: *2244164 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572775692] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572775864] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17443020 kB) free!
[1572776464] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30493296 kB) free.
[1572777064] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10124724 kB) free!
[1572777184] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5855640 kB) free!
[1572777304] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2136064 kB) free!
[1572777751] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1572777871] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1572779929] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19656700 kB) free!
[1572779932] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34210952 kB) free.
[1572780096] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178724 kB) free!
[1572780216] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182408 kB) free!
[1572780336] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187260 kB) free!
[1572780914] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.3% (23034212 kB) free.
[1572781784] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12041504 kB) free!
[1572781904] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8296804 kB) free!
[1572782024] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4867152 kB) free!
[1572782624] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14758812 kB) free!
[1572782786] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9726112 kB) free!
[1572782906] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9725128 kB) free!
[1572783026] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (9726244 kB) free!
[1572783224] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27775704 kB) free.
[1572783824] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7483108 kB) free!
[1572783944] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3626100 kB) free!
[1572784065] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32218692 kB) free.
[1572784664] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8472912 kB) free!
[1572784784] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4360688 kB) free!
[1572784904] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20095048 kB) free!
[1572785504] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (12003308 kB) free!
[1572786104] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23857584 kB) free!
[1572786595] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4907080 kB) free!
[1572786704] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3004024 kB) free!
[1572786715] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4860276 kB) free!
[1572786835] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4803988 kB) free!
[1572787304] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13940920 kB) free!
[1572788127] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189356 kB) free!
[1572788247] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178216 kB) free!
[1572788367] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172580 kB) free!
[1572788504] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4381860 kB) free!
[1572789104] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15850272 kB) free!
[1572789704] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27755548 kB) free.
[1572790141] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22814768 kB) free.
[1572790304] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7269704 kB) free!
[1572790424] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3603516 kB) free!
[1572790544] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31333692 kB) free.
[1572791144] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9365584 kB) free!
[1572791265] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5607196 kB) free!
[1572791385] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31790448 kB) free.
[1572791985] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11822776 kB) free!
[1572792105] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8121416 kB) free!
[1572792225] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4527148 kB) free!
[1572792235] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2409068 kB) free!
[1572792386] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4850628 kB) free!
[1572792506] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4836564 kB) free!
[1572792546] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199088 kB) free!
[1572792627] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (4817128 kB) free!
[1572792666] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188432 kB) free!
[1572792786] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179824 kB) free!
[1572792825] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14101768 kB) free!
[1572793429] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22841100 kB) free.
[1572793466] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 76.7% (18798208 kB) free.
[1572794025] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5891292 kB) free!
[1572794625] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15133976 kB) free!
[1572795233] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197644 kB) free!
[1572795353] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172528 kB) free!
[1572795473] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171252 kB) free!
[1572795825] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6482788 kB) free!
[1572796425] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15250260 kB) free!
[1572796804] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22504504 kB) free.
[1572797416] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175220 kB) free!
[1572797536] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199524 kB) free!
[1572797625] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (7088628 kB) free!
[1572797656] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189152 kB) free!
[1572798225] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16665756 kB) free!
[1572798826] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27380328 kB) free.
[1572799426] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10532684 kB) free!
[1572799546] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7369188 kB) free!
[1572799666] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31324268 kB) free.
[1572800266] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15182692 kB) free!
[1572800386] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12312684 kB) free!
[1572800506] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9530864 kB) free!
[1572801106] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20675312 kB) free!
[1572801706] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30647148 kB) free.
[1572802306] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14229976 kB) free!
[1572802426] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11095216 kB) free!
[1572802546] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (8104576 kB) free!
[1572802850] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 972 processes
[1572803146] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17888192 kB) free!
[1572803450] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1648 processes
[1572803571] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1647 processes
[1572803691] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1647 processes
[1572803746] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (28003128 kB) free.
[1572804271] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1572804346] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11556568 kB) free!
[1572804391] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1572804466] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8468924 kB) free!
[1572804587] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29906756 kB) free.
[1572805186] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12409580 kB) free!
[1572805306] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9111564 kB) free!
[1572805426] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (29992108 kB) free.