[1580841092] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46817796 kB) free.
[1580841202] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.9% (29016772 kB) free.
[1580841567] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188720 kB) free!
[1580841688] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169624 kB) free!
[1580841692] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24946452 kB) free!
[1580841808] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187324 kB) free!
[1580841812] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.5% (48143808 kB) free.
[1580841817] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189552 kB) free!
[1580841937] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182116 kB) free!
[1580842057] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184836 kB) free!
[1580842509] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1580842629] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1580842798] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4902764 kB) free!
[1580843725] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.9% (14432688 kB) free!
[1580844451] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.3% (9948488 kB) free.
[1580844925] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (11112712 kB) free!
[1580845051] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.3% (2779188 kB) free!
[1580845171] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (1226136 kB) free!
[1580845291] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184972 kB) free!
[1580845697] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.3% (7071652 kB) free!
[1580845817] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (6061984 kB) free!
[1580845937] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (5188740 kB) free!
[1580846125] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1580846125] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.1% (14679876 kB) free!
[1580846537] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222456 kB) free!
[1580849008] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.7% (10217344 kB) free.
[1580849529] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5421808 kB) free!
[1580849649] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (1735316 kB) free!
[1580849769] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (1729788 kB) free!
[1580850369] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 77.8% (38335996 kB) free.
[1580850598] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 80.4% (39633052 kB) free.
[1580851938] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(226) < 103.16.101.72 [05/Feb/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1580852018] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(290) < 2020/02/05 03:01:32 [error] 10208#0: *48981638 upstream timed out (110: Connection timed out) while connecting to upstream, client: 182.18.144.246, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580852027] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1966) < 172.105.59.173 [05/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1580852060] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580852138] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580852147] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580856925] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1580859208] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168244 kB) free!
[1580859329] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183244 kB) free!
[1580859338] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(34) < 2020/02/05 04:56:57 [error] 10218#0: *49005097 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580859449] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168764 kB) free!
[1580859458] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580860057] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.6% (20744024 kB) free.
[1580860291] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.2% (11141752 kB) free.
[1580863525] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 34.8% (45808428 kB) free.
[1580864126] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.3% (21499436 kB) free!
[1580864221] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.9% (39852844 kB) free.
[1580864246] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.1% (46263552 kB) free.
[1580864284] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.0% (22048916 kB) free.
[1580865092] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192188 kB) free!
[1580865212] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24683060 kB) free!
[1580865212] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185972 kB) free!
[1580865332] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45223636 kB) free.
[1580865332] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174916 kB) free!
[1580866614] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.0% (22212680 kB) free.
[1580870148] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(16) < 203.212.70.170 [05/Feb/2020:08:02:35 +0530] TCP 502 0 0 63.009 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 3.008, 60.000, 0.000"
[1580870258] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(122) < 2020/02/05 08:02:32 [error] 10222#0: *49049322 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580870268] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580870295] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1580870378] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580870415] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1580871579] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2020/02/05 08:21:16 [error] 10216#0: *49055177 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580871699] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580875068] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 203.212.70.170 [05/Feb/2020:09:22:12 +0530] TCP 502 0 0 135.029 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 15.028, 0.000"
[1580875188] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580878606] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1580878726] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1580880429] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1580880550] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1580880932] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25092628 kB) free!
[1580881052] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19948124 kB) free!
[1580881173] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15391404 kB) free!
[1580881620] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.0% (3453396 kB) free!
[1580881740] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.0% (3217732 kB) free!
[1580881748] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.9% (4151012 kB) free!
[1580881860] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (3002096 kB) free!
[1580881868] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.8% (3868280 kB) free!
[1580881989] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.6% (3576804 kB) free!
[1580882287] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188224 kB) free!
[1580882407] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164324 kB) free!
[1580882415] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1580882460] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (428184 kB) free!
[1580882528] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164352 kB) free!
[1580882535] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1580882655] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 339 processes
[1580882973] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39655396 kB) free.
[1580883131] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (3768412 kB) free!
[1580883189] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1922000 kB) free!
[1580883251] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.5% (3476136 kB) free!
[1580883371] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3153604 kB) free!
[1580883971] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1620796 kB) free!
[1580884455] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1580884571] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.6% (15846252 kB) free.
[1580884773] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.4% (21809804 kB) free.
[1580884788] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [05/Feb/2020:12:01:54 +0530] TCP 502 0 0 60.004 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.003, 0.000"
[1580884858] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.3% (21892588 kB) free.
[1580884909] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580884940] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21623692 kB) free.
[1580885598] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1580885637] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (1129476 kB) free!
[1580885718] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1580885757] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (1127472 kB) free!
[1580885838] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1580885877] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1127616 kB) free!
[1580886163] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.9% (14120020 kB) free.
[1580886438] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1580887173] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14737168 kB) free!
[1580887293] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39526508 kB) free.
[1580887893] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14071588 kB) free!
[1580888013] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40119436 kB) free.
[1580888228] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (249600 kB) free!
[1580888348] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241804 kB) free!
[1580888381] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192256 kB) free!
[1580888388] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39418180 kB) free.
[1580888424] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.4% (21321512 kB) free.
[1580888614] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13970896 kB) free!
[1580888734] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39522000 kB) free.
[1580888999] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9320000 kB) free!
[1580889029] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192620 kB) free!
[1580889119] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9039704 kB) free!
[1580889150] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165640 kB) free!
[1580889239] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.7% (8742712 kB) free!
[1580889270] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187920 kB) free!
[1580889598] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (211776 kB) free!
[1580889718] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230104 kB) free!
[1580889837] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39229056 kB) free.
[1580889934] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18326420 kB) free!
[1580890054] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40264224 kB) free.
[1580890269] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176140 kB) free!
[1580890308] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(596) < 103.255.146.164 [05/Feb/2020:13:31:59 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1580890389] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170824 kB) free!
[1580890428] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580890510] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187556 kB) free!
[1580890572] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1580890572] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21797252 kB) free.
[1580890654] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16592516 kB) free!
[1580890691] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.8% (43276420 kB) free.
[1580890774] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38786632 kB) free.
[1580891257] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4797348 kB) free!
[1580891374] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (14998300 kB) free!
[1580891377] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4763988 kB) free!
[1580891494] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33629032 kB) free.
[1580891497] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (4557264 kB) free!
[1580893749] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1580893869] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580894171] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.7% (21746860 kB) free.
[1580894172] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4403796 kB) free!
[1580894206] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.5% (15834680 kB) free.
[1580894234] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.8% (21269012 kB) free.
[1580894778] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4891324 kB) free!
[1580894838] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 712 processes
[1580894898] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4740512 kB) free!
[1580894958] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 715 processes
[1580895019] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4594384 kB) free!
[1580895078] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 713 processes
[1580896219] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (979652 kB) free!
[1580896269] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1580896278] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1580896278] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1580896294] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20124824 kB) free!
[1580896390] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580896414] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43640956 kB) free.
[1580896521] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.7% (21645392 kB) free.
[1580897014] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19873380 kB) free!
[1580897028] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(939) < 103.255.146.164 [05/Feb/2020:15:30:51 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1580897134] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43241076 kB) free.
[1580897148] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580897499] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/05 15:39:51 [error] 10222#0: *49272730 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1580897620] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580897637] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.9% (7339960 kB) free!
[1580897655] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1580897727] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (1293988 kB) free!
[1580897734] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19555352 kB) free!
[1580897757] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.8% (6317764 kB) free!
[1580897775] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1580897848] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (1187108 kB) free!
[1580897854] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (43000240 kB) free.
[1580897877] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (6318608 kB) free!
[1580897895] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1580897968] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (1117860 kB) free!
[1580898439] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191436 kB) free!
[1580898454] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19272272 kB) free!
[1580898559] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163364 kB) free!
[1580898574] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42416992 kB) free.
[1580898679] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169196 kB) free!
[1580898819] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(131) < 2020/02/05 16:02:06 [error] 10208#0: *49286042 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580898939] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580898948] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(289) < 92.42.104.166 [05/Feb/2020:16:04:35 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1580898978] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165200 kB) free!
[1580899068] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580899098] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198140 kB) free!
[1580899174] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18296880 kB) free!
[1580899218] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167048 kB) free!
[1580899229] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21875468 kB) free.
[1580899261] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1580899261] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.4% (22165392 kB) free.
[1580899294] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42416912 kB) free.
[1580899303] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1580899303] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.8% (21530580 kB) free.
[1580899332] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39417812 kB) free.
[1580899539] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/05 16:12:11 [error] 10222#0: *49292663 open() "/usr/share/nginx/html/js/tmpl/login.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/login.ejs?0.6131470382866637 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1580899659] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580899894] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17846376 kB) free!
[1580900014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41627448 kB) free.
[1580900439] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (364436 kB) free!
[1580900508] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190476 kB) free!
[1580900559] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194852 kB) free!
[1580900614] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17571212 kB) free!
[1580900628] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179216 kB) free!
[1580900679] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (162668 kB) free!
[1580900734] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40750512 kB) free.
[1580900748] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182788 kB) free!
[1580901335] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17693216 kB) free!
[1580901455] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13766052 kB) free!
[1580901495] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1580901575] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37934608 kB) free.
[1580901770] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (2307356 kB) free!
[1580901885] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.5% (20859232 kB) free.
[1580901890] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.2% (2314184 kB) free!
[1580901909] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.9% (21800692 kB) free.
[1580901943] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.8% (14108932 kB) free.
[1580902175] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14475536 kB) free!
[1580902271] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.8% (2642524 kB) free!
[1580902295] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38577340 kB) free.
[1580902391] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.1% (2478904 kB) free!
[1580902512] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2195504 kB) free!
[1580902519] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.7% (2630564 kB) free!
[1580902639] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.5% (2579280 kB) free!
[1580902668] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.3% (22121512 kB) free.
[1580902721] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.8% (21769800 kB) free.
[1580902727] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1580902847] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1580902895] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14182240 kB) free!
[1580902967] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1580903015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38359584 kB) free.
[1580903087] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (577528 kB) free!
[1580903207] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (500224 kB) free!
[1580903271] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.4% (3048016 kB) free!
[1580903327] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184568 kB) free!
[1580903327] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (2267168 kB) free!
[1580903391] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.9% (2919388 kB) free!
[1580903447] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (2205524 kB) free!
[1580903511] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2815796 kB) free!
[1580903568] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2145596 kB) free!
[1580903615] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14750124 kB) free!
[1580903735] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10719520 kB) free!
[1580903823] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.8% (21773920 kB) free.
[1580903850] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.3% (22285700 kB) free.
[1580903855] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34773116 kB) free.
[1580903893] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.1% (22090276 kB) free.
[1580903897] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9087564 kB) free!
[1580903948] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.0% (43362444 kB) free.
[1580904153] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.7% (4819000 kB) free!
[1580904273] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.3% (4741800 kB) free!
[1580904393] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.1% (4670540 kB) free!
[1580904455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10985280 kB) free!
[1580904575] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36256312 kB) free.
[1580905175] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11696872 kB) free!
[1580905295] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35836572 kB) free.
[1580905895] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12799292 kB) free!
[1580906016] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8603744 kB) free!
[1580906136] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33662920 kB) free.
[1580906295] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1580906415] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1580906535] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1580906736] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9487764 kB) free!
[1580906856] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34398300 kB) free.
[1580907078] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1580907394] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2426480 kB) free!
[1580907427] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184320 kB) free!
[1580907456] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9953288 kB) free!
[1580907547] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171792 kB) free!
[1580907576] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6052872 kB) free!
[1580907667] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186276 kB) free!
[1580907696] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31793084 kB) free.
[1580908281] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21651556 kB) free.
[1580908296] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7917308 kB) free!
[1580908317] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.9% (21539540 kB) free.
[1580908416] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31287520 kB) free.
[1580908935] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1580908935] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1580912071] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21920816 kB) free!
[1580912191] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18020912 kB) free!
[1580912311] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44171212 kB) free.
[1580912911] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20589692 kB) free!
[1580913031] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45598732 kB) free.
[1580913460] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4574720 kB) free!
[1580913468] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 203.212.70.170 [05/Feb/2020:20:02:55 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1580913580] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4559000 kB) free!
[1580913588] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580913631] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23599296 kB) free!
[1580913700] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.1% (3477304 kB) free!
[1580913736] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9873492 kB) free!
[1580913751] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19728572 kB) free!
[1580913856] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 20.6% (10162292 kB) free.
[1580913871] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45098776 kB) free.
[1580914471] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22248040 kB) free!
[1580914591] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45148668 kB) free.
[1580914708] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1580914900] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2394104 kB) free!
[1580915056] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9020028 kB) free!
[1580915176] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9010140 kB) free!
[1580915191] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23643388 kB) free!
[1580915296] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (9008088 kB) free!
[1580915311] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19896688 kB) free!
[1580915431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44476120 kB) free.
[1580915908] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201024 kB) free!
[1580916028] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173044 kB) free!
[1580916031] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21569044 kB) free!
[1580916117] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (230768 kB) free!
[1580916148] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202292 kB) free!
[1580916151] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44820912 kB) free.
[1580916237] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (225352 kB) free!
[1580916357] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (226152 kB) free!
[1580916752] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22273196 kB) free!
[1580916872] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18129184 kB) free!
[1580916992] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43718672 kB) free.
[1580917592] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21047700 kB) free!
[1580917712] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17431604 kB) free!
[1580917832] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41566960 kB) free.
[1580918432] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17727612 kB) free!
[1580918535] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1580918552] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (44046908 kB) free.
[1580919152] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19907272 kB) free!
[1580919272] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43893852 kB) free.
[1580919872] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22485860 kB) free!
[1580919992] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18791148 kB) free!
[1580920112] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42445900 kB) free.
[1580920712] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19645832 kB) free!
[1580920832] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43643692 kB) free.
[1580921433] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21246140 kB) free!
[1580921553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43515380 kB) free.
[1580922153] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22348252 kB) free!
[1580922189] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1580922273] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18826804 kB) free!
[1580922310] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580922357] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1580922393] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40614708 kB) free.
[1580922496] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4880184 kB) free!
[1580922993] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19613964 kB) free!
[1580923113] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41016520 kB) free.
[1580923713] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19901124 kB) free!
[1580923833] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41473468 kB) free.
[1580924433] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20213816 kB) free!
[1580924553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40716500 kB) free.
[1580925700] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1580925753] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22423052 kB) free!
[1580925873] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43343392 kB) free.
[1580926404] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220256 kB) free!
[1580926473] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22236664 kB) free!
[1580926524] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219096 kB) free!
[1580926593] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42651652 kB) free.
[1580926644] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238608 kB) free!
[1580926646] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1200348 kB) free!
[1580926766] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (466180 kB) free!
[1580926886] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446160 kB) free!