[1569090979] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569091076] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569091099] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 373 processes
[1569091535] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192404 kB) free!
[1569091600] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198056 kB) free!
[1569091655] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183352 kB) free!
[1569091676] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569091720] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189080 kB) free!
[1569091775] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183396 kB) free!
[1569091840] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177048 kB) free!
[1569092994] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (985296 kB) free!
[1569093114] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (985032 kB) free!
[1569093234] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (988472 kB) free!
[1569093357] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1569093477] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1569093950] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569094077] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569094197] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569094214] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 487 processes
[1569094274] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 481 processes
[1569094335] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 479 processes
[1569094551] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 279 processes
[1569095040] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (385060 kB) free!
[1569095132] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (230468 kB) free!
[1569095148] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 449 processes
[1569095160] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (379628 kB) free!
[1569095178] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8123120 kB) free!
[1569095253] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (253624 kB) free!
[1569095265] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 185516 MB (17% inode=99%):
[1569095280] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385280 kB) free!
[1569095298] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.0% (6917112 kB) free!
[1569095373] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (247960 kB) free!
[1569095385] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 180012 MB (16% inode=99%):
[1569095418] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (6293228 kB) free!
[1569095505] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 180012 MB (16% inode=99%):
[1569095751] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 274 processes
[1569096018] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4160700 kB) free!
[1569096476] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569097676] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569099064] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229968 kB) free!
[1569100264] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.8% (42271564 kB) free.
[1569100677] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569101277] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569101877] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569102568] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 45.2% (11072600 kB) free.
[1569102640] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 46.4% (11370468 kB) free.
[1569103077] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569104278] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569104878] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569105478] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569106678] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569107279] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569108479] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569109079] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569110279] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569110880] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569112080] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569112557] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 89.11, 32.71, 12.30
[1569112677] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 105.14, 57.51, 23.86
[1569112680] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569112774] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 94.6% (33145680 kB) free.
[1569112797] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 107.34, 74.09, 34.02
[1569112797] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 408 processes
[1569112918] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 409 processes
[1569113152] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 94.58, 33.75, 12.67
[1569113235] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 494 processes
[1569113272] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 113.17, 60.99, 25.21
[1569113295] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 492 processes
[1569113392] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 115.49, 79.13, 36.19
[1569113997] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.34, 4.15, 16.10
[1569115080] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 56.3% (74190572 kB) free.
[1569115175] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 46.7% (11532732 kB) free.
[1569115192] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.93, 1.33, 10.87
[1569115305] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 260732 MB (24% inode=99%):
[1569115680] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569115797] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.44, 2.96
[1569116280] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1569116368] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.5% (3790840 kB) free!
[1569116375] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (1509140 kB) free!
[1569116392] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.44, 1.01, 3.62
[1569116440] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192940 kB) free!
[1569116496] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201656 kB) free!
[1569116497] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187892 kB) free!
[1569116561] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174504 kB) free!
[1569116616] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178292 kB) free!
[1569116618] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189620 kB) free!
[1569116681] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196404 kB) free!
[1569117481] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569118081] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569119281] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569119881] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569120117] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569120237] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569120767] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 987.09, 316.82, 113.22
[1569120887] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1290.83, 654.74, 261.74
[1569121007] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1331.84, 880.91, 392.27
[1569121081] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569121681] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569124646] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.3% (23511484 kB) free.
[1569124681] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569124695] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.8% (23245380 kB) free.
[1569124744] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.4% (40096808 kB) free.
[1569124790] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.9% (23020480 kB) free.
[1569124831] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.6% (14404196 kB) free.
[1569125207] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 0.74, 15.14
[1569125281] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569125637] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1569125757] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1569126407] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.63, 0.75, 4.71
[1569126481] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569127081] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569128281] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569128807] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.03, 6.60, 5.55
[1569128881] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569128928] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.44, 4.56, 4.94
[1569129481] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569130681] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569131881] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569132481] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569133681] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569134282] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569134757] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569134877] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1569135482] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569136082] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569136307] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.02, 6.02, 5.22
[1569136427] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.53, 6.20, 5.38
[1569136549] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.70, 6.05, 5.43
[1569137149] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.24, 3.34, 4.45
[1569137282] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569137883] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569138455] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (4184084 kB) free!
[1569138483] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569138575] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (4161896 kB) free!
[1569138695] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4139300 kB) free!
[1569139683] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569140495] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (680324 kB) free!
[1569140884] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 605 processes
[1569140954] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (248712 kB) free!
[1569141075] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (250972 kB) free!
[1569141195] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (249244 kB) free!
[1569141484] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569142684] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569143284] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569144485] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569145085] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569145631] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.4% (23056428 kB) free.
[1569145673] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.6% (40206860 kB) free.
[1569146277] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569146285] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569146397] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 404 processes
[1569146518] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1569146885] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569148086] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569148549] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 346 processes
[1569148686] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1569148864] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9767116 kB) free!
[1569148927] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (2276080 kB) free!
[1569148984] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9703396 kB) free!
[1569149047] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (1725556 kB) free!
[1569149104] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (9644944 kB) free!
[1569149167] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1718520 kB) free!
[1569149517] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569149637] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1569149750] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1569149767] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.0% (3675240 kB) free!
[1569149886] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569150486] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569150950] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1569151550] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1569151550] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1569151567] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2115064 kB) free!
[1569151686] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569152167] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1569152167] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.3% (3984520 kB) free!
[1569152637] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1569152757] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 403 processes
[1569152877] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 402 processes
[1569152886] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569153477] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569153486] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569153967] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (627704 kB) free!
[1569154087] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569155167] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.3% (4473848 kB) free!
[1569155287] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 605 processes
[1569155750] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 346 processes
[1569155887] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569156350] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1569156367] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 24.4% (5972612 kB) free.
[1569156594] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4751344 kB) free!
[1569156715] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4720504 kB) free!
[1569156835] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4690560 kB) free!
[1569156967] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181248 kB) free!
[1569157087] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 25.8% (6320132 kB) free.
[1569157087] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569157505] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4772600 kB) free!
[1569157687] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569157835] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 55.5% (27345176 kB) free.
[1569158287] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569160127] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 367.68, 234.76, 102.75
[1569160247] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 369.19, 279.36, 135.10
[1569160367] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 206.37, 268.45, 149.90
[1569162597] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 401.31, 239.07, 101.32
[1569162717] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 404.86, 294.12, 138.19
[1569162767] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.90, 0.95, 12.06
[1569162837] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 149.08, 257.24, 145.25
[1569163967] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.95, 0.94, 3.98
[1569164392] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 347.56, 274.52, 137.53
[1569164512] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 347.77, 298.74, 163.02
[1569164567] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.68, 23.01, 12.24
[1569164632] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 71.65, 220.36, 151.64
[1569164687] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.92, 15.57, 10.82
[1569164807] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1569164807] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.88, 10.66, 9.59
[1569165238] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.43, 5.86, 13.68
[1569166007] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.47, 0.78, 3.08
[1569166438] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.66, 0.72, 4.19
[1569167032] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 0.76, 12.13
[1569167287] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569167750] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569168232] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.61, 0.89, 4.03
[1569169687] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569171235] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2410432 kB) free!
[1569172678] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569172704] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4880820 kB) free!
[1569172798] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1569172824] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4868788 kB) free!
[1569172944] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4858280 kB) free!
[1569176807] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)