[1572806026] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12401904 kB) free!
[1572806112] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1572806146] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31467284 kB) free.
[1572806232] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1572806434] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197040 kB) free!
[1572806554] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185480 kB) free!
[1572806674] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182340 kB) free!
[1572806674] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186356 kB) free!
[1572806746] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15017888 kB) free!
[1572806794] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (162540 kB) free!
[1572806866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (11523008 kB) free!
[1572806914] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174112 kB) free!
[1572806986] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31485224 kB) free.
[1572807256] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1572807586] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13423144 kB) free!
[1572807706] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9847832 kB) free!
[1572807826] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29873572 kB) free.
[1572808426] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12643732 kB) free!
[1572808546] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26771900 kB) free.
[1572809146] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15207048 kB) free!
[1572809266] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11747904 kB) free!
[1572809387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31506416 kB) free.
[1572809986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13230292 kB) free!
[1572810026] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2774336 kB) free!
[1572810106] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31520996 kB) free.
[1572810706] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14803744 kB) free!
[1572810826] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11178560 kB) free!
[1572810946] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30407608 kB) free.
[1572811546] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11242008 kB) free!
[1572811666] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30274844 kB) free.
[1572812266] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11513076 kB) free!
[1572812386] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30550648 kB) free.
[1572812656] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.8% (5617096 kB) free.
[1572812986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12112684 kB) free!
[1572813106] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31044912 kB) free.
[1572813256] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168816 kB) free!
[1572813376] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163432 kB) free!
[1572813496] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170296 kB) free!
[1572813706] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14266240 kB) free!
[1572813826] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10929404 kB) free!
[1572813946] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30247216 kB) free.
[1572813991] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1572814112] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1572814445] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.0% (6877064 kB) free!
[1572814546] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13425572 kB) free!
[1572814565] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (341876 kB) free!
[1572814666] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31586444 kB) free.
[1572814685] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223664 kB) free!
[1572814945] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219228 kB) free!
[1572815065] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216824 kB) free!
[1572815185] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.1% (44872836 kB) free.
[1572815267] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15145132 kB) free!
[1572815387] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11769576 kB) free!
[1572815507] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30612776 kB) free.
[1572815885] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.4% (46038220 kB) free.
[1572816107] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12689748 kB) free!
[1572816227] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31335616 kB) free.
[1572816827] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12771752 kB) free!
[1572816947] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31084632 kB) free.
[1572817096] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 60.7% (14982148 kB) free.
[1572817547] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12501224 kB) free!
[1572817667] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30704456 kB) free.
[1572817692] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/11/04 03:13:12 [error] 24134#0: *2374435 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572817812] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572818267] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12486804 kB) free!
[1572818387] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572818387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30653360 kB) free.
[1572818832] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1572818952] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 449 processes
[1572818987] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12334432 kB) free!
[1572819013] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2019/11/04 03:39:36 [error] 24130#0: *2375177 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572819107] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30330500 kB) free.
[1572819133] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572819733] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(20) < 2019/11/04 03:43:52 [error] 24130#0: *2375379 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1572819853] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(6) < 2019/11/04 03:54:11 [error] 24134#0: *2375916 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572819973] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(4) < 2019/11/04 03:56:11 [error] 24130#0: *2376130 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572820307] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14876032 kB) free!
[1572820427] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31635184 kB) free.
[1572820474] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2776948 kB) free!
[1572821028] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14417176 kB) free!
[1572821148] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31652988 kB) free.
[1572821173] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572821748] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13085116 kB) free!
[1572821868] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30759432 kB) free.
[1572823068] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15137292 kB) free!
[1572823188] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31680616 kB) free.
[1572823788] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14159520 kB) free!
[1572823908] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31693288 kB) free.
[1572824508] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14008624 kB) free!
[1572824628] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31691732 kB) free.
[1572825228] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13301308 kB) free!
[1572825348] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30856392 kB) free.
[1572825948] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12764600 kB) free!
[1572826068] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30430140 kB) free.
[1572826474] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2453468 kB) free!
[1572826668] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12638592 kB) free!
[1572826789] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30050796 kB) free.
[1572827388] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14112564 kB) free!
[1572827508] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31667648 kB) free.
[1572827773] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/04 06:00:40 [error] 24130#0: *2398946 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1572827893] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572827896] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187980 kB) free!
[1572828016] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167588 kB) free!
[1572828027] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.6% (21224780 kB) free.
[1572828108] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16119616 kB) free!
[1572828136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196560 kB) free!
[1572828221] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.6% (19238820 kB) free!
[1572828228] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13284532 kB) free!
[1572828274] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.5% (21208252 kB) free.
[1572828341] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.2% (18703608 kB) free!
[1572828348] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31629408 kB) free.
[1572828461] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.7% (18103348 kB) free!
[1572828514] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 66.8% (16369628 kB) free.
[1572828948] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16346176 kB) free!
[1572829061] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 36.3% (47799312 kB) free.
[1572829068] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12901592 kB) free!
[1572829188] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31537576 kB) free.
[1572829278] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22954304 kB) free.
[1572829323] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.3% (40546764 kB) free.
[1572829661] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.6% (23197472 kB) free!
[1572829693] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2019/11/04 06:31:05 [error] 24130#0: *2404679 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572829781] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.1% (18641292 kB) free!
[1572829788] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15634560 kB) free!
[1572829814] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572829901] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (13756068 kB) free!
[1572829908] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12163100 kB) free!
[1572830028] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31444988 kB) free.
[1572830413] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/04 06:45:38 [error] 24130#0: *2357570 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:9844/9809, bytes from/to upstream:9809/9844
[1572830501] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1572830501] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9256092 kB) free!
[1572830533] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572830628] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14268604 kB) free!
[1572830748] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11163904 kB) free!
[1572830868] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30278984 kB) free.
[1572831101] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 54.3% (71617872 kB) free.
[1572831468] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14897924 kB) free!
[1572831589] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11798816 kB) free!
[1572831709] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30281608 kB) free.
[1572832309] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14743732 kB) free!
[1572832429] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11736404 kB) free!
[1572832549] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30055000 kB) free.
[1572833149] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15201388 kB) free!
[1572833269] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12086336 kB) free!
[1572833389] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30590128 kB) free.
[1572833989] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15684032 kB) free!
[1572834109] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12755144 kB) free!
[1572834229] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31090228 kB) free.
[1572834829] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14570344 kB) free!
[1572834949] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11235444 kB) free!
[1572835069] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30403532 kB) free.
[1572835669] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13779656 kB) free!
[1572835790] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30903192 kB) free.
[1572836389] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15833924 kB) free!
[1572836398] SERVICE ALERT: webservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: /home 105950 MB (47% inode=99%):
[1572836509] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12454168 kB) free!
[1572836629] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30823616 kB) free.
[1572837229] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14037828 kB) free!
[1572837349] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10676452 kB) free!
[1572837469] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29394284 kB) free.
[1572838069] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12585988 kB) free!
[1572838189] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30666844 kB) free.
[1572838789] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11294236 kB) free!
[1572838909] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7495084 kB) free!
[1572839029] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28786756 kB) free.
[1572839629] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9146992 kB) free!
[1572839749] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48591708 kB) free.
[1572840958] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24019444 kB) free!
[1572841078] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20337632 kB) free!
[1572841198] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16696172 kB) free!
[1572841306] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 298 processes
[1572841426] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1572842501] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572842998] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31111384 kB) free.
[1572844198] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14338456 kB) free!
[1572844318] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35560140 kB) free.
[1572844918] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15065944 kB) free!
[1572845038] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10868636 kB) free!
[1572845158] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34434392 kB) free.
[1572845482] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193656 kB) free!
[1572845602] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176928 kB) free!
[1572845722] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192436 kB) free!
[1572845758] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11259416 kB) free!
[1572845878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34058652 kB) free.
[1572846478] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12529960 kB) free!
[1572846598] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8080028 kB) free!
[1572846718] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33182416 kB) free.
[1572846922] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22076076 kB) free.
[1572847318] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9055692 kB) free!
[1572847439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4702056 kB) free!
[1572847559] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30474840 kB) free.
[1572848159] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6516060 kB) free!
[1572848279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31441692 kB) free.
[1572848355] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.7% (3221288 kB) free!
[1572848476] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.6% (2967428 kB) free!
[1572848596] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.6% (2729272 kB) free!
[1572848879] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9165996 kB) free!
[1572848999] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4780304 kB) free!
[1572849119] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30698332 kB) free.
[1572849196] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.4% (22714508 kB) free.
[1572849719] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6222524 kB) free!
[1572849839] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1837024 kB) free!
[1572849959] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28078364 kB) free.
[1572850474] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4828168 kB) free!
[1572850559] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3334680 kB) free!
[1572850594] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4659344 kB) free!
[1572850679] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29272396 kB) free.
[1572850714] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4512580 kB) free!
[1572851140] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (762436 kB) free!
[1572851261] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190232 kB) free!
[1572851279] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6055836 kB) free!
[1572851381] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189696 kB) free!
[1572851399] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1743360 kB) free!
[1572851519] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28712808 kB) free.
[1572851914] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4394040 kB) free!
[1572852034] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (4256944 kB) free!
[1572852120] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4132256 kB) free!
[1572852154] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4126792 kB) free!
[1572852240] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (410020 kB) free!
[1572852360] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27847180 kB) free.
[1572852447] SERVICE ALERT: sig3-vm2;Total Processes;OK;HARD;1;PROCS OK: 568 processes
[1572852514] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1636288 kB) free!
[1572852598] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1798 processes
[1572852666] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 378 processes
[1572852786] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1572852960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3087716 kB) free!
[1572853080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1385284 kB) free!
[1572853109] SERVICE ALERT: secondary-server;Nginx Access Log Status;WARNING;SOFT;1;NRPE: Unable to read output
[1572853127] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4908716 kB) free!
[1572853200] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26105252 kB) free!
[1572853229] SERVICE ALERT: secondary-server;Nginx Access Log Status;WARNING;SOFT;2;NRPE: Unable to read output
[1572853247] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 20.0% (4892612 kB) free!
[1572853349] SERVICE ALERT: secondary-server;Nginx Access Log Status;WARNING;HARD;3;NRPE: Unable to read output
[1572853367] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.9% (4876280 kB) free!
[1572853800] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1572853800] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1108048 kB) free!
[1572853954] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1698516 kB) free!
[1572854515] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.95, 5.49, 5.17
[1572854635] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.33, 5.27, 5.12
[1572854755] SERVICE ALERT: sig2-vm2;CPU Load;OK;HARD;1;OK - load average: 2.40, 4.09, 4.70
[1572854777] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572855127] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8431800 kB) free!
[1572855207] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 92.6% (22690156 kB) free.
[1572855247] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8431224 kB) free!
[1572855287] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 71.3% (17585952 kB) free.
[1572855337] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.1% (22318192 kB) free.
[1572855367] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (8432872 kB) free!
[1572855375] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43815272 kB) free.
[1572855412] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.2% (40503844 kB) free.
[1572855459] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 70.8% (17359044 kB) free.
[1572855492] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.3% (14673984 kB) free.
[1572855894] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173200 kB) free!
[1572856014] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189392 kB) free!
[1572856134] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1572856134] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170004 kB) free!
[1572856575] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24106516 kB) free!
[1572856695] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32915208 kB) free.
[1572856721] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.5% (21355260 kB) free.
[1572857331] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.9% (3191336 kB) free!
[1572857451] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (2229584 kB) free!
[1572857571] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2149184 kB) free!
[1572858027] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4830348 kB) free!
[1572858147] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4738000 kB) free!
[1572858196] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (4077188 kB) free!
[1572858268] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (4638316 kB) free!
[1572858316] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.7% (3702724 kB) free!
[1572858436] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.9% (3269328 kB) free!
[1572858512] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12694736 kB) free!
[1572858632] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36273648 kB) free.
[1572859036] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1572859036] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1935996 kB) free!
[1572859232] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16606976 kB) free!
[1572859352] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12700600 kB) free!
[1572859472] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35799896 kB) free.
[1572860072] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14939856 kB) free!
[1572860192] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11129788 kB) free!
[1572860237] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.2% (22661332 kB) free.
[1572860312] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34698548 kB) free.
[1572860912] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14158448 kB) free!
[1572861032] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10219376 kB) free!
[1572861152] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33402796 kB) free.
[1572861485] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9373748 kB) free!
[1572861605] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9198184 kB) free!
[1572861725] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9005120 kB) free!
[1572861752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11638132 kB) free!
[1572861872] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34282280 kB) free.
[1572862468] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2100784 kB) free!
[1572862472] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14388076 kB) free!
[1572862593] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10369832 kB) free!
[1572862713] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33650728 kB) free.
[1572862926] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.5% (47053540 kB) free.
[1572862933] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(22) < 2019/11/04 15:52:08 [error] 24131#0: *2507526 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572862997] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(6) < 2019/11/04 15:52:51 [error] 24131#0: *2507588 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572863024] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 422 processes
[1572863038] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(6) < 2019/11/04 15:53:56 [error] 24132#0: *2507546 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572863043] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572863046] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/04 15:54:06 [error] 24130#0: *2507564 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572863053] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572863112] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1572863144] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 417 processes
[1572863162] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21987508 kB) free.
[1572863212] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9499600 kB) free!
[1572863230] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.2% (22108072 kB) free.
[1572863264] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 416 processes
[1572863270] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.2% (40495532 kB) free.
[1572863313] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13488752 kB) free!
[1572863433] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9644212 kB) free!
[1572863553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32746660 kB) free.
[1572863586] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1572863658] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/11/04 15:56:54 [error] 24134#0: *2507688 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572863706] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1572863770] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171388 kB) free!
[1572863778] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572863890] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188888 kB) free!
[1572864010] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193588 kB) free!
[1572864153] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11873460 kB) free!
[1572864273] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (7976288 kB) free!
[1572864393] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31361304 kB) free.
[1572864479] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8291668 kB) free!
[1572864599] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (8291384 kB) free!
[1572864719] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (8291908 kB) free!
[1572864993] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11309280 kB) free!
[1572865059] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4442956 kB) free!
[1572865113] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26698528 kB) free.
[1572865179] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (4289700 kB) free!
[1572865299] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4127736 kB) free!
[1572865579] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2019/11/04 16:32:55 [error] 24134#0: *2508672 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1572865699] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572865713] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12889908 kB) free!
[1572865833] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8847780 kB) free!
[1572865919] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4447840 kB) free!
[1572865953] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31476612 kB) free.
[1572866299] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2019/11/04 16:39:04 [error] 24132#0: *2508985 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572866419] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572866553] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12223652 kB) free!
[1572866673] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8354356 kB) free!
[1572866794] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1572866794] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30952436 kB) free.
[1572867099] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2355476 kB) free!
[1572867337] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4642412 kB) free!
[1572867393] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10840740 kB) free!
[1572867457] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4521500 kB) free!
[1572867513] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6859248 kB) free!
[1572867529] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22843976 kB) free.
[1572867578] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.5% (22905704 kB) free.
[1572867633] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30041756 kB) free.
[1572867849] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1572867969] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1572868069] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.2% (40499484 kB) free.
[1572868121] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 70.8% (17363284 kB) free.
[1572868233] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9554632 kB) free!
[1572868353] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5680288 kB) free!
[1572868473] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28760000 kB) free.
[1572868729] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (204588 kB) free!
[1572868849] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173348 kB) free!
[1572868969] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1572868969] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182840 kB) free!
[1572869073] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8130296 kB) free!
[1572869193] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4348184 kB) free!
[1572869313] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27646072 kB) free.
[1572869913] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7880112 kB) free!
[1572870033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26514036 kB) free.
[1572870633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10347324 kB) free!
[1572870698] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.9% (22669436 kB) free.
[1572870726] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20998540 kB) free!
[1572870730] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33876792 kB) free.
[1572871037] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1572871169] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(188) < 2019/11/04 18:03:15 [error] 15230#0: *1305183 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1572871289] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572871819] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2019/11/04 18:13:19 [error] 24127#0: *2526199 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572871939] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572873101] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171008 kB) free!
[1572873129] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13337400 kB) free!
[1572873222] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189780 kB) free!
[1572873249] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9114672 kB) free!
[1572873342] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (162908 kB) free!
[1572873355] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.38, 6.30, 5.62
[1572873369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34194372 kB) free.
[1572873385] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8917864 kB) free!
[1572873475] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.39, 5.79, 5.53
[1572873505] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (8705980 kB) free!
[1572873596] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.57, 6.33, 5.77
[1572873625] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (8494884 kB) free!
[1572873969] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12479940 kB) free!
[1572874069] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4314072 kB) free!
[1572874089] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8380908 kB) free!
[1572874189] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (1564520 kB) free!
[1572874209] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32897464 kB) free.
[1572874309] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (1250596 kB) free!
[1572874809] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12277416 kB) free!
[1572874929] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7971696 kB) free!
[1572875049] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26821128 kB) free.
[1572875649] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11403380 kB) free!
[1572875769] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7361612 kB) free!
[1572875890] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3407560 kB) free!
[1572875996] SERVICE ALERT: sig2-vm2;CPU Load;OK;HARD;1;OK - load average: 2.33, 3.97, 4.93
[1572876625] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4428880 kB) free!
[1572877090] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17663828 kB) free!
[1572878290] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407856 kB) free!
[1572878328] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4512144 kB) free!
[1572878448] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4423808 kB) free!
[1572878568] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4336128 kB) free!
[1572879490] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1572879490] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14872508 kB) free!
[1572879581] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4441120 kB) free!
[1572879701] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (4338032 kB) free!
[1572879821] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (4272720 kB) free!
[1572880690] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5435548 kB) free!
[1572881237] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4498324 kB) free!
[1572881290] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14139228 kB) free!
[1572881357] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4424256 kB) free!
[1572881477] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.2% (22669308 kB) free.
[1572881569] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2409428 kB) free!
[1572882342] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1572882490] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4167384 kB) free!
[1572883690] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20007972 kB) free!
[1572884021] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2434800 kB) free!
[1572884265] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1572884291] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2693336 kB) free!
[1572884385] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 335 processes
[1572884505] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1572885491] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17455532 kB) free!
[1572886691] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7551952 kB) free!
[1572887291] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15511132 kB) free!
[1572888491] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5939604 kB) free!
[1572889091] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14310808 kB) free!
[1572889901] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (25194124 kB) free!
[1572890021] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28021876 kB) free.
[1572890291] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6434084 kB) free!
[1572890621] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.5% (25663968 kB) free!
[1572890742] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (28169872 kB) free.
[1572890891] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15660308 kB) free!
[1572891491] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1621320 kB) free!