[1552069835] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.1% (4957756 kB) free!
[1552069865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1552069955] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.6% (12618472 kB) free.
[1552070184] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12890420 kB) free.
[1552070555] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.5% (5680504 kB) free!
[1552070675] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (3070112 kB) free!
[1552070784] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5810568 kB) free!
[1552070795] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.7% (12646276 kB) free.
[1552070904] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12906060 kB) free.
[1552071395] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (1249076 kB) free!
[1552071504] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6011676 kB) free!
[1552071515] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10858180 kB) free.
[1552071624] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12854428 kB) free.
[1552071665] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1552071909] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21622 MB (10% inode=99%):
[1552071942] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1552072002] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1552072062] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1552072115] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8299604 kB) free!
[1552072122] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1552072122] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 19.8% (1455600 kB) free!
[1552072235] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.1% (5454852 kB) free!
[1552072242] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 19.7% (1448532 kB) free!
[1552072250] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552072355] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2571896 kB) free!
[1552072362] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 19.4% (1426228 kB) free!
[1552072370] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1552072490] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1552072510] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 32550 MB (15% inode=99%):
[1552072824] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4760980 kB) free!
[1552072944] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12135064 kB) free.
[1552073556] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11103280 kB) free.
[1552074065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1552074156] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9314928 kB) free!
[1552074185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552074276] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6811004 kB) free!
[1552074305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1552074396] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4076512 kB) free!
[1552074522] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 57 users currently logged in
[1552074890] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1552074905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1552075344] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7851856 kB) free!
[1552075464] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6551196 kB) free!
[1552075505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1552075584] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12779376 kB) free.
[1552075596] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1552075596] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5103520 kB) free!
[1552075625] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552076184] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7918876 kB) free!
[1552076226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552076304] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6615664 kB) free!
[1552076346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552076424] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12790848 kB) free.
[1552076466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1552076796] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11169556 kB) free.
[1552077024] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7984120 kB) free!
[1552077066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1552077144] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6697464 kB) free!
[1552077186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552077264] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12781952 kB) free.
[1552077306] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552077864] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8055996 kB) free!
[1552077906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1552077985] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6776604 kB) free!
[1552077996] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2022276 kB) free!
[1552078026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552078105] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12796900 kB) free.
[1552078116] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11181280 kB) free.
[1552078147] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552078705] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8174416 kB) free!
[1552078716] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2601016 kB) free!
[1552078746] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1552078825] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6878096 kB) free!
[1552078837] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11189496 kB) free.
[1552078866] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1552078945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12784552 kB) free.
[1552078986] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552079436] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3214708 kB) free!
[1552079545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8299540 kB) free!
[1552079556] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11193372 kB) free.
[1552079586] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552079665] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7015184 kB) free!
[1552079706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552079785] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5131084 kB) free!
[1552079826] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552080156] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3725572 kB) free!
[1552080276] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.7% (11205284 kB) free.
[1552080426] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552080547] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552080667] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1552080876] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4130828 kB) free!
[1552080985] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12019808 kB) free.
[1552080996] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11212684 kB) free.
[1552081267] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552081387] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552081507] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552081596] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4615360 kB) free!
[1552081716] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (1912192 kB) free!
[1552081836] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11239464 kB) free.
[1552082107] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1552082186] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7371188 kB) free!
[1552082227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552082306] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5766412 kB) free!
[1552082347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552082426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12778528 kB) free.
[1552082436] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2503452 kB) free!
[1552082556] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11260636 kB) free.
[1552082947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552083026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7516476 kB) free!
[1552083068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552083146] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5982520 kB) free!
[1552083157] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3038832 kB) free!
[1552083188] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552083266] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12797884 kB) free.
[1552083277] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.9% (11283416 kB) free.
[1552083866] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7597596 kB) free!
[1552083877] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3479588 kB) free!
[1552083986] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6136340 kB) free!
[1552083997] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11243404 kB) free.
[1552084106] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12774328 kB) free.
[1552084597] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3703232 kB) free!
[1552084706] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7724744 kB) free!
[1552084717] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11260288 kB) free.
[1552084826] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6338512 kB) free!
[1552084946] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12782816 kB) free.
[1552084988] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552085317] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3884236 kB) free!
[1552085437] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.9% (11277032 kB) free.
[1552085547] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7831064 kB) free!
[1552085588] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552085667] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6535512 kB) free!
[1552085708] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552085787] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12764100 kB) free.
[1552085828] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552086037] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4111928 kB) free!
[1552086157] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.9% (11294644 kB) free.
[1552086387] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7933012 kB) free!
[1552086429] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552086507] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6632460 kB) free!
[1552086549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552086627] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12773620 kB) free.
[1552086669] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552086757] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4444548 kB) free!
[1552086877] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.9% (11309876 kB) free.
[1552087227] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8022908 kB) free!
[1552087269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552087347] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6730752 kB) free!
[1552087389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552087467] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12766196 kB) free.
[1552087477] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4669764 kB) free!
[1552087509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552087598] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11328896 kB) free.
[1552088067] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8118964 kB) free!
[1552088109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552088187] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6823588 kB) free!
[1552088197] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4895540 kB) free!
[1552088229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552088308] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12779380 kB) free.
[1552088317] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.9% (7354744 kB) free!
[1552088350] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552088437] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11345368 kB) free.
[1552088907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8243180 kB) free!
[1552088949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552089027] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.1% (6932004 kB) free!
[1552089037] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2689144 kB) free!
[1552089069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552089147] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12766856 kB) free.
[1552089157] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11336704 kB) free.
[1552089189] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552089747] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8345904 kB) free!
[1552089757] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2911412 kB) free!
[1552089789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552089867] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7040560 kB) free!
[1552089877] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552089877] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11359296 kB) free.
[1552089909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552089987] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5184640 kB) free!
[1552090029] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552090477] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3487804 kB) free!
[1552090597] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.1% (11403744 kB) free.
[1552090629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552090750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552090870] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552091187] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12190236 kB) free.
[1552091197] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3752012 kB) free!
[1552091317] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.3% (11463068 kB) free.
[1552091470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1552091590] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552091710] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552091917] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3433588 kB) free!
[1552092038] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.3% (11504288 kB) free.
[1552092310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552092387] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7391756 kB) free!
[1552092430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552092508] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.8% (5816088 kB) free!
[1552092550] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552092628] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12756444 kB) free.
[1552092638] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4712960 kB) free!
[1552092758] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.6% (11613888 kB) free.
[1552092830] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1552093151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552093228] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7461808 kB) free!
[1552093271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552093348] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5919236 kB) free!
[1552093358] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.1% (4957480 kB) free!
[1552093391] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552093468] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12785260 kB) free.
[1552093478] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.6% (11650396 kB) free.
[1552094068] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7513832 kB) free!
[1552094078] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.3% (5051764 kB) free!
[1552094188] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5982112 kB) free!
[1552094198] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.5% (11569496 kB) free.
[1552094308] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12769528 kB) free.
[1552094798] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4886928 kB) free!
[1552094908] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7622028 kB) free!
[1552094918] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.5% (11597964 kB) free.
[1552095028] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6143412 kB) free!
[1552095148] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12782160 kB) free.
[1552095191] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552095518] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4922476 kB) free!
[1552095638] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.5% (11601320 kB) free.
[1552095749] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8217108 kB) free!
[1552095791] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552095869] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6921664 kB) free!
[1552095911] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552095989] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12760180 kB) free.
[1552096032] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552096239] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.4% (5109288 kB) free!
[1552096359] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.6% (11628464 kB) free.
[1552096589] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8328424 kB) free!
[1552096631] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552096709] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7032296 kB) free!
[1552096751] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552096829] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5172544 kB) free!
[1552096871] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552096959] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.4% (5133604 kB) free!
[1552097079] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.6% (11650040 kB) free.
[1552097471] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552097591] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552097679] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.6% (5235668 kB) free!
[1552097711] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552097799] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11673720 kB) free.
[1552098029] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12153028 kB) free.
[1552098311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552098399] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.7% (5280016 kB) free!
[1552098432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552098519] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11692516 kB) free.
[1552098552] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552099119] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.6% (5229272 kB) free!
[1552099152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552099229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7412984 kB) free!
[1552099239] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11713420 kB) free.
[1552099272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552099349] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5840868 kB) free!
[1552099392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552099469] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12760720 kB) free.
[1552099839] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.2% (5042880 kB) free!
[1552099959] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11729676 kB) free.
[1552100070] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7508032 kB) free!
[1552100190] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5969168 kB) free!
[1552100310] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12773624 kB) free.
[1552100559] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.4% (5111416 kB) free!
[1552100680] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11744928 kB) free.
[1552100910] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7570524 kB) free!
[1552101030] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6083548 kB) free!
[1552101150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12757960 kB) free.
[1552101192] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552101279] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.9% (5373592 kB) free!
[1552101399] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11758956 kB) free.
[1552101750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7638912 kB) free!
[1552101792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552101870] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6207612 kB) free!
[1552101913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552101990] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12770948 kB) free.
[1552101999] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.4% (5610996 kB) free!
[1552102033] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552102119] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11771044 kB) free.
[1552102590] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7719428 kB) free!
[1552102633] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552102710] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6343256 kB) free!
[1552102719] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.4% (5628680 kB) free!
[1552102753] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552102831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12752108 kB) free.
[1552102839] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11778656 kB) free.
[1552102873] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552103430] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7729876 kB) free!
[1552103439] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.4% (5607980 kB) free!
[1552103473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552103550] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6419312 kB) free!
[1552103559] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11788848 kB) free.
[1552103593] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552103670] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12771388 kB) free.
[1552103713] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552104159] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.3% (5562556 kB) free!
[1552104270] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7818020 kB) free!
[1552104279] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11797952 kB) free.
[1552104314] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552104390] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6518852 kB) free!
[1552104434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552104510] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12755980 kB) free.
[1552104554] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552104879] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.2% (5536860 kB) free!
[1552105000] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.0% (11806256 kB) free.
[1552105110] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7957632 kB) free!
[1552105154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552105230] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6674348 kB) free!
[1552105274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552105350] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12772416 kB) free.
[1552105394] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552105599] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.5% (5646324 kB) free!
[1552105719] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.0% (11810400 kB) free.
[1552105950] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8155192 kB) free!
[1552105994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552106070] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6853708 kB) free!
[1552106114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552106191] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12751672 kB) free.
[1552106235] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552106319] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.1% (5483668 kB) free!
[1552106439] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.5% (11586564 kB) free.
[1552106790] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5632416 kB) free!
[1552106910] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12769152 kB) free.
[1552107434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552107510] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7420528 kB) free!
[1552107554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552107630] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5843008 kB) free!
[1552107674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552107750] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12753656 kB) free.
[1552108239] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7722960 kB) free!
[1552108350] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7546204 kB) free!
[1552108359] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.6% (5210708 kB) free!
[1552108470] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6047904 kB) free!
[1552108479] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.0% (11804208 kB) free.
[1552108590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12761744 kB) free.
[1552109615] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4735564 kB) free!
[1552109679] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8090060 kB) free!
[1552109735] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4680180 kB) free!
[1552109800] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.8% (5304456 kB) free!
[1552109855] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4620884 kB) free!
[1552109920] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2570004 kB) free!
[1552110391] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5952788 kB) free!
[1552110511] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12712788 kB) free.
[1552110520] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.6% (11150336 kB) free.
[1552110934] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (718752 kB) free!
[1552111111] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5512340 kB) free!
[1552111120] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8959956 kB) free!
[1552111231] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12554972 kB) free.
[1552111240] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.8% (6294612 kB) free!
[1552111360] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3879672 kB) free!
[1552111831] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552112030] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552112150] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1552112270] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1552112560] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1552112560] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.7% (10714036 kB) free.
[1552113031] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4943988 kB) free!
[1552113074] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552113151] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12037900 kB) free.
[1552113160] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8284828 kB) free!
[1552113280] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.6% (5722160 kB) free!
[1552113400] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (3015816 kB) free!
[1552113674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552113794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552113914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552114000] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10453552 kB) free.
[1552114600] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9378892 kB) free!
[1552114721] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.1% (6933760 kB) free!
[1552114841] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4196156 kB) free!
[1552114951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5094124 kB) free!
[1552115072] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12202316 kB) free.
[1552115856] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2453436 kB) free!
[1552116041] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10184512 kB) free.
[1552116162] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (708296 kB) free!
[1552116641] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8534572 kB) free!
[1552116761] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.0% (5932612 kB) free!
[1552116881] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3272376 kB) free!
[1552116914] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552117471] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4967356 kB) free!
[1552117591] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12142604 kB) free.
[1552118081] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9500340 kB) free!
[1552118114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552118235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552118355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552118642] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (1468572 kB) free!
[1552118762] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 19.8% (1459264 kB) free!
[1552118882] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 19.7% (1451824 kB) free!
[1552119391] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4583144 kB) free!
[1552119511] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11699624 kB) free.
[1552119881] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4286024 kB) free!
[1552119910] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22604 MB (10% inode=99%):
[1552121357] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552121710] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 28428 MB (13% inode=99%):
[1552121911] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4774504 kB) free!
[1552122031] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12088396 kB) free.
[1552122281] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.2% (9447732 kB) free!
[1552122557] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552122677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552122798] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552123831] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4306388 kB) free!
[1552123951] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11664772 kB) free.
[1552124081] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4709168 kB) free!
[1552124110] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20149 MB (9% inode=99%):
[1552125910] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1552125910] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 37619 MB (18% inode=99%):
[1552126352] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4630156 kB) free!
[1552126398] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552126472] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11485772 kB) free.
[1552126481] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.2% (5015552 kB) free!
[1552127081] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4271692 kB) free!
[1552128199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552128272] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5109096 kB) free!
[1552128319] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552128392] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12317580 kB) free.
[1552128919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552128992] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6933728 kB) free!
[1552129039] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552129112] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5388064 kB) free!
[1552129159] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552129232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12169476 kB) free.
[1552129481] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9502500 kB) free!
[1552131032] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4443960 kB) free!
[1552131153] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11565780 kB) free.
[1552131281] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4911744 kB) free!
[1552133681] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.8% (9740360 kB) free!
[1552134152] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6647592 kB) free!
[1552134272] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5111264 kB) free!
[1552134392] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12053232 kB) free.
[1552136081] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3694612 kB) free!
[1552136110] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1552136359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552136793] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7545756 kB) free!
[1552136913] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6272048 kB) free!
[1552136960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552137033] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12228860 kB) free.
[1552137080] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552137281] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.5% (5670132 kB) free!
[1552137633] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7706548 kB) free!
[1552137680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552137753] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6398660 kB) free!
[1552137800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552137873] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4564652 kB) free!
[1552137920] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1552138473] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7820512 kB) free!
[1552138520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552138640] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552138760] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552139073] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11571180 kB) free.
[1552139361] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552139481] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552139601] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552139681] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3146832 kB) free!
[1552140201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552140273] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6829252 kB) free!
[1552140281] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.6% (6684604 kB) free!
[1552140321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552140393] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5225268 kB) free!
[1552140441] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552140513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12230228 kB) free.
[1552141041] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552141114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6973024 kB) free!
[1552141161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552141234] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.1% (5489500 kB) free!
[1552141282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552141354] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12214340 kB) free.
[1552141954] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7129644 kB) free!
[1552142074] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5759052 kB) free!
[1552142081] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (771224 kB) free!
[1552142194] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12224836 kB) free.
[1552142794] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7248220 kB) free!
[1552142914] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5956212 kB) free!
[1552143034] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12209668 kB) free.
[1552143082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552143281] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.6% (7674600 kB) free!
[1552143634] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7406356 kB) free!
[1552143682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1552143754] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6123056 kB) free!
[1552143802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552143875] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12228372 kB) free.
[1552143922] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552144474] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7575752 kB) free!
[1552144522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552144594] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6276868 kB) free!
[1552144643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552144714] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12212508 kB) free.
[1552144763] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552145081] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (1493780 kB) free!
[1552145314] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7701412 kB) free!
[1552145363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552145434] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6386420 kB) free!
[1552145483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552145554] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1552145554] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5329352 kB) free!
[1552145603] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552146203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552146281] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.7% (7735540 kB) free!
[1552146323] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552146443] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552146754] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11472992 kB) free.
[1552147044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552147164] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552147284] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552147884] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552147954] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6692820 kB) free!
[1552148004] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552148074] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.2% (5036760 kB) free!
[1552148081] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1202036 kB) free!
[1552148124] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552148195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12206268 kB) free.
[1552148724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552148794] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6885356 kB) free!
[1552148844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552148914] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5348452 kB) free!
[1552148965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552149034] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12218388 kB) free.
[1552149282] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.4% (7604532 kB) free!
[1552149634] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7002608 kB) free!
[1552149754] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5543504 kB) free!
[1552149874] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12210788 kB) free.
[1552149882] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9955588 kB) free.
[1552150474] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7109380 kB) free!
[1552150594] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5786132 kB) free!
[1552150714] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12222748 kB) free.
[1552150765] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552151082] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1026212 kB) free!
[1552151202] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9974304 kB) free.
[1552151314] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7280436 kB) free!
[1552151365] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552151435] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5983112 kB) free!
[1552151485] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1552151555] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12209208 kB) free.
[1552151605] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552151802] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (1390144 kB) free!
[1552151922] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.3% (9989624 kB) free.
[1552152155] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7446808 kB) free!
[1552152205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552152275] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6171772 kB) free!
[1552152326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552152395] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12220948 kB) free.
[1552152446] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552152522] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1753440 kB) free!
[1552152642] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.3% (9996892 kB) free.
[1552152995] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7649716 kB) free!
[1552153046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552153115] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6359868 kB) free!
[1552153166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552153235] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5838712 kB) free!
[1552153242] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2029664 kB) free!
[1552153286] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552153362] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.3% (10010320 kB) free.
[1552153886] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1552153962] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (2583320 kB) free!
[1552154006] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552154082] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.3% (10021176 kB) free.
[1552154126] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552154435] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11524976 kB) free.
[1552154683] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3089692 kB) free!
[1552154727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1552154803] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10033940 kB) free.
[1552154847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552154967] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552155403] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3263700 kB) free!
[1552155523] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10048352 kB) free.
[1552155567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552155635] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6887580 kB) free!
[1552155687] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552155756] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5361172 kB) free!
[1552155807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552155876] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12221260 kB) free.
[1552156123] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3730748 kB) free!