[1581273091] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.9% (30003364 kB) free.
[1581273620] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167416 kB) free!
[1581273740] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202500 kB) free!
[1581273860] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204720 kB) free!
[1581273968] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (246444 kB) free!
[1581274088] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (238024 kB) free!
[1581274208] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (233688 kB) free!
[1581274512] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1581274632] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1581275273] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 180.179.218.150 [10/Feb/2020:00:36:52 +0530] TCP 502 0 0 60.001 "10.147.212.155:13000, 10.147.212.143:13000, 10.147.212.154:13000, 10.147.212.142:13000" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 0.000, 60.000"
[1581275318] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4477632 kB) free!
[1581275393] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581275918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184116 kB) free!
[1581276518] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 63.5% (15664744 kB) free.
[1581277718] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (370608 kB) free!
[1581277838] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184636 kB) free!
[1581277958] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1581277958] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165756 kB) free!
[1581278707] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4499064 kB) free!
[1581278827] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4495248 kB) free!
[1581278947] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4495208 kB) free!
[1581281408] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.6% (10191680 kB) free.
[1581283827] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(210) < 103.16.101.72 [10/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581283948] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(7) < 103.16.101.72 [10/Feb/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581284069] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581284144] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(270) < 2020/02/10 03:01:32 [error] 10222#0: *51463768 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581284264] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581284394] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2236) < 172.105.59.173 [10/Feb/2020:03:01:41 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1581284514] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581288758] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581291608] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197220 kB) free!
[1581291728] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191448 kB) free!
[1581291758] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.5% (11240140 kB) free.
[1581291848] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200728 kB) free!
[1581291860] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.0% (20590856 kB) free.
[1581293864] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/10 05:44:11 [error] 10218#0: *51498156 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1581293984] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581299207] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1581299327] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1581300700] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39718692 kB) free.
[1581300730] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.1% (21842104 kB) free.
[1581302384] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2020/02/10 08:01:03 [error] 10221#0: *51533517 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581302504] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581302558] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (2098876 kB) free!
[1581302678] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (2088124 kB) free!
[1581302798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2073924 kB) free!
[1581303704] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2020/02/10 08:25:40 [error] 10221#0: *51542953 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581303825] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581306114] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(23) < 49.50.86.152 [10/Feb/2020:09:09:36 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581306234] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581307930] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.4% (21576044 kB) free.
[1581311232] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1581311352] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 401 processes
[1581311472] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 402 processes
[1581312072] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1581313643] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581313763] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581314363] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1581314483] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 706 processes
[1581314538] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (4213112 kB) free!
[1581314603] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1581314658] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4124052 kB) free!
[1581314778] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (4002884 kB) free!
[1581315203] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1581315993] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4024648 kB) free!
[1581316113] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.4% (3909016 kB) free!
[1581316233] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.4% (15581764 kB) free.
[1581316403] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1581316424] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(86) < 2020/02/10 12:01:16 [error] 10221#0: *51619725 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581316434] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [10/Feb/2020:12:02:38 +0530] TCP 502 0 0 60.002 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.001, 0.000"
[1581316523] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1581316544] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581316554] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581316578] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1807204 kB) free!
[1581316643] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1581317243] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1581317243] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581318220] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1581318220] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.3% (21295256 kB) free.
[1581318397] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4773436 kB) free!
[1581318516] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4909424 kB) free!
[1581318517] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.2% (4697588 kB) free!
[1581318636] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4892876 kB) free!
[1581318637] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.7% (4585832 kB) free!
[1581318756] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.9% (4873920 kB) free!
[1581322237] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1873856 kB) free!
[1581323208] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21656744 kB) free.
[1581323261] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.4% (22152324 kB) free.
[1581324738] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4376972 kB) free!
[1581324858] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (4269536 kB) free!
[1581324978] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4146200 kB) free!
[1581325428] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.1% (2495952 kB) free!
[1581325460] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4672620 kB) free!
[1581325548] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (2398112 kB) free!
[1581325580] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4579180 kB) free!
[1581325668] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2303000 kB) free!
[1581325700] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4496312 kB) free!
[1581327354] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(170) < 92.42.104.166 [10/Feb/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581327378] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2212508 kB) free!
[1581327437] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/whoami"
[1581327474] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581327501] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186964 kB) free!
[1581327502] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1581327558] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1581327622] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1581327709] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8485148 kB) free!
[1581327742] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1581327829] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8946148 kB) free!
[1581327949] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (8945996 kB) free!
[1581328068] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.4% (21577216 kB) free.
[1581328117] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39697316 kB) free.
[1581328141] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.6% (21969900 kB) free.
[1581328178] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.2% (21381844 kB) free.
[1581328342] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1581328643] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581329137] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9569656 kB) free!
[1581329257] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9408736 kB) free!
[1581329284] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4291164 kB) free!
[1581329320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26265284 kB) free!
[1581329377] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9243556 kB) free!
[1581329404] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (4182852 kB) free!
[1581329440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43899504 kB) free.
[1581329524] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (4083248 kB) free!
[1581330142] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1581330262] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1581330382] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1581330724] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1741676 kB) free!
[1581330945] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(119) < 2020/02/10 16:01:23 [error] 10212#0: *51710674 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581330982] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1581330982] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1581331065] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581332243] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1581332363] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1581332484] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 707 processes
[1581333577] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4286460 kB) free!
[1581334884] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581335265] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/02/10 17:10:11 [error] 10215#0: *51738911 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.29930637666113946 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1581335330] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9710500 kB) free!
[1581335386] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581335432] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.6% (21615788 kB) free.
[1581335450] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9620496 kB) free!
[1581335467] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 24.1% (11897424 kB) free.
[1581335484] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1581335570] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9542688 kB) free!
[1581335604] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1581335724] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1581336325] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1581336325] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581336382] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1581336502] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1581336591] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171360 kB) free!
[1581336623] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1581336712] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185324 kB) free!
[1581336832] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178352 kB) free!
[1581336925] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581337045] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1581337148] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (264024 kB) free!
[1581337165] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1581337268] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191164 kB) free!
[1581337389] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181248 kB) free!
[1581337823] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 354 processes
[1581338187] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 83.3% (41027608 kB) free.
[1581338260] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.7% (21752392 kB) free.
[1581338302] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.8% (20284712 kB) free.
[1581338365] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1581338423] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1581338909] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171928 kB) free!
[1581339030] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (228204 kB) free!
[1581339150] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192924 kB) free!
[1581340223] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 354 processes
[1581341268] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4855776 kB) free!
[1581341388] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (4755256 kB) free!
[1581341423] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 354 processes
[1581341508] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (4266064 kB) free!
[1581341821] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.1% (4679168 kB) free!
[1581341941] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.7% (4571504 kB) free!
[1581342038] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1581342061] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.3% (4474080 kB) free!
[1581343075] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1581343075] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.7% (22226352 kB) free.
[1581343097] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1581343217] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1581343237] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (4207088 kB) free!
[1581343337] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1581343357] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.7% (3626520 kB) free!
[1581343477] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (3552352 kB) free!
[1581343908] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1581343908] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2335760 kB) free!
[1581344461] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2100640 kB) free!
[1581344537] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1581345109] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21664528 kB) free.
[1581345144] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21937608 kB) free.
[1581345277] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2268448 kB) free!
[1581345585] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(166) < 2020/02/10 20:03:29 [error] 10219#0: *51809670 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581345623] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1581345705] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581346117] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8990112 kB) free!
[1581346238] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8989920 kB) free!
[1581346272] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1581346358] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (8990328 kB) free!
[1581346392] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581349165] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581350277] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1581351469] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175940 kB) free!
[1581351478] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195780 kB) free!
[1581351589] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176740 kB) free!
[1581351598] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187664 kB) free!
[1581351710] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176636 kB) free!
[1581351718] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (255752 kB) free!
[1581351749] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191424 kB) free!
[1581351870] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206036 kB) free!
[1581351990] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184584 kB) free!
[1581353423] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1581355992] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581356113] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1581357991] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1581358291] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (242956 kB) free!
[1581358411] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (224432 kB) free!
[1581358531] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240748 kB) free!
[1581359059] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.4% (21572736 kB) free!
[1581359179] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.4% (21587828 kB) free!
[1581359299] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.9% (20897168 kB) free!