[1563301857] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.66, 1.71, 13.40
[1563302059] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1563302115] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8323392 kB) free!
[1563302179] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563302235] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12435820 kB) free.
[1563302372] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180208 kB) free!
[1563302404] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182516 kB) free!
[1563302439] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (2108436 kB) free!
[1563302457] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 17.42, 137.41, 91.04
[1563302492] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (205648 kB) free!
[1563302524] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181164 kB) free!
[1563302559] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (2105264 kB) free!
[1563302612] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199360 kB) free!
[1563302644] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195136 kB) free!
[1563302679] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2103164 kB) free!
[1563302779] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563302835] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7909872 kB) free!
[1563302899] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563302955] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12439172 kB) free.
[1563303499] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563303555] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7488560 kB) free!
[1563303620] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563303675] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12418720 kB) free.
[1563304219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563304339] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563304857] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.52, 3.07, 17.57
[1563304875] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8955896 kB) free!
[1563304939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563304995] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12430204 kB) free.
[1563305059] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563305400] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 72.95, 91.12, 41.65
[1563305462] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.55, 58.33, 38.52
[1563305484] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 480 processes
[1563305498] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (1838844 kB) free!
[1563305520] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.48, 62.00, 36.97
[1563305544] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 479 processes
[1563305595] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9000424 kB) free!
[1563305640] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.09, 42.39, 32.82
[1563305659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563305679] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.5% (3071400 kB) free!
[1563305715] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12422528 kB) free.
[1563305779] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563306240] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.63, 8.09, 18.54
[1563306315] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9079264 kB) free!
[1563306379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563306435] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7385632 kB) free!
[1563306499] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563306555] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12428576 kB) free.
[1563306662] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.62, 1.63, 10.98
[1563307100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563307155] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7418228 kB) free!
[1563307220] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563307262] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.16, 35.08, 24.78
[1563307275] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12419648 kB) free.
[1563307820] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563307862] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.52, 5.17, 13.23
[1563307875] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7772392 kB) free!
[1563307940] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563307995] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12426380 kB) free.
[1563308462] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 57.28, 51.92, 28.28
[1563308540] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563308596] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7930164 kB) free!
[1563308640] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.80, 2.85, 3.99
[1563308660] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563308716] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12420232 kB) free.
[1563309062] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.58, 7.38, 15.05
[1563309098] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (149252 kB) free!
[1563309260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563309316] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8014560 kB) free!
[1563309380] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563309436] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12427952 kB) free.
[1563309662] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16.08, 37.23, 25.55
[1563309981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563310036] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8084204 kB) free!
[1563310101] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563310156] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12429548 kB) free.
[1563310262] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 5.42, 13.64
[1563310701] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563310756] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8249516 kB) free!
[1563310821] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563310862] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.95, 24.69, 23.17
[1563310876] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12422052 kB) free.
[1563311421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563311462] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.58, 5.67, 14.66
[1563311476] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8370308 kB) free!
[1563311541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563311596] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12428480 kB) free.
[1563312141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563312196] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8488616 kB) free!
[1563312262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563312317] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12419708 kB) free.
[1563312382] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563312662] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.26, 21.82, 20.52
[1563312916] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8540908 kB) free!
[1563312982] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563313036] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12425352 kB) free.
[1563313102] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563313636] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8591400 kB) free!
[1563313702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563313756] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12420364 kB) free.
[1563313822] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563314356] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8687264 kB) free!
[1563314422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563314462] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.58, 5.41, 14.81
[1563314476] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12412492 kB) free.
[1563314542] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563315076] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8828696 kB) free!
[1563315142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563315187] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39757256 kB) free.
[1563315196] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12417228 kB) free.
[1563315212] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.8% (21170028 kB) free.
[1563315244] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.4% (21434284 kB) free.
[1563315263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563315662] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.48, 19.83, 21.03
[1563315698] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 70.7% (11484996 kB) free.
[1563315701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (13030980 kB) free!
[1563315749] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.9% (14773056 kB) free.
[1563315796] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9006456 kB) free!
[1563315821] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (18990064 kB) free.
[1563315862] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563315880] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2448224 kB) free!
[1563315916] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12304784 kB) free.
[1563315982] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563316263] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 3.05, 11.25
[1563316516] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9163876 kB) free!
[1563316582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563316637] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7579164 kB) free!
[1563316702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563316757] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12417716 kB) free.
[1563316863] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.70, 22.96, 21.61
[1563317302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563317357] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7808204 kB) free!
[1563317422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563317463] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 3.47, 11.55
[1563317477] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12410032 kB) free.
[1563318023] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563318077] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8015900 kB) free!
[1563318143] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563318197] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12415888 kB) free.
[1563318663] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.38, 0.47, 3.47
[1563318743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563318797] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8179724 kB) free!
[1563318863] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563318917] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12410160 kB) free.
[1563319263] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2.80, 65.58, 50.02
[1563319383] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.94, 44.06, 44.02
[1563319463] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563319503] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.57, 29.68, 38.76
[1563319517] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8340660 kB) free!
[1563319583] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563319637] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12415264 kB) free.
[1563319703] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563320237] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8455448 kB) free!
[1563320303] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563320358] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12416000 kB) free.
[1563320424] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563320957] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8523644 kB) free!
[1563321023] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563321077] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12408664 kB) free.
[1563321143] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563321303] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.67, 2.20, 12.59
[1563321677] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8507368 kB) free!
[1563321743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563321797] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12416304 kB) free.
[1563321863] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563322397] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8485780 kB) free!
[1563322463] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563322503] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.12, 0.75, 3.90
[1563322517] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12407888 kB) free.
[1563322583] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563323117] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8501068 kB) free!
[1563323183] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563323237] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12412400 kB) free.
[1563323304] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563323335] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.53, 18.19, 16.60
[1563323455] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.19, 12.20, 14.60
[1563323575] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.12, 8.23, 12.86
[1563323837] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8572760 kB) free!
[1563323903] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563323957] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12415748 kB) free.
[1563324023] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563324303] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1168.05, 336.47, 116.65
[1563324423] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 157.90, 225.27, 102.55
[1563324543] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.08, 151.04, 90.23
[1563324557] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8537824 kB) free!
[1563324623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563324678] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12406108 kB) free.
[1563324743] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563324775] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.40, 3.71
[1563325277] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8531776 kB) free!
[1563325343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563325397] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12411324 kB) free.
[1563325463] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563325583] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563326021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (12613356 kB) free!
[1563326142] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (19647764 kB) free.
[1563326597] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9523440 kB) free!
[1563326717] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8066952 kB) free!
[1563326837] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12238896 kB) free.
[1563328037] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9258844 kB) free!
[1563328157] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7732084 kB) free!
[1563328277] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11965588 kB) free.
[1563328375] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1253.53, 920.92, 435.28
[1563328495] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 217.29, 659.47, 400.47
[1563328615] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 29.31, 441.27, 351.91
[1563328743] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.66, 5.21, 14.97
[1563329183] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563329343] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.74, 25.47, 21.91
[1563329477] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9043404 kB) free!
[1563329598] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7439652 kB) free!
[1563329718] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183380 kB) free.
[1563329943] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 3.80, 11.70
[1563330318] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7478240 kB) free!
[1563330341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12391524 kB) free!
[1563330438] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12188228 kB) free.
[1563330461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (18614936 kB) free.
[1563330543] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.80, 53.20, 30.26
[1563330983] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563331038] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7589116 kB) free!
[1563331061] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12254996 kB) free!
[1563331104] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563331143] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.29, 7.93, 16.23
[1563331158] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12189952 kB) free.
[1563331181] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (19048232 kB) free.
[1563331615] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1563331615] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.15, 0.23, 14.08
[1563331703] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563331743] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.94, 17.49, 21.87
[1563331758] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7734664 kB) free!
[1563331781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12792196 kB) free!
[1563331823] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563331878] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12180264 kB) free.
[1563331901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (19254984 kB) free.
[1563332343] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.27, 2.74, 11.69
[1563332423] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563332478] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7858828 kB) free!
[1563332543] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563332550] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (170076 kB) free!
[1563332587] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217360 kB) free!
[1563332598] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12186656 kB) free.
[1563332670] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (152448 kB) free!
[1563332707] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (211712 kB) free!
[1563332790] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164232 kB) free!
[1563332815] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.20, 0.28, 4.09
[1563332827] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238920 kB) free!
[1563333143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563333198] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7956976 kB) free!
[1563333263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563333318] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12054148 kB) free.
[1563333543] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.27, 92.56, 50.22
[1563335343] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 3.34, 14.72
[1563335663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563335783] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563335943] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.55, 26.44, 24.85
[1563336318] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8299472 kB) free!
[1563336383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563336438] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12178788 kB) free.
[1563336503] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563336543] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.67, 4.12, 13.35
[1563337038] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8417208 kB) free!
[1563337103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563337158] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12180600 kB) free.
[1563337223] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563337743] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 129.05, 110.36, 49.42
[1563337758] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8501976 kB) free!
[1563337824] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563337878] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12172540 kB) free.
[1563337901] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12370592 kB) free!
[1563337944] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563338021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (18983240 kB) free.
[1563338478] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8268836 kB) free!
[1563338544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563338599] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12175704 kB) free.
[1563338621] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (12910032 kB) free!
[1563338664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563338741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (18644604 kB) free.
[1563338784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563338943] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.42, 4.17, 15.09
[1563339341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13057332 kB) free!
[1563339461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (9237876 kB) free!
[1563339543] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.69, 15.14, 24.45
[1563339581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (16735996 kB) free.
[1563339798] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9322832 kB) free!
[1563339918] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7924816 kB) free!
[1563339925] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2823304 kB) free!
[1563339984] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563340038] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12165644 kB) free.
[1563340046] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.4% (2658608 kB) free!
[1563340143] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.86, 2.59, 13.15
[1563340166] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.0% (2433940 kB) free!
[1563340181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (9372328 kB) free!
[1563340302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (16880480 kB) free.
[1563340766] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (958664 kB) free!
[1563340802] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.14, 6.43, 5.83
[1563340901] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (9081936 kB) free!
[1563340922] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.76, 6.82, 6.04
[1563341021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (17144080 kB) free.
[1563341042] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.59, 7.42, 6.35
[1563341343] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.43, 3.97
[1563341621] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (9525728 kB) free!
[1563341741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (16722128 kB) free.
[1563341784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563341838] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8584196 kB) free!
[1563341904] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563341958] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7132016 kB) free!
[1563342024] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563342078] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11645984 kB) free.
[1563342341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (10289404 kB) free!
[1563342461] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (6457104 kB) free!
[1563342581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (14489192 kB) free.
[1563342624] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563342744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563342864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563343181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (7325088 kB) free!
[1563343301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (15096648 kB) free.
[1563343615] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563343901] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (7688152 kB) free!
[1563344021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (15027660 kB) free.
[1563344409] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.8% (1164140 kB) free!
[1563344529] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.1% (891232 kB) free!
[1563344621] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (6258872 kB) free!
[1563344649] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (541792 kB) free!
[1563344741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (20368084 kB) free.
[1563344975] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.2% (21131612 kB) free.
[1563345064] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39840604 kB) free.
[1563345116] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 80.7% (13116032 kB) free.
[1563345175] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.6% (14885060 kB) free.
[1563345230] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.0% (6763692 kB) free.
[1563345341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12843304 kB) free!
[1563345462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (21051480 kB) free.
[1563346278] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8576552 kB) free!
[1563346398] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12147788 kB) free.
[1563346464] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1563346464] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563347064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563347184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563347261] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (11731620 kB) free!
[1563347304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563347381] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (19362296 kB) free.
[1563347576] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1563347696] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1563347904] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563348212] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179696 kB) free!
[1563348332] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170404 kB) free!
[1563348452] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191868 kB) free!
[1563348505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563348625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563348745] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563348798] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9789720 kB) free!
[1563348918] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8392300 kB) free!
[1563349039] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12134512 kB) free.
[1563349181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12164616 kB) free!
[1563349301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (8341980 kB) free!
[1563349326] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (3201384 kB) free!
[1563349421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (17127372 kB) free.
[1563349446] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (2940496 kB) free!
[1563349566] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.7% (2716392 kB) free!
[1563349975] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.9% (1938184 kB) free!
[1563350021] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (10318860 kB) free!
[1563350095] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (1513268 kB) free!
[1563350141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (17393240 kB) free.
[1563350166] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1563350166] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1584608 kB) free!
[1563350215] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1071344 kB) free!
[1563350741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (11934496 kB) free!
[1563350861] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (7728096 kB) free!
[1563350981] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (16720580 kB) free.
[1563351582] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (11022520 kB) free!
[1563351702] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (7078980 kB) free!
[1563351745] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563351822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (15828228 kB) free.
[1563352143] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563352345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563352422] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (9211624 kB) free!
[1563352438] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.5% (769820 kB) free!
[1563352466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563352542] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (5151784 kB) free!
[1563352558] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (551340 kB) free!
[1563352662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (14179528 kB) free.
[1563352678] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (290328 kB) free!
[1563353252] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.7% (3132560 kB) free!
[1563353269] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (7513912 kB) free!
[1563353389] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (6051648 kB) free!
[1563353509] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (16149404 kB) free.
[1563354109] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (9826908 kB) free!
[1563354229] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (17103048 kB) free.
[1563354265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563354385] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563354452] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2428152 kB) free!
[1563354672] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8223972 kB) free!
[1563354792] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.7% (7236104 kB) free!
[1563354829] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12341068 kB) free!
[1563354912] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2018820 kB) free!
[1563354949] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (8307216 kB) free!
[1563354985] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563355069] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (16564896 kB) free.
[1563355105] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1563355639] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8582196 kB) free!
[1563355670] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (11013044 kB) free!
[1563355705] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563355759] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12120288 kB) free.
[1563355790] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (6768948 kB) free!
[1563355826] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563355910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (16000216 kB) free.
[1563356278] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.9% (6758624 kB) free.
[1563356359] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8634828 kB) free!
[1563356425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563356479] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12110096 kB) free.
[1563356510] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (9851440 kB) free!
[1563356545] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563356630] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (5966452 kB) free!
[1563356750] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (15195696 kB) free.
[1563357079] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8494708 kB) free!
[1563357145] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563357199] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12115300 kB) free.
[1563357265] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563357350] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (6180696 kB) free!
[1563357470] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (19814016 kB) free.
[1563357806] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1563357865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563357926] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1563357985] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563358046] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1563358070] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12547092 kB) free!
[1563358105] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563358190] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (8605336 kB) free!
[1563358310] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (18276696 kB) free.
[1563358910] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11555468 kB) free!
[1563359030] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (15917184 kB) free.
[1563359645] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4833364 kB) free!
[1563359765] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (4756736 kB) free!
[1563359846] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1563359885] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (4696612 kB) free!
[1563360230] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (8766124 kB) free!
[1563360350] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (18475604 kB) free.
[1563360951] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (11773800 kB) free!
[1563361071] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (7587616 kB) free!
[1563361191] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (17522268 kB) free.
[1563361566] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563361791] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10895952 kB) free!
[1563361911] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (6745468 kB) free!
[1563362031] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (16976836 kB) free.
[1563362375] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4428316 kB) free!
[1563362495] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (4300580 kB) free!
[1563362615] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (4164864 kB) free!
[1563362631] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (10061860 kB) free!
[1563362751] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (5753752 kB) free!
[1563362871] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (16402320 kB) free.
[1563362878] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (1355408 kB) free!
[1563362998] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (1146716 kB) free!
[1563363118] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.2% (894880 kB) free!
[1563363199] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9284236 kB) free!
[1563363319] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7770732 kB) free!
[1563363439] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11906144 kB) free.
[1563363471] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (9264584 kB) free!
[1563363485] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2025272 kB) free!
[1563363591] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (4917152 kB) free!
[1563363711] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (15766848 kB) free.
[1563363718] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1563363718] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (130172 kB) free!
[1563364105] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563364311] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (9384772 kB) free!
[1563364415] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2416280 kB) free!
[1563364431] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (5268528 kB) free!
[1563364551] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (15454904 kB) free.
[1563365151] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (8983580 kB) free!
[1563365239] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9427260 kB) free!
[1563365272] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (4798124 kB) free!
[1563365359] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12087192 kB) free.
[1563365392] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (14951368 kB) free.
[1563365674] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2840424 kB) free!
[1563365794] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.8% (2728772 kB) free!
[1563365914] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.1% (2612680 kB) free!
[1563365992] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (8480724 kB) free!
[1563366112] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (4310484 kB) free!
[1563366232] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (14653352 kB) free.
[1563366718] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.1% (6770704 kB) free.
[1563366832] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (8352128 kB) free!
[1563366952] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (4216924 kB) free!
[1563367072] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (14065864 kB) free.
[1563367115] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.8% (13460008 kB) free.
[1563367672] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (8373692 kB) free!
[1563367792] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (3973160 kB) free!
[1563367912] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (13561580 kB) free.
[1563368112] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.2% (39011644 kB) free.
[1563368151] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1563368166] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 81.5% (13251088 kB) free.
[1563368215] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 39.7% (6454584 kB) free.
[1563368252] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 62.4% (15389356 kB) free.
[1563368512] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (8985568 kB) free!
[1563368615] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.7% (20998012 kB) free.
[1563368632] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (4703952 kB) free!
[1563368712] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3941488 kB) free!
[1563368751] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1563368752] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (11974560 kB) free!
[1563368815] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169592 kB) free!
[1563368832] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231584 kB) free!
[1563368935] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165100 kB) free!
[1563368952] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237232 kB) free!
[1563369055] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (186248 kB) free!
[1563369351] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1563369953] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (3483704 kB) free!
[1563370105] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563370280] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 79.9% (19587240 kB) free.
[1563370350] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14652104 kB) free.
[1563370390] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39419392 kB) free.
[1563370454] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 51.3% (33758064 kB) free.
[1563370551] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1563370551] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1563370955] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151012 kB) free!
[1563370998] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.9% (6864672 kB) free!
[1563371065] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.7% (4824552 kB) free!
[1563371075] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (177324 kB) free!
[1563371118] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (263708 kB) free!
[1563371185] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.5% (4791468 kB) free!
[1563371195] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1563371195] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175248 kB) free!
[1563371238] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1563371238] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220096 kB) free!
[1563371305] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.4% (4758744 kB) free!
[1563371852] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4681488 kB) free!
[1563371972] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4588516 kB) free!
[1563372092] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (4507576 kB) free!
[1563372859] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (9130736 kB) free!
[1563372980] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (19066792 kB) free.
[1563373292] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200492 kB) free!
[1563373318] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (1435404 kB) free!
[1563373438] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (1263184 kB) free!
[1563373558] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 15.0% (1104192 kB) free!
[1563374158] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.6% (6734792 kB) free.
[1563374179] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (7823660 kB) free!
[1563374299] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (18141324 kB) free.
[1563374498] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178160 kB) free!
[1563374618] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (147548 kB) free!
[1563374738] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (232800 kB) free!
[1563374766] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1163536 kB) free!
[1563374886] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (182716 kB) free!
[1563374899] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12198500 kB) free!
[1563375006] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156528 kB) free!
[1563375019] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (8000788 kB) free!
[1563375139] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (17537808 kB) free.
[1563375739] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13092080 kB) free!
[1563375859] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.9% (9174016 kB) free!
[1563375896] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1563375979] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1563375979] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (5260352 kB) free!
[1563376016] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1563376579] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (13540240 kB) free.
[1563376806] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 51.5% (8365388 kB) free.
[1563377179] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (8860384 kB) free!
[1563377299] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (5197408 kB) free!
[1563377404] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.6% (21474596 kB) free.
[1563377418] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1563377418] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21801080 kB) free.
[1563377419] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 54.4% (35795760 kB) free.
[1563377438] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 72.2% (35592844 kB) free.
[1563377458] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 53.6% (8719944 kB) free.
[1563377477] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 64.6% (10502776 kB) free.
[1563378047] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215648 kB) free!
[1563378068] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (170904 kB) free!
[1563378167] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (262648 kB) free!
[1563378188] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (172148 kB) free!
[1563378287] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239456 kB) free!
[1563378308] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165788 kB) free!
[1563378416] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1563378536] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1563378843] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.22, 2.74, 3.90
[1563381351] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563384027] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1563385558] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)