[1592418671] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.56, 41.25, 41.16
[1592419271] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 23.80, 33.48, 36.84
[1592422503] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201280 kB) free!
[1592422563] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180660 kB) free!
[1592422625] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (169804 kB) free!
[1592422685] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (181364 kB) free!
[1592423285] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 85.0% (20839672 kB) free.
[1592423892] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11719108 kB) free!
[1592424012] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5978412 kB) free!
[1592424132] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393812 kB) free!
[1592424406] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238680 kB) free!
[1592424526] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234824 kB) free!
[1592424646] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238392 kB) free!
[1592424781] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1592424901] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 336 processes
[1592425021] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1592425622] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1592425846] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.0% (7896508 kB) free!
[1592426446] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.3% (27764904 kB) free.
[1592427671] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 57.88, 39.36, 33.27
[1592428270] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 76.8% (37861212 kB) free.
[1592428271] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1592428271] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 37.32, 34.38, 33.67
[1592428871] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 67.45, 39.82, 35.17
[1592429471] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 34.81, 34.04, 34.00
[1592429555] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.8% (41769488 kB) free.
[1592429576] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(255) < 2020/06/18 03:01:22 [error] 12477#0: *94830324 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.62.98, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1592429578] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(184) < 172.105.62.98 [18/Jun/2020:03:00:12 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1592429678] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1173) < 172.105.62.98 [18/Jun/2020:03:01:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592429696] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592429699] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592429799] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592433378] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (3564992 kB) free!
[1592433671] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 22.98, 27.73, 29.82
[1592433978] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192780 kB) free!
[1592434271] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 32.43, 36.68, 34.59
[1592434391] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 19.31, 30.34, 32.52
[1592434511] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 30.96, 30.18, 32.15
[1592435712] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 30.85, 24.75, 28.20
[1592435778] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.6% (10194316 kB) free.
[1592435822] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1592435942] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1592436912] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 49.67, 41.25, 35.62
[1592436978] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166624 kB) free!
[1592437032] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 38.12, 40.43, 36.03
[1592437098] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (161176 kB) free!
[1592437152] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 52.70, 42.43, 37.13
[1592437218] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1592437218] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (149736 kB) free!
[1592437477] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 5.7% (1406176 kB) free!
[1592437664] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.0% (18126344 kB) free.
[1592438077] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.5% (9742256 kB) free.
[1592440752] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.34, 49.64, 43.12
[1592441239] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 75.8% (18585928 kB) free.
[1592441352] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 32.90, 40.74, 40.82
[1592441532] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.0% (17098684 kB) free!
[1592441585] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.6% (4056932 kB) free!
[1592441645] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.6% (4061288 kB) free!
[1592441705] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.5% (4052436 kB) free!
[1592441765] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.6% (4056592 kB) free!
[1592441952] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.47, 50.57, 45.01
[1592442043] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592442171] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=8.71% system=27.91% iowait=4.54% idle=58.83%
[1592442946] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1592443064] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180700 kB) free!
[1592443067] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1592443184] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189508 kB) free!
[1592443185] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1592443187] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1592443304] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196340 kB) free!
[1592443305] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1592443787] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1592444077] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196688 kB) free!
[1592444197] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189224 kB) free!
[1592444317] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201316 kB) free!
[1592444342] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 20.9% (5113148 kB) free.
[1592444387] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1592444507] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1592444942] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (256868 kB) free!
[1592445062] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180096 kB) free!
[1592445107] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1592445183] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 79.9% (19592412 kB) free.
[1592445227] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1592446142] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1592446262] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1592448018] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1592450402] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.42, 5.63, 5.36
[1592450522] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.68, 6.44, 5.69
[1592450642] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.06, 5.93, 5.60
[1592451773] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592451899] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=11.63% system=27.35% iowait=0.01% idle=61.01%
[1592453352] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592456839] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (3741120 kB) free!
[1592456959] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3492172 kB) free!
[1592457079] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (3234396 kB) free!
[1592457679] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1800064 kB) free!
[1592458279] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3331960 kB) free!
[1592458383] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.8% (4846620 kB) free!
[1592458428] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (3675340 kB) free!
[1592458503] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.8% (4619116 kB) free!
[1592458548] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.4% (3234240 kB) free!
[1592458623] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.8% (4359328 kB) free!
[1592458668] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (2878852 kB) free!
[1592458879] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1592458879] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1762624 kB) free!
[1592459268] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (472384 kB) free!
[1592460423] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1632264 kB) free!
[1592462103] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592462233] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1592462268] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 68.6% (15410404 kB) free.
[1592462363] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1592462969] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=17.37% system=28.19% iowait=0.53% idle=53.92%
[1592463155] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8796908 kB) free!
[1592463275] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8471596 kB) free!
[1592463395] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.5% (8107532 kB) free!
[1592463662] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1592463782] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1592464595] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4370572 kB) free!
[1592467870] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9665672 kB) free!
[1592467983] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1592467990] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9510796 kB) free!
[1592468103] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1592468110] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (9335236 kB) free!
[1592468223] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 329 processes
[1592468341] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 17.9% (5862200 kB) free!
[1592468461] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 17.5% (5739600 kB) free!
[1592468582] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 17.1% (5593804 kB) free!
[1592468823] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1592468868] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (3707416 kB) free!
[1592468988] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.2% (3404644 kB) free!
[1592469108] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3037744 kB) free!
[1592469182] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 21.1% (6923360 kB) free.
[1592469423] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1592469543] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1592469708] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1592469708] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1430376 kB) free!
[1592469782] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.3% (6317656 kB) free!
[1592469902] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 18.8% (6143164 kB) free!
[1592470022] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.2% (5962224 kB) free!
[1592470279] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592471710] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4426740 kB) free!
[1592472709] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.4% (15585804 kB) free.
[1592472999] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(271) < 103.13.97.241 [18/Jun/2020:15:01:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1592473119] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592473743] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1592473863] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1592473983] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1592474222] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1592474222] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (3278840 kB) free!
[1592475296] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/06/18 15:44:17 [error] 12489#0: *94927539 connect() failed (111: Connection refused) while connecting to upstream, client: 15.206.135.98, server: 0.0.0.0:14001, upstream: "10.147.212.61:14001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1592475298] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 15.206.135.98 [18/Jun/2020:15:44:17 +0530] TCP 502 0 0 0.001 "10.147.212.61:14001" "0" "0" "0.001"
[1592475416] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592475419] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592476016] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/06/18 15:54:18 [error] 12489#0: *94930245 connect() failed (111: Connection refused) while connecting to upstream, client: 15.206.135.98, server: 0.0.0.0:14001, upstream: "10.147.212.61:14001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1592476018] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 15.206.135.98 [18/Jun/2020:15:54:18 +0530] TCP 502 0 0 0.001 "10.147.212.61:14001" "0" "0" "0.001"
[1592476136] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/06/18 15:57:48 [error] 12483#0: *94931193 connect() failed (111: Connection refused) while connecting to upstream, client: 15.206.135.98, server: 0.0.0.0:14000, upstream: "10.147.212.49:5000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1592476138] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 15.206.135.98 [18/Jun/2020:15:57:48 +0530] TCP 502 0 0 0.000 "10.147.212.49:5000" "0" "0" "0.000"
[1592476256] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2020/06/18 15:59:17 [error] 12477#0: *94931555 connect() failed (111: Connection refused) while connecting to upstream, client: 15.206.135.98, server: 0.0.0.0:14001, upstream: "10.147.212.61:14001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1592476259] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;HARD;3;(1) < 15.206.135.98 [18/Jun/2020:15:59:17 +0530] TCP 502 0 0 0.001 "10.147.212.61:14001" "0" "0" "0.001"
[1592479910] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.2% (3195900 kB) free!
[1592480030] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.4% (2774356 kB) free!
[1592480150] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.2% (2298000 kB) free!
[1592480511] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1190_n"
[1592480631] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1592480750] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (224220 kB) free!
[1592482022] SERVICE FLAPPING ALERT: mrq4;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1592482865] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 21.4% (5254332 kB) free.
[1592483750] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.4% (15577844 kB) free.
[1592490183] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1592490303] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1592492643] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.49, 4.19, 4.81
[1592492750] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1592493132] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388592 kB) free!
[1592493843] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.95, 5.46, 5.06
[1592493963] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.03, 5.28, 5.05
[1592494083] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.03, 5.33, 5.10
[1592494622] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 68.3% (22370412 kB) free.
[1592494684] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.32, 4.57, 4.89
[1592495284] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.36, 5.90, 5.47
[1592495404] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.73, 5.71, 5.45
[1592495524] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1592495524] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.88, 5.72, 5.48
[1592495979] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592496109] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1592496124] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.15, 4.43, 4.89
[1592496234] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=9.71% system=27.49% iowait=0.03% idle=62.77%
[1592496839] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592496968] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=8.49% system=27.33% iowait=9.43% idle=54.75%
[1592497569] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592497699] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=9.03% system=28.45% iowait=1.66% idle=60.86%
[1592502552] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 33.88, 38.35, 44.49