[1576175796] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195368 kB) free!
[1576175805] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4673708 kB) free!
[1576175918] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187956 kB) free!
[1576175925] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177944 kB) free!
[1576175984] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174076 kB) free!
[1576176038] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (209968 kB) free!
[1576176045] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196748 kB) free!
[1576176104] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191328 kB) free!
[1576176225] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177448 kB) free!
[1576176861] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 320 processes
[1576176981] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 322 processes
[1576177101] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1576178997] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.4% (3803304 kB) free!
[1576179597] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183824 kB) free!
[1576179674] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3264608 kB) free!
[1576184287] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237140 kB) free!
[1576184329] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 406 processes
[1576184407] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (248288 kB) free!
[1576184449] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1576184518] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(67) < 2019/12/13 02:31:58 [error] 24128#0: *16540744 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1576184527] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235636 kB) free!
[1576184577] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 303 processes
[1576184639] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/13 02:32:15 [error] 24128#0: *16540728 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576184759] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576185127] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.5% (45110524 kB) free.
[1576186829] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3005) < 103.209.99.7 [13/Dec/2019:03:10:24 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576186838] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (2473252 kB) free!
[1576186847] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.7% (3114308 kB) free!
[1576186949] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(28) < 203.212.70.170 [13/Dec/2019:03:11:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1576187069] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576187159] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3416) < 2019/12/13 03:12:31 [error] 24128#0: *16550715 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:86/0, bytes from/to upstream:0/86
[1576187248] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(176) < 103.16.101.72 [13/Dec/2019:03:10:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1576187279] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576187305] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(233) < 2019/12/13 03:11:32 [error] 28870#0: *20272919 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576187376] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576187425] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576187438] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2295912 kB) free!
[1576191649] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 406 processes
[1576191769] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1576192847] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2448884 kB) free!
[1576198238] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.7% (20277508 kB) free.
[1576198247] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.2% (20874848 kB) free.
[1576198369] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 410 processes
[1576198425] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.5% (21933868 kB) free.
[1576198489] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 411 processes
[1576199279] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(79) < 2019/12/13 06:37:50 [error] 24128#0: *16584329 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576199399] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/13 06:38:00 [error] 24128#0: *16584273 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576199519] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576202089] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576202209] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576210609] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1576210729] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576211519] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(70) < 2019/12/13 10:01:48 [error] 24130#0: *16616689 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576211639] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(25) < 2019/12/13 10:03:56 [error] 24123#0: *16617046 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1576211759] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(5) < 2019/12/13 10:04:26 [error] 24120#0: *16617032 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576212359] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576214874] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22861108 kB) free.
[1576214906] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.4% (39130516 kB) free.
[1576214929] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1576214942] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 91.7% (22466284 kB) free.
[1576215049] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1576216074] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180008 kB) free!
[1576216194] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190224 kB) free!
[1576216238] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4744768 kB) free!
[1576216314] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172176 kB) free!
[1576216359] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4549148 kB) free!
[1576216479] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (4344572 kB) free!
[1576217469] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4386268 kB) free!
[1576217589] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4067136 kB) free!
[1576217679] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2102724 kB) free!
[1576217709] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.6% (3727408 kB) free!
[1576218049] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1576218170] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 403 processes
[1576218290] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576218910] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 79.4% (17821284 kB) free.
[1576220901] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 327 processes
[1576220931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24494188 kB) free!
[1576221021] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1576221051] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.8% (45851988 kB) free.
[1576222251] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21395664 kB) free!
[1576222371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.8% (51106932 kB) free.
[1576222580] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1576222700] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1576222820] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1576222847] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4721048 kB) free!
[1576222967] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4560344 kB) free!
[1576223087] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4405648 kB) free!
[1576224359] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(131) < 2019/12/13 13:34:09 [error] 24124#0: *16650687 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576224479] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576224887] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2332072 kB) free!
[1576227136] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9243404 kB) free!
[1576227256] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9053064 kB) free!
[1576227376] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.9% (8831452 kB) free!
[1576229025] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4772000 kB) free!
[1576229145] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4726736 kB) free!
[1576229266] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (4678528 kB) free!
[1576230021] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1576230141] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1576230376] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4782008 kB) free!
[1576230516] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9222632 kB) free!
[1576230637] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9221428 kB) free!
[1576230757] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (9221220 kB) free!
[1576230771] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23196264 kB) free!
[1576230891] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.0% (46172140 kB) free.
[1576231315] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.6% (10760804 kB) free.
[1576231350] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 52.9% (12973208 kB) free.
[1576231383] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 39.2% (19334900 kB) free.
[1576231411] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 48.3% (11848156 kB) free.
[1576231491] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22606636 kB) free!
[1576231611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46325124 kB) free.
[1576232211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22415744 kB) free!
[1576232331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.2% (46390784 kB) free.
[1576232866] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2286984 kB) free!
[1576232931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21733208 kB) free!
[1576233051] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45644604 kB) free.
[1576233625] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/12/13 16:04:23 [error] 28863#0: *20574768 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: sservercluster2.server, request: "GET /favicon.ico HTTP/1.1", upstream: "http://10.147.212.61:80/favicon.ico", host: "59.91.63.148:8091", referrer: "http://59.91.63.148:8091/httpsms/bulksms.php?esmeid=1001&username=smppclient1&password=pasword1&source_addr=BP-TESTMS&message=HTTPTest"
[1576233651] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20643412 kB) free!
[1576233745] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576233772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44570264 kB) free.
[1576234371] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22441568 kB) free!
[1576234491] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18709868 kB) free!
[1576234611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42150732 kB) free.
[1576234917] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.7% (2637232 kB) free!
[1576234957] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4926060 kB) free!
[1576235037] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.3% (2542408 kB) free!
[1576235157] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2459044 kB) free!
[1576235211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17353804 kB) free!
[1576235331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42451328 kB) free.
[1576235931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18095592 kB) free!
[1576236051] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42168104 kB) free.
[1576236479] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(91) < 2019/12/13 16:57:46 [error] 20757#0: *16684368 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1576236599] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(6) < 2019/12/13 16:59:45 [error] 20760#0: *16684592 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576236651] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19429576 kB) free!
[1576236719] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(27) < 2019/12/13 17:01:07 [error] 20756#0: *16684594 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576236771] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15940704 kB) free!
[1576236891] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36591100 kB) free.
[1576237319] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576237491] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15525052 kB) free!
[1576237611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.0% (50067456 kB) free.
[1576238550] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 377 processes
[1576238556] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189132 kB) free!
[1576238617] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176976 kB) free!
[1576238671] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1576238676] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184620 kB) free!
[1576238737] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178292 kB) free!
[1576238796] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178836 kB) free!
[1576238857] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168908 kB) free!
[1576239192] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.0% (22702084 kB) free.
[1576239232] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.2% (22103528 kB) free.
[1576239274] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 67.4% (33217164 kB) free.
[1576239308] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.7% (22714700 kB) free.
[1576239342] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.2% (20142904 kB) free.
[1576241211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24427576 kB) free!
[1576241331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (45968776 kB) free.
[1576241356] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9876112 kB) free!
[1576241476] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (9791516 kB) free!
[1576241597] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (9709904 kB) free!
[1576241597] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (2146248 kB) free!
[1576241717] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (1771472 kB) free!
[1576241837] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164236 kB) free!
[1576241931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26307260 kB) free!
[1576242051] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.0% (22437616 kB) free!
[1576242171] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.2% (49017036 kB) free.
[1576242290] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576242410] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1576242771] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23301960 kB) free!
[1576242891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19589636 kB) free!
[1576243011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.0% (46153988 kB) free.
[1576243611] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (21018452 kB) free!
[1576243731] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (47543236 kB) free.
[1576244036] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195328 kB) free!
[1576244115] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168564 kB) free!
[1576244156] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (221800 kB) free!
[1576244164] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166980 kB) free!
[1576244235] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170988 kB) free!
[1576244276] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195832 kB) free!
[1576244284] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173228 kB) free!
[1576244332] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22927876 kB) free!
[1576244355] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195480 kB) free!
[1576244404] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (243540 kB) free!
[1576244452] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (18945772 kB) free!
[1576244572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.6% (46869048 kB) free.
[1576245172] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20402216 kB) free!
[1576245292] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16914880 kB) free!
[1576245412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44393404 kB) free.
[1576245593] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1576245620] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1576246012] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19661804 kB) free!
[1576246132] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40206324 kB) free.
[1576246732] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21596800 kB) free!
[1576246852] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17874064 kB) free!
[1576246972] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45275772 kB) free.
[1576247572] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18616828 kB) free!
[1576247692] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.4% (21672932 kB) free!
[1576247812] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40869120 kB) free.
[1576248412] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15159196 kB) free!
[1576248533] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44126928 kB) free.
[1576248719] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2019/12/13 20:21:54 [error] 20756#0: *16717787 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576248839] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(32) < 2019/12/13 20:23:48 [error] 20760#0: *16717892 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576248870] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 375 processes
[1576248959] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(53) < 2019/12/13 20:25:52 [error] 20757#0: *16718138 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576248990] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 447 processes
[1576249132] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19058000 kB) free!
[1576249252] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15499860 kB) free!
[1576249372] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41845792 kB) free.
[1576249397] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4701048 kB) free!
[1576249559] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576249972] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19215408 kB) free!
[1576250092] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43305672 kB) free.
[1576250692] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18989980 kB) free!
[1576250812] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15150628 kB) free!
[1576250932] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40211460 kB) free.
[1576251532] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17173276 kB) free!
[1576251652] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42593424 kB) free.
[1576252252] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19257128 kB) free!
[1576252372] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15790536 kB) free!
[1576252492] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39882588 kB) free.
[1576253092] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17584684 kB) free!
[1576253213] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42205852 kB) free.
[1576253541] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 541 processes
[1576253661] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 542 processes
[1576253812] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19060412 kB) free!
[1576253932] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42040012 kB) free.
[1576254005] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 551 processes
[1576254125] SERVICE ALERT: webservercluster1;Total Processes;OK;HARD;1;PROCS OK: 550 processes
[1576254433] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.6% (13970064 kB) free!
[1576254532] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20427756 kB) free!
[1576254553] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9238624 kB) free!
[1576254652] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16864936 kB) free!
[1576254673] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3165604 kB) free!
[1576254716] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 68.8% (16980268 kB) free.
[1576254725] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 568 processes
[1576254745] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 71.6% (17542924 kB) free.
[1576254772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40345280 kB) free.
[1576254774] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 47.5% (23408176 kB) free.
[1576254801] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.0% (23778880 kB) free!
[1576254804] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.7% (37894640 kB) free.
[1576254837] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.7% (22725676 kB) free.
[1576254845] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 566 processes
[1576254865] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.3% (22369072 kB) free.
[1576254870] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(43) < 103.16.101.72 [13/Dec/2019:22:02:31 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576254960] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(62) < 2019/12/13 22:02:31 [error] 20756#0: *16734724 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "backend_13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576254965] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 568 processes
[1576254990] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576255080] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576255345] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(34) < 2019/12/13 22:03:02 [error] 28871#0: *20720876 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:1714/10568, bytes from/to upstream:10568/1714
[1576255372] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22030580 kB) free!
[1576255465] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576255492] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41933240 kB) free.
[1576256092] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21786948 kB) free!
[1576256212] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41844104 kB) free.
[1576256766] SERVICE ALERT: webservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 610 processes
[1576256812] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19644760 kB) free!
[1576256932] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15566448 kB) free!
[1576257052] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39808752 kB) free.
[1576257652] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18404120 kB) free!
[1576257772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41182740 kB) free.
[1576258373] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18590228 kB) free!
[1576258493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41051256 kB) free.
[1576259093] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18834976 kB) free!
[1576259213] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40921968 kB) free.
[1576259813] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19609176 kB) free!
[1576259933] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39965312 kB) free.
[1576260274] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9790912 kB) free!
[1576260394] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9787928 kB) free!
[1576260514] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9788108 kB) free!
[1576260533] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20960368 kB) free!
[1576260653] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17591600 kB) free!
[1576260773] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39402380 kB) free.
[1576261373] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18524124 kB) free!
[1576261493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40299496 kB) free.