[1587926347] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10574120 kB) free!
[1587926467] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32674008 kB) free.
[1587927067] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11265812 kB) free!
[1587927187] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32679908 kB) free.
[1587927367] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (2644080 kB) free!
[1587927487] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3375088 kB) free.
[1587927586] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1587927750] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1587927787] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12352924 kB) free!
[1587927791] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195548 kB) free!
[1587927851] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191892 kB) free!
[1587927907] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32647500 kB) free.
[1587927911] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (173008 kB) free!
[1587927971] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (196820 kB) free!
[1587928272] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.2% (17456948 kB) free.
[1587928350] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1587928507] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13101024 kB) free!
[1587928627] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32674380 kB) free.
[1587928688] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (353772 kB) free!
[1587928808] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1203616 kB) free!
[1587928928] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1200660 kB) free!
[1587929227] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13891396 kB) free!
[1587929347] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10259832 kB) free!
[1587929467] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30405484 kB) free.
[1587930667] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13194932 kB) free!
[1587930728] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (2914764 kB) free!
[1587930787] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32651872 kB) free.
[1587930973] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3246232 kB) free!
[1587931093] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231752 kB) free!
[1587931213] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216668 kB) free!
[1587931253] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1375240 kB) free!
[1587931373] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (530864 kB) free!
[1587931387] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12429200 kB) free!
[1587931493] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2271672 kB) free!
[1587931507] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32087368 kB) free.
[1587932107] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12191544 kB) free!
[1587932227] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30248704 kB) free.
[1587932827] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12081156 kB) free!
[1587932947] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37652180 kB) free.
[1587933613] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 11.8% (5799524 kB) free!
[1587934565] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 72.7% (35840984 kB) free.
[1587934747] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24060048 kB) free!
[1587934867] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42590336 kB) free.
[1587935467] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23319992 kB) free!
[1587935581] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.6% (44638196 kB) free.
[1587935587] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41473988 kB) free.
[1587936188] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22379112 kB) free!
[1587936308] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40405064 kB) free.
[1587936613] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.5% (44577008 kB) free.
[1587936751] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1587936914] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(318) < 103.16.101.72 [27/Apr/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1587937035] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587937089] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2508) < 203.212.70.170 [27/Apr/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1587937160] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(411) < 2020/04/27 03:01:32 [error] 22126#0: *85904396 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1587937210] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587937281] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587937508] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23907940 kB) free!
[1587937628] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41120540 kB) free.
[1587938828] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24512176 kB) free!
[1587938948] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42439504 kB) free.
[1587939249] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.7% (10456644 kB) free.
[1587939548] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22988160 kB) free!
[1587939668] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40174372 kB) free.
[1587940328] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3618288 kB) free.
[1587940449] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.6% (2838436 kB) free!
[1587940569] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (272492 kB) free!
[1587940689] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (277420 kB) free!
[1587940868] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23272436 kB) free!
[1587940988] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40216996 kB) free.
[1587941889] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.5% (9681324 kB) free.
[1587943389] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25418464 kB) free!
[1587943509] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41078444 kB) free.
[1587944289] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168796 kB) free!
[1587944409] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175892 kB) free!
[1587944529] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1587944529] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173544 kB) free!
[1587945111] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 53.0% (13080284 kB) free.
[1587945143] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.9% (19090852 kB) free.
[1587945238] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 55.3% (13545800 kB) free.
[1587945309] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22662624 kB) free!
[1587945429] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40748428 kB) free.
[1587946328] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1587948893] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 37.0% (48738748 kB) free.
[1587949311] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180332 kB) free!
[1587949431] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170200 kB) free!
[1587949551] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167952 kB) free!
[1587954129] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1587955786] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.18, 11.12, 7.90
[1587955906] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.85, 9.41, 7.66
[1587956001] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1587956026] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.39, 7.69, 7.25
[1587956121] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1587956481] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/27 08:25:21 [error] 22120#0: *85929792 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1587956601] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587957429] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23779284 kB) free!
[1587957549] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (18593932 kB) free!
[1587957669] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14563892 kB) free!
[1587957751] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1587957871] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 343 processes
[1587957991] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1587958438] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4848824 kB) free!
[1587958558] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4633516 kB) free!
[1587958678] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4468516 kB) free!
[1587961079] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2416848 kB) free!
[1587962791] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 346 processes
[1587963043] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1587963069] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12367488 kB) free!
[1587963143] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192964 kB) free!
[1587963163] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1587963263] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170960 kB) free!
[1587963284] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1587963383] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174184 kB) free!
[1587963391] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1587964269] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17041812 kB) free!
[1587965791] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1587965791] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 348 processes
[1587966069] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1587966069] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34769572 kB) free.
[1587966338] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 54.6% (13377020 kB) free.
[1587967242] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.1% (2727828 kB) free!
[1587967328] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3248156 kB) free!
[1587967362] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.4% (2324080 kB) free!
[1587967448] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3549056 kB) free.
[1587967482] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1784764 kB) free!
[1587968470] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16828816 kB) free!
[1587968590] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38241348 kB) free.
[1587968648] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3111232 kB) free!
[1587968768] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (3088372 kB) free!
[1587968791] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1587968888] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (2290872 kB) free!
[1587969009] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(687) < 103.255.146.164 [27/Apr/2020:11:55:22 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587969130] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587969282] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16067844 kB) free.
[1587969488] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3379264 kB) free.
[1587969790] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14868648 kB) free!
[1587969910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36768024 kB) free.
[1587970089] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3183276 kB) free!
[1587970209] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3137872 kB) free!
[1587970329] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (2824112 kB) free!
[1587971110] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15000180 kB) free!
[1587971230] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36673484 kB) free.
[1587971738] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.5% (4299780 kB) free!
[1587971830] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12894620 kB) free!
[1587971858] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.8% (4105892 kB) free!
[1587971950] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35695168 kB) free.
[1587971978] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.0% (3922944 kB) free!
[1587972550] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11547100 kB) free!
[1587972670] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34139072 kB) free.
[1587973178] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2350616 kB) free!
[1587973270] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10314148 kB) free!
[1587973390] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32558056 kB) free.
[1587973990] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8501556 kB) free!
[1587974111] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31186128 kB) free.
[1587974710] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7943972 kB) free!
[1587974830] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29738900 kB) free.
[1587975729] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(699) < 103.255.146.164 [27/Apr/2020:13:47:38 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587975849] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587975921] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1587976030] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8977500 kB) free!
[1587976041] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1587976150] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31067956 kB) free.
[1587976750] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7124200 kB) free!
[1587976870] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30219924 kB) free.
[1587977470] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6102996 kB) free!
[1587977529] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (762184 kB) free!
[1587977581] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9710384 kB) free!
[1587977590] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1587977590] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29274452 kB) free.
[1587977701] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9537748 kB) free!
[1587977791] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1587977822] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.0% (9351856 kB) free!
[1587978129] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1587978129] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (2951352 kB) free!
[1587978190] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5659860 kB) free!
[1587978249] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(131) < 103.255.146.164 [27/Apr/2020:14:31:18 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587978282] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4392912 kB) free!
[1587978310] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26624700 kB) free.
[1587978369] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587978402] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4010180 kB) free!
[1587978522] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.7% (3534752 kB) free!
[1587978910] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (429340 kB) free!
[1587979031] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39064900 kB) free.
[1587979122] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1587979122] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1312840 kB) free!
[1587979569] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(585) < 103.255.146.164 [27/Apr/2020:14:49:54 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587979630] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13952436 kB) free!
[1587979690] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587979750] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38422396 kB) free.
[1587980093] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24359504 kB) free!
[1587980191] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1587980213] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.0% (17138436 kB) free!
[1587980222] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4119232 kB) free!
[1587980289] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 185.255.8.58 [27/Apr/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1587980322] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16075596 kB) free.
[1587980333] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.9% (26212656 kB) free!
[1587980350] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14140096 kB) free!
[1587980409] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587980470] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9709644 kB) free!
[1587980590] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36100424 kB) free.
[1587980933] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.7% (29860188 kB) free.
[1587981190] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11210684 kB) free!
[1587981310] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6803336 kB) free!
[1587981430] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34193408 kB) free.
[1587981533] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (21258524 kB) free!
[1587981653] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26476724 kB) free.
[1587981966] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9374032 kB) free!
[1587982030] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7758652 kB) free!
[1587982086] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9146072 kB) free!
[1587982150] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3263748 kB) free!
[1587982206] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8895432 kB) free!
[1587982209] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(175) < 103.255.146.164 [27/Apr/2020:15:31:16 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587982253] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.2% (21412100 kB) free!
[1587982270] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30094376 kB) free.
[1587982329] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3428344 kB) free.
[1587982330] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587982373] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.3% (29388192 kB) free.
[1587982870] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4699600 kB) free!
[1587982929] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2891236 kB) free!
[1587982973] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (22702224 kB) free!
[1587982990] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (387116 kB) free!
[1587983049] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2812276 kB) free!
[1587983093] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.0% (15767220 kB) free!
[1587983110] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28939168 kB) free.
[1587983169] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (2682056 kB) free!
[1587983213] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.0% (23744928 kB) free!
[1587983711] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3137864 kB) free!
[1587983831] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395976 kB) free!
[1587983951] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30804080 kB) free.
[1587984370] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (317628 kB) free!
[1587984551] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (393368 kB) free!
[1587984671] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (390604 kB) free!
[1587984791] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382432 kB) free!
[1587985013] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1587985013] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (10458136 kB) free!
[1587985570] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2578020 kB) free!
[1587985806] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4525736 kB) free!
[1587985991] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25161508 kB) free!
[1587986170] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (605420 kB) free!
[1587986322] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.5% (3022500 kB) free!
[1587986442] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.9% (2664304 kB) free!
[1587986562] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (2261364 kB) free!
[1587986591] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396648 kB) free!
[1587986770] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (2463280 kB) free!
[1587987162] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (474684 kB) free!
[1587987413] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.5% (25753400 kB) free!
[1587987791] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1587987791] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19938892 kB) free!
[1587988391] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28978552 kB) free.
[1587988991] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2194100 kB) free!
[1587989111] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403164 kB) free!
[1587989213] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5175364 kB) free!
[1587989231] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400348 kB) free!
[1587990413] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27196296 kB) free.
[1587990431] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22144556 kB) free!
[1587990762] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.6% (16302424 kB) free.
[1587990970] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1409660 kB) free!
[1587991013] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (25996096 kB) free!
[1587991031] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32460556 kB) free.
[1587991133] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.9% (19649808 kB) free!
[1587991253] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (27341372 kB) free.
[1587991402] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/27 18:03:50 [error] 22135#0: *85999838 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1587991522] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587991631] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7771120 kB) free!
[1587991751] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3374364 kB) free!
[1587991853] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25800496 kB) free!
[1587991871] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23093888 kB) free!
[1587991973] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.0% (18429708 kB) free!
[1587992093] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.7% (25933444 kB) free!
[1587992471] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8863020 kB) free!
[1587993071] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22191476 kB) free!
[1587993370] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (2823784 kB) free!
[1587993671] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33408516 kB) free.
[1587994041] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1587994161] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1587994271] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6607796 kB) free!
[1587994391] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2484880 kB) free!
[1587994511] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391824 kB) free!
[1587995026] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.51, 4.68, 4.87
[1587995711] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25372308 kB) free!
[1587995770] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1515088 kB) free!
[1587996293] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (28330228 kB) free.
[1587996311] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1629436 kB) free!
[1587997493] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.3% (25396372 kB) free!
[1587997512] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23894748 kB) free!
[1587997613] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.9% (19586440 kB) free!
[1587997733] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27219564 kB) free.
[1587998112] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391236 kB) free!
[1587998333] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.9% (23577780 kB) free!
[1587998453] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.7% (18054508 kB) free!
[1587998573] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.0% (25027896 kB) free!
[1587999162] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1587999312] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24478032 kB) free!
[1587999370] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2778980 kB) free!
[1587999912] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (823512 kB) free!
[1588000522] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/27 20:44:28 [error] 22133#0: *86017252 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1588000642] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588001112] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23785132 kB) free!
[1588001712] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384728 kB) free!
[1588002773] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.0% (34295820 kB) free.
[1588002912] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24854720 kB) free!
[1588002970] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1335124 kB) free!
[1588004112] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12863272 kB) free!
[1588004170] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2691860 kB) free!
[1588004712] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23489260 kB) free!
[1588005312] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34101724 kB) free.
[1588005370] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1280504 kB) free!
[1588005912] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10618076 kB) free!
[1588005970] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2776556 kB) free!
[1588006032] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6563828 kB) free!
[1588006153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33986324 kB) free.
[1588006752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11561784 kB) free!
[1588006872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7581536 kB) free!
[1588006992] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3717816 kB) free!
[1588007170] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (720248 kB) free!
[1588007770] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.3% (4768264 kB) free.
[1588008192] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20168360 kB) free!
[1588008370] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (2499296 kB) free!
[1588008490] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (2726680 kB) free!
[1588008610] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2948096 kB) free!
[1588008792] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31684480 kB) free.
[1588009210] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.7% (4177640 kB) free.
[1588009392] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8126120 kB) free!
[1588009512] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4434432 kB) free!
[1588009632] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32484708 kB) free.
[1588010232] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8407820 kB) free!
[1588010352] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4708580 kB) free!
[1588010362] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1588010472] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33270504 kB) free.
[1588010482] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1588010573] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1588011072] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8583144 kB) free!
[1588011192] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4869920 kB) free!
[1588011312] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33181988 kB) free.
[1588011913] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9107908 kB) free!
[1588012033] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5510848 kB) free!
[1588012153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33070484 kB) free.