[1568056307] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568056427] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1568056876] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.9% (10332352 kB) free.
[1568058076] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4848928 kB) free!
[1568058196] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4128516 kB) free!
[1568058316] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.8% (3657384 kB) free!
[1568058917] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194760 kB) free!
[1568063719] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 94.2% (46422372 kB) free.
[1568065497] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.7% (10227088 kB) free.
[1568065563] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 44.9% (11015472 kB) free.
[1568077526] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.3% (9706372 kB) free.
[1568077940] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 95.67, 40.97, 16.03
[1568078060] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 103.53, 62.04, 26.77
[1568078180] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 104.68, 76.19, 36.23
[1568079380] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.62, 3.07, 14.73
[1568080197] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 380 processes
[1568080257] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 380 processes
[1568080317] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 374 processes
[1568080499] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182712 kB) free!
[1568080526] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198140 kB) free!
[1568080572] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193256 kB) free!
[1568080619] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188792 kB) free!
[1568080647] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186508 kB) free!
[1568080692] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199492 kB) free!
[1568080740] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195216 kB) free!
[1568080767] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193720 kB) free!
[1568080812] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182392 kB) free!
[1568081180] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.33, 2.97, 4.38
[1568091099] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1568094012] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (2754880 kB) free!
[1568094612] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1908388 kB) free!
[1568094711] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.07, 9.97, 8.11
[1568094831] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.68, 6.82, 7.19
[1568094951] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.97, 4.88, 6.43
[1568095060] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1568095180] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1568095551] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.91, 2.17, 4.12
[1568099737] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 303 processes
[1568099857] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1568100581] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1568100701] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1568102266] SERVICE ALERT: rq-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1568102386] SERVICE ALERT: rq-2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1568102986] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.67, 7.60, 5.77
[1568103106] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.91, 7.50, 5.94
[1568103226] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.16, 6.61, 5.81
[1568104426] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.62, 4.25, 4.78
[1568106326] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8650176 kB) free!
[1568106447] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8304156 kB) free!
[1568106567] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.2% (7975940 kB) free!
[1568106826] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.04, 6.17, 5.68
[1568106946] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.73, 5.71, 5.57
[1568107066] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.80, 5.79, 5.61
[1568107767] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4512376 kB) free!
[1568109101] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1568109221] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1568112951] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.71, 5.37, 5.50
[1568112980] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.38, 5.34, 5.19
[1568113071] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.31, 4.38, 5.12
[1568113100] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.14, 6.02, 5.45
[1568113191] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.21, 3.63, 4.76
[1568113220] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.39, 6.13, 5.56
[1568114384] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 144.18, 123.95, 68.05
[1568114391] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 109.39, 41.63, 17.27
[1568114420] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.31, 2.44, 3.81
[1568114504] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 68.63, 110.67, 70.53
[1568114511] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 125.31, 70.22, 30.71
[1568114596] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.5% (1703620 kB) free!
[1568114624] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.65, 74.58, 62.17
[1568114631] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 127.30, 89.30, 42.50
[1568114716] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.4% (1695420 kB) free!
[1568114836] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (1696024 kB) free!
[1568115459] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.43, 5.95, 5.24
[1568115579] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.18, 5.43, 5.15
[1568115699] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.61, 5.46, 5.22
[1568115824] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.06, 3.02, 18.25
[1568116220] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 174.74, 134.78, 67.58
[1568116300] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 4.84, 4.16, 4.51
[1568116340] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 98.51, 128.94, 74.20
[1568116431] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.72, 7.56, 15.96
[1568116460] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.92, 87.26, 65.56
[1568116666] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.25, 4.50, 4.84
[1568117021] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1568117141] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1568117266] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.62, 7.14, 6.16
[1568117386] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.35, 7.25, 6.31
[1568117506] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.24, 7.21, 6.41
[1568117508] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (6648960 kB) free!
[1568117624] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.17, 1.31, 3.81
[1568117628] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (6216616 kB) free!
[1568117660] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1568117660] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.47, 3.71, 19.99
[1568117748] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 16.0% (5613056 kB) free!
[1568117863] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3498920 kB) free!
[1568117983] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.9% (3041484 kB) free!
[1568118103] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.1% (2623100 kB) free!
[1568118231] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.40, 2.85, 4.49
[1568118349] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1831976 kB) free!
[1568118704] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.9% (23077272 kB) free.
[1568118831] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.22, 8.12, 6.22
[1568118949] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 94.7% (33157752 kB) free.
[1568118951] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.05, 8.78, 6.69
[1568119071] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1568119071] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.81, 6.94, 6.28
[1568119300] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.42, 5.66, 5.16
[1568119420] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.40, 6.34, 5.47
[1568119541] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.54, 5.95, 5.43
[1568120271] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.54, 2.95, 4.80
[1568120660] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.42, 1.85, 3.67
[1568120741] SERVICE FLAPPING ALERT: sig5-vm2;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1568120741] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 3.64, 4.53, 4.92
[1568122672] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.59, 8.93, 5.98
[1568122792] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.97, 9.21, 6.43
[1568122912] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.45, 8.67, 6.56
[1568125906] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.33, 3.83, 4.71
[1568126060] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.71, 11.56, 7.21
[1568126180] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.93, 8.37, 6.58
[1568126300] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.57, 6.04, 5.94
[1568126900] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.92, 1.84, 3.71
[1568127097] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1568127217] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1568127337] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1568127938] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1568131541] SERVICE FLAPPING ALERT: sig5-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568132512] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1568137187] SERVICE ALERT: rq-2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1568137307] SERVICE ALERT: rq-2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1568138300] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)