[1550773922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550774043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550774642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550774762] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550775362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550775482] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550776082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550776202] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550776802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550776923] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1550777522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550777642] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550778242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550778362] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550778962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550779082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550779682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550779802] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1550779802] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550781002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550781122] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550781722] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550781842] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1550781962] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550782562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550782682] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550783283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1550783403] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550784003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550784123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1550784723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550784843] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550785443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550785563] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550786163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550786284] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550786883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550787003] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550787603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550787723] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550788323] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550788443] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1550789043] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550789164] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550789763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550789883] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550790483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550790603] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550791203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550791323] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1550791923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550792044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550792643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550792763] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550793363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550793483] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550794083] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550794203] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550795403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550795524] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550796123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550796243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550796843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550796963] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550797563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550797683] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1550798283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550798404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550799003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550799123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550799723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550799843] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550800443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550800563] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550801163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550801283] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1550801884] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550802004] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550802604] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550802724] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550803324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550803444] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1550804044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550804164] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550804765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550804885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550805005] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550807837] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3248376 kB) free!
[1550807957] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.9% (3242204 kB) free!
[1550808077] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.9% (3232516 kB) free!
[1550809205] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550809805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550809925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550810045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550810813] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.9% (14773996 kB) free.
[1550810840] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (14913092 kB) free.
[1550813114] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4821292 kB) free!
[1550813234] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4605672 kB) free!
[1550813354] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4394524 kB) free!
[1550813761] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23858 MB (10% inode=99%):
[1550814199] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24010 MB (10% inode=99%):
[1550814245] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1550814245] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550814554] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2345676 kB) free!
[1550814686] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 19.3% (3131304 kB) free!
[1550814754] SERVICE ALERT: webserver2;Current Users;CRITICAL;SOFT;1;USERS CRITICAL - 67 users currently logged in
[1550814806] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 18.9% (3079696 kB) free!
[1550814814] SERVICE ALERT: webserver2;Current Users;CRITICAL;SOFT;2;USERS CRITICAL - 67 users currently logged in
[1550814845] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550814874] SERVICE ALERT: webserver2;Current Users;CRITICAL;SOFT;3;USERS CRITICAL - 67 users currently logged in
[1550814926] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 18.6% (3016304 kB) free!
[1550814934] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 67 users currently logged in
[1550814965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550815085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550815564] SERVICE ALERT: smscapp1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1550815685] SERVICE ALERT: smscapp1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1550815805] SERVICE ALERT: smscapp1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1550818685] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1550819285] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550819405] SERVICE ALERT: smscapp1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1550819405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550819525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550820026] SERVICE ALERT: sig1;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'sig1' was out of bounds)
[1550820029] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1550820033] SERVICE ALERT: sig1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig1' was out of bounds)
[1550820033] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig1' was out of bounds)
[1550820086] SERVICE ALERT: sig1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig1' was out of bounds)
[1550820093] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1550820127] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1550820146] SERVICE ALERT: sig1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 33098 MB (64% inode=98%):
[1550820153] SERVICE ALERT: sig1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 142538 MB (63% inode=99%):
[1550820153] SERVICE ALERT: sig1;Sensors Status;OK;HARD;1;SENSORS OK
[1550820206] SERVICE ALERT: sig1;Total Processes;OK;HARD;1;PROCS OK: 329 processes
[1550820248] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1550820326] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.6% (15368756 kB) free.
[1550820385] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550820505] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550820625] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 306 processes
[1550820625] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550820745] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1550821226] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550822349] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1550822387] SERVICE ALERT: sig4;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4' was out of bounds)
[1550822428] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;HARD;1;CPU CRITICAL: 1 crit, 0 warn out of 261 processes
[1550822429] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1550822507] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 94.7% (15385696 kB) free.
[1550822548] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1550823361] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 165086 MB (73% inode=99%):
[1550823555] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 28.5% (6989904 kB) free.
[1550827900] SERVICE ALERT: sig4;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550827911] SERVICE ALERT: sig4;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550827911] SERVICE ALERT: sig4;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550827913] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550827917] SERVICE ALERT: sig4;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550827930] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1550827948] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4' was out of bounds)
[1550828020] SERVICE ALERT: sig4;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1550828024] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1550828031] SERVICE ALERT: sig4;CPU Load;OK;HARD;1;OK - load average: 1.04, 0.28, 0.10
[1550828032] SERVICE ALERT: sig4;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 180929 MB (80% inode=99%):
[1550828033] SERVICE ALERT: sig4;HDD_Root status;OK;HARD;1;DISK OK - free space: / 32981 MB (64% inode=98%):
[1550828037] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 94.8% (15402636 kB) free.
[1550828069] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 300 processes
[1550828355] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (214436 kB) free!
[1550828475] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (206712 kB) free!
[1550828595] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201832 kB) free!
[1550828599] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29228 MB (13% inode=99%):
[1550828668] SERVICE FLAPPING ALERT: sig4;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1550828668] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550828805] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.9% (23020548 kB) free.
[1550830325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550830622] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2557608 kB) free!
[1550830742] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.4% (2509044 kB) free!
[1550830862] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.2% (2465432 kB) free!
[1550830925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550831045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550831166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550831766] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550832366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550832486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550832606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550832831] SERVICE ALERT: sig4;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'sig4' was out of bounds)
[1550832831] SERVICE ALERT: sig4;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig4' was out of bounds)
[1550832833] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4' was out of bounds)
[1550832847] SERVICE ALERT: sig4;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550832877] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1550832878] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550832951] SERVICE ALERT: sig4;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 180929 MB (80% inode=99%):
[1550832951] SERVICE ALERT: sig4;CPU Load;OK;HARD;1;OK - load average: 0.41, 0.10, 0.03
[1550832953] SERVICE ALERT: sig4;HDD_Root status;OK;HARD;1;DISK OK - free space: / 32981 MB (64% inode=98%):
[1550832957] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1550832967] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 94.9% (15420324 kB) free.
[1550832999] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 304 processes
[1550833044] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (3031564 kB) free!
[1550833164] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (2984044 kB) free!
[1550833284] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (2932560 kB) free!
[1550833598] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1550833806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550834406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550834527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550834647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550835662] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1559020 kB) free!
[1550835684] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1432852 kB) free!
[1550835961] SERVICE ALERT: sig4;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550835961] SERVICE ALERT: sig4;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550835963] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550835963] HOST ALERT: sig4;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.56)
[1550835970] SERVICE ALERT: sig4;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4' was out of bounds)
[1550836001] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4' was out of bounds)
[1550836081] SERVICE ALERT: sig4;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 180929 MB (80% inode=99%):
[1550836081] SERVICE ALERT: sig4;CPU Load;OK;HARD;1;OK - load average: 0.47, 0.16, 0.06
[1550836083] SERVICE ALERT: sig4;HDD_Root status;OK;HARD;1;DISK OK - free space: / 32980 MB (64% inode=98%):
[1550836085] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1550836090] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 94.9% (15425736 kB) free.
[1550836121] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1550836321] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1550836441] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1550836447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550836561] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1550837047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550837167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550837287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550837954] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.7% (14748900 kB) free.
[1550837978] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.1% (14636592 kB) free.
[1550838361] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1550839420] SERVICE ALERT: sig4;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550839540] SERVICE ALERT: sig4;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550839660] SERVICE ALERT: sig4;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550841799] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24700 MB (10% inode=99%):
[1550842234] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1550845234] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 67 users currently logged in
[1550845534] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1550845660] SERVICE ALERT: sig4;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550845687] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550846287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550846407] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550846527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550847522] SERVICE FLAPPING ALERT: sig4;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550847727] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550848327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550848447] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550848450] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1550848514] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1550848567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550849140] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1550849260] SERVICE ALERT: sig4;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig4' was out of bounds)
[1550849283] SERVICE ALERT: sig4;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig4' was out of bounds)
[1550849283] SERVICE ALERT: sig4;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig4' was out of bounds)
[1550849283] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4' was out of bounds)
[1550849290] SERVICE ALERT: sig4;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4' was out of bounds)
[1550849323] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4' was out of bounds)
[1550849427] SERVICE ALERT: sig4;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4' was out of bounds)
[1550849427] SERVICE ALERT: sig4;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig4' was out of bounds)
[1550849427] SERVICE ALERT: sig4;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig4' was out of bounds)
[1550849694] HOST ALERT: sig4;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.56)
[1550850160] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550850162] HOST ALERT: sig5;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.57)
[1550850280] SERVICE ALERT: sig5;CPU_Procs;WARNING;HARD;1;CPU WARNING: 1 warn out of 314 processes
[1550850289] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1550850880] SERVICE ALERT: sig5;CPU_Procs;OK;HARD;1;CPU OK: 286 processes
[1550852167] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550854315] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1550854348] SERVICE ALERT: sig1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig1' was out of bounds)
[1550854353] SERVICE ALERT: sig1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig1' was out of bounds)
[1550854353] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig1' was out of bounds)
[1550854380] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1550854468] SERVICE ALERT: sig1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37124 MB (72% inode=98%):
[1550854473] SERVICE ALERT: sig1;Sensors Status;OK;HARD;1;SENSORS OK
[1550854473] SERVICE ALERT: sig1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 141397 MB (62% inode=99%):
[1550854625] SERVICE ALERT: sig2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162008 kB) free!
[1550854745] SERVICE ALERT: sig2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (159460 kB) free!
[1550854826] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550854865] SERVICE ALERT: sig2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155196 kB) free!
[1550854947] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550855067] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1550855768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550855888] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550856008] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550856065] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 58.3% (9470468 kB) free.
[1550856426] SERVICE ALERT: sig3;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550856427] SERVICE ALERT: sig3;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550856442] SERVICE ALERT: sig3;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550856442] SERVICE ALERT: sig3;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550856456] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1550856546] SERVICE ALERT: sig3;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1550856547] SERVICE ALERT: sig3;HDD_Root status;OK;HARD;1;DISK OK - free space: / 30484 MB (59% inode=98%):
[1550856550] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1550856562] SERVICE ALERT: sig3;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550856562] SERVICE ALERT: sig3;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 178202 MB (79% inode=99%):
[1550858408] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes