[1560019442] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 133.91, 46.11, 16.69
[1560019493] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 388 processes
[1560019562] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 163.55, 86.49, 35.05
[1560019660] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (1449776 kB) free!
[1560019682] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 167.48, 113.48, 51.18
[1560019719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560019780] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10149676 kB) free.
[1560020013] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (2889336 kB) free!
[1560020319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560020440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560020560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560020971] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560020980] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1759544 kB) free!
[1560021100] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10249324 kB) free.
[1560021482] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 0.68, 10.66
[1560022494] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 414 processes
[1560022682] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.56, 0.38, 3.16
[1560022864] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1560022900] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (827636 kB) free!
[1560022960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560023013] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2416412 kB) free!
[1560023021] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10062952 kB) free.
[1560023094] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 388 processes
[1560023560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560023680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560023801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560023882] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 188.51, 311.83, 171.10
[1560024002] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 25.48, 208.68, 150.37
[1560024122] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1560024122] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.56, 139.70, 132.16
[1560024213] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (3127404 kB) free!
[1560024220] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2375600 kB) free!
[1560024267] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560024340] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10741284 kB) free.
[1560024397] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1560024795] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560024807] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (3326724 kB) free.
[1560024940] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1567104 kB) free!
[1560025001] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1560025060] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10024360 kB) free.
[1560025922] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.27, 0.63, 19.24
[1560026007] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (223748 kB) free!
[1560026127] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (896712 kB) free!
[1560026201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560026247] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179580 kB) free!
[1560026321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560026441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560026860] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1568176 kB) free!
[1560026980] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9970060 kB) free.
[1560027060] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1560027294] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 409 processes
[1560027447] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.0% (1944160 kB) free!
[1560027722] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.30, 3.01
[1560028494] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 403 processes
[1560028647] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1390924 kB) free!
[1560028781] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (1516092 kB) free!
[1560028840] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 219.56, 141.89, 63.03
[1560028841] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1560028901] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9916588 kB) free.
[1560028960] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 220.41, 167.94, 82.14
[1560029080] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 220.98, 185.54, 98.99
[1560029085] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5450 MB (10% inode=97%):
[1560029374] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (2728860 kB) free!
[1560029441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560029561] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560029681] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560029685] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6987 MB (13% inode=98%):
[1560029847] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.7% (2714812 kB) free!
[1560030447] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (292284 kB) free!
[1560030574] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2428388 kB) free!
[1560030701] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2082292 kB) free!
[1560030821] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10357744 kB) free.
[1560030880] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.79, 2.36, 17.66
[1560031047] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.5% (3970420 kB) free.
[1560032013] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2427060 kB) free!
[1560032021] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (1885172 kB) free!
[1560032081] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560032094] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 368 processes
[1560032141] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10126840 kB) free.
[1560032680] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.49, 1.46, 3.83
[1560032694] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 384 processes
[1560032847] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.0% (2428164 kB) free!
[1560032968] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.9% (2409252 kB) free!
[1560033088] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.4% (2499716 kB) free!
[1560033281] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560033401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1560033521] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560033941] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8246360 kB) free!
[1560034062] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.8% (9778256 kB) free!
[1560034182] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.0% (7891512 kB) free!
[1560034288] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.3% (4262144 kB) free.
[1560036088] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3067356 kB) free!
[1560036208] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.2% (4412164 kB) free.
[1560036521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560036582] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10198008 kB) free.
[1560036894] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 361 processes
[1560037121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560037182] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9662452 kB) free!
[1560037241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1560037302] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7811268 kB) free!
[1560037361] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560037408] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.8% (2081712 kB) free!
[1560037422] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6254968 kB) free!
[1560037480] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 514.24, 366.54, 170.54
[1560037528] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.0% (3570240 kB) free.
[1560037600] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 515.06, 415.75, 212.34
[1560037720] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1560037720] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 288.09, 391.95, 230.09
[1560037774] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (2820936 kB) free!
[1560037922] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560038022] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1560038022] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (1414340 kB) free!
[1560038094] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 392 processes
[1560038128] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.0% (2764656 kB) free!
[1560038248] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.4% (2174324 kB) free!
[1560038368] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (3415036 kB) free.
[1560038374] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1560038374] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2454568 kB) free!
[1560038622] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5251908 kB) free!
[1560038968] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.4% (2657188 kB) free!
[1560039088] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.8% (2082372 kB) free!
[1560039209] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 20.0% (3237568 kB) free!
[1560039761] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560040121] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.91, 2.03, 19.12
[1560040361] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1560040409] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171864 kB) free!
[1560040481] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560040601] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1560041022] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (695328 kB) free!
[1560041609] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.4% (4606676 kB) free.
[1560041801] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1560041921] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.78, 1.52, 4.12
[1560042122] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15.86, 57.20, 35.62
[1560042213] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2823328 kB) free!
[1560042242] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.48, 38.41, 31.35
[1560042362] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.60, 25.83, 27.60
[1560042402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1560042522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560042642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560042895] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 398 processes
[1560042962] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.11, 3.67, 14.58
[1560043422] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5451244 kB) free!
[1560044163] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.28, 4.18
[1560044695] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 390 processes
[1560045209] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3181348 kB) free!
[1560045329] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.0% (4544288 kB) free.
[1560046395] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 172.05, 123.37, 57.55
[1560046422] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (563960 kB) free!
[1560046515] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 33.05, 90.73, 53.94
[1560046635] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.92, 60.87, 47.45
[1560046863] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2705644 kB) free!
[1560047013] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2349968 kB) free!
[1560047835] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 1.45, 13.33
[1560048063] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2435772 kB) free!
[1560048929] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.7% (2542388 kB) free!
[1560048984] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 43.6% (10685492 kB) free.
[1560049036] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.55, 0.45, 3.95
[1560049049] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (3796676 kB) free.
[1560049260] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1560049284] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4716532 kB) free!
[1560049344] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4716632 kB) free!
[1560049404] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.3% (4722520 kB) free!
[1560049422] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5492536 kB) free!
[1560049464] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.3% (4722456 kB) free!
[1560049783] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560049887] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 591 processes
[1560049948] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 585 processes
[1560050836] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 512.92, 374.49, 176.89
[1560050956] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 75.61, 256.53, 157.93
[1560051043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1560051076] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1560051076] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.45, 171.73, 138.81
[1560051147] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 582 processes
[1560051207] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 576 processes
[1560051449] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.8% (2567500 kB) free!
[1560051569] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (3366264 kB) free.
[1560051643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560051763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560051883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560051886] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (1701740 kB) free!
[1560052121] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560052169] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3170944 kB) free!
[1560052289] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.2% (3121904 kB) free!
[1560052409] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (3074044 kB) free!
[1560053023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10168924 kB) free.
[1560053066] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 374 processes
[1560053086] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174104 kB) free!
[1560053186] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 375 processes
[1560053306] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 1 warn out of 377 processes
[1560053476] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.97, 0.82, 11.04
[1560053609] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (3500736 kB) free.
[1560053921] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 91.78, 25.75, 9.99
[1560054041] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 161.53, 74.30, 29.69
[1560054161] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 170.65, 106.67, 46.96
[1560054210] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.8% (2083920 kB) free!
[1560054223] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (1667804 kB) free!
[1560054330] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (3262940 kB) free.
[1560054343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10138168 kB) free.
[1560054676] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.19, 0.92, 3.60
[1560054895] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560054930] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.1% (2126044 kB) free!
[1560054943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8053012 kB) free!
[1560055050] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (1455084 kB) free!
[1560055063] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 18.2% (8980968 kB) free!
[1560055170] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.4% (2823172 kB) free!
[1560055183] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6756312 kB) free!
[1560055961] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.74, 3.02, 18.65
[1560056095] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 410 processes
[1560056696] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 453 processes
[1560056983] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10292032 kB) free.
[1560057761] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.22, 2.08, 4.45
[1560058171] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (1987120 kB) free!
[1560058183] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (311472 kB) free!
[1560058304] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9454408 kB) free!
[1560058424] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7373544 kB) free!
[1560058496] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 409 processes
[1560058771] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156212 kB) free!
[1560059013] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 63.2% (15502640 kB) free.
[1560059096] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1560059096] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 444 processes
[1560059306] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1560059383] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 68.6% (16803600 kB) free.
[1560059460] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1560059464] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 70.4% (17375908 kB) free.
[1560059906] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 375 processes
[1560060027] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1560060224] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10663024 kB) free.
[1560060824] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8234880 kB) free!
[1560060883] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560060944] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9490332 kB) free!
[1560061064] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7392628 kB) free!
[1560061171] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180092 kB) free!
[1560061483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560061496] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 408 processes
[1560061596] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (2959276 kB) free!
[1560061603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560061716] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (2854132 kB) free!
[1560061723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560061771] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.5% (1873444 kB) free!
[1560061836] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.0% (2767428 kB) free!
[1560062096] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 415 processes
[1560062163] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 471.73, 203.12, 77.07
[1560062283] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 507.97, 305.84, 130.00
[1560062403] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 513.44, 374.82, 176.60
[1560062864] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10707844 kB) free.
[1560063636] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1270908 kB) free!
[1560063896] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 395 processes
[1560064065] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (910916 kB) free!
[1560064185] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10878016 kB) free.
[1560064226] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 373 processes
[1560064346] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 373 processes
[1560064466] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1560064496] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 399 processes
[1560064771] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (947504 kB) free!
[1560064771] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (2500972 kB) free!
[1560064803] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 0.47, 14.48
[1560065371] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.3% (2325460 kB) free!
[1560065371] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1560065371] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156772 kB) free!
[1560065385] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (1359464 kB) free!
[1560065476] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560065505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11050928 kB) free.
[1560066003] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.49, 0.40, 4.26
[1560066296] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 419 processes
[1560066705] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (1176036 kB) free!
[1560066825] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11178208 kB) free.
[1560066896] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 406 processes
[1560067496] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 435 processes
[1560068025] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (591544 kB) free!
[1560068097] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 385 processes
[1560068145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10971648 kB) free.
[1560068971] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (319880 kB) free!
[1560069346] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4676512 kB) free!
[1560069466] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10041920 kB) free.
[1560070172] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (3437488 kB) free.
[1560070771] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (1980900 kB) free!
[1560070772] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171608 kB) free!
[1560070892] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (1392836 kB) free!
[1560071012] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (3442588 kB) free.
[1560071266] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (381240 kB) free!
[1560071323] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560071371] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162220 kB) free!
[1560071386] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11162116 kB) free.
[1560071612] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151828 kB) free!
[1560071666] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1560071732] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.2% (1651856 kB) free!
[1560071786] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1560071852] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.8% (3057248 kB) free!
[1560071894] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (1441160 kB) free!
[1560071923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560072003] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 15.79, 3.62, 1.59
[1560072014] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 18.8% (1380524 kB) free!
[1560072043] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560072123] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 170.85, 66.97, 25.05
[1560072134] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 17.5% (1289144 kB) free!
[1560072164] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560072243] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1560072243] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 192.03, 109.44, 45.70
[1560072586] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (772932 kB) free!
[1560072706] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10380232 kB) free.
[1560072734] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.2% (6851380 kB) free.
[1560073052] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (672216 kB) free!
[1560073276] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 148.29, 60.47, 22.76
[1560073396] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 162.89, 95.11, 40.03
[1560073516] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 164.98, 118.37, 55.23
[1560073652] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.2% (2949668 kB) free!
[1560073906] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (546024 kB) free!
[1560073964] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560074026] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560074026] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10133592 kB) free.
[1560074043] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.21, 0.91, 12.65
[1560074252] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169604 kB) free!
[1560074716] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.36, 2.95, 16.82
[1560074852] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.5% (2843580 kB) free!
[1560075164] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560075227] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237480 kB) free!
[1560075243] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.38, 3.75
[1560075284] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560075347] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10714364 kB) free.
[1560075453] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (717960 kB) free!
[1560075572] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (2377828 kB) free!
[1560075884] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560076004] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560076125] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560076172] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155604 kB) free!
[1560076516] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.92, 0.70, 2.98
[1560076547] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2089992 kB) free!
[1560076653] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.0% (2915412 kB) free!
[1560076667] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10353916 kB) free.
[1560076772] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (2636444 kB) free!
[1560077372] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (220540 kB) free!
[1560077853] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (784272 kB) free!
[1560077867] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237612 kB) free!
[1560077925] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560077987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10272420 kB) free.
[1560078761] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 99.45, 79.77, 41.20
[1560078881] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 77.58, 81.47, 46.69
[1560079001] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.03, 55.12, 41.26
[1560079125] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560079245] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560079365] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560079497] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560079653] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.9% (2579596 kB) free!
[1560079787] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (647492 kB) free!
[1560079907] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10169704 kB) free.
[1560080116] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.50, 217.79, 151.67
[1560080201] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.64, 2.71, 12.57
[1560080236] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.44, 145.76, 133.29
[1560080253] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1237200 kB) free!
[1560080356] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1560080356] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.99, 97.65, 117.17
[1560080372] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (1750128 kB) free!
[1560080643] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.23, 12.34, 8.58
[1560080763] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.52, 8.40, 7.59
[1560080883] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.19, 5.66, 6.69
[1560080972] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (204736 kB) free!
[1560081402] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.19, 2.05, 4.89
[1560081707] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (883004 kB) free!
[1560081765] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1560081827] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10866740 kB) free.
[1560081897] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 437 processes
[1560082260] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1560082365] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560082485] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560082497] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 443 processes
[1560082605] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560082654] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.7% (2554244 kB) free!
[1560082683] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.57, 3.58
[1560082757] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 0.53, 10.73
[1560083098] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 421 processes
[1560083383] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4682064 kB) free!
[1560083503] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4577632 kB) free!
[1560083623] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4506352 kB) free!
[1560083627] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (935204 kB) free!
[1560083698] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1560083698] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 453 processes
[1560083747] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11448892 kB) free.
[1560083854] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (645052 kB) free!
[1560083957] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.64, 0.68, 3.40
[1560084298] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 435 processes
[1560084454] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.0% (2273356 kB) free!
[1560084898] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 439 processes
[1560084947] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (701920 kB) free!
[1560085006] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560085054] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (170768 kB) free!
[1560085067] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10642260 kB) free.
[1560085606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560085726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560085846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1560086254] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.6% (2533920 kB) free!
[1560086267] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (905368 kB) free!
[1560086372] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (1727092 kB) free!
[1560086387] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11285416 kB) free.
[1560086972] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179820 kB) free!
[1560087454] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155484 kB) free!
[1560087587] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (927288 kB) free!
[1560087646] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1560087707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11086716 kB) free.
[1560087823] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2136440 kB) free!
[1560088246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560088367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560088487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560088907] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (812820 kB) free!
[1560089028] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10562740 kB) free.
[1560089254] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.9% (2415544 kB) free!
[1560089854] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.0% (4048324 kB) free.
[1560090298] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 414 processes
[1560090455] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.2% (2310056 kB) free!
[1560090575] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (172436 kB) free!
[1560090695] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (825312 kB) free!
[1560090827] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1054612 kB) free!
[1560090887] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560090898] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 409 processes
[1560090947] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10501528 kB) free.
[1560091487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560091607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1560091728] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560092147] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (1227976 kB) free!
[1560092267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10647364 kB) free.
[1560092495] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.3% (2478216 kB) free!
[1560092883] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560093298] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 397 processes
[1560093664] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4849144 kB) free!
[1560093784] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4833788 kB) free!
[1560093898] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 421 processes
[1560093904] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4811136 kB) free!
[1560094067] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1590344 kB) free!
[1560094083] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 49.66, 122.95, 72.22
[1560094128] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560094187] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10659260 kB) free.
[1560094203] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.03, 82.38, 63.51
[1560094295] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (852200 kB) free!
[1560094323] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.27, 55.22, 55.84
[1560094504] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.2% (5481260 kB) free.
[1560094757] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560094895] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.8% (2721844 kB) free!
[1560095328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560095448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1560095523] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.22, 1.22, 15.50
[1560095569] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1560096095] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (734100 kB) free!
[1560096587] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (1927768 kB) free!
[1560096696] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (3276656 kB) free.
[1560096707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10830148 kB) free.
[1560096723] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.39, 4.51
[1560097296] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (158072 kB) free!
[1560097416] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1206584 kB) free!
[1560097498] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 416 processes
[1560097536] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (3771144 kB) free.
[1560097907] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1793964 kB) free!
[1560097969] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560098027] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10709528 kB) free.
[1560098136] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3171728 kB) free!
[1560098202] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 507.22, 324.44, 142.39
[1560098256] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (458760 kB) free!
[1560098322] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 510.37, 386.12, 187.07
[1560098372] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560098376] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.1% (1956320 kB) free!
[1560098442] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 432.21, 410.51, 220.83
[1560098698] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 383 processes
[1560099169] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560099289] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560099298] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 416 processes
[1560099410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1560099576] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (854336 kB) free!
[1560099827] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (1886692 kB) free!
[1560099947] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10514780 kB) free.
[1560100176] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (3506788 kB) free.
[1560100504] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4946912 kB) free!
[1560100624] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4945796 kB) free!
[1560100744] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (4937736 kB) free!
[1560100776] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3160864 kB) free!
[1560100842] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.34, 1.65, 18.04
[1560100896] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (716172 kB) free!
[1560101016] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.7% (2867804 kB) free!
[1560101698] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 369 processes
[1560102298] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 391 processes
[1560102347] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8213844 kB) free!
[1560102410] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560102467] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10100220 kB) free.
[1560102642] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.67, 1.59, 3.90
[1560102816] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181156 kB) free!
[1560102899] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 415 processes
[1560104211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1560104331] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560104442] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 51.80, 13.97, 5.95
[1560104451] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1560104562] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 106.86, 47.58, 19.25
[1560104682] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1560104682] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 114.22, 70.01, 30.90