[1579977359] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 368 processes
[1579977480] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1579977481] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1579977489] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.9% (30483716 kB) free.
[1579977507] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191524 kB) free!
[1579977628] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187064 kB) free!
[1579977748] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180404 kB) free!
[1579978681] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579978801] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579978921] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1579979548] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 49.7% (12267388 kB) free.
[1579980697] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163640 kB) free!
[1579980721] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579980841] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579980961] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1579981349] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (1328380 kB) free!
[1579981380] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579981469] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176516 kB) free!
[1579981561] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579981589] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (230548 kB) free!
[1579981681] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579981768] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2020/01/26 01:17:44 [error] 18368#0: *43833196 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1579981801] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1579981888] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579982178] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9642980 kB) free!
[1579982298] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9647084 kB) free!
[1579982418] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (9626092 kB) free!
[1579983001] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1579983121] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579983697] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 45.4% (11127884 kB) free.
[1579984321] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579984441] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579984562] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579985162] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579985668] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 59.1% (29116632 kB) free.
[1579985829] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9695724 kB) free!
[1579985949] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.6% (44144172 kB) free.
[1579987562] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579987683] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1579987888] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(431) < 2020/01/26 03:00:32 [error] 18361#0: *43851655 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579987890] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2301) < 172.105.59.173 [26/Jan/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579988008] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579988011] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(162) < 49.50.86.152 [26/Jan/2020:03:02:03 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579988131] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579988160] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(351) < 103.16.101.72 [26/Jan/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1579988281] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579989482] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579989602] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579989722] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579990322] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1579990773] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1579990893] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1579990923] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579991043] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579991789] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579994643] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579994763] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579994883] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579995098] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189508 kB) free!
[1579995218] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187992 kB) free!
[1579995338] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195912 kB) free!
[1579995483] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1579995603] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1579995989] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 51.1% (12613712 kB) free.
[1579996203] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 354 processes
[1579996323] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1579996924] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1579997044] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1579997644] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579997764] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579997884] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1579998484] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579999262] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.6% (22204812 kB) free.
[1579999302] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.6% (21216144 kB) free.
[1579999910] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (222544 kB) free!
[1580000030] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187756 kB) free!
[1580000150] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188960 kB) free!
[1580001484] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1580002684] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580002805] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1580002925] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1580004125] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1580004245] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580004365] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580004965] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1580006165] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580006230] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.3% (21405084 kB) free.
[1580006286] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1580006838] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187056 kB) free!
[1580006885] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1580006958] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168904 kB) free!
[1580007005] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1580007078] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1580007078] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (213012 kB) free!
[1580007125] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1580007725] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1580008925] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580009046] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1580010208] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2020/01/26 09:06:54 [error] 18364#0: *43931727 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580010328] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580010925] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.7% (21017212 kB) free.
[1580011445] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1580011529] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178984 kB) free!
[1580011565] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1580011650] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169328 kB) free!
[1580011770] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178964 kB) free!
[1580012165] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580012285] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1580012886] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580013006] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1580013126] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1580015189] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4652800 kB) free!
[1580015309] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4601324 kB) free!
[1580015429] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4532636 kB) free!
[1580016911] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.8% (21514332 kB) free.
[1580017219] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4645500 kB) free!
[1580017326] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1580017519] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166408 kB) free!
[1580017639] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191464 kB) free!
[1580017736] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167360 kB) free!
[1580017759] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174824 kB) free!
[1580017829] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2096032 kB) free!
[1580017856] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180672 kB) free!
[1580017926] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580017976] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181168 kB) free!
[1580020692] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.8% (21412604 kB) free.
[1580020725] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.2% (21864316 kB) free.
[1580020763] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.6% (40219136 kB) free.
[1580020795] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.9% (21792304 kB) free.
[1580021398] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175008 kB) free!
[1580021518] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180688 kB) free!
[1580021638] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183984 kB) free!
[1580025120] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.9% (21552152 kB) free.
[1580025126] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1580025731] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179380 kB) free!
[1580025851] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186504 kB) free!
[1580025971] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182392 kB) free!
[1580026098] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.4% (3548728 kB) free!
[1580026218] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.3% (3522232 kB) free!
[1580026326] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1580026338] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.1% (3490240 kB) free!
[1580026446] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1580026566] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1580026985] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1580027105] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1580027166] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1580028138] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1580028138] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1631980 kB) free!
[1580028305] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580028425] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1580029026] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1580029146] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580031068] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (4404476 kB) free!
[1580031188] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (4287396 kB) free!
[1580031308] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4211100 kB) free!
[1580032177] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.6% (21121684 kB) free.
[1580032205] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.2% (23095564 kB) free.
[1580032573] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1580032693] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1580032813] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1580034042] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1580034162] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1580034308] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2110172 kB) free!
[1580037010] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1580037278] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9644676 kB) free!
[1580037398] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9583196 kB) free!
[1580037518] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (9518464 kB) free!
[1580038566] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1580039977] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1580041149] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9808896 kB) free!
[1580041269] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9704576 kB) free!
[1580041389] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.5% (9590420 kB) free!
[1580041569] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 79.4% (19471348 kB) free.
[1580041727] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4874408 kB) free!
[1580041847] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.5% (4777300 kB) free!
[1580041967] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.0% (4660752 kB) free!
[1580042131] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(247) < 185.255.8.58 [26/Jan/2020:18:00:57 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1580042251] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580044146] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580044266] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580045567] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461652 kB) free!
[1580045766] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1580045851] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 203.212.70.170 [26/Jan/2020:19:00:09 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1580045972] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580047177] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4391360 kB) free!
[1580047297] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.9% (3675324 kB) free!
[1580047418] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (3664676 kB) free!
[1580047718] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4939176 kB) free!
[1580051018] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2393048 kB) free!
[1580052189] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4925068 kB) free!
[1580052666] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1580052786] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 401 processes
[1580052906] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580053813] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1580055789] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.6% (21614644 kB) free.
[1580055820] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 30.8% (15164576 kB) free.
[1580055852] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 33.9% (16685700 kB) free.
[1580055884] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.2% (21864128 kB) free.
[1580062090] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (392232 kB) free!
[1580062170] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9863088 kB) free!
[1580062210] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217596 kB) free!
[1580062290] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9862480 kB) free!
[1580062330] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (221908 kB) free!
[1580062410] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9862492 kB) free!
[1580062520] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.7% (24710364 kB) free!
[1580062640] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.7% (24709796 kB) free!
[1580062760] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.7% (24707384 kB) free!