[1559759466] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.3% (2805872 kB) free!
[1559760036] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 124.08, 53.22, 23.15
[1559760666] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1217360 kB) free!
[1559761054] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (2788504 kB) free!
[1559761266] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.8% (3045476 kB) free!
[1559761318] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559761598] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 394 processes
[1559761836] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1559761836] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.53, 2.66, 16.37
[1559761866] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175868 kB) free!
[1559762466] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.9% (1768056 kB) free!
[1559762798] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 411 processes
[1559763398] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1559763636] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.33, 1.70, 3.69
[1559764165] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559764598] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 387 processes
[1559764718] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 398 processes
[1559764838] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 395 processes
[1559764866] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1194104 kB) free!
[1559765466] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.7% (4014804 kB) free.
[1559766066] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.6% (2533404 kB) free!
[1559766186] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.4% (4448652 kB) free.
[1559766786] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (165248 kB) free!
[1559766906] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (1031488 kB) free!
[1559767026] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1032840 kB) free!
[1559767271] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 197.55, 122.77, 53.05
[1559767392] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 168.81, 141.63, 68.68
[1559767512] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.06, 94.91, 60.41
[1559767620] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2767320 kB) free!
[1559768254] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2444108 kB) free!
[1559768712] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.25, 1.98, 16.77
[1559769312] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 74.88, 253.48, 163.46
[1559770026] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (1692128 kB) free!
[1559771227] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.6% (3664092 kB) free.
[1559771712] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1559771712] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.32, 0.31, 12.53
[1559771827] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.2% (2630948 kB) free!
[1559771947] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (3767148 kB) free.
[1559772912] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.25, 0.28, 3.64
[1559773147] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (3148996 kB) free!
[1559773267] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.8% (4021908 kB) free.
[1559773765] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5556612 kB) free!
[1559773836] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559773918] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 503.79, 286.84, 119.01
[1559774038] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 511.14, 361.41, 166.68
[1559774128] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (2954428 kB) free!
[1559774158] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 285.58, 354.72, 189.65
[1559775038] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1559775638] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 396 processes
[1559775758] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 404 processes
[1559775878] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 390 processes
[1559776478] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 389 processes
[1559776558] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.05, 0.31, 14.61
[1559776572] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1559776598] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1559776718] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1559777172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559777292] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559777758] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.23, 4.17
[1559777892] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559777919] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 391 processes
[1559778012] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559778039] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 387 processes
[1559778132] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559778159] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1559778759] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 395 processes
[1559778879] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 391 processes
[1559778999] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 394 processes
[1559779165] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1289196 kB) free!
[1559780220] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2413628 kB) free!
[1559780799] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1559781067] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.9% (2260200 kB) free!
[1559781187] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 37.0% (6002100 kB) free.
[1559781399] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1559781520] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 419 processes
[1559781640] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 420 processes
[1559781742] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 22.6% (5544064 kB) free.
[1559781928] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (873580 kB) free!
[1559782387] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559782840] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 391 processes
[1559783440] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 385 processes
[1559783560] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 384 processes
[1559783680] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 384 processes
[1559783712] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559783842] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4387316 kB) free!
[1559783902] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4393180 kB) free!
[1559783962] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 17.9% (4397240 kB) free!
[1559784023] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 18.0% (4401096 kB) free!
[1559784454] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 47.9% (11734292 kB) free.
[1559784880] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1559784971] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 1 warn out of 380 processes
[1559785091] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 1 warn out of 380 processes
[1559785211] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 1 warn out of 381 processes
[1559785481] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1559785601] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1559785721] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559786411] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1559786921] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 391 processes
[1559787036] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 135.07, 97.96, 46.57
[1559787156] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 97.21, 101.63, 54.48
[1559787276] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.10, 68.40, 48.01
[1559787521] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 379 processes
[1559787641] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1559787732] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559788241] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 398 processes
[1559788361] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 403 processes
[1559788476] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.56, 2.95, 14.42
[1559788481] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 405 processes
[1559790276] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.78, 1.77, 3.56
[1559791476] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 242.80, 356.04, 201.85
[1559791512] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 101.06, 35.64, 13.05
[1559791596] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 34.20, 238.76, 177.57
[1559791632] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 121.13, 64.98, 26.54
[1559791716] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1559791716] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.86, 160.25, 156.21
[1559791752] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 123.73, 84.57, 38.39
[1559791884] SERVICE ALERT: pservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 577 processes
[1559792004] SERVICE ALERT: pservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 577 processes
[1559792124] SERVICE ALERT: pservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 576 processes
[1559792220] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (2904396 kB) free!
[1559792724] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 314 processes
[1559792820] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2377568 kB) free!
[1559792952] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 2.64, 13.62
[1559793187] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (3055688 kB) free!
[1559793307] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.7% (3679596 kB) free.
[1559794116] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.16, 2.15, 13.57
[1559794152] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.53, 0.78, 4.36
[1559794507] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.6% (2859720 kB) free!
[1559794627] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.6% (2685684 kB) free!
[1559794670] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559794747] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.5% (2521900 kB) free!
[1559794790] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559794911] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559795347] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (841588 kB) free!
[1559795846] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.25, 6.76, 5.85
[1559795854] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (2795696 kB) free!
[1559795916] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.26, 1.95, 3.59
[1559795948] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (3276852 kB) free.
[1559795966] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.10, 6.62, 5.91
[1559795974] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2359060 kB) free!
[1559796086] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.33, 6.56, 5.97
[1559796094] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2310492 kB) free!
[1559796132] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559796252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559796300] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 7 users currently logged in
[1559796372] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559796548] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.8% (1759020 kB) free!
[1559796668] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.3% (1677564 kB) free!
[1559796788] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1559796788] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (1682064 kB) free!
[1559797528] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 34.6% (8485640 kB) free.
[1559799281] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559799788] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.4% (4929588 kB) free.
[1559799928] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (957804 kB) free!
[1559800049] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (693488 kB) free!
[1559800152] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 422.89, 152.32, 55.63
[1559800169] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (489836 kB) free!
[1559800272] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 499.95, 271.35, 110.97
[1559800388] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.3% (1995464 kB) free!
[1559800392] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1559800392] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 526.11, 357.85, 162.13
[1559800508] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173176 kB) free!
[1559800628] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.3% (1832216 kB) free!
[1559801124] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (5277836 kB) free.
[1559801228] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157200 kB) free!
[1559801407] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (290736 kB) free!
[1559801527] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (120500 kB) free!
[1559801647] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (130420 kB) free!
[1559801724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175516 kB) free!
[1559801844] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176452 kB) free!
[1559801964] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157400 kB) free!
[1559801965] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5158736 kB) free!
[1559802792] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.03, 1.19, 16.59
[1559802848] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.8% (6747972 kB) free.
[1559803054] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1559803164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (4848788 kB) free.
[1559803765] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176976 kB) free!
[1559803885] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (219380 kB) free!
[1559804005] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1559804005] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (148632 kB) free!
[1559804172] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559804592] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.78, 1.27, 3.42
[1559804828] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (3333892 kB) free.
[1559805205] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 633 processes
[1559805266] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 636 processes
[1559805326] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 630 processes
[1559805428] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.0% (2432992 kB) free!
[1559805548] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (160952 kB) free!
[1559805668] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (694248 kB) free!
[1559806116] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559806573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559806693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559806813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559806868] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.4% (2658040 kB) free!
[1559807468] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (148692 kB) free!
[1559807965] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4685364 kB) free!
[1559808358] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 501.37, 275.51, 112.95
[1559808478] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 510.87, 353.85, 161.37
[1559808566] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6007760 kB) free!
[1559808569] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2859220 kB) free!
[1559808598] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 222.58, 327.47, 177.12
[1559808669] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.1% (2942896 kB) free!
[1559809448] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.5% (1284240 kB) free!
[1559809568] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.1% (1113756 kB) free!
[1559809688] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (954012 kB) free!
[1559809769] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2170944 kB) free!
[1559809869] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166200 kB) free!
[1559810005] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (4001116 kB) free.
[1559810288] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1559810288] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (130420 kB) free!
[1559810366] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1559810366] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9968556 kB) free.
[1559810486] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.85, 4.72, 4.84
[1559810605] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157172 kB) free!
[1559810725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (198256 kB) free!
[1559810846] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176996 kB) free!
[1559810966] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (250476 kB) free!
[1559810969] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1559810969] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (2473768 kB) free!
[1559811013] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559811069] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.1% (2605684 kB) free!
[1559811086] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9406972 kB) free!
[1559811206] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7233660 kB) free!
[1559811488] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 55.1% (4052796 kB) free.
[1559811569] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1285424 kB) free!
[1559811613] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559811669] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (150836 kB) free!
[1559811686] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.21, 5.55, 5.35
[1559811733] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559811806] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.82, 5.66, 5.40
[1559811853] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1559811853] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559811926] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.13, 5.98, 5.54
[1559812198] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 0.79, 12.74
[1559812269] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.6% (3825232 kB) free.
[1559812500] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1559812620] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1559812740] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1559812869] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (152828 kB) free!
[1559812989] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.6% (4151208 kB) free.
[1559813006] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10136784 kB) free.
[1559813288] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (1408844 kB) free!
[1559813398] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.40, 3.80
[1559813408] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.3% (6716820 kB) free.
[1559813589] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (156796 kB) free!
[1559813709] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.2% (2296724 kB) free!
[1559813830] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168804 kB) free!
[1559814254] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559814430] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.3% (1837212 kB) free!
[1559814792] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559814806] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (405352 kB) free!
[1559814926] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9352872 kB) free!
[1559815030] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156276 kB) free!
[1559815046] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7117624 kB) free!
[1559815630] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.4% (2170500 kB) free!
[1559816054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559816174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559816230] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (195320 kB) free!
[1559816294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559816681] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1559816830] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.8% (2393244 kB) free!
[1559817282] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1559817430] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155816 kB) free!
[1559817516] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 79.96, 22.87, 9.22
[1559817636] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 139.96, 64.72, 26.21
[1559817756] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 148.38, 92.81, 41.17
[1559817790] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45798 MB (20% inode=99%):
[1559817910] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45527 MB (20% inode=99%):
[1559818012] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1559818030] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45514 MB (20% inode=99%):
[1559818046] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10344184 kB) free.
[1559818072] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1559818132] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1559818140] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1559818192] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1559818211] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1559818260] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1559818331] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1559818380] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1559818451] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1559818980] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1559819051] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 378 processes
[1559819171] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 377 processes
[1559819230] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.7% (2879240 kB) free!
[1559819246] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233352 kB) free!
[1559819291] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 374 processes
[1559819294] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559819366] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9511384 kB) free!
[1559819408] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.0% (1250908 kB) free!
[1559819486] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7203840 kB) free!
[1559819528] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.0% (1026592 kB) free!
[1559819556] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.49, 3.12, 15.32
[1559819648] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.7% (783540 kB) free!
[1559819831] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.3% (3945156 kB) free.
[1559819894] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559819992] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1559820014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559820134] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559820248] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (122112 kB) free!
[1559820295] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.6% (2695772 kB) free!
[1559820415] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2577040 kB) free!
[1559820431] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (1232272 kB) free!
[1559820536] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.1% (2460400 kB) free!
[1559820551] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169224 kB) free!
[1559820671] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.2% (3107840 kB) free!
[1559821286] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10475664 kB) free.
[1559821357] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.21, 2.13, 4.10
[1559821871] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (153772 kB) free!
[1559821886] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (248524 kB) free!
[1559821934] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559822006] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9916992 kB) free.
[1559822246] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559822336] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1134988 kB) free!
[1559822392] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1559822534] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559822580] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1559822655] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559822683] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1559822775] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559822970] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559823206] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (656992 kB) free!
[1559823292] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1559823292] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1559823326] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10373708 kB) free.
[1559823883] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1559823883] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 354 processes
[1559824448] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.4% (6794512 kB) free.
[1559824526] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (702880 kB) free!
[1559824575] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559824646] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10326996 kB) free.
[1559824736] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.5% (13414084 kB) free.
[1559825083] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1559825126] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.69, 4.14, 4.56
[1559825175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559825295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559825415] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559825683] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 343 processes
[1559825846] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244220 kB) free!
[1559825966] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9555412 kB) free!
[1559825998] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 41.29, 8.90, 3.18
[1559826087] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7339704 kB) free!
[1559826119] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 447.23, 174.87, 64.69
[1559826239] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 501.69, 285.84, 118.72
[1559826780] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1559826780] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1559827073] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559827194] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559827215] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1559827314] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559827646] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (2454860 kB) free!
[1559827815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559827887] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10556340 kB) free.
[1559827914] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1559827935] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559828055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559828083] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 400 processes
[1559828246] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171304 kB) free!
[1559828392] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559828487] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9432684 kB) free!
[1559828607] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9531368 kB) free!
[1559828683] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 439 processes
[1559828727] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7478648 kB) free!
[1559829072] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.7% (3852252 kB) free.
[1559829239] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.22, 0.26, 12.29
[1559829284] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 404 processes
[1559829672] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162820 kB) free!
[1559829792] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.9% (3391092 kB) free.
[1559829855] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559829884] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 440 processes
[1559830392] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (152752 kB) free!
[1559830439] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.32, 3.61
[1559830456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559830512] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.8% (3544560 kB) free.
[1559830527] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10468072 kB) free.
[1559830576] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559830696] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559831084] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 436 processes
[1559831112] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (779368 kB) free!
[1559831127] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (351028 kB) free!
[1559831232] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (159676 kB) free!
[1559831248] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9928100 kB) free.
[1559831352] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.1% (2936340 kB) free!
[1559831684] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 448 processes
[1559832111] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1559832248] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1559832284] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 436 processes
[1559832496] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559832552] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (3293568 kB) free.
[1559832714] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559832834] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559832884] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 362 processes
[1559832954] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559833047] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (1092620 kB) free!
[1559833153] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3178608 kB) free!
[1559833167] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10276644 kB) free.
[1559833273] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (159420 kB) free!
[1559833393] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157428 kB) free!
[1559833696] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559833816] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559833937] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559834367] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (1157856 kB) free!
[1559834487] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10591180 kB) free.
[1559834580] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1559834754] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1559834754] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1559834820] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (2579572 kB) free!
[1559835087] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (770964 kB) free!
[1559835137] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559835192] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 196.87, 114.77, 48.47
[1559835193] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.6% (2688244 kB) free!
[1559835207] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9703044 kB) free!
[1559835239] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 148.37, 90.44, 38.80
[1559835312] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 62.11, 101.23, 52.33
[1559835327] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7502848 kB) free!
[1559835359] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 138.21, 107.77, 51.52
[1559835432] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.81, 67.90, 46.04
[1559835479] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1559835479] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 18.72, 72.13, 45.27
[1559835569] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2636924 kB) free!
[1559835694] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (3243188 kB) free!
[1559835737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559835793] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (202724 kB) free!
[1559835857] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559835977] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559836020] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2478196 kB) free!
[1559836393] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.4% (1850668 kB) free!
[1559836484] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 423 processes
[1559836632] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.73, 2.00, 13.21
[1559836680] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.72, 1.65, 12.62
[1559836769] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194576 kB) free!
[1559836894] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172096 kB) free!
[1559836994] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152820 kB) free!
[1559837084] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 403 processes
[1559837127] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10488520 kB) free.
[1559837727] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (569620 kB) free!
[1559837778] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559837832] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.48, 1.11, 4.22
[1559837847] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9576668 kB) free!
[1559837880] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.07, 0.89, 4.60
[1559837967] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.6% (7691352 kB) free!
[1559838378] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559838498] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559838618] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559838794] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.8% (2400536 kB) free!
[1559839032] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 307.29, 86.45, 30.96
[1559839152] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 483.66, 226.98, 89.18
[1559839167] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (9987196 kB) free.
[1559839272] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1559839272] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 507.49, 321.01, 140.35
[1559839394] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169600 kB) free!
[1559839484] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 401 processes
[1559839524] SERVICE ALERT: pservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 588 processes
[1559839554] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559839620] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2753480 kB) free!
[1559839644] SERVICE ALERT: pservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 588 processes
[1559839674] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559839764] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 327 processes
[1559839794] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559840084] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 390 processes
[1559840367] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (1077936 kB) free!
[1559840418] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559840487] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 20.0% (9853224 kB) free!
[1559840607] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.8% (7775844 kB) free!
[1559841018] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559841139] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559841194] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.7% (2219476 kB) free!
[1559841259] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559841672] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 0.71, 18.28
[1559841794] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159052 kB) free!
[1559841884] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 407 processes
[1559842407] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10402852 kB) free.
[1559842485] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 426 processes
[1559842620] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1559842620] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2452880 kB) free!
[1559842995] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.7% (2062624 kB) free!
[1559843059] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559843472] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.32, 2.92
[1559843595] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (449360 kB) free!
[1559843607] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (1646724 kB) free!
[1559843659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559843727] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10186832 kB) free.
[1559843779] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559843821] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2801940 kB) free!
[1559843899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559844285] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 421 processes
[1559844421] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202928 kB) free!
[1559844885] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 389 processes
[1559844927] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1039372 kB) free!
[1559845021] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2766140 kB) free!
[1559845047] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9619344 kB) free!
[1559845167] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7545260 kB) free!