[1561228391] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561228511] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561228608] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3239524 kB) free!
[1561228728] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.9% (3232916 kB) free!
[1561228739] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1561228739] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199312 kB) free!
[1561228795] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188524 kB) free!
[1561228848] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.8% (3225652 kB) free!
[1561228915] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194356 kB) free!
[1561228942] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561229036] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187620 kB) free!
[1561229111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561229231] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561229831] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561229951] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561230551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561230577] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9215556 kB) free!
[1561230671] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561230697] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13202484 kB) free.
[1561231272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561231298] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9192388 kB) free!
[1561231392] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561231396] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561231418] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13208380 kB) free.
[1561231992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561232018] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9125080 kB) free!
[1561232055] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (1581172 kB) free!
[1561232112] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561232115] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (1584584 kB) free!
[1561232138] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13198600 kB) free.
[1561232175] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 6.5% (1584264 kB) free!
[1561232235] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 6.5% (1585448 kB) free!
[1561232712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561232738] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9012208 kB) free!
[1561232832] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561232858] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13202752 kB) free.
[1561233432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561233458] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8882624 kB) free!
[1561233552] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561233578] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13205516 kB) free.
[1561234152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561234178] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8725304 kB) free!
[1561234273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561234298] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13194840 kB) free.
[1561234872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561234898] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8582932 kB) free!
[1561234992] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561234996] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 40.21, 10.91, 4.77
[1561235018] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13203372 kB) free.
[1561235116] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 92.30, 40.49, 16.35
[1561235236] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 98.53, 60.04, 26.44
[1561235592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561235712] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561235817] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1106128 kB) free!
[1561235881] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2533160 kB) free!
[1561235937] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (159304 kB) free!
[1561236001] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 77.5% (12603300 kB) free.
[1561236057] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (187144 kB) free!
[1561236312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561236432] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561236436] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.63, 5.30, 17.71
[1561236539] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1561237632] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561237656] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561237752] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561237942] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 23.68, 5.56, 1.99
[1561238062] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 75.11, 31.23, 11.83
[1561238182] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 82.09, 48.42, 20.48
[1561238236] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.61, 1.58, 3.94
[1561238352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561238472] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561238511] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2750608 kB) free!
[1561238782] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.18, 13.04, 17.65
[1561239072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561239192] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561239792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561239818] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8784748 kB) free!
[1561239913] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561239938] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13182096 kB) free.
[1561239982] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.43, 4.98
[1561240512] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561240632] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561241138] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9825936 kB) free!
[1561241232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561241258] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13191720 kB) free.
[1561241352] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561241782] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 45.20, 10.19, 4.20
[1561241858] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9564924 kB) free!
[1561241902] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 181.93, 74.04, 28.32
[1561241979] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13194028 kB) free.
[1561242022] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1561242022] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 200.41, 116.79, 49.53
[1561242111] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (995024 kB) free!
[1561242552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561242578] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9117520 kB) free!
[1561242672] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561242698] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13184304 kB) free.
[1561242711] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.8% (3392556 kB) free!
[1561243273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561243298] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8733676 kB) free!
[1561243393] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561243418] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13190432 kB) free.
[1561243822] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 0.91, 12.82
[1561243993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561244113] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561244618] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9779044 kB) free!
[1561244738] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13182212 kB) free.
[1561245022] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.33, 3.74
[1561245313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561245338] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9510220 kB) free!
[1561245433] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561245458] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13187824 kB) free.
[1561246033] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561246058] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9146232 kB) free!
[1561246154] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561246178] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13188084 kB) free.
[1561246753] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561246779] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8785900 kB) free!
[1561246873] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561246899] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12533600 kB) free.
[1561246993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561247225] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9875584 kB) free!
[1561247345] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9875660 kB) free!
[1561247465] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9877220 kB) free!
[1561247856] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.87, 8.83, 7.29
[1561247939] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2882792 kB) free!
[1561247976] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.39, 6.02, 6.44
[1561248096] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 4.18, 5.72
[1561248696] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.58, 0.99, 3.23
[1561248699] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8783508 kB) free!
[1561248819] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12691556 kB) free.
[1561249565] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2289704 kB) free!
[1561250019] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9832468 kB) free!
[1561250139] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13130784 kB) free.
[1561250457] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (2734260 kB) free!
[1561251057] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179836 kB) free!
[1561251658] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 62.2% (10104904 kB) free.
[1561251939] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9474256 kB) free!
[1561252060] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12974388 kB) free.
[1561252248] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1536576 kB) free!
[1561252993] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561253339] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2431036 kB) free!
[1561253593] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561253713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561253834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561253859] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8950552 kB) free!
[1561253979] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12706828 kB) free.
[1561254965] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (2481008 kB) free!
[1561255822] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561256236] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 146.61, 61.57, 24.11
[1561256356] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 159.85, 94.74, 40.81
[1561256379] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8592520 kB) free!
[1561256476] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 162.37, 117.20, 55.57
[1561256499] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12945012 kB) free.
[1561256511] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2443848 kB) free!
[1561256688] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5162 MB (10% inode=97%):
[1561257099] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8424628 kB) free!
[1561257220] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12953572 kB) free.
[1561257288] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6999 MB (13% inode=98%):
[1561258034] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1561258034] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561258035] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 28.0% (6857908 kB) free.
[1561258258] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (189432 kB) free!
[1561258276] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.38, 2.18, 15.50
[1561258378] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (188700 kB) free!
[1561258419] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9647348 kB) free!
[1561258499] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1561258499] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (193124 kB) free!
[1561258539] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12946128 kB) free.
[1561259139] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9428644 kB) free!
[1561259259] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12953600 kB) free.
[1561259834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561259859] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9225456 kB) free!
[1561259954] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561259979] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12952628 kB) free.
[1561260076] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.81, 1.79, 3.81
[1561260554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561260579] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9084280 kB) free!
[1561260674] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561260699] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12940692 kB) free.
[1561261274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561261299] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8847488 kB) free!
[1561261394] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561261419] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12950748 kB) free.
[1561261565] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2100616 kB) free!
[1561261995] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561262020] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8990952 kB) free!
[1561262115] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561262140] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12939252 kB) free.
[1561262715] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561262740] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8345388 kB) free!
[1561262835] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561262860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12946580 kB) free.
[1561263435] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561263555] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561263696] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 242.45, 173.66, 81.04
[1561263816] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 54.25, 133.45, 78.25
[1561263936] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7.79, 89.49, 68.83
[1561264139] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2742644 kB) free!
[1561264155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561264248] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.0% (13482584 kB) free.
[1561264275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561264499] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1561264740] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180240 kB) free!
[1561264876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561264996] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561265136] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.28, 2.80, 19.69
[1561265334] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561265454] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1561265511] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 26.0% (6373728 kB) free.
[1561265596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561265716] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561266111] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174064 kB) free!
[1561266231] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200844 kB) free!
[1561266316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561266351] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171624 kB) free!
[1561266436] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561266540] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.3% (18082532 kB) free.
[1561266937] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.86, 0.66, 3.56
[1561267036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561267060] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8698192 kB) free!
[1561267156] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1561267156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561267180] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12950096 kB) free.
[1561267757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561267781] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8740368 kB) free!
[1561267877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561267901] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.3% (12940784 kB) free.
[1561268151] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1561268151] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 75.7% (18553492 kB) free.
[1561268477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561268501] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8542664 kB) free!
[1561268597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561268621] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12324416 kB) free.
[1561268717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561269659] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3188004 kB) free!
[1561269779] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (3192596 kB) free!
[1561269899] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (3191720 kB) free!
[1561269965] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 14.2% (3487192 kB) free!
[1561270276] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 242.12, 181.50, 87.89
[1561270396] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 242.21, 201.59, 106.61
[1561270488] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5456 MB (10% inode=97%):
[1561270516] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 147.82, 192.12, 115.42
[1561271088] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6933 MB (13% inode=98%):
[1561271865] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.25, 6.02, 5.22
[1561271985] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.61, 6.77, 5.60
[1561272105] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.23, 6.37, 5.59
[1561272316] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.83, 2.31, 18.26
[1561272665] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4899732 kB) free!
[1561273306] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.29, 3.80, 4.33
[1561273517] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561273822] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 225.93, 161.88, 75.39
[1561273942] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 50.57, 124.39, 72.80
[1561274062] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7.24, 83.41, 64.04
[1561274116] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.55, 2.48, 4.59
[1561274448] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.7% (2873760 kB) free!
[1561274568] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.9% (2746324 kB) free!
[1561274621] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9612496 kB) free!
[1561274688] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.2% (2636536 kB) free!
[1561274717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561274741] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12925252 kB) free.
[1561274837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561275262] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.22, 1.78, 17.79
[1561275341] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9548240 kB) free!
[1561275461] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12933116 kB) free.
[1561275654] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1561275774] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561275889] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1315932 kB) free!
[1561276037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561276061] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9457392 kB) free!
[1561276157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561276181] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12931352 kB) free.
[1561276537] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 26.57, 19.90, 9.74
[1561276565] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2323100 kB) free!
[1561276657] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.79, 16.75, 9.93
[1561276757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561276777] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.48, 11.34, 8.77
[1561276781] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9315232 kB) free!
[1561276877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561276901] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12926592 kB) free.
[1561277062] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.80, 3.81
[1561277165] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.6% (3094504 kB) free!
[1561277377] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 1.88, 4.80
[1561277477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561277501] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9214256 kB) free!
[1561277597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561277621] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12933784 kB) free.
[1561278197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561278221] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9008968 kB) free!
[1561278299] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1554204 kB) free!
[1561278316] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 175.66, 132.68, 64.80
[1561278318] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561278342] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12920328 kB) free.
[1561278351] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561278436] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 175.78, 146.96, 78.26
[1561278556] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1561278556] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 77.99, 129.93, 81.06
[1561278836] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (2569232 kB) free!
[1561278917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561278941] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8746244 kB) free!
[1561279037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561279061] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12931376 kB) free.
[1561279374] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1561279436] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2414000 kB) free!
[1561279494] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561279637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561279661] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8550644 kB) free!
[1561279757] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561279781] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12932088 kB) free.
[1561280356] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.66, 2.17, 13.29
[1561280357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561280381] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8730572 kB) free!
[1561280477] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561280501] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12923700 kB) free.
[1561281077] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1561281101] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8766988 kB) free!
[1561281166] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (1470564 kB) free!
[1561281198] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561281221] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12929960 kB) free.
[1561281286] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 18.8% (1384268 kB) free!
[1561281406] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 17.6% (1292368 kB) free!
[1561281797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561281821] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8769612 kB) free!
[1561281894] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1561281917] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561281941] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12918596 kB) free.
[1561282014] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1561282156] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.51, 1.74, 3.46
[1561282517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561282541] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8713460 kB) free!
[1561282606] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.5% (6804200 kB) free.
[1561282637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561282662] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12924904 kB) free.
[1561283237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561283357] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561283814] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1561283861] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9580944 kB) free!
[1561283934] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1561283957] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561283981] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12924156 kB) free.
[1561284078] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561284677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561284797] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561285397] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561285517] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561286117] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561286237] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561286837] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561286958] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561287557] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561287677] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561287722] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.6% (21712824 kB) free.
[1561287838] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39282032 kB) free.
[1561287942] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.8% (14599536 kB) free.
[1561288027] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.2% (14826616 kB) free.
[1561288097] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 68.6% (11143768 kB) free.
[1561288181] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9697848 kB) free!
[1561288277] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561288301] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12917656 kB) free.
[1561288397] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561288902] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9626172 kB) free!
[1561288997] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561289022] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12923776 kB) free.
[1561289117] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1561289622] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9576616 kB) free!
[1561289742] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12915516 kB) free.
[1561289934] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561290054] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1561290318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561290342] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9511296 kB) free!
[1561290438] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561290462] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12920236 kB) free.
[1561291038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561291062] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9456316 kB) free!
[1561291158] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561291182] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12922276 kB) free.
[1561291758] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561291782] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9392480 kB) free!
[1561291878] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561291902] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12915060 kB) free.
[1561292356] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561292478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561292502] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9385340 kB) free!
[1561292598] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561292622] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12921684 kB) free.
[1561293199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561293223] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9389576 kB) free!
[1561293319] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561293343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12910780 kB) free.
[1561293919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561293943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9448268 kB) free!
[1561294039] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561294063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12919152 kB) free.
[1561294639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561294663] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9398844 kB) free!
[1561294759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561294783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12919664 kB) free.
[1561295359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561295383] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9401484 kB) free!
[1561295480] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561295503] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12917584 kB) free.
[1561295956] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 293.03, 223.00, 108.89
[1561296076] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 292.95, 246.14, 131.21
[1561296079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561296103] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9398828 kB) free!
[1561296196] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 100.19, 204.74, 131.15
[1561296199] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561296223] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12920476 kB) free.
[1561296262] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 88.22, 29.05, 10.46
[1561296383] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 112.71, 57.99, 23.32
[1561296503] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 115.51, 77.19, 34.56
[1561296799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561296823] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9303716 kB) free!
[1561296919] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561296943] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12910472 kB) free.
[1561297177] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 23.10, 48.27, 28.53
[1561297297] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.33, 32.38, 25.10
[1561297417] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.98, 21.85, 22.12
[1561297519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561297543] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9340292 kB) free!
[1561297639] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561297664] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12918600 kB) free.
[1561297703] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.16, 2.20, 12.19
[1561298017] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 3.32, 11.80
[1561298239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561298263] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9223636 kB) free!
[1561298360] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561298383] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12918776 kB) free.
[1561298596] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.02, 1.95, 11.54
[1561298903] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.31, 3.53
[1561298959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561298983] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9271448 kB) free!
[1561299079] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561299103] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12914088 kB) free.
[1561299217] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.50, 3.55
[1561299679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561299703] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8635112 kB) free!
[1561299796] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.03, 1.88, 4.52
[1561299799] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561299823] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12916108 kB) free.
[1561300399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561300519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561301023] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9307040 kB) free!
[1561301143] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12909364 kB) free.
[1561301719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561301743] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8592048 kB) free!
[1561301840] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561301863] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12913560 kB) free.
[1561302439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561302559] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561303063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9279864 kB) free!
[1561303183] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12915536 kB) free.
[1561303566] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.1% (4926860 kB) free.
[1561303759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561303783] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8593160 kB) free!
[1561303879] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561303903] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12913444 kB) free.
[1561304479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561304599] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561305103] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9244108 kB) free!
[1561305223] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12912964 kB) free.
[1561305799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561305823] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8519276 kB) free!
[1561305919] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561305943] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12905360 kB) free.
[1561306519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561306639] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561306996] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 23.91, 9.08, 4.27
[1561307116] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 29.38, 16.04, 7.41
[1561307143] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9325308 kB) free!
[1561307236] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 29.87, 20.61, 10.13
[1561307263] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12912148 kB) free.
[1561307839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561307863] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9313168 kB) free!
[1561307959] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561307983] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12913972 kB) free.
[1561308560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561308583] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9336108 kB) free!
[1561308680] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561308703] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12906740 kB) free.
[1561309036] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1561309036] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.78, 1.83, 3.95
[1561309280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561309303] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9610392 kB) free!
[1561309400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561309423] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12916108 kB) free.
[1561309520] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1561310023] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9577748 kB) free!
[1561310120] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561310143] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12907192 kB) free.
[1561310240] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561310304] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 211.06, 147.04, 67.34
[1561310424] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 66.04, 124.30, 69.50
[1561310544] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.98, 83.20, 61.08
[1561310743] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9465020 kB) free!
[1561310863] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12914720 kB) free.
[1561311441] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561311463] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9341444 kB) free!
[1561311561] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561311584] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12911396 kB) free.
[1561311655] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1561311744] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.09, 1.72, 16.94
[1561311775] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1561312161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561312183] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9195304 kB) free!
[1561312281] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561312303] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12912792 kB) free.
[1561312881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561312903] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8967696 kB) free!
[1561312944] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.34, 4.86
[1561313001] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561313023] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12911976 kB) free.
[1561313601] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561313623] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8686904 kB) free!
[1561313721] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561313743] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12903260 kB) free.
[1561314322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1561314343] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8558956 kB) free!
[1561314442] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561314463] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12913168 kB) free.