[1576002761] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sservercluster1' was out of bounds)
[1576002817] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1576002820] HOST ALERT: sservercluster1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.49)
[1576002834] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sservercluster1' was out of bounds)
[1576002854] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sservercluster1' was out of bounds)
[1576002854] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sservercluster1' was out of bounds)
[1576002883] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sservercluster1' was out of bounds)
[1576002908] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(122) < 2019/12/11 00:05:05 [error] 24130#0: *15608624 connect() failed (113: No route to host) while connecting to upstream, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576002937] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sservercluster1' was out of bounds)
[1576002947] HOST ALERT: sservercluster1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1576002954] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1576002974] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.22, 0.08
[1576002974] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 297 processes
[1576002982] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1576002985] HOST ALERT: sservercluster2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.61)
[1576002989] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sservercluster2' was out of bounds)
[1576003028] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(271) < 2019/12/11 00:07:08 [error] 24130#0: *15609674 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576003039] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 67.3% (33164200 kB) free.
[1576003042] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sservercluster2' was out of bounds)
[1576003049] HOST ALERT: sservercluster2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1576003102] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 324 processes
[1576003109] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 95.3% (23348568 kB) free.
[1576003148] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(9526) < 2019/12/11 00:09:08 [error] 24127#0: *15619680 no live upstreams while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "backend_3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576003162] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.42, 0.17
[1576003205] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183988 kB) free!
[1576003206] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1814) < 2019/12/11 00:10:05 [error] 26846#0: *18825646 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576003278] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(29867) < 103.255.146.164 [11/Dec/2019:00:11:18 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576003314] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192528 kB) free!
[1576003325] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176076 kB) free!
[1576003326] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2034) < 2019/12/11 00:12:06 [error] 26847#0: *18828647 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576003398] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(18278) < 103.255.146.164 [11/Dec/2019:00:13:18 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576003434] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4147) < 103.67.237.219 [11/Dec/2019:00:13:52 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1576003434] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194552 kB) free!
[1576003445] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186128 kB) free!
[1576003446] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(2024) < 2019/12/11 00:14:05 [error] 26838#0: *18831680 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576003481] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1576003519] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(18989) < 103.255.146.164 [11/Dec/2019:00:15:18 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576003537] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 96.1% (23550872 kB) free.
[1576003555] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197396 kB) free!
[1576003555] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1607) < 103.16.101.72 [11/Dec/2019:00:15:54 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1576003559] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'webservercluster2' was out of bounds)
[1576003573] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'webservercluster2' was out of bounds)
[1576003640] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'webservercluster2' was out of bounds)
[1576003670] SERVICE ALERT: webservercluster2;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'webservercluster2' was out of bounds)
[1576003670] SERVICE ALERT: webservercluster2;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'webservercluster2' was out of bounds)
[1576003670] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;SOFT;1;(Return code of 255 for service 'Sensors Status' on host 'webservercluster2' was out of bounds)
[1576003670] SERVICE ALERT: webservercluster2;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'webservercluster2' was out of bounds)
[1576003677] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;HARD;3;(1564) < 103.16.101.72 [11/Dec/2019:00:17:55 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1576003679] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;(Return code of 255 for service 'CPU_Procs' on host 'webservercluster2' was out of bounds)
[1576003683] SERVICE ALERT: webservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'webservercluster2' was out of bounds)
[1576003684] SERVICE ALERT: webservercluster2;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'webservercluster2' was out of bounds)
[1576003693] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;SOFT;2;(Return code of 255 for service 'Current Users' on host 'webservercluster2' was out of bounds)
[1576003760] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;(Return code of 255 for service 'Memory' on host 'webservercluster2' was out of bounds)
[1576003790] SERVICE ALERT: webservercluster2;HDD_Home status;CRITICAL;SOFT;2;(Return code of 255 for service 'HDD_Home status' on host 'webservercluster2' was out of bounds)
[1576003790] SERVICE ALERT: webservercluster2;CPU Load;CRITICAL;SOFT;2;(Return code of 255 for service 'CPU Load' on host 'webservercluster2' was out of bounds)
[1576003790] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;SOFT;2;(Return code of 255 for service 'Sensors Status' on host 'webservercluster2' was out of bounds)
[1576003790] SERVICE ALERT: webservercluster2;Total Processes;CRITICAL;SOFT;2;(Return code of 255 for service 'Total Processes' on host 'webservercluster2' was out of bounds)
[1576003799] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;(Return code of 255 for service 'CPU_Procs' on host 'webservercluster2' was out of bounds)
[1576003803] SERVICE ALERT: webservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'webservercluster2' was out of bounds)
[1576003804] SERVICE ALERT: webservercluster2;HDD_Root status;CRITICAL;SOFT;2;(Return code of 255 for service 'HDD_Root status' on host 'webservercluster2' was out of bounds)
[1576003813] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;(Return code of 255 for service 'Current Users' on host 'webservercluster2' was out of bounds)
[1576003880] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;(Return code of 255 for service 'Memory' on host 'webservercluster2' was out of bounds)
[1576003910] SERVICE ALERT: webservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 223681 MB (99% inode=99%):
[1576003910] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.07, 0.19, 0.13
[1576003910] SERVICE ALERT: webservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1576003910] SERVICE ALERT: webservercluster2;Total Processes;OK;HARD;1;PROCS OK: 298 processes
[1576003923] SERVICE ALERT: webservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1576003924] SERVICE ALERT: webservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 40198 MB (78% inode=99%):
[1576004399] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1576004413] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1576004480] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 84.1% (13675328 kB) free.
[1576004835] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.3293.1357> (10.147.212.158:33732 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1576004955] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1576005325] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576005505] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576005845] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.1% (10635364 kB) free.
[1576006401] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4714812 kB) free!
[1576006445] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169452 kB) free!
[1576006547] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576006565] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188244 kB) free!
[1576006578] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576006685] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171248 kB) free!
[1576009547] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/12/11 01:55:43 [error] 24128#0: *15968624 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576009668] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(107) < 2019/12/11 01:57:45 [error] 24120#0: *15968930 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576009788] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2019/12/11 01:58:11 [error] 24128#0: *15968774 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576010388] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1576010388] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576011497] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (228856 kB) free!
[1576011617] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227872 kB) free!
[1576011737] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229692 kB) free!
[1576012337] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.8% (44764120 kB) free.
[1576012728] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576012848] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576013389] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(121) < 2019/12/11 02:55:03 [error] 24121#0: *15984270 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576013510] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576014109] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1736) < 2019/12/11 03:11:16 [error] 24133#0: *15990807 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576014229] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576014389] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1592) < 169.38.76.212 [11/Dec/2019:03:10:13 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576014509] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576021429] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(123) < 2019/12/11 05:12:19 [error] 24128#0: *16023563 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576021549] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576022149] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576023648] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576023769] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576024812] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1259 processes
[1576024932] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1259 processes
[1576025052] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 368 processes
[1576025137] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187636 kB) free!
[1576025257] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.6% (22219032 kB) free.
[1576025755] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.3% (20899828 kB) free.
[1576028546] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1576028568] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576028666] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1576028688] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576028749] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(125) < 2019/12/11 07:13:10 [error] 24128#0: *16056043 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576028870] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576031088] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1576031208] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576031809] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576031929] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576036069] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(169) < 2019/12/11 09:14:48 [error] 24128#0: *16088443 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576036189] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576040418] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1576040538] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1576041201] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 24.0% (11839832 kB) free.
[1576041211] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.2% (23004592 kB) free.
[1576042789] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/11 11:09:46 [error] 24121#0: *16119834 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576042909] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(45) < 2019/12/11 11:11:37 [error] 24128#0: *16119874 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1576042938] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 325 processes
[1576043014] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181456 kB) free!
[1576043029] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(7) < 2019/12/11 11:13:48 [error] 24120#0: *16120518 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576043058] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1576043134] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (1054344 kB) free!
[1576043254] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (432336 kB) free!
[1576044229] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576045429] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(19) < 2019/12/11 11:52:34 [error] 24120#0: *16130557 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1576045549] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576046755] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4354492 kB) free!
[1576046875] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4169332 kB) free!
[1576046995] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (3994280 kB) free!
[1576048195] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2150532 kB) free!
[1576050949] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(44) < 2019/12/11 13:19:38 [error] 24120#0: *16153854 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576051069] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576051105] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(19) < 2019/12/11 13:22:27 [error] 26844#0: *19043599 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:51769/733332, bytes from/to upstream:733332/51769
[1576051225] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576051396] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.5% (22660752 kB) free.
[1576051427] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22857208 kB) free.
[1576051799] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1576051919] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 297 processes
[1576052040] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 296 processes
[1576052666] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1576052786] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1576052857] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (4125088 kB) free!
[1576052870] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(111) < 2019/12/11 13:56:15 [error] 24122#0: *16163730 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576052977] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.1% (3955852 kB) free!
[1576052990] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576053097] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.8% (3626236 kB) free!
[1576053258] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 328 processes
[1576053378] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 324 processes
[1576053498] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1576053801] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9828304 kB) free!
[1576053921] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 21.1% (10404108 kB) free.
[1576054297] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2127828 kB) free!
[1576054729] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1576054849] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1576054937] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9092220 kB) free!
[1576055057] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8892584 kB) free!
[1576055121] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8119488 kB) free!
[1576055177] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.7% (8701592 kB) free!
[1576055241] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.5% (8118256 kB) free!
[1576055361] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (8118444 kB) free!
[1576055637] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (1708824 kB) free!
[1576055757] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (1579080 kB) free!
[1576055877] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1450392 kB) free!
[1576056058] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3603088 kB) free!
[1576056178] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3345772 kB) free!
[1576056298] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3162604 kB) free!
[1576056899] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 79.4% (17821316 kB) free.
[1576057577] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4052460 kB) free!
[1576057790] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(54) < 2019/12/11 15:15:26 [error] 24128#0: *16186169 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576057910] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576058961] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4623808 kB) free!
[1576059169] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25293820 kB) free!
[1576059289] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.7% (47077948 kB) free.
[1576059710] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2019/12/11 15:51:49 [error] 24121#0: *16196443 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1576059831] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(58) < 2019/12/11 15:53:49 [error] 24128#0: *16196993 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1576059951] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(52) < 2019/12/11 15:55:50 [error] 24123#0: *16197344 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576061151] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576061298] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1576061418] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1576061909] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4671416 kB) free!
[1576062029] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4589060 kB) free!
[1576062149] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (4528516 kB) free!
[1576062653] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.4% (20571952 kB) free.
[1576062690] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 66.9% (32961992 kB) free.
[1576062723] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 23.3% (11479184 kB) free.
[1576062751] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 62.6% (15332684 kB) free.
[1576063411] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4684076 kB) free!
[1576063531] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4553300 kB) free!
[1576063651] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4374144 kB) free!
[1576064689] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16114768 kB) free!
[1576064809] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.2% (48981088 kB) free.
[1576065056] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.5% (3335192 kB) free!
[1576065176] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.8% (3166036 kB) free!
[1576065297] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2193580 kB) free!
[1576065451] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2243676 kB) free!
[1576065750] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2357392 kB) free!
[1576067151] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(31) < 2019/12/11 17:52:35 [error] 24128#0: *16222125 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576067271] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576067425] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/11 17:50:45 [error] 26853#0: *19296145 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:1087/17, bytes from/to upstream:17/1087
[1576067545] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576068649] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576068737] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9588040 kB) free!
[1576068769] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1576068857] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9495988 kB) free!
[1576068977] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (9393816 kB) free!
[1576069609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25025896 kB) free!
[1576069729] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.3% (49220000 kB) free.
[1576069818] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1576069938] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1576070058] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1576070329] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24431928 kB) free!
[1576070449] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.8% (48501716 kB) free.
[1576070733] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 366 processes
[1576070757] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.5% (20367548 kB) free.
[1576070803] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 63.6% (15578664 kB) free.
[1576070832] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 57.9% (14194504 kB) free.
[1576070853] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 413 processes
[1576070871] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(69) < 2019/12/11 18:57:42 [error] 24128#0: *16233098 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576070991] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2019/12/11 18:58:18 [error] 24124#0: *16233084 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576071049] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25282740 kB) free!
[1576071111] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576071169] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48592464 kB) free.
[1576071712] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2019/12/11 19:03:01 [error] 24128#0: *16233817 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576071769] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (24984416 kB) free!
[1576071832] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576071889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.2% (47757688 kB) free.
[1576072489] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25711300 kB) free!
[1576072609] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (47547044 kB) free.
[1576073209] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21663156 kB) free!
[1576073329] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.6% (46936020 kB) free.
[1576073765] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174064 kB) free!
[1576073886] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168680 kB) free!
[1576073930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20962616 kB) free!
[1576074006] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184888 kB) free!
[1576074050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.2% (46364496 kB) free.
[1576074650] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20152128 kB) free!
[1576074770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45806156 kB) free.
[1576075370] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22546568 kB) free!
[1576075490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45452232 kB) free.
[1576075577] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4828188 kB) free!
[1576076090] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21160784 kB) free!
[1576076210] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44997140 kB) free.
[1576076658] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1576076778] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1576076810] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19395736 kB) free!
[1576076898] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1576076930] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30662664 kB) free.
[1576077530] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20103888 kB) free!
[1576077650] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16507660 kB) free!
[1576077770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41373452 kB) free.
[1576078370] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15713208 kB) free!
[1576078490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41736640 kB) free.
[1576078957] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4850352 kB) free!
[1576079032] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/12/11 21:13:32 [error] 24121#0: *16255958 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576079077] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 21.6% (5298880 kB) free.
[1576079091] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17014296 kB) free!
[1576079152] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576079211] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41581188 kB) free.
[1576079811] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18492048 kB) free!
[1576079931] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41927228 kB) free.
[1576080095] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9110460 kB) free!
[1576080215] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9107328 kB) free!
[1576080240] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1576080267] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.4% (22317036 kB) free.
[1576080277] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4858864 kB) free!
[1576080314] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 47.0% (23165360 kB) free.
[1576080336] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (9106644 kB) free!
[1576080398] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4822964 kB) free!
[1576080518] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4782124 kB) free!
[1576080531] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19345372 kB) free!
[1576080651] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41931056 kB) free.
[1576081251] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19594032 kB) free!
[1576081371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41892164 kB) free.
[1576081971] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19538744 kB) free!
[1576082091] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40941800 kB) free.
[1576082676] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192176 kB) free!
[1576082691] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18517228 kB) free!
[1576082786] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1576082796] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189268 kB) free!
[1576082811] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39241124 kB) free.
[1576082906] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 391 processes
[1576082916] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181888 kB) free!
[1576083352] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(80) < 2019/12/11 22:25:19 [error] 24128#0: *16267511 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576083473] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576084011] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22460288 kB) free!
[1576084131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42047968 kB) free.
[1576084731] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20684112 kB) free!
[1576084852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40683356 kB) free.
[1576086051] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21209076 kB) free!
[1576086171] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41596940 kB) free.
[1576086771] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21060344 kB) free!
[1576086891] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41542152 kB) free.
[1576087118] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2314644 kB) free!
[1576087491] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20616188 kB) free!
[1576087611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40015440 kB) free.
[1576088239] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216756 kB) free!
[1576088359] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242836 kB) free!
[1576088421] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 207815 MB (19% inode=99%):
[1576088479] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (249508 kB) free!
[1576088539] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.6% (19278200 kB) free!
[1576088541] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 200936 MB (18% inode=99%):
[1576088659] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11849152 kB) free!
[1576088661] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 193875 MB (18% inode=99%):
[1576088779] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11866828 kB) free!
[1576088811] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21934724 kB) free!
[1576088931] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39147432 kB) free.