[1565548225] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11534672 kB) free.
[1565548230] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1565548313] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 383 processes
[1565548551] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.22, 26.79, 20.14
[1565548723] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194296 kB) free!
[1565548762] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.10, 4.87, 11.20
[1565548825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7850532 kB) free!
[1565548829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565548843] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (398180 kB) free!
[1565548945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11549596 kB) free.
[1565548949] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565548963] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (378324 kB) free!
[1565549513] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 3 warn out of 379 processes
[1565549545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8153004 kB) free!
[1565549549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565549666] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6533752 kB) free!
[1565549669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565549786] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11541080 kB) free.
[1565549789] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565550351] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.79, 7.91, 13.47
[1565550386] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7007836 kB) free!
[1565550389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565550506] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11549004 kB) free.
[1565550509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565551106] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7460020 kB) free!
[1565551110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565551226] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11538600 kB) free.
[1565551230] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565551776] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 67.10, 34.89, 20.35
[1565551826] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7814132 kB) free!
[1565551830] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565551832] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.5% (10002944 kB) free.
[1565551946] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11546552 kB) free.
[1565551950] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565552362] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 45.28, 24.69, 16.32
[1565552546] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7935500 kB) free!
[1565552550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565552666] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11549860 kB) free.
[1565552670] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565552962] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.78, 3.92, 8.88
[1565552976] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 5.08, 10.98
[1565553266] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8003412 kB) free!
[1565553270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565553386] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10847564 kB) free.
[1565553390] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565553951] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 87.20, 35.82, 19.17
[1565553987] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8082996 kB) free!
[1565553991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565554107] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6175948 kB) free!
[1565554111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565554227] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11538532 kB) free.
[1565554231] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565554551] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.64, 7.08, 11.27
[1565554827] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6267620 kB) free!
[1565554831] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565554947] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11546696 kB) free.
[1565554951] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565555362] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 79.42, 32.84, 17.39
[1565555547] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6279492 kB) free!
[1565555551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565555667] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11549944 kB) free.
[1565555671] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565555976] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.50, 0.87, 4.45
[1565556113] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 397 processes
[1565556267] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6376348 kB) free!
[1565556271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565556351] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 21.85, 28.47, 20.90
[1565556387] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11537600 kB) free.
[1565556391] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565556562] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 4.19, 9.50
[1565556576] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.25, 3.91, 6.01
[1565556632] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164808 kB) free!
[1565556697] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 131.18, 31.78, 14.84
[1565556713] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 5 warn out of 393 processes
[1565556752] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185776 kB) free!
[1565556817] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.03, 28.44, 16.03
[1565556872] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197016 kB) free!
[1565556931] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 12.7% (6252256 kB) free!
[1565556987] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6478696 kB) free!
[1565556992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565557051] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.0% (45837240 kB) free.
[1565557107] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11548968 kB) free.
[1565557112] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565557707] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6543564 kB) free!
[1565557712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565557828] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11536192 kB) free.
[1565557832] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565558017] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 4.39, 9.52
[1565558427] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6502304 kB) free!
[1565558432] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565558547] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547508 kB) free.
[1565558552] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565558751] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.11, 6.94, 13.05
[1565559032] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4550900 kB) free!
[1565559147] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6777608 kB) free!
[1565559152] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197520 kB) free!
[1565559152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565559267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11548016 kB) free.
[1565559272] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201656 kB) free!
[1565559273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565559562] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 39.46, 24.12, 15.81
[1565559763] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.7% (10470976 kB) free.
[1565559867] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7274136 kB) free!
[1565559872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565559872] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 44.3% (10862624 kB) free.
[1565559987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11539100 kB) free.
[1565559992] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565560162] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 3.67, 8.53
[1565560417] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 44.24, 25.30, 16.13
[1565560551] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.57, 25.84, 17.00
[1565560587] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7446868 kB) free!
[1565560592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1565560707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11544812 kB) free.
[1565560712] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565560762] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.98, 4.74
[1565561017] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.98, 3.98, 8.74
[1565561151] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.52, 5.74, 10.17
[1565561307] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7585696 kB) free!
[1565561312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565561362] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.45, 10.79, 9.26
[1565561427] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11535244 kB) free.
[1565561432] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565561482] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.86, 7.45, 8.22
[1565561602] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.62, 5.20, 7.31
[1565562027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7787192 kB) free!
[1565562032] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565562148] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547776 kB) free.
[1565562152] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565562747] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7873848 kB) free!
[1565562753] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565562867] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547712 kB) free.
[1565562873] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565563467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7853812 kB) free!
[1565563473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565563587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11541412 kB) free.
[1565563593] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565563913] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 378 processes
[1565564151] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 107.95, 38.79, 23.00
[1565564187] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7901548 kB) free!
[1565564193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565564307] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547832 kB) free.
[1565564313] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565564513] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 374 processes
[1565564617] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 69.02, 27.55, 17.52
[1565564907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7967908 kB) free!
[1565564913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1565565027] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11538796 kB) free.
[1565565033] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565565202] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 97.89, 35.70, 19.72
[1565565217] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 4.24, 9.47
[1565565627] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7990164 kB) free!
[1565565634] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565565747] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10486300 kB) free.
[1565565754] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565565802] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 5.31, 10.62
[1565565951] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.71, 6.61, 12.62
[1565566347] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8010488 kB) free!
[1565566354] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565566468] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11548792 kB) free.
[1565566474] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565567067] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8222224 kB) free!
[1565567074] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565567187] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6545360 kB) free!
[1565567194] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565567307] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547940 kB) free.
[1565567314] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565567907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6935876 kB) free!
[1565567914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565568027] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11537588 kB) free.
[1565568034] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1565568627] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7185984 kB) free!
[1565568635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565568747] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11546248 kB) free.
[1565568755] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565568802] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 46.87, 12.42, 11.85
[1565568817] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 54.47, 25.20, 14.22
[1565568951] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 70.09, 37.55, 22.90
[1565569347] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7482896 kB) free!
[1565569355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565569402] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 4.82, 9.76
[1565569467] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547056 kB) free.
[1565569475] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565569703] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1003 processes
[1565569823] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1017 processes
[1565569943] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 373 processes
[1565570067] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7716040 kB) free!
[1565570075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565570152] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.79, 8.37, 14.74
[1565570187] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11540164 kB) free.
[1565570195] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565570617] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 4.60, 10.17
[1565570788] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8015588 kB) free!
[1565570795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565570908] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6252384 kB) free!
[1565570915] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565570999] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1565571028] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11538204 kB) free.
[1565571036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565571159] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 7 users currently logged in
[1565571628] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7371600 kB) free!
[1565571635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565571748] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11544876 kB) free.
[1565571755] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565572003] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565572123] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1565572243] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1565572348] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7948296 kB) free!
[1565572355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565572402] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.94, 27.49, 20.00
[1565572468] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547676 kB) free.
[1565572475] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565573068] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8321424 kB) free!
[1565573075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565573188] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6784280 kB) free!
[1565573195] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565573308] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11545264 kB) free.
[1565573315] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565573513] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 354 processes
[1565573602] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 4.16, 10.36
[1565573752] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 38.32, 29.97, 17.34
[1565573908] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7209096 kB) free!
[1565573916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565574028] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11534740 kB) free.
[1565574036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565574113] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 353 processes
[1565574217] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 45.36, 10.10, 6.93
[1565574352] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.77, 7.02, 10.60
[1565574629] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7712004 kB) free!
[1565574636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565574749] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11547068 kB) free.
[1565574756] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565574763] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173436 kB) free!
[1565574817] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 5.79, 8.77
[1565574890] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177128 kB) free!
[1565575010] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167140 kB) free!
[1565575349] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7972348 kB) free!
[1565575356] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565575469] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6372512 kB) free!
[1565575476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565575589] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11535288 kB) free.
[1565575596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565576189] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7500164 kB) free!
[1565576196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565576309] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6119684 kB) free!
[1565576317] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565576429] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11535048 kB) free.
[1565576603] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 125.02, 37.68, 20.39
[1565576752] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.44, 27.73, 16.31
[1565576916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565577029] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6481080 kB) free!
[1565577036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565577113] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 383 processes
[1565577149] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11543576 kB) free.
[1565577156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565577203] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.24, 5.84, 11.06
[1565577352] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.79, 6.10, 9.85
[1565577749] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6741528 kB) free!
[1565577756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565577869] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11545000 kB) free.
[1565577876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565578313] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 355 processes
[1565578417] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 45.71, 12.35, 11.26
[1565578470] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6956156 kB) free!
[1565578476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565578590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11536284 kB) free.
[1565578596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565579017] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 4.69, 9.35
[1565579190] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7735816 kB) free!
[1565579196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565579310] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6226904 kB) free!
[1565579317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1565579430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11536632 kB) free.
[1565579437] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565580030] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7128628 kB) free!
[1565580037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565580140] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1565580150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11541664 kB) free.
[1565580157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565580203] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 64.76, 27.28, 16.24
[1565580260] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565580352] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.09, 29.64, 22.08
[1565580750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7630344 kB) free!
[1565580757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565580803] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 4.31, 8.88
[1565580870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11544988 kB) free.
[1565580877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565581470] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7974592 kB) free!
[1565581477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565581552] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.91, 8.49, 14.54
[1565581590] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6231824 kB) free!
[1565581597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565581710] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11541052 kB) free.
[1565581718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565581913] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 366 processes
[1565582017] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.09, 26.19, 18.70
[1565582310] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6923792 kB) free!
[1565582317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565582431] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11536556 kB) free.
[1565582437] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565582514] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 356 processes
[1565582603] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.08, 31.09, 18.04
[1565582617] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.67, 4.06, 10.11
[1565583030] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7463556 kB) free!
[1565583037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565583150] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10998072 kB) free.
[1565583157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565583203] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.84, 4.89, 9.81
[1565583352] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 93.56, 38.73, 23.60
[1565583750] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8231276 kB) free!
[1565583757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565583870] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6747588 kB) free!
[1565583877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565583952] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.98, 7.52, 13.63
[1565583990] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11535196 kB) free.
[1565583997] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565584343] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 635 processes
[1565584463] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 377 processes
[1565584590] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7919168 kB) free!
[1565584597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565584710] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6683452 kB) free!
[1565584718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565584830] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11531684 kB) free.
[1565584838] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565585430] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7794776 kB) free!
[1565585438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565585462] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.5% (14228668 kB) free.
[1565585525] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40307996 kB) free.
[1565585550] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6383148 kB) free!
[1565585558] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565585570] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.9% (22776528 kB) free.
[1565585614] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22120016 kB) free.
[1565585671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11540672 kB) free.
[1565585678] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565585752] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 50.91, 34.13, 22.30
[1565586073] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (184860 kB) free!
[1565586131] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235848 kB) free!
[1565586193] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174172 kB) free!
[1565586251] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214396 kB) free!
[1565586260] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565586270] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7874828 kB) free!
[1565586278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565586313] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169848 kB) free!
[1565586352] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.68, 6.89, 12.95
[1565586371] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217540 kB) free!
[1565586380] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1565586390] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6612060 kB) free!
[1565586398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565586510] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11540556 kB) free.
[1565586518] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565586872] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4209724 kB) free!
[1565586952] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1279.23, 1040.51, 539.10
[1565586992] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3834208 kB) free!
[1565587110] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7781584 kB) free!
[1565587112] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 26.2% (6415084 kB) free.
[1565587119] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565587230] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6421604 kB) free!
[1565587239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565587350] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11540964 kB) free.
[1565587359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565587417] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 38.47, 31.64, 19.19
[1565587950] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8146464 kB) free!
[1565587959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565588003] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 70.29, 34.71, 20.56
[1565588070] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7033328 kB) free!
[1565588079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565588190] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11522484 kB) free.
[1565588199] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565588312] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.2% (2742824 kB) free!
[1565588336] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 106416 MB (77% inode=99%):
[1565588432] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.5% (2562816 kB) free!
[1565588552] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (2531464 kB) free!
[1565588559] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1565588617] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.36, 7.08, 11.73
[1565588624] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172912 kB) free!
[1565588679] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1565588744] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178748 kB) free!
[1565588790] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8383480 kB) free!
[1565588799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565588799] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1565588864] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165580 kB) free!
[1565588910] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7313172 kB) free!
[1565588919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565589031] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11525480 kB) free.
[1565589040] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565589152] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2370952 kB) free!
[1565589599] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565589630] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8317772 kB) free!
[1565589639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565589750] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7179852 kB) free!
[1565589759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565589870] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11527400 kB) free.
[1565589879] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565590065] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 24.8% (6120700 kB) free.
[1565590352] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (2854084 kB) free!
[1565590403] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.62, 15.17, 16.80
[1565590470] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8486448 kB) free!
[1565590479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565590590] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7374616 kB) free!
[1565590599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565590665] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195572 kB) free!
[1565590710] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11529028 kB) free.
[1565590719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565590785] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (329412 kB) free!
[1565590905] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1565590905] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (292076 kB) free!
[1565591199] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 13 users currently logged in
[1565591310] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8599780 kB) free!
[1565591319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565591430] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7549892 kB) free!
[1565591440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565591505] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.5% (19872936 kB) free.
[1565591550] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6110444 kB) free!
[1565591552] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1565591552] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452192 kB) free!
[1565591560] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565591617] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.62, 1.72, 4.55
[1565591752] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.78, 6.02, 15.53
[1565592114] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 390 processes
[1565592160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565592217] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.81, 8.73, 10.08
[1565592280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565592337] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.54, 6.00, 8.92
[1565592395] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565592400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565592457] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.51, 4.19, 7.90
[1565592515] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1565592635] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1565592646] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1565592750] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11083232 kB) free.
[1565593057] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.08, 1.11, 4.38
[1565593235] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565593305] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190764 kB) free!
[1565593315] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 4 warn out of 364 processes
[1565593425] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188940 kB) free!
[1565593545] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195076 kB) free!
[1565593652] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244904 kB) free!
[1565593772] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236228 kB) free!
[1565593892] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220276 kB) free!
[1565593951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7837916 kB) free!
[1565594071] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6492092 kB) free!
[1565594191] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11520540 kB) free.
[1565594200] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565594318] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 362 processes
[1565594438] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 367 processes
[1565594558] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1565594791] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7893580 kB) free!
[1565594801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565594911] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6615544 kB) free!
[1565594921] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565595031] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11516768 kB) free.
[1565595041] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565595203] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.97, 0.97, 3.14
[1565595352] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 52.68, 27.55, 15.96
[1565595457] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.11, 14.56, 8.79
[1565595577] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.35, 10.02, 7.82
[1565595631] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7794056 kB) free!
[1565595641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565595697] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.37, 6.78, 6.90
[1565595715] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 361 processes
[1565595751] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6444876 kB) free!
[1565595758] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1565595761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565595803] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.12, 7.35, 9.16
[1565595871] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11523016 kB) free.
[1565595881] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565595923] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 75.45, 29.78, 16.76
[1565596043] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.45, 20.06, 14.78
[1565596297] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.48, 1.38, 3.85
[1565596315] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 362 processes
[1565596471] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7917728 kB) free!
[1565596481] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565596552] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.56, 7.81, 12.58
[1565596591] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6572972 kB) free!
[1565596601] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565596712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11519808 kB) free.
[1565596721] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565596892] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.9% (42841312 kB) free.
[1565596958] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 364 processes
[1565597078] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 364 processes
[1565597198] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 369 processes
[1565597309] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.65, 6.35, 5.15
[1565597311] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8071192 kB) free!
[1565597322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565597429] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.61, 6.65, 5.40
[1565597431] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6779456 kB) free!
[1565597442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565597497] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.53, 3.73, 5.25
[1565597515] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 360 processes
[1565597549] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.47, 6.62, 5.55
[1565597551] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11523612 kB) free.
[1565597562] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565597617] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.61, 2.70, 4.69
[1565597799] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1565597799] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1565598149] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.68, 4.03, 4.67
[1565598151] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8193976 kB) free!
[1565598162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565598217] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.57, 8.99, 6.58
[1565598271] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (6999308 kB) free!
[1565598282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565598338] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.95, 6.15, 5.84
[1565598391] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11515432 kB) free.
[1565598402] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565598443] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 35.33, 16.78, 11.84
[1565598458] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.30, 4.47, 5.25
[1565598716] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 372 processes
[1565598991] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8293368 kB) free!
[1565599002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565599044] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.24, 8.18, 9.75
[1565599111] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7115672 kB) free!
[1565599122] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565599231] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11514216 kB) free.
[1565599242] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565599552] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.03, 27.69, 20.12
[1565599644] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 87.63, 25.34, 15.10
[1565599658] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.98, 12.08, 10.78
[1565599831] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8377800 kB) free!
[1565599843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565599951] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7263408 kB) free!
[1565599963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565600071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11520908 kB) free.
[1565600083] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565600152] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.56, 6.11, 11.88
[1565600244] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.44, 9.94, 11.74
[1565600258] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.00, 12.78, 12.59
[1565600672] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8511872 kB) free!
[1565600683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565600792] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7398164 kB) free!
[1565600803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565600912] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11521172 kB) free.
[1565600923] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565600973] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1565601033] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1565601093] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1565601153] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1565601512] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8627692 kB) free!
[1565601523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565601632] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7553224 kB) free!
[1565601643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565601752] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6103908 kB) free!
[1565601753] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 59 users currently logged in
[1565601763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565601952] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.51, 2.52, 3.96
[1565602353] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565602552] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 26.32, 12.13, 6.84
[1565602644] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.52, 14.60, 10.87
[1565602672] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.50, 8.90, 6.29
[1565602792] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.70, 6.73, 5.82
[1565602916] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 388 processes
[1565602953] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1565603013] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1565603073] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1565603133] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1565603244] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.64, 6.43, 8.68
[1565603258] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 79.82, 23.94, 15.21
[1565603392] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.65, 3.05, 4.23
[1565603517] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 374 processes
[1565603552] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1565603552] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11335216 kB) free.
[1565603745] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1565603858] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.70, 3.77, 8.30
[1565604163] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1565604163] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565604190] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 457 processes
[1565604310] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 456 processes
[1565604752] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8152104 kB) free!
[1565604763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565604873] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6921732 kB) free!
[1565604883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565604993] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11501620 kB) free.
[1565605003] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565605044] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.79, 3.48
[1565605193] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.13, 5.63, 6.05
[1565605313] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.67, 4.67, 5.65
[1565605433] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.59, 3.99, 5.28
[1565605593] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8256308 kB) free!
[1565605603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565605713] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7049028 kB) free!
[1565605724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565605833] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11500652 kB) free.
[1565605844] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565606033] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.95, 2.94, 4.07
[1565606258] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.67, 0.74, 3.03
[1565606433] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8412660 kB) free!
[1565606444] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565606553] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7288332 kB) free!
[1565606564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565606673] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11509896 kB) free.
[1565606684] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565606844] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.93, 20.78, 17.62
[1565606858] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.87, 10.06, 6.40
[1565606964] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 18.23, 24.18, 19.43
[1565606978] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 28.87, 17.06, 9.31
[1565607084] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.69, 16.63, 17.24
[1565607098] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.61, 11.99, 8.38
[1565607233] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.07, 14.33, 11.06
[1565607273] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8358796 kB) free!
[1565607284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565607353] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.34, 10.62, 10.10
[1565607393] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7229600 kB) free!
[1565607404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565607473] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.05, 8.14, 9.26
[1565607513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11506604 kB) free.
[1565607524] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565607699] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.76, 2.33, 4.79
[1565608113] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8470580 kB) free!
[1565608125] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565608233] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7332220 kB) free!
[1565608245] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565608354] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11503636 kB) free.
[1565608365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565608533] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 57 users currently logged in
[1565608673] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.44, 2.69, 4.45
[1565608884] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.67, 0.70, 3.07
[1565608953] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8229964 kB) free!
[1565608965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565609073] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7005100 kB) free!
[1565609085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565609193] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11494964 kB) free.
[1565609199] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565609205] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565610085] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 28.14, 11.36, 5.18
[1565610099] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.92, 15.44, 7.44
[1565610205] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.32, 7.76, 4.61
[1565610219] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.24, 10.52, 6.60
[1565610339] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.61, 7.13, 5.83
[1565610393] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8480300 kB) free!
[1565610405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565610513] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7410696 kB) free!
[1565610526] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565610633] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11500796 kB) free.
[1565610646] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565610939] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.03, 1.68, 3.44
[1565611233] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8616116 kB) free!
[1565611246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565611352] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.8% (23002996 kB) free.
[1565611353] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7548260 kB) free!
[1565611366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565611385] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 68.1% (11071980 kB) free.
[1565611405] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.19, 10.60, 7.28
[1565611408] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.9% (36912520 kB) free.
[1565611447] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.3% (22284960 kB) free.
[1565611473] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.2% (8007012 kB) free!
[1565611486] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565611525] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 20.85, 18.32, 10.68
[1565611539] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.64, 17.32, 10.75
[1565611645] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.33, 12.45, 9.46
[1565611659] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.82, 11.72, 9.50
[1565611673] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.14, 8.99, 8.79
[1565611779] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.59, 7.96, 8.38
[1565611793] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.67, 6.88, 8.04
[1565611913] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.62, 5.47, 7.39
[1565611993] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181588 kB) free!
[1565612014] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (266388 kB) free!
[1565612086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565612113] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (163060 kB) free!
[1565612134] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214608 kB) free!
[1565612206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565612234] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174424 kB) free!
[1565612255] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (247256 kB) free!
[1565612326] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565612673] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10239924 kB) free.
[1565612845] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.26, 22.72, 13.28
[1565612927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565613047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565613113] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.53, 2.77, 4.04
[1565613167] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565613445] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.21, 10.05, 10.43
[1565613635] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565613755] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1565613767] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565613873] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6394748 kB) free!
[1565613875] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1565613887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565613907] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196296 kB) free!
[1565613993] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11496844 kB) free.
[1565614007] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565614028] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (235184 kB) free!
[1565614148] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (219748 kB) free!
[1565614313] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.52, 4.87, 5.20
[1565614433] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 77.20, 20.66, 10.30
[1565614553] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 12.55, 14.66, 9.36
[1565614593] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7887068 kB) free!
[1565614607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565614713] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6563376 kB) free!
[1565614727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565614833] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11494880 kB) free.
[1565614847] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565614917] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565615245] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.79, 1.34, 3.51
[1565615379] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 43.37, 31.86, 15.56
[1565615394] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565615399] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1565615433] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7957732 kB) free!
[1565615448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565615553] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6667592 kB) free!
[1565615568] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565615673] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11503692 kB) free.
[1565615688] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565615753] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 41.86, 18.27, 11.27
[1565615845] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.50, 6.89, 5.37
[1565615965] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.66, 4.78, 4.78
[1565616273] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8103868 kB) free!
[1565616275] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1565616288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565616353] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.66, 9.68, 9.98
[1565616393] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6850836 kB) free!
[1565616408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565616513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11502272 kB) free.
[1565616528] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565616565] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.15, 12.07, 9.72
[1565616579] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.17, 11.73, 12.79
[1565616685] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.16, 8.23, 8.59
[1565616805] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.06, 5.81, 7.66
[1565616980] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565617101] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1565617114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7502828 kB) free!
[1565617128] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565617234] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11498780 kB) free.
[1565617248] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1565617248] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565617834] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8037480 kB) free!
[1565617849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565617954] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6393740 kB) free!
[1565617969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565618074] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11504104 kB) free.
[1565618089] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565618674] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7192852 kB) free!
[1565618689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565618794] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11489416 kB) free.
[1565618809] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565619394] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7824488 kB) free!
[1565619409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565619514] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11496180 kB) free.
[1565619529] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565620114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8173028 kB) free!
[1565620129] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565620234] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6471940 kB) free!
[1565620250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565620354] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11492960 kB) free.
[1565620370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565620955] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7464384 kB) free!
[1565620970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565621075] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11494312 kB) free.
[1565621090] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565621379] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 86.49, 24.34, 15.54
[1565621675] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7923880 kB) free!
[1565621690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565621795] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8379660 kB) free!
[1565621810] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565621915] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11492608 kB) free.
[1565621979] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 3.70, 8.36
[1565622263] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 908 processes
[1565622383] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 368 processes
[1565622410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565622515] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6393864 kB) free!
[1565622530] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565622635] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11483676 kB) free.
[1565622651] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565623235] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7019112 kB) free!
[1565623250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565623355] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11494772 kB) free.
[1565623370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1565623553] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 13.24, 23.58, 20.83
[1565623955] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7290408 kB) free!
[1565623970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565624075] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11492908 kB) free.
[1565624090] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565624153] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.99, 10.60, 16.26
[1565624675] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7493548 kB) free!
[1565624690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565624796] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11483580 kB) free.
[1565624810] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565625395] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7869968 kB) free!
[1565625410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565625515] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6115736 kB) free!
[1565625531] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565625579] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.93, 20.18, 15.50
[1565625635] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11482876 kB) free.
[1565625805] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 72.57, 25.51, 17.00
[1565626130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565626180] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.56, 12.19, 13.25
[1565626235] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6777640 kB) free!
[1565626250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565626355] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565626355] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11489700 kB) free.
[1565626370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565626405] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 3.87, 9.14
[1565626955] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7052196 kB) free!
[1565626970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565627075] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11493788 kB) free.
[1565627090] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565627380] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.67, 0.67, 3.95
[1565627605] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.78, 1.96, 4.69
[1565627675] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7472856 kB) free!
[1565627690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565627795] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11484920 kB) free.
[1565627811] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565628395] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7758704 kB) free!
[1565628410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565628515] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11493160 kB) free.
[1565628530] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565629115] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8112776 kB) free!
[1565629130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565629236] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6444892 kB) free!
[1565629250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565629356] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11491236 kB) free.
[1565629370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565629956] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7122796 kB) free!
[1565629970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565630005] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 56.11, 17.75, 6.83
[1565630076] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11492796 kB) free.
[1565630090] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565630125] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.67, 11.94, 6.02
[1565630246] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.48, 8.18, 5.36
[1565630380] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.10, 6.00, 5.06
[1565630500] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.43, 4.16, 4.50
[1565630676] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7599404 kB) free!
[1565630690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565630796] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11481164 kB) free.
[1565630811] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565630846] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 1.48, 3.04
[1565631100] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.51, 4.32, 5.27
[1565631220] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 2.98, 4.66
[1565631396] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7902644 kB) free!
[1565631410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565631516] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6151088 kB) free!
[1565631530] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565631636] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11482444 kB) free.
[1565631820] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.14, 11.60, 9.05
[1565631940] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.54, 7.86, 7.99
[1565632060] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 5.34, 7.05
[1565632130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565632236] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6449292 kB) free!
[1565632250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565632356] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11489596 kB) free.
[1565632370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565632553] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.88, 2.83, 4.53
[1565632646] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 40.24, 11.34, 5.75
[1565632660] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.95, 3.82
[1565632766] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.34, 7.88, 5.16
[1565632886] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.07, 5.41, 4.59
[1565632956] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6895900 kB) free!
[1565632970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565633077] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11489892 kB) free.
[1565633090] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565633486] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.67, 4.48, 5.01
[1565633606] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 3.07, 4.43
[1565633676] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7430572 kB) free!
[1565633691] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565633796] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11479768 kB) free.
[1565633811] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565634206] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.66, 13.61, 7.92
[1565634326] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 33.67, 15.79, 9.13
[1565634353] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.63, 17.19, 11.19
[1565634396] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7721028 kB) free!
[1565634411] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565634446] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.58, 15.11, 9.92
[1565634460] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.70, 16.85, 12.09
[1565634473] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.85, 12.40, 10.16
[1565634516] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11489276 kB) free.
[1565634531] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565634580] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.58, 11.42, 10.68
[1565634593] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.80, 9.18, 9.25