[1560624560] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 20.4% (10076912 kB) free.
[1560625166] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560625286] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1560625436] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9977656 kB) free.
[1560626036] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3761468 kB) free!
[1560626156] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (2993824 kB) free!
[1560626277] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2351756 kB) free!
[1560626612] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560627560] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9479648 kB) free!
[1560627680] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9863796 kB) free!
[1560627800] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9858016 kB) free!
[1560627967] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1352628 kB) free!
[1560628286] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560628407] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1560628726] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 396 processes
[1560628846] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 396 processes
[1560628966] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 394 processes
[1560628995] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171536 kB) free!
[1560629000] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 20.3% (9982584 kB) free.
[1560629115] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (179312 kB) free!
[1560629235] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168764 kB) free!
[1560629277] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6193924 kB) free!
[1560629583] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 219142 MB (20% inode=99%):
[1560629703] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 212029 MB (19% inode=99%):
[1560629823] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 209936 MB (19% inode=99%):
[1560629877] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (1885284 kB) free!
[1560630166] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 390 processes
[1560630767] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 396 processes
[1560630887] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 395 processes
[1560631007] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 394 processes
[1560631014] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 154.39, 108.00, 50.20
[1560631118] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.0% (2924348 kB) free!
[1560631134] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 154.52, 123.41, 62.85
[1560631238] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 78.4% (12734832 kB) free.
[1560631254] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 35.91, 93.99, 59.87
[1560631406] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560631526] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1560631607] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1560631607] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 390 processes
[1560632167] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (5236344 kB) free!
[1560632454] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.75, 3.22, 17.53
[1560632612] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 60.51, 16.83, 6.05
[1560632732] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 100.43, 46.56, 18.25
[1560632852] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 105.34, 66.29, 28.91
[1560633295] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2772676 kB) free!
[1560634052] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 2.36, 11.72
[1560634077] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.6% (5708136 kB) free!
[1560634254] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.80, 1.67, 3.88
[1560634400] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9876740 kB) free!
[1560634450] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 508.05, 364.20, 169.41
[1560634485] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 582 processes
[1560634520] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9872132 kB) free!
[1560634556] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2786772 kB) free!
[1560634570] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 104.15, 273.45, 161.08
[1560634605] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 581 processes
[1560634640] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1560634640] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9872820 kB) free!
[1560634678] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1368904 kB) free!
[1560634690] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.17, 183.02, 141.57
[1560634725] SERVICE ALERT: sservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 580 processes
[1560635252] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.83, 0.57, 3.56
[1560635325] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 325 processes
[1560636295] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2467692 kB) free!
[1560636967] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1335012 kB) free!
[1560637090] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 0.32, 10.91
[1560638278] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (5997108 kB) free!
[1560638290] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 0.29, 3.21
[1560638756] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2433336 kB) free!
[1560638878] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223708 kB) free!
[1560639343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560639479] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6055856 kB) free!
[1560639943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560640064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1560640079] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219308 kB) free!
[1560640184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1560640679] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6077972 kB) free!
[1560640784] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1560641240] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1560641279] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222216 kB) free!
[1560641384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560641504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560641624] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1560641624] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560641767] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (5218284 kB) free!
[1560641840] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 21.0% (10334716 kB) free.
[1560643007] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560643640] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9266768 kB) free!
[1560643679] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7456468 kB) free!
[1560643760] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 21.2% (10424736 kB) free.
[1560644279] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (221940 kB) free!
[1560645435] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 69.8% (11319392 kB) free.
[1560645967] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (582560 kB) free!
[1560646007] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 396 processes
[1560646035] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.6% (2696064 kB) free!
[1560646127] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 394 processes
[1560646155] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164860 kB) free!
[1560646247] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 394 processes
[1560646275] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 73.8% (11968952 kB) free.
[1560647280] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.0% (7906356 kB) free!
[1560647853] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 504.42, 319.38, 138.96
[1560647880] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2549448 kB) free!
[1560647973] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 508.01, 382.00, 183.79
[1560648093] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 81.18, 267.32, 166.39
[1560648295] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2787048 kB) free!
[1560648490] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 129.94, 128.41, 65.14
[1560648610] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.91, 86.06, 57.30
[1560648730] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.85, 57.73, 50.40
[1560648815] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1560648935] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1560649055] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1560649411] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 556 processes
[1560649471] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 543 processes
[1560649568] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (5215768 kB) free!
[1560649760] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9823272 kB) free!
[1560649880] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9829012 kB) free!
[1560649930] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.39, 1.34, 14.06
[1560650000] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9818264 kB) free!
[1560650280] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.1% (6455452 kB) free!
[1560650493] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 0.50, 12.98
[1560650695] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2474980 kB) free!
[1560650880] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (266140 kB) free!
[1560651130] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.28, 4.06
[1560651693] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.51, 0.45, 3.91
[1560651801] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4223012 kB) free!
[1560653024] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560653168] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1560653168] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (720920 kB) free!
[1560654824] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1560655293] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 5.50, 56.98, 42.97
[1560655413] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.98, 38.23, 37.79
[1560655424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560655495] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2869644 kB) free!
[1560655533] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1560655533] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.34, 25.65, 33.24
[1560655544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560655664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560655800] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1560655920] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1560656040] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1560656133] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.66, 3.86, 17.61
[1560656695] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1560656695] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2137232 kB) free!
[1560656768] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 10.2% (5024076 kB) free!
[1560657933] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.65, 3.12
[1560658080] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.2% (5512636 kB) free!
[1560658680] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (854928 kB) free!
[1560659516] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 25.2% (4090500 kB) free.
[1560659768] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216188 kB) free!
[1560659880] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.2% (8001752 kB) free!
[1560660116] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (3054648 kB) free!
[1560660236] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.4% (2832988 kB) free!
[1560660356] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.1% (2622636 kB) free!
[1560660480] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3729760 kB) free!
[1560661556] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1468616 kB) free!
[1560661602] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1560661722] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1560662168] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (6017100 kB) free!
[1560662831] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.5% (23322736 kB) free.
[1560662871] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.0% (22304740 kB) free.
[1560662915] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1560662915] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.2% (14815436 kB) free.
[1560662945] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.2% (22356120 kB) free.
[1560662984] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.0% (14961384 kB) free.
[1560663015] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 54.3% (26749208 kB) free.
[1560663060] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5578188 kB) free!
[1560663588] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168860 kB) free!
[1560663664] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1278700 kB) free!
[1560663709] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166292 kB) free!
[1560663816] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.09, 6.64, 5.14
[1560663829] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155156 kB) free!
[1560663936] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.57, 6.89, 5.41
[1560663968] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212672 kB) free!
[1560664056] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.31, 7.31, 5.73
[1560664722] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 368 processes
[1560664842] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1560664854] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 514.76, 349.12, 158.00
[1560664974] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 516.61, 404.64, 201.52
[1560665094] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 402.92, 416.57, 231.51
[1560665443] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1560665464] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5137472 kB) free!
[1560665563] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1560665768] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 12.3% (6056548 kB) free!
[1560665856] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.75, 4.42, 4.96
[1560666037] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (723188 kB) free!
[1560666064] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230076 kB) free!
[1560666157] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (513636 kB) free!
[1560666163] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1560666277] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (345112 kB) free!
[1560666283] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1560666403] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1560666456] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.81, 5.40, 5.20
[1560666576] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.45, 5.06, 5.10
[1560666589] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.0% (14961980 kB) free.
[1560666637] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.1% (6845540 kB) free.
[1560666657] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7339416 kB) free!
[1560666696] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.98, 4.71, 4.97
[1560666968] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.3% (10014512 kB) free.
[1560667198] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166412 kB) free!
[1560667318] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165792 kB) free!
[1560667438] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1560667438] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164268 kB) free!
[1560667494] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.50, 1.80, 19.09
[1560667568] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (867612 kB) free!
[1560667688] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.6% (10130196 kB) free.
[1560667865] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (823164 kB) free!
[1560667869] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1560667947] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.0% (14956064 kB) free.
[1560667987] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.9% (8348820 kB) free!
[1560667989] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1560668109] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1560668133] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560668204] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1560668288] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (1079576 kB) free!
[1560668408] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.3% (9981556 kB) free.
[1560668418] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8184676 kB) free!
[1560668538] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (8222684 kB) free!
[1560668555] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (170068 kB) free!
[1560668598] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4724864 kB) free!
[1560668658] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (8453448 kB) free!
[1560668675] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (191712 kB) free!
[1560668795] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165928 kB) free!
[1560668831] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1560668864] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1560669008] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (695396 kB) free!
[1560669128] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.7% (10217236 kB) free.
[1560669294] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.78, 1.79, 4.30
[1560669465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560669586] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560669706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560669728] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1043244 kB) free!
[1560669848] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.8% (10247068 kB) free.
[1560669924] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.0% (14949764 kB) free.
[1560670389] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.9% (8306960 kB) free!
[1560670715] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (3199760 kB) free!
[1560670835] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (3101164 kB) free!
[1560670930] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 366.25, 363.22, 184.18
[1560670955] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.5% (3010652 kB) free!
[1560670996] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (1955468 kB) free!
[1560671050] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 49.60, 243.11, 161.88
[1560671090] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9874128 kB) free!
[1560671129] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (158908 kB) free!
[1560671170] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7.16, 162.84, 142.31
[1560671204] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 368 processes
[1560671210] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8076320 kB) free!
[1560671249] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (170528 kB) free!
[1560671324] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1560671330] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 33.7% (16615316 kB) free.
[1560671369] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181192 kB) free!
[1560671930] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8304496 kB) free!
[1560672051] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 33.5% (16484788 kB) free.
[1560672650] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8204844 kB) free!
[1560672770] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 33.2% (16338680 kB) free.
[1560673370] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7836428 kB) free!
[1560673490] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 32.9% (16228580 kB) free.
[1560673570] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 5.42, 18.16
[1560673955] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1310864 kB) free!
[1560673996] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4990384 kB) free!
[1560674090] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7732872 kB) free!
[1560674210] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 32.6% (16072836 kB) free.
[1560674324] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 368 processes
[1560674444] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1560674596] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219860 kB) free!
[1560674733] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 128.46, 69.85, 28.74
[1560674810] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7775952 kB) free!
[1560674853] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 131.37, 90.31, 41.22
[1560674930] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 32.4% (15977044 kB) free.
[1560674973] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 52.87, 82.04, 44.69
[1560675370] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.39, 2.97
[1560675530] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7649904 kB) free!
[1560675650] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 32.2% (15859808 kB) free.
[1560675858] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (1436432 kB) free!
[1560675885] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5164 MB (10% inode=97%):
[1560675894] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 114.00, 98.52, 54.63
[1560676014] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.94, 67.70, 48.73
[1560676134] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.66, 45.83, 43.02
[1560676173] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.86, 2.22, 12.81
[1560676250] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7500052 kB) free!
[1560676371] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 32.0% (15765612 kB) free.
[1560676458] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (5077396 kB) free!
[1560676485] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6973 MB (13% inode=98%):
[1560676970] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7537308 kB) free!
[1560677058] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1560677058] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2932332 kB) free!
[1560677090] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 31.8% (15676940 kB) free.
[1560677334] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1560677334] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.87, 2.47, 13.02
[1560677373] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.48, 0.57, 3.96
[1560677690] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7301272 kB) free!
[1560677810] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 11.6% (5698816 kB) free!
[1560677930] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 30.0% (14785584 kB) free.
[1560678045] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 369 processes
[1560678165] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1560678196] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5300796 kB) free!
[1560678530] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5610656 kB) free!
[1560678534] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.60, 1.57, 4.70
[1560678650] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.8% (14665116 kB) free.
[1560678796] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3379300 kB) free!
[1560679250] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5360880 kB) free!
[1560679365] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 371 processes
[1560679370] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560679370] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.6% (14564836 kB) free.
[1560679485] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 370 processes
[1560679605] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 370 processes
[1560679773] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 78.58, 21.51, 7.89
[1560679893] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.95, 14.54, 6.99
[1560679970] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5729088 kB) free!
[1560680013] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1560680013] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.25, 10.08, 6.28
[1560680090] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 30.1% (14829616 kB) free.
[1560680205] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1560680613] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.74, 1.90, 3.58
[1560680690] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5598136 kB) free!
[1560680806] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 366 processes
[1560680811] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.8% (14673448 kB) free.
[1560680926] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1560681410] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5306192 kB) free!
[1560681526] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 370 processes
[1560681530] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.6% (14562976 kB) free.
[1560681569] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560681646] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 5 crit, 0 warn out of 372 processes
[1560681766] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 369 processes
[1560682130] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5310776 kB) free!
[1560682250] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.2% (14395996 kB) free.
[1560682850] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5217768 kB) free!
[1560682955] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1560682970] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 30.0% (14800324 kB) free.
[1560682996] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5949312 kB) free!
[1560683566] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1560683566] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1560683570] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5161000 kB) free!
[1560683596] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1573452 kB) free!
[1560683690] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.9% (14717324 kB) free.
[1560684290] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5388596 kB) free!
[1560684410] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.4% (14507576 kB) free.
[1560685010] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5030136 kB) free!
[1560685047] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 634 processes
[1560685091] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.1% (6952072 kB) free!
[1560685130] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.1% (14323412 kB) free.
[1560685134] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.37, 16.85, 8.48
[1560685139] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.1% (36514788 kB) free.
[1560685167] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 634 processes
[1560685182] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (14920372 kB) free.
[1560685215] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.0% (14959680 kB) free.
[1560685254] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.66, 11.93, 7.69
[1560685287] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 638 processes
[1560685366] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 373 processes
[1560685374] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.81, 8.49, 6.94
[1560685486] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1560685509] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1560685696] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11351552 kB) free.
[1560685731] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5315932 kB) free!
[1560685815] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (153824 kB) free!
[1560685851] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.5% (14534348 kB) free.
[1560685887] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 367 processes
[1560685935] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (188744 kB) free!
[1560685974] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.70, 2.50, 4.38
[1560686040] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1560686055] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (191636 kB) free!
[1560686110] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1560686296] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (537424 kB) free!
[1560686405] SERVICE ALERT: db2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1560686419] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (223072 kB) free!
[1560686451] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5241836 kB) free!
[1560686530] SERVICE ALERT: db2;Sensors Status;OK;HARD;1;SENSORS OK
[1560686539] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228912 kB) free!
[1560686571] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.3% (14448664 kB) free.
[1560687171] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4977520 kB) free!
[1560687291] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 29.1% (14337176 kB) free.
[1560687891] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4898736 kB) free!
[1560688011] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.8% (14178404 kB) free.
[1560688339] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5822664 kB) free!
[1560688611] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4083340 kB) free!
[1560688731] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.4% (14001276 kB) free.
[1560688939] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (1127204 kB) free!
[1560689331] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4390244 kB) free!
[1560689451] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.5% (14033428 kB) free.
[1560689507] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 339 processes
[1560689948] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1560690051] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4760376 kB) free!
[1560690107] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560690172] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.3% (13952572 kB) free.
[1560690771] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4786596 kB) free!
[1560690891] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 28.2% (13898208 kB) free.
[1560691491] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4387932 kB) free!
[1560691611] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.6% (13603568 kB) free.
[1560691907] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1560692013] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560692211] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3987340 kB) free!
[1560692331] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.1% (13371868 kB) free.
[1560692931] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4272620 kB) free!
[1560693051] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.9% (13244864 kB) free.
[1560693107] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1560693107] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1560693139] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11677708 kB) free.
[1560693651] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4297364 kB) free!
[1560693707] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560693739] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4481668 kB) free!
[1560693771] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.7% (13665700 kB) free.
[1560693827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560693860] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (3610000 kB) free!
[1560693886] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560693947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560693980] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2829200 kB) free!
[1560694371] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4194396 kB) free!
[1560694491] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 27.6% (13607520 kB) free.
[1560695091] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4039500 kB) free!
[1560695212] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.8% (12245740 kB) free.
[1560695811] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5529924 kB) free!
[1560695931] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7204468 kB) free!
[1560696051] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (5281820 kB) free!
[1560696509] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 374 processes
[1560696630] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1560696980] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.6% (5716552 kB) free!
[1560697375] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560697580] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (1504100 kB) free!
[1560697851] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (252512 kB) free!
[1560699174] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.27, 12.34, 6.20
[1560699294] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.36, 8.82, 5.66
[1560699414] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.09, 6.51, 5.19
[1560699651] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (5076820 kB) free!
[1560699980] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7249516 kB) free!
[1560700014] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.65, 2.19, 3.42
[1560700580] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (2428016 kB) free!
[1560702051] SERVICE FLAPPING ALERT: sigqscluster2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1560702051] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.5% (10092848 kB) free.
[1560702651] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (1416360 kB) free!
[1560702771] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.4% (10036612 kB) free.
[1560703371] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (914860 kB) free!
[1560703491] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9418776 kB) free!
[1560703611] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (7512072 kB) free!
[1560704811] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.5% (10109840 kB) free.
[1560705347] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560705411] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (868532 kB) free!
[1560705531] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (9133868 kB) free!
[1560705651] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (7256312 kB) free!
[1560707451] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 20.6% (10162560 kB) free.
[1560707747] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560707780] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6604868 kB) free!
[1560708051] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (1213108 kB) free!
[1560708171] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9331092 kB) free!
[1560708291] SERVICE ALERT: sigqscluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (7548536 kB) free!
[1560708347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560708380] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224720 kB) free!
[1560708467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1560708587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1560708981] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5027672 kB) free!
[1560709581] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (1800624 kB) free!
[1560709787] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 339 processes