[1582309818] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.7% (29403956 kB) free.
[1582310570] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175252 kB) free!
[1582310690] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174096 kB) free!
[1582310691] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178696 kB) free!
[1582310810] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169544 kB) free!
[1582310812] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178636 kB) free!
[1582310932] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (207272 kB) free!
[1582311706] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (2207888 kB) free!
[1582311768] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193332 kB) free!
[1582311828] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (194240 kB) free!
[1582311891] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (188320 kB) free!
[1582312191] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 67.2% (16482836 kB) free.
[1582312219] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [22/Feb/2020:00:36:59 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1582312339] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582312502] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25725276 kB) free!
[1582312622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22528124 kB) free!
[1582312742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.2% (49004252 kB) free.
[1582314089] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8689404 kB) free!
[1582314209] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8323056 kB) free!
[1582314329] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.1% (7950420 kB) free!
[1582314929] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4735256 kB) free!
[1582315142] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24794840 kB) free!
[1582315262] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.4% (50679052 kB) free.
[1582315429] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2824632 kB) free!
[1582316029] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185228 kB) free!
[1582317662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16573892 kB) free!
[1582317783] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.8% (49838696 kB) free.
[1582318383] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22409244 kB) free!
[1582318503] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18016864 kB) free!
[1582318623] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.1% (50221168 kB) free.
[1582318930] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235124 kB) free!
[1582319050] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236276 kB) free!
[1582319170] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238932 kB) free!
[1582319223] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24040800 kB) free!
[1582319343] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18678636 kB) free!
[1582319463] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582319463] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.3% (49233364 kB) free.
[1582319541] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(146) < 92.42.104.166 [22/Feb/2020:02:35:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582319661] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582319770] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.8% (40808120 kB) free.
[1582320063] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24972620 kB) free!
[1582320183] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18909096 kB) free!
[1582320303] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.7% (48426924 kB) free.
[1582320410] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.3% (10359288 kB) free.
[1582320769] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(249) < 103.16.101.72 [22/Feb/2020:03:00:31 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582320837] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(330) < 2020/02/22 03:01:32 [error] 12255#0: *57677740 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582320861] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1879) < 203.212.70.170 [22/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582320890] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582320903] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24549288 kB) free!
[1582320957] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582320981] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582321023] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.2% (50316868 kB) free.
[1582328212] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201964 kB) free!
[1582328332] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173540 kB) free!
[1582328452] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177416 kB) free!
[1582328932] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.2% (20874872 kB) free.
[1582329229] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 65.2% (16096748 kB) free.
[1582332202] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 87.2% (114931004 kB) free.
[1582332253] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39485084 kB) free.
[1582332291] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.9% (21786248 kB) free.
[1582333170] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 408 processes
[1582333290] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1582334030] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (4022184 kB) free!
[1582334150] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (4013304 kB) free!
[1582334270] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (4008184 kB) free!
[1582337270] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179320 kB) free!
[1582338981] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 172.105.59.173 [22/Feb/2020:08:02:37 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582339101] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582340157] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2020/02/22 08:17:19 [error] 12255#0: *57755158 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582340277] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582344202] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21780116 kB) free.
[1582345131] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 334 processes
[1582345251] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1582345481] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (208608 kB) free!
[1582345601] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (217572 kB) free!
[1582345721] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (212652 kB) free!
[1582346610] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.3% (3040268 kB) free!
[1582346730] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.4% (2804172 kB) free!
[1582346823] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23959184 kB) free!
[1582346850] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (617304 kB) free!
[1582346943] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19956040 kB) free!
[1582347063] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43658928 kB) free.
[1582347223] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.4% (3780672 kB) free!
[1582347343] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.6% (3586176 kB) free!
[1582347464] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.7% (3357592 kB) free!
[1582347663] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20513276 kB) free!
[1582347783] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39339576 kB) free.
[1582348064] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2162612 kB) free!
[1582348383] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10750856 kB) free!
[1582348503] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6601600 kB) free!
[1582348624] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37075572 kB) free.
[1582349223] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13918592 kB) free!
[1582349343] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39587960 kB) free.
[1582349943] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17520968 kB) free!
[1582349999] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.8% (21511148 kB) free.
[1582350063] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13265896 kB) free!
[1582350063] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.9% (20445644 kB) free.
[1582350117] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.9% (14119404 kB) free.
[1582350153] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (3989940 kB) free!
[1582350183] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9240136 kB) free!
[1582350273] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.8% (3558464 kB) free!
[1582350393] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3151136 kB) free!
[1582350540] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 52.0% (68604572 kB) free.
[1582350946] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582350993] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1273508 kB) free!
[1582351012] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582351066] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1582351132] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582351132] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195516 kB) free!
[1582351186] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1582351200] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165068 kB) free!
[1582351252] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1582351252] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196860 kB) free!
[1582351321] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187072 kB) free!
[1582351372] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190820 kB) free!
[1582351441] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1582351441] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170164 kB) free!
[1582351456] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8465760 kB) free!
[1582351576] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8459160 kB) free!
[1582351696] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (8459268 kB) free!
[1582351818] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9196240 kB) free!
[1582351852] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582351938] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8599140 kB) free!
[1582352058] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.2% (7994764 kB) free!
[1582352121] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4883508 kB) free!
[1582352242] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4526140 kB) free!
[1582352362] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.9% (4169972 kB) free!
[1582352386] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1582353067] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4681256 kB) free!
[1582353187] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4115364 kB) free!
[1582353258] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.6% (30360280 kB) free.
[1582353307] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1582353307] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (3933948 kB) free!
[1582353394] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.2% (16208936 kB) free.
[1582353477] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(102) < 2020/02/22 12:01:22 [error] 12252#0: *57843702 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1582353501] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 172.105.59.173 [22/Feb/2020:12:02:49 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582353562] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (562140 kB) free!
[1582353574] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21614492 kB) free.
[1582353597] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582353621] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582353626] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.3% (21398040 kB) free.
[1582353866] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582353907] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197484 kB) free!
[1582353986] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582354106] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582354162] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.5% (22843468 kB) free.
[1582354252] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582354372] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582354492] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582354696] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4240500 kB) free!
[1582354740] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16680504 kB) free!
[1582354831] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (1167084 kB) free!
[1582354861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.3% (47813828 kB) free.
[1582354951] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182980 kB) free!
[1582355071] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187144 kB) free!
[1582355170] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8952512 kB) free!
[1582355290] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8663120 kB) free!
[1582355410] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.9% (8347496 kB) free!
[1582355460] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21923768 kB) free!
[1582355580] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17657216 kB) free!
[1582355698] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4473172 kB) free!
[1582355700] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (47952960 kB) free.
[1582355819] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (4354240 kB) free!
[1582355939] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (4191940 kB) free!
[1582356300] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23357860 kB) free!
[1582356420] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18793008 kB) free!
[1582356540] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14538572 kB) free!
[1582356892] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1582356892] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582357029] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21782436 kB) free.
[1582357073] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39473788 kB) free.
[1582357127] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1582357127] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.2% (21371232 kB) free.
[1582357139] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1534104 kB) free!
[1582357199] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.4% (21432976 kB) free.
[1582357210] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4585396 kB) free!
[1582357633] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.1% (2498348 kB) free!
[1582357706] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582357740] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1582357740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35621984 kB) free.
[1582357753] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.7% (3143876 kB) free!
[1582357873] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.0% (2957248 kB) free!
[1582358281] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (241428 kB) free!
[1582358306] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582358340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11483488 kB) free!
[1582358401] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233380 kB) free!
[1582358426] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582358460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30396016 kB) free.
[1582358473] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168020 kB) free!
[1582358522] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1582358522] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238088 kB) free!
[1582358546] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582358932] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1613372 kB) free!
[1582358985] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < operation basic.publish caused a connection exception channel_error: "expected 'channel.open'"
[1582359022] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(24) < 203.212.70.170 [22/Feb/2020:13:40:17 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1582359053] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185768 kB) free!
[1582359060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (14972460 kB) free!
[1582359106] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1582359142] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582359173] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190824 kB) free!
[1582359180] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10328428 kB) free!
[1582359300] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6043420 kB) free!
[1582359742] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(517) < 103.255.146.164 [22/Feb/2020:13:45:07 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1582359862] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582360501] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24963012 kB) free!
[1582361101] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401580 kB) free!
[1582361701] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14815340 kB) free!
[1582361794] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4380872 kB) free!
[1582361914] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4119220 kB) free!
[1582362034] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.3% (3884712 kB) free!
[1582362301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30208292 kB) free.
[1582362582] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582362901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5323644 kB) free!
[1582363021] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (923376 kB) free!
[1582363141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34237488 kB) free.
[1582363234] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1846712 kB) free!
[1582363741] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7582260 kB) free!
[1582363821] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9599000 kB) free!
[1582363835] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1582363835] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16197268 kB) free.
[1582363861] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3120224 kB) free!
[1582363941] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9460936 kB) free!
[1582363981] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376988 kB) free!
[1582364061] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (9336532 kB) free!
[1582364062] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(150) < 92.42.104.166 [22/Feb/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582364182] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582364382] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.6% (4550728 kB) free!
[1582364502] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.9% (4390596 kB) free!
[1582364546] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1582364622] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.3% (4245424 kB) free!
[1582364666] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1582364692] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 363 processes
[1582364734] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.2% (19062312 kB) free.
[1582364786] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1582364812] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1582364866] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.3% (50512384 kB) free.
[1582364921] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 50.8% (25049068 kB) free.
[1582364932] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1582364952] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 39.2% (19298864 kB) free.
[1582364978] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.0% (18143920 kB) free.
[1582365583] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186800 kB) free!
[1582365703] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175028 kB) free!
[1582365823] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184940 kB) free!
[1582365940] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164668 kB) free!
[1582366060] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170956 kB) free!
[1582366069] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20944320 kB) free!
[1582366180] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196636 kB) free!
[1582366189] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44135536 kB) free.
[1582366422] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1930108 kB) free!
[1582366789] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12487568 kB) free!
[1582366909] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8281952 kB) free!
[1582367029] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4211240 kB) free!
[1582367630] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18877172 kB) free!
[1582367997] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2020/02/22 16:00:51 [error] 12255#0: *57954178 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582368117] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582368230] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44939748 kB) free.
[1582368261] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4774512 kB) free!
[1582368830] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11404604 kB) free!
[1582368950] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7393420 kB) free!
[1582369070] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3908712 kB) free!
[1582369129] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1582369423] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.6% (5040740 kB) free.
[1582369670] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23390396 kB) free!
[1582370023] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197296 kB) free!
[1582370143] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168980 kB) free!
[1582370263] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163312 kB) free!
[1582370270] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39377240 kB) free.
[1582370400] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4559692 kB) free!
[1582370520] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4432832 kB) free!
[1582370640] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4340512 kB) free!
[1582370863] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 25.1% (6155472 kB) free.
[1582370870] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11273900 kB) free!
[1582370932] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1582370990] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7067352 kB) free!
[1582371110] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45160048 kB) free.
[1582371386] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582371463] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3994356 kB) free!
[1582371532] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1582371583] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.2% (3964260 kB) free!
[1582371703] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 47.1% (11534792 kB) free.
[1582371710] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15769136 kB) free!
[1582371830] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10834980 kB) free!
[1582371950] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6362728 kB) free!
[1582372129] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8767880 kB) free!
[1582372132] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582372173] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (2724432 kB) free!
[1582372249] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8768228 kB) free!
[1582372252] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582372294] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3057396 kB) free!
[1582372304] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178092 kB) free!
[1582372369] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (8768836 kB) free!
[1582372373] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582372414] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3343760 kB) free.
[1582372424] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3673320 kB) free!
[1582372440] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2135884 kB) free!
[1582372447] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 76.9% (18853608 kB) free.
[1582372476] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.4% (14658492 kB) free.
[1582372509] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 37.1% (18271164 kB) free.
[1582372536] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42555944 kB) free.
[1582372544] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1537600 kB) free!
[1582372565] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 61.6% (15088364 kB) free.
[1582372586] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582372594] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.6% (17558800 kB) free.
[1582372706] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582372826] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582373802] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (841540 kB) free!
[1582373922] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170668 kB) free!
[1582374042] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179068 kB) free!
[1582374773] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1582374773] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1582374941] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11503596 kB) free!
[1582375062] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7351784 kB) free!
[1582375182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46817004 kB) free.
[1582375235] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582375369] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4394760 kB) free!
[1582375414] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3046864 kB) free!
[1582375534] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3271064 kB) free.
[1582375782] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17575244 kB) free!
[1582375902] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13081968 kB) free!
[1582376022] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8813316 kB) free!
[1582376141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 53.2% (70172592 kB) free.
[1582376182] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.5% (39194136 kB) free.
[1582376214] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.5% (21701308 kB) free.
[1582376680] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166452 kB) free!
[1582376800] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172120 kB) free!
[1582376920] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177752 kB) free!
[1582377951] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10326592 kB) free!
[1582378071] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26806048 kB) free.
[1582378671] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (394284 kB) free!
[1582378791] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386344 kB) free!
[1582378911] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377032 kB) free!
[1582379512] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15345016 kB) free!
[1582380112] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36347700 kB) free.
[1582380712] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6773448 kB) free!
[1582380832] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2662448 kB) free!
[1582380952] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (421016 kB) free!
[1582381552] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16876720 kB) free!
[1582382152] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36267860 kB) free.
[1582382752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11453340 kB) free!
[1582382789] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9336628 kB) free!
[1582382872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7412432 kB) free!
[1582382909] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9336680 kB) free!
[1582382992] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13522976 kB) free!
[1582383026] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1582383029] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (9339436 kB) free!
[1582383520] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1582383522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1582383642] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 339 processes
[1582383762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1582384192] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32865896 kB) free.
[1582384362] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1582384373] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582384622] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582384792] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9711264 kB) free!
[1582384912] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5878656 kB) free!
[1582385032] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40089912 kB) free.
[1582385633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12895880 kB) free!
[1582385753] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8966828 kB) free!
[1582385873] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (5238792 kB) free!
[1582386473] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17988128 kB) free!
[1582387073] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32763408 kB) free.
[1582387583] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 78.2% (19285240 kB) free.
[1582387673] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7011000 kB) free!
[1582387793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (2981724 kB) free!
[1582387913] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40164428 kB) free.
[1582388513] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13161724 kB) free!
[1582388633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9250520 kB) free!
[1582388753] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5439536 kB) free!
[1582389353] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16534596 kB) free!
[1582389953] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28468108 kB) free.
[1582390553] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5887248 kB) free!
[1582390673] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39778648 kB) free.
[1582391274] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14953848 kB) free!
[1582391394] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10896276 kB) free!
[1582391514] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (7219492 kB) free!
[1582392114] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17172004 kB) free!
[1582392714] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27572048 kB) free.
[1582393229] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4853884 kB) free!
[1582393314] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7386668 kB) free!
[1582393317] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(11) < 2020/02/22 23:10:06 [error] 12257#0: *58161153 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582393434] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36042472 kB) free.
[1582393437] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582394034] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12982332 kB) free!
[1582394154] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9508316 kB) free!
[1582394274] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39448420 kB) free.
[1582394874] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14750544 kB) free!
[1582394994] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10916860 kB) free!
[1582395114] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39375680 kB) free.
[1582395259] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227972 kB) free!
[1582395379] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235188 kB) free!
[1582395499] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242584 kB) free!
[1582395714] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16729604 kB) free!
[1582395813] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9871116 kB) free!
[1582395834] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (12984508 kB) free!
[1582395933] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2249540 kB) free!
[1582395954] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9516600 kB) free!
[1582396053] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2831268 kB) free!