[1589826647] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15903552 kB) free!
[1589826834] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589827195] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3322008 kB) free.
[1589827247] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35847108 kB) free.
[1589827847] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8870748 kB) free!
[1589827967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4625612 kB) free!
[1589828087] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (693372 kB) free!
[1589828687] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15766080 kB) free!
[1589829287] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34823144 kB) free.
[1589829595] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1367852 kB) free!
[1589829715] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (1252680 kB) free!
[1589829835] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (998892 kB) free!
[1589829887] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9308544 kB) free!
[1589830007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5258084 kB) free!
[1589830127] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1527244 kB) free!
[1589830284] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (477964 kB) free!
[1589830345] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201536 kB) free!
[1589830408] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (191088 kB) free!
[1589830442] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (1926668 kB) free!
[1589830468] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (200304 kB) free!
[1589830728] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15230208 kB) free!
[1589831112] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 29812 MB (13.26% inode=100%):
[1589831172] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 25972 MB (11.55% inode=100%):
[1589831232] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;3;DISK WARNING - free space: /home 27570 MB (12.26% inode=100%):
[1589831292] SERVICE ALERT: webserver2;HDD_Home status;WARNING;HARD;4;DISK WARNING - free space: /home 27569 MB (12.26% inode=100%):
[1589831328] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30301328 kB) free.
[1589831368] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 82.8% (20299668 kB) free.
[1589831592] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 78685 MB (35.01% inode=100%):
[1589831642] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (644960 kB) free!
[1589831788] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2248592 kB) free!
[1589831908] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243068 kB) free!
[1589831928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7339888 kB) free!
[1589832028] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242488 kB) free!
[1589832048] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3824072 kB) free!
[1589832168] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36214540 kB) free.
[1589832242] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (1898016 kB) free!
[1589832584] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (485136 kB) free!
[1589832704] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (449932 kB) free!
[1589832768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11783440 kB) free!
[1589832824] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2444012 kB) free!
[1589832888] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7815616 kB) free!
[1589833008] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4241380 kB) free!
[1589833442] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1406420 kB) free!
[1589833608] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15015928 kB) free!
[1589833828] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.5% (5186680 kB) free!
[1589834042] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3317992 kB) free.
[1589834208] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27893244 kB) free.
[1589834808] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6586712 kB) free!
[1589834928] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29905780 kB) free.
[1589835528] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13916848 kB) free!
[1589835648] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10071208 kB) free!
[1589835768] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6491272 kB) free!
[1589836368] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16631760 kB) free!
[1589836424] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (44862320 kB) free.
[1589836581] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.0% (44374584 kB) free.
[1589836945] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4117216 kB) free!
[1589836969] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28033256 kB) free.
[1589837428] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.1% (44391532 kB) free.
[1589837545] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1538744 kB) free!
[1589837569] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7737524 kB) free!
[1589837689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35848672 kB) free.
[1589837723] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1589837829] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2619) < 172.105.59.173 [19/May/2020:03:03:32 +0530] TCP 502 0 0 180.001 "10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 60.000, 0.000"
[1589837864] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(360) < 49.50.86.152 [19/May/2020:03:01:32 +0530] TCP 502 0 0 60.000 "10.147.212.38:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1589837949] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589837984] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(462) < 2020/05/19 03:01:32 [error] 22126#0: *89214775 no live upstreams while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589837985] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589838105] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589838242] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3257392 kB) free!
[1589838289] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12388364 kB) free!
[1589838362] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2789732 kB) free!
[1589838409] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8608288 kB) free!
[1589838482] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (3254400 kB) free!
[1589838529] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23136944 kB) free!
[1589839682] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1238736 kB) free!
[1589839945] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.0% (9796844 kB) free.
[1589840282] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3495060 kB) free.
[1589840329] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35441916 kB) free.
[1589840929] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12485144 kB) free!
[1589841049] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8691280 kB) free!
[1589841169] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35400304 kB) free.
[1589841745] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (1864896 kB) free!
[1589841769] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5206924 kB) free!
[1589841865] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (1372820 kB) free!
[1589841889] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1042108 kB) free!
[1589841985] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1589841985] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (532980 kB) free!
[1589842009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29923040 kB) free.
[1589842610] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7861636 kB) free!
[1589842730] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (3872704 kB) free!
[1589842850] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (445204 kB) free!
[1589843785] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 36.4% (8920004 kB) free.
[1589844050] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22235996 kB) free!
[1589844650] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35359824 kB) free.
[1589844985] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180620 kB) free!
[1589845106] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174768 kB) free!
[1589845226] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174472 kB) free!
[1589845250] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10719084 kB) free!
[1589845370] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7224748 kB) free!
[1589845490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37144352 kB) free.
[1589845904] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 05:14:56 [error] 22129#0: *89231503 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:141/17, bytes from/to upstream:17/141
[1589846024] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589846090] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16024560 kB) free!
[1589846210] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12372560 kB) free!
[1589846307] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 54.8% (13429244 kB) free.
[1589846322] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 72.7% (17821628 kB) free.
[1589846330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38172272 kB) free.
[1589846518] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.9% (10345052 kB) free.
[1589846930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18259748 kB) free!
[1589847050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39377352 kB) free.
[1589847650] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19939964 kB) free!
[1589847770] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15793992 kB) free!
[1589847891] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39153688 kB) free.
[1589848490] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16679872 kB) free!
[1589848610] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12676452 kB) free!
[1589848730] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36605880 kB) free.
[1589849024] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 06:12:49 [error] 22127#0: *89238860 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:194/17, bytes from/to upstream:17/194
[1589849144] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589849330] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14586532 kB) free!
[1589849450] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38461208 kB) free.
[1589850050] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17474820 kB) free!
[1589850170] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13353644 kB) free!
[1589850290] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37054568 kB) free.
[1589850568] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3618104 kB) free!
[1589850628] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3620584 kB) free!
[1589850688] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 14.8% (3619440 kB) free!
[1589850748] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 14.8% (3617464 kB) free!
[1589850890] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15980272 kB) free!
[1589851010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37889764 kB) free.
[1589851107] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (373192 kB) free!
[1589851227] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (328792 kB) free!
[1589851318] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167608 kB) free!
[1589851319] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 7.9% (1939084 kB) free!
[1589851347] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (292184 kB) free!
[1589851424] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 71.7% (94521984 kB) free.
[1589851438] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165864 kB) free!
[1589851558] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194880 kB) free!
[1589851610] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18248444 kB) free!
[1589851683] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589851730] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14312196 kB) free!
[1589851850] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37297904 kB) free.
[1589852450] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15289900 kB) free!
[1589852570] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37283808 kB) free.
[1589853171] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17306140 kB) free!
[1589853291] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13039960 kB) free!
[1589853411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36892892 kB) free.
[1589854011] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15042744 kB) free!
[1589854131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29491380 kB) free.
[1589854731] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16874788 kB) free!
[1589854851] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12695924 kB) free!
[1589854971] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36166176 kB) free.
[1589855571] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13887508 kB) free!
[1589855691] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9635648 kB) free!
[1589855811] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33488768 kB) free.
[1589856030] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589856411] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11586100 kB) free!
[1589856531] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35074524 kB) free.
[1589856944] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 08:21:41 [error] 22125#0: *89254722 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:233/17, bytes from/to upstream:17/233
[1589857064] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589857131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12842636 kB) free!
[1589857222] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.92, 7.24, 5.74
[1589857251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8570552 kB) free!
[1589857342] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.74, 8.00, 6.19
[1589857371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33250048 kB) free.
[1589857462] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.79, 7.48, 6.22
[1589857972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10728504 kB) free!
[1589858092] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33729112 kB) free.
[1589858662] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.48, 3.41, 4.13
[1589858692] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12767176 kB) free!
[1589858812] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8532156 kB) free!
[1589858932] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32975592 kB) free.
[1589859262] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.99, 6.93, 5.38
[1589859383] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.02, 6.32, 5.36
[1589859503] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.92, 6.16, 5.41
[1589859532] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8802520 kB) free!
[1589859652] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4204996 kB) free!
[1589859772] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589859772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31473212 kB) free.
[1589860372] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8208508 kB) free!
[1589860492] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3671712 kB) free!
[1589860612] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30092064 kB) free.
[1589860664] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 09:20:28 [error] 22130#0: *89262238 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1589860784] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589861123] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1589861212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5037544 kB) free!
[1589861218] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1589861243] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589861322] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196464 kB) free!
[1589861332] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (481072 kB) free!
[1589861338] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1589861363] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1589861442] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169176 kB) free!
[1589861452] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20680040 kB) free!
[1589861459] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1589861562] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187840 kB) free!
[1589861903] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1589861903] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 35.61, 31.81, 20.57
[1589861963] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1589862052] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5916040 kB) free!
[1589862059] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1589862503] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.00, 15.78, 16.96
[1589862653] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16121208 kB) free!
[1589863073] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 60.9% (14935112 kB) free.
[1589863103] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.45, 22.46, 20.92
[1589863163] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589863253] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1589863253] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406628 kB) free!
[1589863283] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1589863403] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1589863703] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.69, 12.25, 16.05
[1589863785] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 10:10:31 [error] 22120#0: *89268325 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:243/17, bytes from/to upstream:17/243
[1589863905] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589863949] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (202272 kB) free!
[1589864282] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2991048 kB) free!
[1589864402] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3500188 kB) free.
[1589865002] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2586236 kB) free!
[1589865105] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 10:34:18 [error] 22133#0: *89271884 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:244/17, bytes from/to upstream:17/244
[1589865122] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1900704 kB) free!
[1589865225] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589865241] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2809092 kB) free!
[1589865242] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2840012 kB) free!
[1589865361] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (2252676 kB) free!
[1589865481] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1621188 kB) free!
[1589865653] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18319484 kB) free!
[1589865714] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 332 processes
[1589865834] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 329 processes
[1589865891] SERVICE ALERT: sig1-vm2;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig1-vm2' was out of bounds)
[1589865914] SERVICE ALERT: sig1-vm2;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'sig1-vm2' was out of bounds)
[1589865954] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1589866011] SERVICE ALERT: sig1-vm2;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.29, 0.11
[1589866034] SERVICE ALERT: sig1-vm2;Total Processes;OK;HARD;1;PROCS OK: 267 processes
[1589866073] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.8% (4845980 kB) free!
[1589866103] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 29.21, 22.20, 20.24
[1589866193] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.1% (4436648 kB) free!
[1589866253] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35355940 kB) free.
[1589866314] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.5% (4033440 kB) free!
[1589866853] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9641012 kB) free!
[1589866914] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2090124 kB) free!
[1589866973] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5118184 kB) free!
[1589867025] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/19 11:10:35 [error] 22134#0: *89277920 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:251/17, bytes from/to upstream:17/251
[1589867093] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (779748 kB) free!
[1589867146] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589867303] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 15.18, 18.19, 19.83
[1589867693] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16653316 kB) free!
[1589868113] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (6894108 kB) free!
[1589868203] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1589868203] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1589868233] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.1% (6338064 kB) free!
[1589868243] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (589176 kB) free!
[1589868345] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 11:30:37 [error] 22126#0: *89281340 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:309/17, bytes from/to upstream:17/309
[1589868353] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 16.4% (5753700 kB) free!
[1589868422] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.1% (4222604 kB) free!
[1589868465] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589868542] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.1% (3722496 kB) free!
[1589868662] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.1% (3238024 kB) free!
[1589868803] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589868843] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2613780 kB) free!
[1589868893] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (8008652 kB) free!
[1589868953] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2944684 kB) free!
[1589869262] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (953468 kB) free!
[1589869493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28603292 kB) free.
[1589869665] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 11:54:25 [error] 22135#0: *89285438 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:217/17, bytes from/to upstream:17/217
[1589869703] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 24.32, 23.18, 21.15
[1589869785] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589870024] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9246656 kB) free!
[1589870043] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1589870043] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1120928 kB) free!
[1589870093] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2436556 kB) free!
[1589870144] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8921764 kB) free!
[1589870153] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.5% (31689372 kB) free.
[1589870214] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (430620 kB) free!
[1589870264] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (8547412 kB) free!
[1589870281] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.2% (15770972 kB) free.
[1589870334] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379872 kB) free!
[1589870562] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4473464 kB) free!
[1589870643] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.4% (2184480 kB) free!
[1589870781] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8305372 kB) free!
[1589870901] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7947848 kB) free!
[1589870985] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/19 12:18:11 [error] 22121#0: *89289674 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:127/17, bytes from/to upstream:17/127
[1589871021] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (7628144 kB) free!
[1589871063] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.0% (22966192 kB) free.
[1589871105] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589871503] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.23, 10.83, 17.35
[1589871534] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34796792 kB) free.
[1589872064] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3788872 kB) free!
[1589872134] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6921556 kB) free!
[1589872221] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4791716 kB) free!
[1589872254] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2364560 kB) free!
[1589872362] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1930468 kB) free!
[1589872374] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432304 kB) free!
[1589872974] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14963332 kB) free!
[1589873574] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446168 kB) free!
[1589873850] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 11.7% (2863876 kB) free!
[1589874774] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24274968 kB) free!
[1589875374] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409084 kB) free!
[1589875443] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1132948 kB) free!
[1589875905] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 13:32:04 [error] 22124#0: *89304435 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:922/17, bytes from/to upstream:17/922
[1589875974] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27496748 kB) free.
[1589876025] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/19 13:43:20 [error] 22123#0: *89306815 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1589876145] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589876281] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4453056 kB) free!
[1589876401] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4073852 kB) free!
[1589876521] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (3651596 kB) free!
[1589876574] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (508996 kB) free!
[1589876643] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (2923720 kB) free!
[1589876694] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (440684 kB) free!
[1589876815] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391140 kB) free!
[1589876949] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(447) < 103.255.146.164 [19/May/2020:13:59:03 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589877069] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589877121] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1589877121] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1659012 kB) free!
[1589878870] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(564) < 103.255.146.164 [19/May/2020:14:30:49 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589878990] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589879643] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3448620 kB) free.
[1589879745] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 14:40:54 [error] 22131#0: *89319587 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:198/17, bytes from/to upstream:17/198
[1589879865] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589879972] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 25347 MB (19% inode=99%):
[1589880092] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 25351 MB (19% inode=99%):
[1589880203] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589880212] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 25351 MB (19% inode=99%):
[1589880244] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3122608 kB) free!
[1589880353] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (6823076 kB) free!
[1589880364] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3078952 kB) free!
[1589880465] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 14:54:39 [error] 22131#0: *89322751 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:253/17, bytes from/to upstream:17/253
[1589880474] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.7% (6540088 kB) free!
[1589880484] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157124 kB) free!
[1589880585] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/19 14:58:26 [error] 22130#0: *89323489 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:253/17, bytes from/to upstream:17/253
[1589880594] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 17.9% (6274156 kB) free!
[1589880705] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589880721] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.9% (15929072 kB) free.
[1589880790] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(144) < 185.255.8.58 [19/May/2020:15:00:06 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1589880910] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589881084] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.9% (5020860 kB) free.
[1589881194] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.7% (31749568 kB) free.
[1589881264] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 12.9% (3185540 kB) free!
[1589881384] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 11.5% (2848084 kB) free!
[1589881504] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (2502372 kB) free!
[1589881615] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33157948 kB) free.
[1589881684] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2857564 kB) free!
[1589881703] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1589881804] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.5% (4956932 kB) free.
[1589882104] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.1% (22980984 kB) free.
[1589882215] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7940612 kB) free!
[1589882335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3420240 kB) free!
[1589882404] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.4% (2337816 kB) free!
[1589882455] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411936 kB) free!
[1589882524] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.6% (4814492 kB) free.
[1589883055] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21700596 kB) free!
[1589883125] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.1% (2296620 kB) free!
[1589883245] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3475320 kB) free.
[1589883354] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1589883474] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1589883655] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (462152 kB) free!
[1589883705] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/19 15:49:35 [error] 22135#0: *89333404 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net"
[1589883812] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 112026 MB (87% inode=99%):
[1589883825] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589883845] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.7% (2065092 kB) free!
[1589883965] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3334640 kB) free.
[1589884074] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1589884194] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1589884255] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (22019400 kB) free!
[1589884425] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 15:55:32 [error] 22135#0: *89334625 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net"
[1589884545] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589884846] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 361 processes
[1589884855] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379988 kB) free!
[1589884966] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1589885165] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.7% (1905404 kB) free!
[1589885285] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3356964 kB) free.
[1589885456] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22420072 kB) free!
[1589885885] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (685096 kB) free!
[1589886005] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.3% (2161344 kB) free!
[1589886056] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (507308 kB) free!
[1589886126] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3483756 kB) free.
[1589886656] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20143540 kB) free!
[1589886725] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (1515332 kB) free!
[1589886845] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (2933564 kB) free!
[1589886965] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3406068 kB) free.
[1589887256] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394408 kB) free!
[1589887565] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2113500 kB) free!
[1589887685] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3493772 kB) free.
[1589888058] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589888285] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1855532 kB) free!
[1589888405] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3501160 kB) free.
[1589888521] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.7% (2625592 kB) free!
[1589888641] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.6% (2375616 kB) free!
[1589888761] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2085200 kB) free!
[1589889006] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (182524 kB) free!
[1589889126] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2659076 kB) free!
[1589889246] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (737196 kB) free!
[1589889345] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 17:22:56 [error] 22133#0: *89350371 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589889466] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589890065] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 17:36:40 [error] 22135#0: *89352615 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard/getesmegraph HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard/getesmegraph", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/dashboard"
[1589890185] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/19 17:39:24 [error] 22130#0: *89353103 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589890256] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19670256 kB) free!
[1589890305] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589890446] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3964124 kB) free.
[1589890856] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448004 kB) free!
[1589890905] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/05/19 17:50:31 [error] 22135#0: *89355005 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589891025] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589891047] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2599608 kB) free!
[1589891167] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.2% (3935156 kB) free.
[1589891456] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22928496 kB) free!
[1589891625] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/19 17:59:02 [error] 22135#0: *89356542 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589891746] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589891761] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.0% (16163344 kB) free.
[1589891767] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3249172 kB) free!
[1589891887] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4036448 kB) free.
[1589892056] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (472408 kB) free!
[1589892304] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4844976 kB) free!
[1589892424] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4549760 kB) free!
[1589892544] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.1% (22974656 kB) free.
[1589892656] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31439656 kB) free.
[1589893087] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3187616 kB) free!
[1589893207] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.6% (4323744 kB) free.
[1589893257] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6776272 kB) free!
[1589893377] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2391224 kB) free!
[1589893497] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405172 kB) free!
[1589894407] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (1230588 kB) free!
[1589894528] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (154176 kB) free!
[1589894648] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (2502156 kB) free!
[1589894962] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1589895082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1589895202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1589895248] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3766748 kB) free.
[1589895848] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (2584808 kB) free!
[1589895968] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.9% (4705664 kB) free.
[1589896402] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 333 processes
[1589897003] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1589897058] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 762 processes
[1589897097] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17984604 kB) free!
[1589897168] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3138388 kB) free!
[1589897178] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 759 processes
[1589897288] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.9% (4859524 kB) free.
[1589897298] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 761 processes
[1589897697] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405912 kB) free!
[1589899088] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2825448 kB) free!
[1589899103] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.86, 4.29, 5.00
[1589899209] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.4% (4778696 kB) free.
[1589899703] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.27, 4.92, 5.14
[1589899823] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.27, 5.20, 5.22
[1589899943] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.11, 4.82, 5.09
[1589900097] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28519304 kB) free.
[1589900408] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2717620 kB) free!
[1589900528] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.0% (4555332 kB) free.
[1589900698] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (382508 kB) free!
[1589900761] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1589900818] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389732 kB) free!
[1589900938] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401108 kB) free!
[1589901361] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4004332 kB) free!
[1589901481] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.0% (3820348 kB) free!
[1589901538] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35506620 kB) free.
[1589901601] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (3663412 kB) free!
[1589901728] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (1345144 kB) free!
[1589901848] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3635596 kB) free.
[1589902138] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (384656 kB) free!
[1589902258] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (457148 kB) free!
[1589902378] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (466556 kB) free!
[1589902448] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3100600 kB) free!
[1589902568] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165836 kB) free!
[1589902689] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2811536 kB) free!
[1589902801] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.4% (16042272 kB) free.
[1589902943] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.80, 3.60, 4.49
[1589902978] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32845084 kB) free.
[1589903289] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (803728 kB) free!
[1589903578] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1871436 kB) free!
[1589903698] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394728 kB) free!
[1589903818] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406496 kB) free!
[1589903889] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4195308 kB) free.
[1589904418] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17976392 kB) free!
[1589904489] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (1568328 kB) free!
[1589904609] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3614980 kB) free.
[1589904659] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1589904780] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589904900] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1589905018] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423084 kB) free!
[1589905209] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (915716 kB) free!
[1589905314] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 47.2% (11579856 kB) free.
[1589905329] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (2926128 kB) free!
[1589905449] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4068756 kB) free.
[1589905618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32086856 kB) free.
[1589906049] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.4% (1690356 kB) free!
[1589906169] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3406180 kB) free.
[1589906218] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (405956 kB) free!
[1589906339] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (412336 kB) free!
[1589906459] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (445232 kB) free!
[1589907059] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14968036 kB) free!
[1589907369] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.4% (2335112 kB) free!
[1589907489] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3802936 kB) free.
[1589907599] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1699 processes
[1589907659] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448628 kB) free!
[1589908259] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21330384 kB) free!
[1589908651] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 8.4% (2057916 kB) free!
[1589908690] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (1557816 kB) free!
[1589908810] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2903192 kB) free!
[1589908859] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (471836 kB) free!
[1589908919] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 54.5% (13350028 kB) free.
[1589908930] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3681764 kB) free.
[1589909459] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24627284 kB) free!
[1589909530] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (2257112 kB) free!
[1589909650] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (1367340 kB) free!
[1589909770] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (2825780 kB) free!
[1589910059] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398432 kB) free!
[1589910370] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3503128 kB) free.
[1589910659] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23015840 kB) free!
[1589911199] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1702 processes
[1589911259] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397828 kB) free!
[1589911800] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1696 processes
[1589911859] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15366908 kB) free!
[1589912170] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3213436 kB) free!
[1589912290] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3215160 kB) free!
[1589912410] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1257608 kB) free!
[1589912459] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397428 kB) free!