[1557945172] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11549556 kB) free.
[1557945565] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557946372] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9589936 kB) free!
[1557946492] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8070248 kB) free!
[1557946612] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6167156 kB) free!
[1557948412] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9978540 kB) free.
[1557948565] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557948685] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557949213] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 248.94, 132.17, 54.30
[1557949285] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557949333] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 254.92, 173.13, 78.75
[1557949406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557949453] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 255.64, 200.46, 100.21
[1557949526] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557949616] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2200484 kB) free!
[1557950126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557950246] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557950812] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6899768 kB) free!
[1557950846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557950932] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12184312 kB) free.
[1557950966] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557951532] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6740984 kB) free!
[1557951566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557951652] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12170980 kB) free.
[1557951687] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557951853] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.05, 2.03, 12.57
[1557952252] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6846816 kB) free!
[1557952286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1557952372] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183724 kB) free.
[1557952406] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557952973] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6970996 kB) free!
[1557953006] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557953053] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.04, 1.95, 4.79
[1557953093] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12182912 kB) free.
[1557953126] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1557953693] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7116264 kB) free!
[1557953726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1557953813] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12178292 kB) free.
[1557953846] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557954413] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7372276 kB) free!
[1557954446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557954533] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12182112 kB) free.
[1557954567] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1557954567] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557955133] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7528408 kB) free!
[1557955166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557955253] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12173432 kB) free.
[1557955286] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557955853] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7519016 kB) free!
[1557955886] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1557955973] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12174828 kB) free.
[1557956006] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557956573] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7457804 kB) free!
[1557956606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557956694] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12179044 kB) free.
[1557956726] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1557957293] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7783784 kB) free!
[1557957326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1557957413] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12177560 kB) free.
[1557957447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1557958013] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7903492 kB) free!
[1557958046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557958133] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12180740 kB) free.
[1557958166] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1557958550] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 24.50, 121.81, 79.73
[1557958670] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.44, 81.58, 70.09
[1557958733] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8047484 kB) free!
[1557958766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557958790] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.57, 54.61, 61.60
[1557958853] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12177092 kB) free.
[1557958886] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557959453] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7976284 kB) free!
[1557959486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557959573] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12181772 kB) free.
[1557959606] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1557959990] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 1.32, 17.20
[1557960173] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7948016 kB) free!
[1557960207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557960293] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1557960293] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12185784 kB) free.
[1557960327] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1557960893] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7936540 kB) free!
[1557960927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1557961014] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12178260 kB) free.
[1557961047] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557961191] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.29, 4.94
[1557961613] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7772604 kB) free!
[1557961647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557961733] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183100 kB) free.
[1557961767] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557962333] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7717136 kB) free!
[1557962367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557962453] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12170824 kB) free.
[1557962487] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1557963053] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7721160 kB) free!
[1557963088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557963173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12176888 kB) free.
[1557963208] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557963773] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7709048 kB) free!
[1557963808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557963893] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12173796 kB) free.
[1557963928] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1557964493] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7646908 kB) free!
[1557964528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557964613] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12168844 kB) free.
[1557964648] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557964768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557965813] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9314900 kB) free!
[1557965933] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7714104 kB) free!
[1557966053] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12127600 kB) free.
[1557966568] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557967168] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557967253] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9753772 kB) free!
[1557967288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1557967373] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8211744 kB) free!
[1557967408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557967493] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6534872 kB) free!
[1557968093] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12084052 kB) free.
[1557968608] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1557968608] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1557968653] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 210.92, 71.71, 26.72
[1557968773] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 262.60, 137.56, 56.31
[1557968893] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 269.85, 181.70, 82.34
[1557969047] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5368 MB (10% inode=97%):
[1557969209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557969293] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9544180 kB) free!
[1557969329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557969413] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8049156 kB) free!
[1557969449] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557969534] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12067696 kB) free.
[1557969647] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7005 MB (13% inode=98%):
[1557970049] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557970649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557970733] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9520944 kB) free!
[1557970769] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557970853] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8052596 kB) free!
[1557970889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557970973] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6107256 kB) free!
[1557971293] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.63, 1.75, 12.43
[1557971573] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1557971573] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11784680 kB) free.
[1557972493] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.58, 1.52, 4.48
[1557972773] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9066712 kB) free!
[1557972893] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7558368 kB) free!
[1557973013] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12037392 kB) free.
[1557974214] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9005368 kB) free!
[1557974334] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7465292 kB) free!
[1557974454] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12022336 kB) free.
[1557975654] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8956488 kB) free!
[1557975774] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7385480 kB) free!
[1557975894] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12012992 kB) free.
[1557976890] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1557977094] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9461264 kB) free!
[1557977214] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7970296 kB) free!
[1557977334] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.7% (5771744 kB) free!
[1557977490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1557977610] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557977730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1557977934] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11631888 kB) free.
[1557979134] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8396328 kB) free!
[1557979254] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6603516 kB) free!
[1557979313] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.4% (22897868 kB) free.
[1557979374] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11932296 kB) free.
[1557980574] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8630468 kB) free!
[1557980694] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7034404 kB) free!
[1557980814] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11894044 kB) free.
[1557981330] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1557981930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1557982014] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8846144 kB) free!
[1557982051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557982134] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7163932 kB) free!
[1557982171] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557982254] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5008748 kB) free!
[1557982357] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.1% (2944292 kB) free!
[1557982478] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.0% (2756812 kB) free!
[1557982598] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.1% (2609948 kB) free!
[1557982771] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557982854] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10760576 kB) free.
[1557983371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557983455] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9103180 kB) free!
[1557983491] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557983518] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.0% (2947236 kB) free!
[1557983519] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174940 kB) free!
[1557983575] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7446084 kB) free!
[1557983611] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557983628] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179424 kB) free!
[1557983638] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (1510328 kB) free!
[1557983639] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176900 kB) free!
[1557983695] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.0% (5427660 kB) free!
[1557983709] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.4% (22646984 kB) free.
[1557983748] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172624 kB) free!
[1557983759] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.7% (23617820 kB) free.
[1557983798] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1094712 kB) free!
[1557983805] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.1% (22563856 kB) free.
[1557983840] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.5% (14714328 kB) free.
[1557984295] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11647980 kB) free.
[1557984895] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9789088 kB) free!
[1557985015] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8103788 kB) free!
[1557985135] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6381056 kB) free!
[1557985735] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11534308 kB) free.
[1557986935] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8773524 kB) free!
[1557987055] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7047204 kB) free!
[1557987175] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4962556 kB) free!
[1557987211] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557987775] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11181624 kB) free.
[1557987811] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557987931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557988051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557988376] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9437600 kB) free!
[1557988496] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7852148 kB) free!
[1557988616] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6201700 kB) free!
[1557989216] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11281132 kB) free.
[1557989251] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557989717] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (204712 kB) free!
[1557989838] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189180 kB) free!
[1557989852] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557989958] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165856 kB) free!
[1557989972] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557990088] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 601 processes
[1557990092] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557990148] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 599 processes
[1557990416] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8814468 kB) free!
[1557990536] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.4% (7085320 kB) free!
[1557990656] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5082296 kB) free!
[1557991050] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2828460 kB) free!
[1557991170] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.5% (2686000 kB) free!
[1557991256] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10983796 kB) free.
[1557991290] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.7% (2554096 kB) free!
[1557991730] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.2% (22353888 kB) free.
[1557992456] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8824472 kB) free!
[1557992490] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1557992490] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1323604 kB) free!
[1557992576] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7060056 kB) free!
[1557992697] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5089132 kB) free!
[1557993297] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11079936 kB) free.
[1557993897] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9518752 kB) free!
[1557994017] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7860576 kB) free!
[1557994137] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6181960 kB) free!
[1557994293] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1557994737] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4018724 kB) free!
[1557994893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557995013] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557995133] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557995293] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 290.34, 158.12, 65.45
[1557995337] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10612288 kB) free.
[1557995413] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 296.06, 204.08, 93.53
[1557995533] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 297.02, 234.92, 118.24
[1557995937] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9235396 kB) free!
[1557996057] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7692928 kB) free!
[1557996177] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6100440 kB) free!
[1557996333] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1557996405] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4904168 kB) free!
[1557996526] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4836184 kB) free!
[1557996646] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4761452 kB) free!
[1557996881] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.4% (14694960 kB) free.
[1557996933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557997053] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557997174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557997191] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 37.67, 149.61, 95.65
[1557997246] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 21.6% (5304484 kB) free.
[1557997311] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 5.18, 100.15, 84.07
[1557997377] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9953148 kB) free.
[1557997431] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.12, 67.17, 73.93
[1557997933] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.14, 2.26, 13.87
[1557997978] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8290480 kB) free!
[1557998098] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6460696 kB) free!
[1557998218] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4401380 kB) free!
[1557998446] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4527688 kB) free!
[1557998566] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (4451628 kB) free!
[1557998686] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4381436 kB) free!
[1557998818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10668252 kB) free.
[1557999231] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.18, 0.46, 10.90
[1557999418] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9672548 kB) free!
[1557999538] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7652308 kB) free!
[1557999567] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.23, 6.70, 5.60
[1557999658] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.9% (5863828 kB) free!
[1557999687] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.78, 6.47, 5.64
[1557999733] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.23, 2.27, 3.87
[1557999807] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.76, 5.75, 5.48
[1558000258] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4241380 kB) free!
[1558000407] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.65, 4.24, 4.84
[1558000431] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.24, 3.19
[1558000858] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10495880 kB) free.
[1558001007] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.27, 5.46, 5.22
[1558001127] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.64, 5.27, 5.18
[1558001247] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.67, 5.20, 5.18
[1558001458] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8523588 kB) free!
[1558001578] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6719880 kB) free!
[1558001698] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4842476 kB) free!
[1558001847] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1558001847] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.28, 4.42, 4.82
[1558001912] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3202384 kB) free!
[1558002033] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3074340 kB) free!
[1558002153] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.9% (2916204 kB) free!
[1558002299] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10408816 kB) free.
[1558002888] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1558002899] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8781168 kB) free!
[1558003019] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6910808 kB) free!
[1558003139] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5129728 kB) free!
[1558003488] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (2860704 kB) free!
[1558003608] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.8% (2737684 kB) free!
[1558003728] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.0% (2607272 kB) free!
[1558003739] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3135916 kB) free!
[1558003774] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558003953] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1565128 kB) free!
[1558004339] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10149708 kB) free.
[1558004374] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558004494] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558004614] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1558004929] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1507676 kB) free!
[1558004939] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8346220 kB) free!
[1558005059] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6551412 kB) free!
[1558005098] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176536 kB) free!
[1558005102] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 345 processes
[1558005179] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4702572 kB) free!
[1558005218] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (177848 kB) free!
[1558005222] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 345 processes
[1558005287] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1558005287] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2411892 kB) free!
[1558005338] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152444 kB) free!
[1558005342] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 343 processes
[1558005539] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197232 kB) free!
[1558005659] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186352 kB) free!
[1558005779] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200656 kB) free!
[1558005779] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10213500 kB) free.
[1558005815] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558005831] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1558005831] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.6% (14728000 kB) free.
[1558005883] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.2% (13690624 kB) free.
[1558005922] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.1% (22581100 kB) free.
[1558005942] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1558005966] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.4% (22408444 kB) free.
[1558006002] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.8% (14265180 kB) free.
[1558006379] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8815084 kB) free!
[1558006415] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558006499] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7152928 kB) free!
[1558006535] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558006620] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.2% (5498696 kB) free!
[1558006655] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1558007220] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4121168 kB) free!
[1558007820] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10140840 kB) free.
[1558008420] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8601120 kB) free!
[1558008540] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6902544 kB) free!
[1558008660] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.6% (5202776 kB) free!
[1558009260] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3655732 kB) free!
[1558009542] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558009662] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558009860] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.6% (9659008 kB) free!
[1558012260] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (9994040 kB) free.
[1558012861] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9004528 kB) free!
[1558012981] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7262604 kB) free!
[1558013101] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5668936 kB) free!
[1558013247] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558013701] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4001136 kB) free!
[1558014133] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 242.79, 82.53, 30.54
[1558014253] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 302.63, 158.48, 64.68
[1558014301] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9943144 kB) free.
[1558014373] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 310.43, 209.21, 94.64
[1558014901] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8677736 kB) free!
[1558015021] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7028596 kB) free!
[1558015141] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5352824 kB) free!
[1558015392] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558015440] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558015741] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3570944 kB) free!
[1558016341] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.3% (9495208 kB) free!
[1558016766] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558016766] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4852424 kB) free!
[1558016773] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.61, 1.91, 16.82
[1558016886] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4811036 kB) free!
[1558017007] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4783024 kB) free!
[1558018055] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1558018142] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4230532 kB) free!
[1558018573] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.48, 1.59, 3.86
[1558018742] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.9% (9785308 kB) free!
[1558019255] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1558019855] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558019976] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558020096] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558020771] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 628 processes
[1558020831] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 23.90, 76.03, 40.37
[1558020891] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 627 processes
[1558020951] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.47, 50.96, 35.51
[1558021011] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 322 processes
[1558021071] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.50, 34.12, 31.21
[1558021142] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4442412 kB) free!
[1558021671] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.26, 4.78, 16.47
[1558022342] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.1% (8905592 kB) free!
[1558022871] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.43, 4.77
[1558024142] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4103312 kB) free!
[1558024369] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1558024489] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1558024742] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.6% (9666204 kB) free!
[1558026696] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558027142] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3543288 kB) free!
[1558027296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558027416] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558027537] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1558027537] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558027742] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.7% (9201200 kB) free!
[1558029542] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4890816 kB) free!
[1558030142] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.5% (9634424 kB) free!
[1558031137] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes