[1576521416] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (667304 kB) free!
[1576521489] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176224 kB) free!
[1576521498] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188124 kB) free!
[1576521541] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192408 kB) free!
[1576521610] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193080 kB) free!
[1576521618] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182192 kB) free!
[1576521661] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165836 kB) free!
[1576521730] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176284 kB) free!
[1576521739] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201512 kB) free!
[1576521856] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184748 kB) free!
[1576521976] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172584 kB) free!
[1576522096] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200132 kB) free!
[1576523074] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199232 kB) free!
[1576523134] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 58.3% (14283852 kB) free.
[1576523419] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(22) < 2019/12/17 00:35:45 [error] 20756#0: *17620414 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1576523539] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.5% (9743724 kB) free.
[1576523539] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576523810] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576523930] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 401 processes
[1576524050] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576524739] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (3498836 kB) free!
[1576524859] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198396 kB) free!
[1576524979] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177644 kB) free!
[1576525375] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7414980 kB) free!
[1576525495] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.6% (6220212 kB) free!
[1576525615] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3781256 kB) free!
[1576525851] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1576525971] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576527379] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1576527379] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.7% (10278888 kB) free.
[1576528206] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 593 processes
[1576528806] SERVICE ALERT: webservercluster1;Total Processes;OK;HARD;1;PROCS OK: 500 processes
[1576529541] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2019/12/17 02:22:14 [error] 20756#0: *17637307 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576529615] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8549968 kB) free!
[1576529661] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(31) < 2019/12/17 02:24:11 [error] 20760#0: *17637475 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576529735] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (2484000 kB) free!
[1576529779] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186212 kB) free!
[1576529781] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(34) < 2019/12/17 02:26:16 [error] 20756#0: *17637655 recv() failed (104: Connection reset by peer) while proxying connection, client: 54.254.180.107, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:50/0, bytes from/to upstream:0/50
[1576529855] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229304 kB) free!
[1576529899] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195232 kB) free!
[1576530019] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197168 kB) free!
[1576530456] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.7% (43229968 kB) free.
[1576530981] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576532426] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/17 03:10:01 [error] 28861#0: *21871608 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:37145232/10763895, bytes from/to upstream:10763895/37145232
[1576532546] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576539619] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1576541571] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1576541691] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576542291] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576542412] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576543094] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 67.7% (16695868 kB) free.
[1576543136] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.2% (22350052 kB) free.
[1576543230] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.8% (38857316 kB) free.
[1576543277] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.5% (22671564 kB) free.
[1576543324] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.5% (22663256 kB) free.
[1576543698] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190532 kB) free!
[1576543748] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166380 kB) free!
[1576543818] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176240 kB) free!
[1576543868] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.0% (22316968 kB) free.
[1576543938] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189740 kB) free!
[1576544781] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2019/12/17 06:35:43 [error] 20756#0: *17680657 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576544901] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(51) < 2019/12/17 06:38:17 [error] 20749#0: *17680888 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:44/0, bytes from/to upstream:0/44
[1576545021] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(28) < 2019/12/17 06:39:42 [error] 20746#0: *17680993 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576545621] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576546553] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.1% (22963556 kB) free.
[1576550211] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576550331] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576557621] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(124) < 2019/12/17 10:08:01 [error] 20754#0: *17714089 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576557741] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576558943] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 383 processes
[1576559063] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1576559158] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.0% (3956104 kB) free!
[1576559278] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.6% (3842664 kB) free!
[1576559398] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.1% (3729480 kB) free!
[1576560598] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181648 kB) free!
[1576561016] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.2% (22756000 kB) free.
[1576562468] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4391828 kB) free!
[1576562588] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (4194640 kB) free!
[1576562708] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (3994436 kB) free!
[1576562826] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (1191988 kB) free!
[1576562946] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193120 kB) free!
[1576562955] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3511508 kB) free!
[1576563066] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1576563066] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181404 kB) free!
[1576563075] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3128220 kB) free!
[1576563195] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.3% (2761116 kB) free!
[1576563795] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1128620 kB) free!
[1576563901] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (305536 kB) free!
[1576563908] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2125784 kB) free!
[1576564021] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (354536 kB) free!
[1576564112] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1099) < 103.255.146.164 [17/Dec/2019:11:54:11 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576564141] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (361636 kB) free!
[1576564232] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576564341] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1157) < 2019/12/17 11:54:11 [error] 20755#0: *17733054 no live upstreams while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "backend_3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576564395] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 78.8% (17683708 kB) free.
[1576564461] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576564523] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22661564 kB) free.
[1576564556] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.6% (22698984 kB) free.
[1576564642] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 12.7% (3142608 kB) free!
[1576564678] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.1% (14811432 kB) free.
[1576564762] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 11.1% (2741352 kB) free!
[1576564837] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8788840 kB) free!
[1576564882] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2281212 kB) free!
[1576564957] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (7879912 kB) free!
[1576565061] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/17 12:08:16 [error] 20755#0: *17735472 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:47/0, bytes from/to upstream:0/47
[1576565077] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (7865848 kB) free!
[1576565181] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576565482] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 94.0% (23201056 kB) free.
[1576566326] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4684344 kB) free!
[1576566446] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4579292 kB) free!
[1576566566] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (4488528 kB) free!
[1576566877] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4447476 kB) free!
[1576568077] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (5683512 kB) free!
[1576568366] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2256104 kB) free!
[1576568677] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576568677] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (433468 kB) free!
[1576569085] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4679940 kB) free!
[1576569205] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4554732 kB) free!
[1576569326] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4434788 kB) free!
[1576569458] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8808500 kB) free!
[1576569578] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8594712 kB) free!
[1576569698] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.1% (8406420 kB) free!
[1576569877] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (8188020 kB) free!
[1576569981] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(97) < 2019/12/17 13:36:14 [error] 20756#0: *17749427 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576570101] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(38) < 2019/12/17 13:37:27 [error] 20757#0: *17749461 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576570222] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576571726] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1846816 kB) free!
[1576572621] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/17 14:15:18 [error] 20757#0: *17755984 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576572741] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576573298] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4797740 kB) free!
[1576573341] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/17 14:26:34 [error] 20756#0: *17757821 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576573462] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576575126] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4910744 kB) free!
[1576575246] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4796740 kB) free!
[1576575301] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.6% (22610904 kB) free.
[1576575336] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 35.5% (17492916 kB) free.
[1576575366] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (4698948 kB) free!
[1576575377] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 66.4% (16278696 kB) free.
[1576575877] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4342368 kB) free!
[1576576131] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576576252] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1576577706] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (4262268 kB) free!
[1576577827] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (4175716 kB) free!
[1576577947] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (4077192 kB) free!
[1576578966] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2321424 kB) free!
[1576579232] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(27) < 85.10.207.167 [17/Dec/2019:16:10:32 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576579353] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576579461] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2019/12/17 16:10:32 [error] 20759#0: *17775689 no live upstreams while connecting to upstream, client: 85.10.207.167, server: 0.0.0.0:13000, upstream: "backend_13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576579581] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576579747] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1361376 kB) free!
[1576580768] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4562756 kB) free!
[1576580888] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4424836 kB) free!
[1576581008] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4247436 kB) free!
[1576581877] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1576582495] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22510892 kB) free.
[1576582526] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.9% (38883752 kB) free.
[1576582564] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 70.9% (17370152 kB) free.
[1576582808] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2167864 kB) free!
[1576583096] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (2479100 kB) free!
[1576583216] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.1% (2500116 kB) free!
[1576583336] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2407932 kB) free!
[1576584981] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2019/12/17 17:45:45 [error] 20758#0: *17791972 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576585058] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9401148 kB) free!
[1576585101] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576585178] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9306864 kB) free!
[1576585298] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (9214592 kB) free!
[1576585463] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 397 processes
[1576585583] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 394 processes
[1576586183] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1576586303] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 408 processes
[1576586423] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1576589902] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2019/12/17 18:59:30 [error] 20758#0: *17804733 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1576590022] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576590334] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6191804 kB) free!
[1576590455] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (391328 kB) free!
[1576590575] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235388 kB) free!
[1576590698] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4809896 kB) free!
[1576591222] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(98) < 2019/12/17 19:28:14 [error] 20756#0: *17809232 recv() failed (104: Connection reset by peer) while proxying connection, client: 139.59.21.208, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:57/0, bytes from/to upstream:0/57
[1576591341] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.8% (22397012 kB) free.
[1576591342] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576591373] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 66.1% (16194996 kB) free.
[1576591402] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 32.3% (15911688 kB) free.
[1576591433] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 59.7% (29430200 kB) free.
[1576591946] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.0% (3955396 kB) free!
[1576592066] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.7% (3877820 kB) free!
[1576592186] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.5% (3828296 kB) free!
[1576592546] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/17 19:50:41 [error] 28866#0: *22515567 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1576592666] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576592786] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192620 kB) free!
[1576593051] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1576593172] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576593266] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/17 19:56:50 [error] 28875#0: *22518128 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1576593386] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576593609] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 555 processes
[1576593729] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 556 processes
[1576593835] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8134568 kB) free!
[1576593849] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 558 processes
[1576593955] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.9% (7848140 kB) free!
[1576593985] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4701512 kB) free!
[1576594075] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.8% (7293856 kB) free!
[1576594106] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4646036 kB) free!
[1576594226] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4595000 kB) free!
[1576594676] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1576594676] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2645060 kB) free!
[1576595128] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1576595249] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1576595369] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1576595428] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 1.5% (368752 kB) free!
[1576595542] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/12/17 20:37:22 [error] 20756#0: *17820704 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576595548] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 1.5% (366856 kB) free!
[1576595649] SERVICE ALERT: webservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 607 processes
[1576595662] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576595668] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 1.5% (361528 kB) free!
[1576595786] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/17 20:38:36 [error] 28874#0: *22533107 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:233/17, bytes from/to upstream:17/233
[1576595874] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.8% (22401396 kB) free.
[1576595906] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576595915] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 91.5% (22425012 kB) free.
[1576595938] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 42.7% (21028384 kB) free.
[1576595969] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1576596089] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1576596209] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1576597678] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192928 kB) free!
[1576597706] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/17 21:13:10 [error] 28875#0: *22546321 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:247/17, bytes from/to upstream:17/247
[1576597798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179696 kB) free!
[1576597827] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/17 21:18:55 [error] 28875#0: *22548387 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:443/17, bytes from/to upstream:17/443
[1576597918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187520 kB) free!
[1576597947] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576598347] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8733976 kB) free!
[1576598467] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.7% (7244524 kB) free!
[1576598587] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (6307380 kB) free!
[1576598662] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(56) < 2019/12/17 21:31:18 [error] 20756#0: *17828455 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:350410/132667, bytes from/to upstream:132667/350410
[1576598746] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.0% (22211588 kB) free.
[1576598782] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576599147] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2019/12/17 21:33:15 [error] 28875#0: *22549726 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:2438/973, bytes from/to upstream:973/2438
[1576599187] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2643500 kB) free!
[1576599267] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576600827] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2416512 kB) free!
[1576601782] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2019/12/17 22:20:27 [error] 20758#0: *17837930 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576601902] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576603102] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2019/12/17 22:48:21 [error] 20758#0: *17842258 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576603222] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(27) < 2019/12/17 22:50:18 [error] 20755#0: *17842385 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576603342] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(44) < 2019/12/17 22:52:22 [error] 20756#0: *17842580 connect() failed (111: Connection refused) while connecting to upstream, client: 210.212.235.210, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576605752] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(14) < 103.209.99.7 [17/Dec/2019:23:28:24 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576605872] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(42) < 180.179.218.150 [17/Dec/2019:23:32:58 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576605992] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576606942] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found