[1580322690] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(32) < 2020/01/30 00:01:29 [error] 18339#0: *27198175 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580322810] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(20) < 2020/01/30 00:03:29 [error] 18340#0: *27198250 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580323319] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184420 kB) free!
[1580323349] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.6% (3081292 kB) free!
[1580323440] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192648 kB) free!
[1580323470] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.6% (3082368 kB) free!
[1580323537] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183280 kB) free!
[1580323560] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178060 kB) free!
[1580323590] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (3075140 kB) free!
[1580323657] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177908 kB) free!
[1580323777] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174472 kB) free!
[1580325302] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.6% (21125096 kB) free.
[1580325344] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.7% (21485040 kB) free.
[1580325381] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580325388] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/30 00:46:27 [error] 18343#0: *27199831 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580325508] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(20) < 2020/01/30 00:48:27 [error] 18343#0: *27199908 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580325628] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1580325628] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(20) < 2020/01/30 00:50:27 [error] 18343#0: *27199991 access forbidden by rule while initializing session, client: 54.254.180.107, server: 0.0.0.0:13000
[1580326590] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175796 kB) free!
[1580327067] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9134340 kB) free!
[1580327187] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.9% (7846424 kB) free!
[1580327252] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(139) < 2020/01/30 01:17:30 [error] 10221#0: *45592260 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1580327307] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.3% (6546552 kB) free!
[1580327372] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(7) < 2020/01/30 01:17:57 [error] 10222#0: *45593071 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580327492] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580327645] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 103.16.101.72 [30/Jan/2020:01:16:57 +0530] TCP 502 0 0 60.002 "10.147.212.61:4000, 10.147.212.49:4000, 10.147.212.38:4000, backend3" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 60.001, 0.000"
[1580327708] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (2080972 kB) free!
[1580327766] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580327828] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (376656 kB) free!
[1580327907] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (265072 kB) free!
[1580327948] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169764 kB) free!
[1580329590] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.3% (10849564 kB) free.
[1580331670] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 70.0% (34511324 kB) free.
[1580333492] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(320) < 2020/01/30 03:01:31 [error] 10221#0: *45633826 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580333612] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580333766] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(244) < 103.16.101.72 [30/Jan/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1580333856] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2947) < 172.105.59.173 [30/Jan/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1580333888] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580333976] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580336428] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1580337213] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2020/01/30 04:00:40 [error] 10212#0: *45656392 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580337333] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580338052] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21854844 kB) free.
[1580338097] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.3% (40069040 kB) free.
[1580341060] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.8% (3645428 kB) free!
[1580341181] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172172 kB) free!
[1580341301] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189280 kB) free!
[1580341591] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (218884 kB) free!
[1580341711] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (238536 kB) free!
[1580341831] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (240084 kB) free!
[1580341901] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.1% (18788088 kB) free.
[1580342107] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1580342228] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1580346427] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580346547] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1580346564] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21942156 kB) free.
[1580350748] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580350868] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1580351733] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(47) < 2020/01/30 08:01:22 [error] 10216#0: *45740437 connect() failed (111: Connection refused) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580351853] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580356223] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.4% (14213880 kB) free.
[1580358118] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171380 kB) free!
[1580358238] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195992 kB) free!
[1580358349] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165016 kB) free!
[1580358358] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173228 kB) free!
[1580358469] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172432 kB) free!
[1580358589] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189928 kB) free!
[1580358731] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1580358851] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580358971] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580359571] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1580359963] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1580360083] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1580360374] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.8% (21531368 kB) free.
[1580360424] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21717956 kB) free.
[1580361701] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4568228 kB) free!
[1580361821] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4513272 kB) free!
[1580361941] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4446896 kB) free!
[1580363141] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2236408 kB) free!
[1580363286] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.6% (21379820 kB) free.
[1580365631] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1580365700] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4265912 kB) free!
[1580365820] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.3% (16237620 kB) free.
[1580365887] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.7% (4837480 kB) free!
[1580366007] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.6% (4815060 kB) free!
[1580366127] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.6% (4796524 kB) free!
[1580366253] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(124) < 2020/01/30 12:01:36 [error] 10207#0: *45861954 connect() failed (111: Connection refused) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580366373] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580367567] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164332 kB) free!
[1580367687] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186056 kB) free!
[1580367808] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186960 kB) free!
[1580367831] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 75.6% (18533572 kB) free.
[1580368173] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/30 12:36:49 [error] 10215#0: *45885034 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /main.php?page=cligrptosenderid&action=add HTTP/1.0", upstream: "http://10.147.212.40:80/main.php?page=cligrptosenderid&action=add", host: "altsmsc.tesync.net"
[1580368293] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580368439] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192508 kB) free!
[1580368559] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193520 kB) free!
[1580368679] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195728 kB) free!
[1580369395] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1580369395] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.2% (21627308 kB) free.
[1580369425] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.7% (22226164 kB) free.
[1580369889] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.8% (2671192 kB) free!
[1580370009] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196612 kB) free!
[1580370130] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165244 kB) free!
[1580370169] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1580370169] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.6% (19139584 kB) free.
[1580370309] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214608 kB) free!
[1580370429] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233416 kB) free!
[1580370489] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.1% (42937784 kB) free.
[1580370779] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189600 kB) free!
[1580370899] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173096 kB) free!
[1580371019] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195344 kB) free!
[1580371031] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1580371152] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1580371270] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7799628 kB) free!
[1580371272] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1580371390] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3675760 kB) free!
[1580371510] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212112 kB) free!
[1580371714] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.5% (19136492 kB) free.
[1580371752] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 88.4% (43582464 kB) free.
[1580372918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187400 kB) free!
[1580373038] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180680 kB) free!
[1580373158] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178964 kB) free!
[1580373498] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9879588 kB) free!
[1580373618] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8901428 kB) free!
[1580373738] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (8258824 kB) free!
[1580374161] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.2% (21025248 kB) free.
[1580374199] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.3% (40044088 kB) free.
[1580375370] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (4095520 kB) free!
[1580375490] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (4000628 kB) free!
[1580375610] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (3919968 kB) free!
[1580376576] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(73) < 92.42.104.166 [30/Jan/2020:14:57:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1580376696] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580376810] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2465040 kB) free!
[1580377296] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(320) < 103.255.146.164 [30/Jan/2020:15:02:09 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1580377416] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580378979] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4809632 kB) free!
[1580379099] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4742268 kB) free!
[1580379219] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4652012 kB) free!
[1580379598] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1580380293] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2020/01/30 16:00:52 [error] 10219#0: *46019352 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580380413] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(127) < 2020/01/30 16:03:05 [error] 10222#0: *46020272 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580380533] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580380826] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.2% (4705384 kB) free!
[1580380947] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.7% (4580352 kB) free!
[1580381067] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.0% (4421228 kB) free!
[1580384398] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (4420204 kB) free!
[1580384518] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (4289388 kB) free!
[1580384638] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4140448 kB) free!
[1580385267] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2427668 kB) free!
[1580387038] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2157788 kB) free!
[1580387610] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1580388933] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/30 18:22:42 [error] 10222#0: *46105912 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET / HTTP/1.0", upstream: "http://10.147.212.40:80/", host: "altsmsc.tesync.net"
[1580389053] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580389868] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580389988] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580390083] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1580390203] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1580390619] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2406436 kB) free!
[1580392388] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580392509] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580392814] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.9% (7825272 kB) free!
[1580392934] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (4748160 kB) free!
[1580393054] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3661756 kB) free!
[1580394816] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(42) < 203.212.70.170 [30/Jan/2020:20:02:20 +0530] TCP 502 0 0 91.066 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "31.066, 60.000, 0.000"
[1580394936] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580395053] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(64) < 2020/01/30 20:01:24 [error] 10212#0: *46154707 connect() failed (111: Connection refused) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580395174] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580397308] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580397428] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580398169] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.2% (21523400 kB) free.
[1580398229] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.3% (40044104 kB) free.
[1580398246] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21789732 kB) free.
[1580398319] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.5% (21691848 kB) free.
[1580398336] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.2% (20404012 kB) free.
[1580400577] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.3% (3032664 kB) free!
[1580400697] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (3024868 kB) free!
[1580400817] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (3015772 kB) free!
[1580402617] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2379528 kB) free!
[1580407639] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2288864 kB) free!
[1580407759] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240684 kB) free!
[1580407879] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224492 kB) free!