[1582828205] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6473960 kB) free!
[1582828325] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30234916 kB) free.
[1582828494] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.0% (29067504 kB) free.
[1582828685] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2020/02/28 00:01:17 [error] 12247#0: *61282401 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1582828742] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1582828805] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582828862] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1582828925] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9127520 kB) free!
[1582828983] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1582829045] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32783416 kB) free.
[1582829442] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1582829645] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11910504 kB) free!
[1582829739] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582829765] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8283300 kB) free!
[1582829885] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31501264 kB) free.
[1582830360] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188112 kB) free!
[1582830420] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192264 kB) free!
[1582830433] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (4498504 kB) free!
[1582830480] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.0% (16908324 kB) free.
[1582830485] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10174704 kB) free!
[1582830605] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32746784 kB) free.
[1582830785] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1582831033] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (207672 kB) free!
[1582831205] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12105636 kB) free!
[1582831325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8669012 kB) free!
[1582831445] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30038964 kB) free.
[1582831633] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.0% (9371668 kB) free.
[1582831985] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1582832045] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10352980 kB) free!
[1582832105] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582832165] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32685496 kB) free.
[1582832225] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1582832233] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (2052208 kB) free!
[1582832353] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (516244 kB) free!
[1582832474] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1582832474] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174992 kB) free!
[1582832766] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11719420 kB) free!
[1582832886] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32727224 kB) free.
[1582833486] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12389332 kB) free!
[1582833606] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (9008244 kB) free!
[1582833726] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29588172 kB) free.
[1582834326] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9488864 kB) free!
[1582834446] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30188648 kB) free.
[1582834626] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582834746] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1582835046] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10215836 kB) free!
[1582835166] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30915960 kB) free.
[1582835346] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582835466] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1582835586] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1582835766] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10437128 kB) free!
[1582835886] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30915564 kB) free.
[1582836186] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582836306] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582836486] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24633524 kB) free!
[1582836550] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.6% (9957688 kB) free.
[1582836606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29260228 kB) free.
[1582836839] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5848120 kB) free!
[1582836959] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238916 kB) free!
[1582837080] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237172 kB) free!
[1582837507] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582837627] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582837680] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 84.5% (41656156 kB) free.
[1582837729] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.8% (16894308 kB) free!
[1582837747] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582837806] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12004996 kB) free!
[1582837927] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582837927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32433272 kB) free.
[1582838347] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1582838412] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.2% (43978084 kB) free.
[1582838467] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582838526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11176400 kB) free!
[1582838587] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582838646] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30385400 kB) free.
[1582839148] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3609) < 203.212.70.170 [28/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582839187] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 357 processes
[1582839268] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582839307] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582839427] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1582839541] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(232) < 103.16.101.72 [28/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582839605] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(293) < 2020/02/28 03:00:22 [error] 12243#0: *61362101 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582839663] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582839725] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582839846] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12469048 kB) free!
[1582839966] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29439272 kB) free.
[1582840027] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582840628] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1582841828] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582841948] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582842068] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1582842366] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12869916 kB) free!
[1582842486] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29909512 kB) free.
[1582842668] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1582842788] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1582842908] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1582843508] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582843874] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582844286] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22432512 kB) free!
[1582844406] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29722804 kB) free.
[1582845908] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1582846508] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 355 processes
[1582846628] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582846748] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582846750] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196124 kB) free!
[1582846806] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7225736 kB) free!
[1582846870] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178644 kB) free!
[1582846926] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29986536 kB) free.
[1582846990] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.4% (5488340 kB) free.
[1582847590] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194780 kB) free!
[1582847710] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195604 kB) free!
[1582847830] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191444 kB) free!
[1582848074] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 46.2% (11396784 kB) free.
[1582848126] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9739588 kB) free!
[1582848246] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30115588 kB) free.
[1582848548] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1582849148] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582849269] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582849389] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582849989] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1582850589] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582850709] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582850929] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 29.6% (39005312 kB) free.
[1582851309] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582851429] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582852029] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582852149] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582852269] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1582852869] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582852990] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582853442] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 78.8% (19319344 kB) free.
[1582853479] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39265780 kB) free.
[1582853511] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.7% (21753028 kB) free.
[1582853589] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1582853709] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582854046] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.4% (3043232 kB) free!
[1582854075] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4759696 kB) free!
[1582854166] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.4% (3031892 kB) free!
[1582854195] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (208680 kB) free!
[1582854286] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.3% (3025936 kB) free!
[1582854309] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582854315] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195908 kB) free!
[1582854429] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582854549] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1582854886] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191732 kB) free!
[1582855149] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582856349] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1582858972] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21914648 kB) free.
[1582859002] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1582859002] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.7% (22225172 kB) free.
[1582859068] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [28/Feb/2020:08:32:54 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1582859188] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582859732] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582859852] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582860246] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17218288 kB) free!
[1582860366] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31839612 kB) free.
[1582860768] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 186.89, 48.65, 19.38
[1582860889] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 527.85, 226.90, 88.28
[1582861009] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 683.72, 387.21, 163.94
[1582862809] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.86, 7.45, 30.45
[1582863409] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 18.20, 10.84, 20.74
[1582864010] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 114.83, 36.90, 25.47
[1582864130] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 498.05, 173.90, 74.83
[1582864251] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1582864251] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 555.68, 288.55, 128.05
[1582864566] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22789916 kB) free!
[1582864686] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.3% (24061264 kB) free!
[1582864806] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18897544 kB) free!
[1582864925] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2020/02/28 10:05:14 [error] 12253#0: *61529068 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582865012] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1582865045] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582865728] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1153 processes
[1582865848] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1153 processes
[1582865895] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 25.51, 15.23, 7.28
[1582865969] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1149 processes
[1582866006] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43226300 kB) free.
[1582866015] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 27.11, 18.26, 9.31
[1582866042] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.7% (3854272 kB) free!
[1582866135] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 28.36, 22.63, 12.08
[1582866162] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.2% (3473924 kB) free!
[1582866283] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 12.9% (3161812 kB) free!
[1582866549] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1582866606] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14683576 kB) free!
[1582866669] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 352 processes
[1582866726] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45120132 kB) free.
[1582866735] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.58, 29.72, 20.35
[1582866789] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1582866883] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1391268 kB) free!
[1582867251] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.66, 15.14, 33.28
[1582867310] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582867326] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21555192 kB) free!
[1582867335] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 14.97, 19.28, 19.83
[1582867430] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1582867446] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44731564 kB) free.
[1582867550] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1582867852] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 14.69, 16.76, 25.60
[1582867936] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1582867936] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.73, 23.16, 21.94
[1582868046] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20316476 kB) free!
[1582868167] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16526060 kB) free!
[1582868264] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.7% (2626948 kB) free!
[1582868287] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39009268 kB) free.
[1582868384] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (2190276 kB) free!
[1582868504] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1816720 kB) free!
[1582868581] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (1090652 kB) free!
[1582868589] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582868701] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (842736 kB) free!
[1582868821] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1581276 kB) free!
[1582868887] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15899128 kB) free!
[1582869007] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40186208 kB) free.
[1582869350] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582869607] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16396644 kB) free!
[1582869682] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.6% (6715616 kB) free!
[1582869727] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12777620 kB) free!
[1582869736] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.02, 8.52, 14.71
[1582869802] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.1% (4978792 kB) free!
[1582869836] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21621968 kB) free.
[1582869845] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2020/02/28 11:33:22 [error] 12257#0: *61572944 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: rq-1-1gbps.server, request: "GET /rabbitmq/api/queues/?page=2&page_size=500&name=&use_regex=false&pagination=true HTTP/1.1", upstream: "http://10.147.212.203:15672/api/queues/%2F/?page=2&page_size=500&name=&use_regex=false&pagination=true", host: "rq-1-1gbps.server:8091", referrer: "http://rq-1-1gbps.server:8091/rabbitmq/"
[1582869847] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35339136 kB) free.
[1582869866] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.9% (21687156 kB) free.
[1582869889] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 51.3% (11524060 kB) free.
[1582869920] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.4% (38638148 kB) free.
[1582869966] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582870045] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4476592 kB) free!
[1582870165] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (4016296 kB) free!
[1582870285] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.5% (3573216 kB) free!
[1582870447] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12701324 kB) free!
[1582870567] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12635032 kB) free!
[1582870687] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31687152 kB) free.
[1582870885] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1158160 kB) free!
[1582871012] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3957312 kB) free!
[1582871070] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (4621040 kB) free!
[1582871094] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8278188 kB) free!
[1582871132] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.5% (3800844 kB) free!
[1582871190] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (4384152 kB) free!
[1582871194] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 15.4% (5361124 kB) free!
[1582871214] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7865720 kB) free!
[1582871252] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (3651036 kB) free!
[1582871287] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9718060 kB) free!
[1582871310] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4176140 kB) free!
[1582871314] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (4980792 kB) free!
[1582871334] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.2% (7508464 kB) free!
[1582871407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33384444 kB) free.
[1582871434] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.1% (32077696 kB) free.
[1582871577] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1582871589] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 391 processes
[1582871697] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1582871765] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2020/02/28 12:01:12 [error] 12255#0: *61585932 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582871788] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 172.105.59.173 [28/Feb/2020:12:02:13 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1582871885] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582871908] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582871911] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1582871911] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164796 kB) free!
[1582871934] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.4% (30236820 kB) free.
[1582872007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10987420 kB) free!
[1582872127] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7177484 kB) free!
[1582872248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29859644 kB) free.
[1582872685] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22890164 kB) free.
[1582872847] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4855544 kB) free!
[1582872967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1097188 kB) free!
[1582873052] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2262772 kB) free!
[1582873087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41930816 kB) free.
[1582873687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18827756 kB) free!
[1582873807] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15163876 kB) free!
[1582873812] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9172152 kB) free!
[1582873927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39298616 kB) free.
[1582873932] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8885132 kB) free!
[1582874035] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.5% (21351384 kB) free.
[1582874052] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.4% (8556056 kB) free!
[1582874176] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.1% (22076632 kB) free.
[1582874204] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 50.5% (24893224 kB) free.
[1582874527] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16498324 kB) free!
[1582874647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12879028 kB) free!
[1582874767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38212236 kB) free.
[1582875119] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4368964 kB) free!
[1582875239] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.5% (4043992 kB) free!
[1582875359] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (3739676 kB) free!
[1582875367] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14922580 kB) free!
[1582875487] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11127800 kB) free!
[1582875607] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37563864 kB) free.
[1582875920] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9781992 kB) free!
[1582875960] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2447124 kB) free!
[1582876040] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 21.7% (10696328 kB) free.
[1582876208] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13465380 kB) free!
[1582876328] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9714984 kB) free!
[1582876448] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35353764 kB) free.
[1582876640] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9296136 kB) free!
[1582876760] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (8276656 kB) free!
[1582876880] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (7264240 kB) free!
[1582877040] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173836 kB) free!
[1582877048] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12630004 kB) free!
[1582877160] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169636 kB) free!
[1582877168] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8716120 kB) free!
[1582877280] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190936 kB) free!
[1582877288] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (5067224 kB) free!
[1582878350] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582878470] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1582878480] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8798800 kB) free!
[1582878488] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1582878488] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22142288 kB) free!
[1582878591] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1582878600] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (8622568 kB) free!
[1582878720] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (8350364 kB) free!
[1582878843] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.6% (3824828 kB) free!
[1582878891] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.1% (3386556 kB) free!
[1582878963] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.9% (3642052 kB) free!
[1582879011] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3074260 kB) free!
[1582879030] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.0% (2272828 kB) free!
[1582879083] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.8% (3388420 kB) free!
[1582879088] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (944548 kB) free!
[1582879131] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.2% (2736980 kB) free!
[1582879150] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.9% (2258372 kB) free!
[1582879252] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582879270] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 14.0% (2270996 kB) free!
[1582879280] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4001156 kB) free!
[1582879391] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.1% (21253688 kB) free.
[1582879497] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 335 processes
[1582879617] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1582879683] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1918516 kB) free!
[1582879731] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (756728 kB) free!
[1582879791] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1582879936] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582880288] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37108004 kB) free.
[1582880520] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4537548 kB) free!
[1582880589] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1582880888] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2902408 kB) free!
[1582881008] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403000 kB) free!
[1582881085] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4618572 kB) free!
[1582881128] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402180 kB) free!
[1582881205] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.1% (4214544 kB) free!
[1582881325] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 15.5% (3816120 kB) free!
[1582881339] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (2726648 kB) free!
[1582881459] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2793380 kB) free!
[1582881579] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (2875280 kB) free!
[1582881736] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 2.98, 2.95, 4.24
[1582881925] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (24.3% change >= 20.0% threshold)
[1582881925] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1853004 kB) free!
[1582881925] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.9% (4866140 kB) free!
[1582881980] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4899116 kB) free!
[1582882046] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.8% (4846364 kB) free!
[1582882100] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4849416 kB) free!
[1582882128] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 545 processes
[1582882134] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9123136 kB) free!
[1582882166] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.7% (4825940 kB) free!
[1582882179] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3313184 kB) free.
[1582882220] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.8% (5095784 kB) free.
[1582882264] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882328] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37445760 kB) free.
[1582882387] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.2% (29193144 kB) free.
[1582882509] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882639] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882694] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882708] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(155) < 92.42.104.166 [28/Feb/2020:15:00:12 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582882731] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.4% (15798200 kB) free.
[1582882735] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882769] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882820] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882820] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4714884 kB) free!
[1582882824] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882828] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582882851] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882857] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882858] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1582882893] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12959912 kB) free!
[1582882941] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4678352 kB) free!
[1582882944] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38752 MB (75% inode=98%):
[1582882944] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 143548 MB (63% inode=99%):
[1582882981] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882987] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582882994] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582883003] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (2385336 kB) free!
[1582883023] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582883048] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9398944 kB) free!
[1582883061] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4645480 kB) free!
[1582883111] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582883117] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582883123] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (2161776 kB) free!
[1582883124] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582883125] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 86.0% (21240376 kB) free.
[1582883144] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2149 processes
[1582883168] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5943740 kB) free!
[1582883243] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2747940 kB) free!
[1582883254] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582883380] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (2647752 kB) free!
[1582883462] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582883465] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582883500] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2793044 kB) free!
[1582883586] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 27.63, 29.79, 28.78
[1582883587] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1582883620] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3125060 kB) free!
[1582883722] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582883844] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (354076 kB) free!
[1582884220] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1582884220] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4201444 kB) free.
[1582884368] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32371652 kB) free.
[1582884670] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181764 kB) free!
[1582884750] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582884754] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582884792] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582884795] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582884820] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1849720 kB) free!
[1582884861] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173480 kB) free!
[1582884880] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582884884] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582884922] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582884925] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582884932] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582884940] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.2% (2152396 kB) free!
[1582884968] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6786596 kB) free!
[1582885010] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582885013] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38751 MB (75% inode=98%):
[1582885049] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 28.52, 31.77, 30.67
[1582885049] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1582885060] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2532168 kB) free!
[1582885062] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582885088] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (2975496 kB) free!
[1582885192] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582885208] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408892 kB) free!
[1582885519] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1582885614] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582885617] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 135904 MB (60% inode=99%):
[1582885655] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582885659] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1582885661] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.1% (29641136 kB) free.
[1582885738] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38752 MB (75% inode=98%):
[1582885780] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1582885808] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14862228 kB) free!
[1582885989] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1582886109] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1582886209] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9166548 kB) free!
[1582886229] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582886260] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.4% (4450960 kB) free.
[1582886260] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 16.94, 23.52, 28.59
[1582886285] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(63) < 2020/02/28 16:00:55 [error] 12248#0: *61705544 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582886329] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8857356 kB) free!
[1582886370] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2166 processes
[1582886405] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582886408] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29692356 kB) free.
[1582886428] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(96) < 203.212.70.170 [28/Feb/2020:16:02:54 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1582886450] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (8523360 kB) free!
[1582886549] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582886860] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (2790408 kB) free!
[1582886980] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3288912 kB) free.
[1582887008] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5693472 kB) free!
[1582887129] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1904744 kB) free!
[1582887249] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13651892 kB) free!
[1582887580] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (1428400 kB) free!
[1582887700] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.9% (2099548 kB) free!
[1582887820] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2695268 kB) free!
[1582887849] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (11535248 kB) free!
[1582888250] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4759912 kB) free!
[1582888420] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1195048 kB) free!
[1582888449] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27920488 kB) free.
[1582889020] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3288852 kB) free.
[1582889049] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4421000 kB) free!
[1582889169] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (512164 kB) free!
[1582889289] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32199200 kB) free.
[1582889621] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3113916 kB) free!
[1582889741] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3121572 kB) free!
[1582889817] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1582889861] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3117240 kB) free!
[1582889889] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7216320 kB) free!
[1582889937] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1582890009] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3342936 kB) free!
[1582890129] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392572 kB) free!
[1582890729] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14183356 kB) free!
[1582891061] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.8% (4349116 kB) free.
[1582891329] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33717476 kB) free.
[1582891731] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (3524348 kB) free!
[1582891806] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/28 17:37:03 [error] 12254#0: *61758669 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1582891852] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.7% (3305916 kB) free!
[1582891926] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582891929] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8484260 kB) free!
[1582891972] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3037852 kB) free!
[1582892049] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4521940 kB) free!
[1582892170] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (822080 kB) free!
[1582892572] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1896120 kB) free!
[1582892770] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14767768 kB) free!
[1582892861] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2602284 kB) free!
[1582892981] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3494424 kB) free.
[1582893172] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582893172] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.3% (15793120 kB) free.
[1582893370] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35543356 kB) free.
[1582893581] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (2702480 kB) free!
[1582893701] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3588508 kB) free.
[1582893925] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582893970] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10639772 kB) free!
[1582894090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6433512 kB) free!
[1582894210] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2479972 kB) free!
[1582894301] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (1877816 kB) free!
[1582894421] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2864036 kB) free!
[1582894542] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3431292 kB) free.
[1582894810] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16557344 kB) free!
[1582895245] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582895337] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1582895410] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34974640 kB) free.
[1582895457] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1582895548] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(543) < 172.105.59.173 [28/Feb/2020:18:42:18 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582895668] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(98) < 203.212.70.170 [28/Feb/2020:18:44:20 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582895741] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2712520 kB) free!
[1582895788] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582895829] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1582895861] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3326920 kB) free.
[1582896010] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9883476 kB) free!
[1582896130] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5593604 kB) free!
[1582896250] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1647712 kB) free!
[1582896850] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13656904 kB) free!
[1582898050] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2233508 kB) free!
[1582898261] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3090972 kB) free!
[1582898381] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3423516 kB) free.
[1582898991] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1582899111] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1582899232] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1582899251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26701500 kB) free.
[1582899581] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.2% (1820244 kB) free!
[1582899701] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.3% (1992444 kB) free!
[1582899821] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (2189544 kB) free!
[1582899832] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582899851] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (455740 kB) free!
[1582899971] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (380648 kB) free!
[1582900091] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403644 kB) free!
[1582900327] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(169) < 2020/02/28 20:01:49 [error] 12255#0: *61829689 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1582900421] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3298272 kB) free.
[1582900432] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582900447] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(5) < 2020/02/28 20:02:20 [error] 12252#0: *61829816 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582900552] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1582900567] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582900672] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1582900691] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18083440 kB) free!
[1582901022] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (1724896 kB) free!
[1582901142] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 36.2% (5881340 kB) free.
[1582901291] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38981396 kB) free.
[1582901891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13668992 kB) free!
[1582902011] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9599340 kB) free!
[1582902131] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5704312 kB) free!
[1582902731] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20235908 kB) free!
[1582903072] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1582903072] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582903331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37840692 kB) free.
[1582903672] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1582903792] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1582903912] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1582903931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14437612 kB) free!
[1582904051] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10075044 kB) free!
[1582904171] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6224024 kB) free!
[1582904572] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582904772] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18418612 kB) free!
[1582905372] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38106560 kB) free.
[1582905972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12108932 kB) free!
[1582906092] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7728364 kB) free!
[1582906212] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3704268 kB) free!
[1582906257] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1582906378] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1582906812] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20469512 kB) free!
[1582907412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38084144 kB) free.
[1582908012] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12588344 kB) free!
[1582908132] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8435148 kB) free!
[1582908252] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4646840 kB) free!
[1582908852] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17335184 kB) free!
[1582909452] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36897064 kB) free.
[1582909912] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582910052] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10445244 kB) free!
[1582910172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6263940 kB) free!
[1582910292] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2592816 kB) free!
[1582910512] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582910632] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1582910752] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1582910892] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16724784 kB) free!
[1582911452] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582911493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31593432 kB) free.
[1582911572] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582911942] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582911952] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582912093] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8894828 kB) free!
[1582912213] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5236980 kB) free!
[1582912333] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37041908 kB) free.
[1582912552] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582912672] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1582912772] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1582912792] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1582912893] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1582912933] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13099960 kB) free!
[1582913053] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9137008 kB) free!
[1582913173] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (5525240 kB) free!
[1582913261] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240464 kB) free!
[1582913382] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219088 kB) free!
[1582913502] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240684 kB) free!
[1582913773] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16314328 kB) free!
[1582913929] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (564696 kB) free!
[1582914049] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404700 kB) free!
[1582914169] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4356472 kB) free!
[1582914373] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28557436 kB) free.