[1579890868] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9759576 kB) free!
[1579890988] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (211992 kB) free!
[1579891108] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223884 kB) free!
[1579891249] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (943860 kB) free!
[1579891369] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (941904 kB) free!
[1579891489] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (943176 kB) free!
[1579893144] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.68, 6.21, 5.11
[1579893264] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.75, 6.12, 5.22
[1579893384] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.38, 6.42, 5.43
[1579893984] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 0.47, 2.06, 3.76
[1579894477] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (3552940 kB) free!
[1579895077] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196328 kB) free!
[1579897385] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579897505] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579898932] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.5% (10428952 kB) free.
[1579899857] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 74.4% (36687844 kB) free.
[1579901487] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(305) < 2020/01/25 03:00:22 [error] 18356#0: *43493329 no live upstreams while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579901607] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(59) < 2020/01/25 03:01:32 [error] 18362#0: *43493581 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579901639] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(240) < 103.67.237.219 [25/Jan/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1579901705] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579901727] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579901760] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579901825] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579901969] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2367) < 203.212.70.170 [25/Jan/2020:03:02:32 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1579902089] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579909133] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188064 kB) free!
[1579909253] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199548 kB) free!
[1579909373] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (237944 kB) free!
[1579909477] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.3% (11171476 kB) free.
[1579909490] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.4% (20675728 kB) free.
[1579913803] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40407520 kB) free.
[1579913841] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.2% (22353000 kB) free.
[1579920239] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 356 processes
[1579920278] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4874260 kB) free!
[1579920359] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1579920398] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4824048 kB) free!
[1579920518] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4790596 kB) free!
[1579921025] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579921145] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1579923227] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.3% (1833272 kB) free!
[1579923347] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.3% (1832740 kB) free!
[1579923467] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (1833348 kB) free!
[1579924118] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1577764 kB) free!
[1579925084] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.4% (21816000 kB) free.
[1579925268] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4793648 kB) free!
[1579925389] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4765696 kB) free!
[1579925509] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.3% (4733996 kB) free!
[1579930309] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1879732 kB) free!
[1579930853] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21793448 kB) free.
[1579933890] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(111) < 103.255.146.164 [25/Jan/2020:12:00:48 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579934010] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(235) < 185.255.8.58 [25/Jan/2020:12:03:30 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1579934049] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 133.51, 179.50, 90.67
[1579934130] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(3) < 92.42.104.166 [25/Jan/2020:12:03:35 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1579934170] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 30.58, 125.23, 81.58
[1579934290] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.14, 88.22, 73.33
[1579934461] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.0% (4174376 kB) free!
[1579934519] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579934566] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1579934581] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199156 kB) free!
[1579934639] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579934686] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579934690] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (4373560 kB) free!
[1579934690] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197796 kB) free!
[1579934701] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1579934701] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192472 kB) free!
[1579934759] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1579934806] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579934810] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (4281564 kB) free!
[1579934810] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188068 kB) free!
[1579934930] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (4202652 kB) free!
[1579934930] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198080 kB) free!
[1579935490] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 10.62, 12.97, 28.85
[1579935530] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2251936 kB) free!
[1579935930] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579936530] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(173) < 103.255.146.164 [25/Jan/2020:12:45:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579936650] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579937850] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(68) < 103.255.146.164 [25/Jan/2020:13:00:27 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579937936] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.3% (21882732 kB) free.
[1579937971] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579937976] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.8% (22153428 kB) free.
[1579938007] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.2% (21863692 kB) free.
[1579938359] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579938479] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579938514] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1579938599] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579938634] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1579938754] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1579939207] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178892 kB) free!
[1579939328] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169580 kB) free!
[1579939355] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1579939448] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172956 kB) free!
[1579940399] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579940727] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2020/01/25 13:54:17 [error] 18367#0: *43637524 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579940848] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579940970] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 203.212.70.170 [25/Jan/2020:13:57:36 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1579941090] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579941155] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1579941248] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1579941248] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 66.0% (16183668 kB) free.
[1579941275] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1579941395] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1579942602] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1579942722] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1579946336] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1579946379] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (3690532 kB) free!
[1579946500] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.2% (2517740 kB) free!
[1579946620] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2400048 kB) free!
[1579946850] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4810288 kB) free!
[1579946945] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579946970] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (4748484 kB) free!
[1579947065] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579947090] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4638648 kB) free!
[1579948722] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1579948842] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1579949057] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9484088 kB) free!
[1579949177] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9370864 kB) free!
[1579949298] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9284104 kB) free!
[1579950406] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9441148 kB) free!
[1579950526] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9441372 kB) free!
[1579950646] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (9440664 kB) free!
[1579951448] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579951890] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2214220 kB) free!
[1579953249] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183320 kB) free!
[1579953253] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579953369] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (205364 kB) free!
[1579953373] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579953489] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199720 kB) free!
[1579953493] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579953599] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579953646] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (253548 kB) free!
[1579953690] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (4037476 kB) free!
[1579953719] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579953839] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579954439] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1579954538] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.2% (20049556 kB) free.
[1579954573] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 55.8% (27486016 kB) free.
[1579954600] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.1% (21593920 kB) free.
[1579954693] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579954813] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579954933] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579955039] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579955208] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.5% (3299680 kB) free!
[1579955298] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4792176 kB) free!
[1579955328] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195060 kB) free!
[1579955448] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (210000 kB) free!
[1579957265] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579957333] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579957385] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579957553] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (4306212 kB) free!
[1579957673] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4236588 kB) free!
[1579957793] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4165860 kB) free!
[1579958039] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1579958039] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1579958490] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1579958490] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2419288 kB) free!
[1579958639] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579959239] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1579960193] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2305516 kB) free!
[1579960582] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1579960582] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.5% (21347864 kB) free.
[1579960619] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 37.0% (18256196 kB) free.
[1579960655] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 78.2% (19172400 kB) free.
[1579960679] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1579960679] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.5% (20946772 kB) free.
[1579960933] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579961053] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579961173] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579961289] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.1% (2963940 kB) free!
[1579961409] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199536 kB) free!
[1579961529] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199812 kB) free!
[1579961773] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1579961773] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579962839] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1579962959] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579963559] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579963679] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579963799] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1579964999] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579965119] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579965239] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579965839] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 370 processes
[1579966047] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(64) < 2020/01/25 20:52:54 [error] 18367#0: *43769286 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579966167] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579966228] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9747520 kB) free!
[1579966348] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9728300 kB) free!
[1579966439] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579966468] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (9708260 kB) free!
[1579966667] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (259964 kB) free!
[1579967039] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1579967187] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (246160 kB) free!
[1579967307] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (233520 kB) free!
[1579967427] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (220492 kB) free!
[1579969657] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1579970299] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.6% (21127444 kB) free.
[1579970333] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.4% (21176964 kB) free.
[1579970357] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.2% (14341968 kB) free.
[1579970639] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579970760] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1579970857] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4880320 kB) free!
[1579970939] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4494020 kB) free!
[1579970977] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4864364 kB) free!
[1579971059] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4480492 kB) free!
[1579971097] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4852940 kB) free!
[1579971180] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170996 kB) free!
[1579971359] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579971479] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1579972079] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1579972199] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1579973173] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579973399] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579973519] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1579973639] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579974239] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579974360] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579974480] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579976048] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239632 kB) free!
[1579976168] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234384 kB) free!
[1579976288] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230472 kB) free!