[1553020264] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7952860 kB) free!
[1553020538] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18914 MB (9% inode=99%):
[1553020821] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553020864] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12239792 kB) free.
[1553021138] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23002 MB (11% inode=99%):
[1553021421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553021464] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4276824 kB) free!
[1553021541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553021584] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12260392 kB) free.
[1553022141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553022184] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6697392 kB) free!
[1553022261] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553022305] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5395268 kB) free!
[1553022382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553022425] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12243716 kB) free.
[1553022938] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21475 MB (10% inode=99%):
[1553023025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7371036 kB) free!
[1553023145] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6055640 kB) free!
[1553023265] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4380448 kB) free!
[1553023539] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23028 MB (11% inode=99%):
[1553023582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553023865] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.0% (8358928 kB) free!
[1553024182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553024302] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553024465] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12235716 kB) free.
[1553024808] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167480 kB) free!
[1553024902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1553024928] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 31.2% (5067472 kB) free.
[1553025022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553025065] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5113576 kB) free!
[1553025142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1553025185] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12255900 kB) free.
[1553025339] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21031 MB (10% inode=99%):
[1553025785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7170104 kB) free!
[1553025905] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5882848 kB) free!
[1553025939] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23066 MB (11% inode=99%):
[1553026026] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4144184 kB) free!
[1553026342] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553026626] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7950100 kB) free!
[1553026942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553027062] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553027226] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12252200 kB) free.
[1553027662] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553027739] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19514 MB (9% inode=99%):
[1553027782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553027826] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (4537148 kB) free!
[1553027902] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553027946] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12238376 kB) free.
[1553028339] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23112 MB (11% inode=99%):
[1553028502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553028546] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6819264 kB) free!
[1553028623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553028666] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5506888 kB) free!
[1553028743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553028786] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12253952 kB) free.
[1553029386] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7615304 kB) free!
[1553029506] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6325908 kB) free!
[1553029626] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4898240 kB) free!
[1553029943] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553030139] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21091 MB (10% inode=99%):
[1553030226] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.1% (8449292 kB) free!
[1553030543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553030663] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553030739] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23113 MB (11% inode=99%):
[1553030827] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12251568 kB) free.
[1553031263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553031383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553031427] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5093680 kB) free!
[1553031503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553031547] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12235904 kB) free.
[1553032147] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7181648 kB) free!
[1553032267] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5919416 kB) free!
[1553032387] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4230428 kB) free!
[1553032539] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19511 MB (9% inode=99%):
[1553032703] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553032987] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.2% (7984192 kB) free!
[1553033139] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23104 MB (11% inode=99%):
[1553033303] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553033423] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553033587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12233568 kB) free.
[1553034023] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553034143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1553034187] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4941692 kB) free!
[1553034263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553034307] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12236192 kB) free.
[1553034907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7203912 kB) free!
[1553034939] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18207 MB (8% inode=99%):
[1553035028] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5923964 kB) free!
[1553035148] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4162356 kB) free!
[1553035464] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553035748] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7952980 kB) free!
[1553036064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553036139] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23142 MB (11% inode=99%):
[1553036184] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553036348] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12238980 kB) free.
[1553036784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553036904] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553036948] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4605108 kB) free!
[1553037024] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553037068] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12219420 kB) free.
[1553037624] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553037668] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6931904 kB) free!
[1553037744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553037788] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.4% (5636308 kB) free!
[1553037865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553037908] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3747036 kB) free!
[1553037939] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22726 MB (10% inode=99%):
[1553038508] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.6% (7707096 kB) free!
[1553038539] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23140 MB (11% inode=99%):
[1553039065] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553039108] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12208556 kB) free.
[1553039665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553039708] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4121632 kB) free!
[1553039739] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22146 MB (10% inode=99%):
[1553039785] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553039829] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12228748 kB) free.
[1553040339] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23116 MB (11% inode=99%):
[1553040385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553040428] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6569032 kB) free!
[1553040505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553040548] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5289216 kB) free!
[1553040625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553040668] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12210668 kB) free.
[1553041268] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7296604 kB) free!
[1553041388] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (6010892 kB) free!
[1553041508] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4344148 kB) free!
[1553041539] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21115 MB (10% inode=99%):
[1553041825] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553042108] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.2% (7982108 kB) free!
[1553042139] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27233 MB (13% inode=99%):
[1553042425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553042545] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553042708] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12206624 kB) free.
[1553043145] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553043265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553043308] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4644600 kB) free!
[1553043339] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19075 MB (9% inode=99%):
[1553043385] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553043428] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12213916 kB) free.
[1553043939] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23139 MB (11% inode=99%):
[1553043985] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553044029] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6945044 kB) free!
[1553044106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553044149] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.5% (5664520 kB) free!
[1553044226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553044269] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3791584 kB) free!
[1553044869] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7714208 kB) free!
[1553045426] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553045469] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12209336 kB) free.
[1553045739] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21621 MB (10% inode=99%):
[1553046026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553046069] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4136800 kB) free!
[1553046146] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553046189] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12196788 kB) free.
[1553046339] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23161 MB (11% inode=99%):
[1553046746] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553046789] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6535408 kB) free!
[1553046866] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553046909] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5207396 kB) free!
[1553046986] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553047029] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12213020 kB) free.
[1553047629] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7243828 kB) free!
[1553047749] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5963328 kB) free!
[1553047870] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4287736 kB) free!
[1553048139] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20059 MB (9% inode=99%):
[1553048186] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553048470] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.6% (7705264 kB) free!
[1553048740] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23104 MB (11% inode=99%):
[1553048786] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553048906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553049026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553049626] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553049670] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12184324 kB) free.
[1553050226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553050270] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5328508 kB) free!
[1553050346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553050390] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12203300 kB) free.
[1553050467] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553050540] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19521 MB (9% inode=99%):
[1553050990] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7406796 kB) free!
[1553051066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553051110] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6110592 kB) free!
[1553051140] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23109 MB (11% inode=99%):
[1553051174] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2332272 kB) free!
[1553051186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553051230] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4559636 kB) free!
[1553051306] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553051830] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.6% (8183160 kB) free!
[1553051906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553052026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553052146] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553052430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12197988 kB) free.
[1553052940] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19047 MB (9% inode=99%):
[1553053031] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4680896 kB) free!
[1553053151] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12180664 kB) free.
[1553053751] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3869848 kB) free!
[1553053871] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11679556 kB) free.
[1553054140] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23048 MB (11% inode=99%):
[1553055071] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3190832 kB) free!
[1553055146] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553055191] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11927844 kB) free.
[1553055340] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22693 MB (10% inode=99%):
[1553055746] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553055791] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2753404 kB) free!
[1553055866] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553055911] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11204512 kB) free.
[1553055940] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24958 MB (11% inode=99%):
[1553055986] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553057019] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1553057111] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3082072 kB) free!
[1553057139] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1553057140] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19524 MB (9% inode=99%):
[1553057187] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553057232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11572800 kB) free.
[1553057259] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1553057672] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 34.0% (8333620 kB) free.
[1553057774] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 26.1% (6404892 kB) free.
[1553057787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553057877] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 24.1% (5913340 kB) free.
[1553057907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553058027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553058340] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29201 MB (13% inode=99%):
[1553058431] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3484880 kB) free!
[1553058551] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11946144 kB) free.
[1553058940] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18961 MB (9% inode=99%):
[1553059151] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (2237988 kB) free!
[1553059227] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553059271] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11274268 kB) free.
[1553059827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553059898] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (736272 kB) free!
[1553059948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553060068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553060140] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26023 MB (12% inode=99%):
[1553060471] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3514268 kB) free!
[1553060591] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11846908 kB) free.
[1553061191] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2856712 kB) free!
[1553061311] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11256560 kB) free.
[1553061340] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20890 MB (10% inode=99%):
[1553061374] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1764380 kB) free!
[1553061494] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (1682000 kB) free!
[1553061614] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1599856 kB) free!
[1553061941] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24940 MB (11% inode=99%):
[1553062511] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3623908 kB) free!
[1553062631] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11811944 kB) free.
[1553062686] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.5% (13900248 kB) free.
[1553062853] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 105488 MB (46% inode=99%):
[1553063141] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20798 MB (9% inode=99%):
[1553063231] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2737624 kB) free!
[1553063351] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11135736 kB) free.
[1553063741] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24849 MB (11% inode=99%):
[1553064551] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3010840 kB) free!
[1553064671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11309280 kB) free.
[1553064942] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 16636 MB (7% inode=99%):
[1553065271] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1553065697] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 10.8% (796556 kB) free!
[1553065832] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7152 MB (13% inode=98%):
[1553065871] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4081700 kB) free!
[1553065991] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12407784 kB) free.
[1553066297] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (731836 kB) free!
[1553066592] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3520364 kB) free!
[1553066611] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2522120 kB) free!
[1553066668] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553066712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12169068 kB) free.
[1553067211] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2065700 kB) free!
[1553067268] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553067312] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3924556 kB) free!
[1553067342] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26582 MB (12% inode=99%):
[1553067388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553067432] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11352800 kB) free.
[1553067508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553067528] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151124 kB) free!
[1553067648] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (189832 kB) free!
[1553067768] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167700 kB) free!
[1553068542] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22741 MB (10% inode=99%):
[1553068632] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2829352 kB) free!
[1553068709] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553068752] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11775856 kB) free.
[1553069309] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553069429] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553069549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553069952] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3529516 kB) free!
[1553070072] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12183152 kB) free.
[1553070673] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3600856 kB) free!
[1553070749] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553070793] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12166952 kB) free.
[1553070811] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (3676984 kB) free!
[1553071237] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 336 processes
[1553071349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553071393] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2871508 kB) free!
[1553071469] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553071513] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11529532 kB) free.
[1553072070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553072142] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 323 processes
[1553072190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553072310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553072611] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1553072611] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2469276 kB) free!
[1553072713] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3235436 kB) free!
[1553072833] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11862744 kB) free.
[1553072877] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4669532 kB) free!
[1553072997] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4660696 kB) free!
[1553073117] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4656536 kB) free!
[1553073433] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3444840 kB) free!
[1553073553] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11137964 kB) free.
[1553074110] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553074317] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.9% (5116796 kB) free.
[1553074411] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.6% (3117820 kB) free!
[1553074477] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.3% (14682104 kB) free.
[1553074533] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.0% (14797576 kB) free.
[1553074542] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 324 processes
[1553074645] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.0% (6990208 kB) free.
[1553074653] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1553074710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553074753] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3058172 kB) free!
[1553074773] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1553074830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553074873] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11437528 kB) free.
[1553074881] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 95.2% (6998468 kB) free.
[1553074893] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1553074918] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4750124 kB) free!
[1553074958] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.6% (5707000 kB) free!
[1553075038] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4681920 kB) free!
[1553075096] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.5% (21839336 kB) free.
[1553075158] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4613724 kB) free!
[1553075174] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.3% (22873156 kB) free.
[1553075236] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1553075236] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.6% (22460972 kB) free.
[1553075431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553075437] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 337 processes
[1553075551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553075653] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3554372 kB) free!
[1553075671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553075815] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 93.2% (6853288 kB) free.
[1553076073] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3187768 kB) free!
[1553076193] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11804740 kB) free.
[1553076486] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (2951200 kB) free!
[1553076606] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (2856596 kB) free!
[1553076726] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.9% (2749716 kB) free!
[1553076793] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2419244 kB) free!
[1553076868] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.1% (6445784 kB) free!
[1553076901] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185872 kB) free!
[1553076913] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11130696 kB) free.
[1553076975] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21854504 kB) free.
[1553076995] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1553076995] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4671984 kB) free!
[1553077181] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 94.8% (6971340 kB) free.
[1553077324] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44398 MB (21% inode=99%):
[1553077436] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 94.7% (6963024 kB) free.
[1553077445] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.6% (7029524 kB) free.
[1553077926] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1504916 kB) free!
[1553078134] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4671760 kB) free!
[1553078254] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12565560 kB) free.
[1553078372] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;1;WARNING - 19.6% (1444840 kB) free!
[1553078492] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;2;WARNING - 19.5% (1436424 kB) free!
[1553078612] SERVICE ALERT: sig5-vm1;Memory;WARNING;HARD;3;WARNING - 19.4% (1427736 kB) free!
[1553079272] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4826176 kB) free!
[1553079392] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4748696 kB) free!
[1553079454] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4783936 kB) free!
[1553079513] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4668100 kB) free!
[1553079574] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12798068 kB) free.
[1553079872] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.4% (15022216 kB) free.
[1553079935] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.7% (23199672 kB) free.
[1553080034] SERVICE ALERT: sig5-vm1;Memory;OK;HARD;1;OK - 94.7% (6967296 kB) free.
[1553080174] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3993840 kB) free!
[1553080293] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1553080294] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12425044 kB) free.
[1553080893] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1553081013] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1553081051] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1553081133] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1553081494] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4425964 kB) free!
[1553081614] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12669212 kB) free.
[1553082124] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40023 MB (19% inode=99%):
[1553082214] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3481436 kB) free!
[1553082244] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40021 MB (19% inode=99%):
[1553082271] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553082334] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12115980 kB) free.
[1553082346] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (185148 kB) free!
[1553082364] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 40026 MB (19% inode=99%):
[1553082381] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (1557560 kB) free!
[1553082467] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178380 kB) free!
[1553082501] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (1517696 kB) free!
[1553082587] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151292 kB) free!
[1553082621] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1482668 kB) free!
[1553082871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553082964] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51644 MB (24% inode=99%):
[1553082991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553083111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553083534] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (4011468 kB) free!
[1553083654] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12485348 kB) free.
[1553084012] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 88.9% (14445844 kB) free.
[1553084087] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21796664 kB) free.
[1553084254] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3041192 kB) free!
[1553084312] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553084374] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11686728 kB) free.
[1553084783] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175696 kB) free!
[1553084903] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174332 kB) free!
[1553084912] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553085023] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178260 kB) free!
[1553085032] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553085152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553085575] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3359268 kB) free!
[1553085695] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12015852 kB) free.
[1553085913] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 281 processes
[1553085964] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 35756 MB (17% inode=99%):
[1553086033] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 276 processes
[1553086041] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1553086084] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 35756 MB (17% inode=99%):
[1553086204] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 35756 MB (17% inode=99%):
[1553086445] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (115292 kB) free!
[1553086459] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.1% (22576164 kB) free.
[1553086533] SERVICE FLAPPING ALERT: db1;Current Users;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1553086533] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1553086565] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (120740 kB) free!
[1553086685] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (137472 kB) free!
[1553086793] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;1;WARNING - 19.6% (1442940 kB) free!
[1553086895] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3597280 kB) free!
[1553086913] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;2;WARNING - 19.4% (1427464 kB) free!
[1553086952] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553087015] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12077016 kB) free.
[1553087033] SERVICE ALERT: sig6-vm1;Memory;WARNING;HARD;3;WARNING - 19.2% (1413932 kB) free!
[1553087133] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1553087253] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1553087373] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1553087552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553087615] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2658692 kB) free!
[1553087672] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553087735] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11241272 kB) free.
[1553087887] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.6% (7028476 kB) free.
[1553088043] SERVICE ALERT: db1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553088166] SERVICE ALERT: db1;CPU Load;OK;HARD;1;OK - load average: 2.06, 2.02, 1.83
[1553088273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553088393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553088463] SERVICE ALERT: db1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553088463] SERVICE ALERT: db1;Mariadb status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553088506] SERVICE ALERT: db1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553088513] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553088584] SERVICE ALERT: db1;Mariadb status;OK;HARD;1;PROCS OK: 1 process with command name 'mysqld'
[1553088584] SERVICE ALERT: db1;Sensors Status;OK;HARD;1;SENSORS OK
[1553088635] SERVICE ALERT: db1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 46438 MB (90% inode=99%):
[1553088668] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553088773] SERVICE ALERT: db1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553088798] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1553088903] SERVICE ALERT: db1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1553088928] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553088935] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3716436 kB) free!
[1553089033] SERVICE ALERT: db1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089055] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11915604 kB) free.
[1553089167] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089183] SERVICE ALERT: db1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089193] SERVICE ALERT: db1;Mariadb status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089193] SERVICE ALERT: db1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089214] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089236] SERVICE ALERT: db1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089263] SERVICE ALERT: db1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089297] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089323] SERVICE ALERT: db1;Mariadb status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089323] SERVICE ALERT: db1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089359] SERVICE ALERT: db1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 46438 MB (90% inode=99%):
[1553089388] SERVICE ALERT: db1;Total Processes;OK;HARD;1;PROCS OK: 339 processes
[1553089427] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089453] SERVICE ALERT: db1;Mariadb status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089453] SERVICE ALERT: db1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553089655] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3201104 kB) free!
[1553089696] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.87, 5.88, 5.28
[1553089713] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553089776] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11237440 kB) free.
[1553089816] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.62, 6.09, 5.42
[1553089937] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.85, 5.81, 5.41
[1553090029] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1553090055] SERVICE ALERT: db1;Mariadb status;OK;HARD;1;PROCS OK: 1 process with command name 'mysqld'
[1553090055] SERVICE ALERT: db1;Sensors Status;OK;HARD;1;SENSORS OK
[1553090134] SERVICE ALERT: db1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1553090313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553090433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553090553] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553090683] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (3152628 kB) free!
[1553090803] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (3023640 kB) free!
[1553090923] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (3023032 kB) free!
[1553090975] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3136532 kB) free!
[1553091095] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11625728 kB) free.
[1553091166] SERVICE ALERT: db1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553091193] SERVICE ALERT: db1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553091239] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553091289] SERVICE ALERT: db1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 46441 MB (90% inode=99%):
[1553091315] SERVICE ALERT: db1;Total Processes;OK;HARD;1;PROCS OK: 347 processes
[1553091369] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1553091438] SERVICE ALERT: db1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.71, 14.32, 13.90
[1553091499] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553091695] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3829272 kB) free!
[1553091737] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.41, 4.54, 4.75
[1553091815] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12048892 kB) free.
[1553091923] SERVICE ALERT: db1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553091944] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1553092048] SERVICE ALERT: db1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553092053] SERVICE ALERT: db1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1553092074] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1553092183] SERVICE ALERT: db1;Total Processes;OK;HARD;1;PROCS OK: 345 processes
[1553092204] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1553092337] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.62, 5.48, 5.01
[1553092415] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3850700 kB) free!
[1553092457] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.15, 5.82, 5.20
[1553092535] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11985464 kB) free.
[1553092577] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.34, 4.96, 4.97
[1553092785] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1553092805] SERVICE FLAPPING ALERT: db1;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1553092805] SERVICE ALERT: db1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1553092815] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 37083 MB (17% inode=99%):
[1553093135] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3369672 kB) free!
[1553093249] SERVICE FLAPPING ALERT: db1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1553093249] SERVICE ALERT: db1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.32, 8.26, 10.93
[1553093255] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11651660 kB) free.
[1553093300] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1553093300] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1553093691] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1553093853] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41993 MB (18% inode=99%):
[1553093973] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41993 MB (18% inode=99%):
[1553094093] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41993 MB (18% inode=99%):
[1553094263] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (4318336 kB) free!
[1553094383] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (750120 kB) free!
[1553094455] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4267256 kB) free!
[1553094503] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174960 kB) free!
[1553094576] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11994272 kB) free.
[1553094615] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47381 MB (22% inode=99%):
[1553095050] SERVICE ALERT: db1;CPU Load;OK;HARD;1;OK - load average: 2.81, 2.33, 3.76
[1553095175] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3488676 kB) free!
[1553095295] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11487916 kB) free.
[1553095491] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (333996 kB) free!
[1553095611] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181432 kB) free!
[1553095731] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187116 kB) free!
[1553096495] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4059676 kB) free!
[1553096553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553096615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11992116 kB) free.
[1553097153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553097215] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6506288 kB) free!
[1553097273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553097335] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5168984 kB) free!
[1553097393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553097455] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3204192 kB) free!
[1553097459] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1553098055] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7340496 kB) free!
[1553098593] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553098655] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11995116 kB) free.
[1553099193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553099256] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3861872 kB) free!
[1553099314] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1553099376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11992520 kB) free.
[1553099913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553099976] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6397844 kB) free!
[1553100033] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553100096] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.3% (5068036 kB) free!
[1553100153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553100216] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11914352 kB) free.
[1553100816] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7222260 kB) free!
[1553100936] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5885996 kB) free!
[1553101056] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1553101056] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4403732 kB) free!
[1553101185] SERVICE FLAPPING ALERT: db1;Current Users;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1553101353] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553101656] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.2% (7979140 kB) free!
[1553101953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553102073] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553102256] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11989716 kB) free.
[1553102673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553102794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553102856] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4757188 kB) free!
[1553102914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553102976] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11981744 kB) free.
[1553103415] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.8% (8781568 kB) free!
[1553103501] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1553103577] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6948556 kB) free!
[1553103697] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.5% (5647424 kB) free!
[1553103817] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4036524 kB) free!
[1553104114] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553104205] SERVICE FLAPPING ALERT: db1;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1553104417] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7838988 kB) free!
[1553104616] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1301640 kB) free!
[1553104714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553104815] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1553104834] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553105017] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11982084 kB) free.
[1553105250] SERVICE FLAPPING ALERT: db1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1553105434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553105554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553105617] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4591268 kB) free!
[1553105675] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553105737] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11984352 kB) free.
[1553106337] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6837648 kB) free!
[1553106457] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.2% (5527920 kB) free!
[1553106577] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3769600 kB) free!