[1548268764] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2091368 kB) free!
[1548269733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1548269964] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5072392 kB) free!
[1548272293] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 351 processes
[1548295164] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4267832 kB) free!
[1548295615] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.1% (5470292 kB) free!
[1548295696] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 7 warn out of 335 processes
[1548296265] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1548296385] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548296505] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548297706] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548298164] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5060220 kB) free!
[1548298764] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3162364 kB) free!
[1548299364] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1548299364] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5042704 kB) free!
[1548304646] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1548304752] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.9% (1926868 kB) free!
[1548304766] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1548304873] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3361456 kB) free.
[1548304886] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1548305472] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3186944 kB) free!
[1548305592] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3154068 kB) free!
[1548305712] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3123976 kB) free!
[1548308712] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 64.5% (10481436 kB) free.
[1548308964] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3264572 kB) free!
[1548309333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1548309415] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4346376 kB) free!
[1548309453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1548309573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1548310306] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548310412] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.70, 10.05, 10.02
[1548310426] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548310532] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.10, 10.45, 10.17
[1548310546] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548310652] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.34, 10.38, 10.18
[1548311146] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548311252] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.72, 9.78, 9.97
[1548311471] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 34963 MB (15% inode=99%):
[1548311572] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.7% (2551608 kB) free!
[1548312168] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 9727 MB (19% inode=98%):
[1548312288] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 9727 MB (19% inode=98%):
[1548312347] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548312408] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 11752 MB (22% inode=99%):
[1548312467] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548312587] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548312772] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1571292 kB) free!
[1548313052] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.41, 10.36, 10.02
[1548313172] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.78, 10.17, 10.00
[1548313173] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1548313187] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1548313187] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548313772] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.29, 10.18, 10.12
[1548313892] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.61, 10.37, 10.20
[1548314012] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.08, 10.36, 10.22
[1548314373] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548314494] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1548314808] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10375 MB (20% inode=98%):
[1548314815] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.9% (6360144 kB) free!
[1548314928] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 10250 MB (20% inode=98%):
[1548314964] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1548315048] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 12219 MB (23% inode=99%):
[1548315212] SERVICE FLAPPING ALERT: qservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1548315212] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 10.13, 9.84, 9.95
[1548315587] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548315648] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10145 MB (19% inode=98%):
[1548315693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548315707] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548315768] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 12178 MB (23% inode=99%):
[1548315812] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.85, 10.15, 10.11
[1548315813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548315828] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548315933] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.22, 10.19, 10.13
[1548315933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548316053] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.34, 10.19, 10.14
[1548316368] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10114 MB (19% inode=98%):
[1548316428] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548316488] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 10114 MB (19% inode=98%):
[1548316608] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 11890 MB (23% inode=99%):
[1548317028] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548317148] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548317268] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548317808] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 9795 MB (19% inode=98%):
[1548317815] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4522708 kB) free!
[1548317868] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548317928] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 9795 MB (19% inode=98%):
[1548318048] SERVICE ALERT: db2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 9795 MB (19% inode=98%):
[1548318415] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1548318415] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.1% (4968140 kB) free!
[1548319015] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4671432 kB) free!
[1548319068] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548319188] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548319248] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 11218 MB (21% inode=99%):
[1548319308] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548319653] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.41, 9.80, 10.00
[1548320449] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 9161 MB (17% inode=98%):
[1548320508] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548320569] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 11201 MB (21% inode=99%):
[1548320733] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1548320733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1548321108] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548321228] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548321769] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 8991 MB (17% inode=98%):
[1548321828] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548321889] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 9055 MB (17% inode=98%):
[1548321948] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548322009] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 10980 MB (21% inode=99%):
[1548322068] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548322533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548322653] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548322773] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548322871] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24641 MB (10% inode=99%):
[1548323268] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548323853] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.66, 10.49, 10.15
[1548323869] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548323973] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.26, 10.45, 10.18
[1548323989] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548324093] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.45, 10.46, 10.21
[1548324109] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548325609] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10354 MB (20% inode=98%):
[1548325729] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 10353 MB (20% inode=98%):
[1548325764] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4780972 kB) free!
[1548325849] SERVICE ALERT: db2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 10481 MB (20% inode=98%):
[1548325884] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4727884 kB) free!
[1548325909] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548326004] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (4675340 kB) free!
[1548326293] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 352 processes
[1548326604] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 38.1% (9334204 kB) free.
[1548327049] SERVICE FLAPPING ALERT: db2;HDD_Root status;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1548327049] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 11050 MB (21% inode=99%):
[1548327109] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548327230] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548328249] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10537 MB (20% inode=98%):
[1548328369] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 11043 MB (21% inode=99%):
[1548329029] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548329149] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548330169] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 9902 MB (19% inode=98%):
[1548330290] SERVICE ALERT: db2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 9902 MB (19% inode=98%):
[1548330410] SERVICE ALERT: db2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 10271 MB (20% inode=98%):
[1548330415] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548330949] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548331069] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548331189] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548331893] SERVICE FLAPPING ALERT: qservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1548332389] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548332509] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548332629] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548332973] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1548333093] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.42, 9.58, 9.90
[1548333204] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4477376 kB) free!
[1548333229] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548333324] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4476788 kB) free!
[1548333444] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4476348 kB) free!
[1548333573] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1548333829] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548333949] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548334069] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548334173] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548334294] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548334309] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1548334512] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548334893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1548334909] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1548335013] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1548335029] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 309 processes
[1548335133] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548335149] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 309 processes
[1548335334] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.89, 7.89, 5.89
[1548335454] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.60, 8.54, 6.37
[1548335574] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.15, 8.65, 6.67
[1548335869] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548335871] SERVICE ALERT: db1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 5644 MB (11% inode=98%):
[1548335989] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548336109] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548336709] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548336829] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548336933] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1548338844] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 25.7% (6303448 kB) free.
[1548339410] SERVICE FLAPPING ALERT: db2;HDD_Root status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1548340293] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.27, 10.02, 10.05
[1548340413] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.94, 10.01, 10.04
[1548340512] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168816 kB) free!
[1548340533] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.59, 9.93, 10.01
[1548340632] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180152 kB) free!
[1548340752] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168996 kB) free!
[1548341133] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 10.21, 9.92, 9.96
[1548342933] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.44, 10.29, 10.09
[1548343053] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.36, 10.33, 10.13
[1548343173] SERVICE FLAPPING ALERT: qservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1548343173] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.01, 10.27, 10.14
[1548343773] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.52, 9.68, 9.91
[1548345229] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548345573] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.28, 10.10, 10.01
[1548345693] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.77, 9.98, 9.98
[1548346293] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.60, 10.26, 10.08
[1548346413] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.44, 10.00, 10.01
[1548346533] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.87, 10.00, 10.01
[1548349093] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 351 processes
[1548351333] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.59, 9.75, 9.94
[1548352534] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548352615] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.1% (4995984 kB) free!
[1548352654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548352774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548353374] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548353974] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548354094] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes