[1581791423] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 207913 MB (19% inode=99%):
[1581791540] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12182004 kB) free!
[1581791543] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 207913 MB (19% inode=99%):
[1581791660] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8624060 kB) free!
[1581791663] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 207913 MB (19% inode=99%):
[1581791780] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35567184 kB) free.
[1581791989] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.0% (29089160 kB) free.
[1581792380] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12440484 kB) free!
[1581792500] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8886056 kB) free!
[1581792620] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35516136 kB) free.
[1581793220] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12716540 kB) free!
[1581793340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9144692 kB) free!
[1581793460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35507320 kB) free.
[1581793475] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168516 kB) free!
[1581793535] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180372 kB) free!
[1581793595] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (189364 kB) free!
[1581793655] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 66.9% (16393200 kB) free.
[1581793923] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(150) < 92.42.104.166 [16/Feb/2020:00:36:01 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581794043] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581794061] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12356944 kB) free!
[1581794181] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8788848 kB) free!
[1581794301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34632384 kB) free.
[1581794901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10955596 kB) free!
[1581795021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35483004 kB) free.
[1581795621] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13573408 kB) free!
[1581795741] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10042776 kB) free!
[1581795861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35432000 kB) free.
[1581796067] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4607348 kB) free!
[1581796461] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13019236 kB) free!
[1581796581] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9597164 kB) free!
[1581796667] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192308 kB) free!
[1581796701] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34567296 kB) free.
[1581797301] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12330120 kB) free!
[1581797421] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9036660 kB) free!
[1581797541] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32601096 kB) free.
[1581798141] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11436812 kB) free!
[1581798261] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35440780 kB) free.
[1581798862] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13312784 kB) free!
[1581798982] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9809412 kB) free!
[1581799102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31162324 kB) free.
[1581799702] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4298256 kB) free!
[1581799822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34908992 kB) free.
[1581800077] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230292 kB) free!
[1581800197] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (246080 kB) free!
[1581800305] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (1745312 kB) free!
[1581800318] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218840 kB) free!
[1581800422] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7744388 kB) free!
[1581800425] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 80.0% (39419808 kB) free.
[1581800542] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4156664 kB) free!
[1581800662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34877028 kB) free.
[1581801262] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8147780 kB) free!
[1581801382] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3502748 kB) free!
[1581801502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32688432 kB) free.
[1581801518] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 80.8% (39841928 kB) free.
[1581801677] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.6% (9959528 kB) free.
[1581802102] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8383952 kB) free!
[1581802222] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34925936 kB) free.
[1581802238] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(228) < 182.18.144.246 [16/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581802359] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581802444] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2108) < 203.212.70.170 [16/Feb/2020:03:03:32 +0530] TCP 502 0 0 180.000 "10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1581802564] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581802791] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(290) < 2020/02/16 03:01:32 [error] 19841#0: *54539948 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581802822] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12449724 kB) free!
[1581802911] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581802942] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9053200 kB) free!
[1581803063] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35672880 kB) free.
[1581803662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15555288 kB) free!
[1581803782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37012532 kB) free.
[1581804382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16708864 kB) free!
[1581804502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37027360 kB) free.
[1581805102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18432104 kB) free!
[1581805222] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581805222] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37040560 kB) free.
[1581805822] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19828456 kB) free!
[1581805911] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2020/02/16 04:01:01 [error] 19830#0: *54553075 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581805942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (16610080 kB) free!
[1581806031] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581806062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36787480 kB) free.
[1581807262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22026784 kB) free!
[1581807382] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39212108 kB) free.
[1581807982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21665024 kB) free!
[1581808102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39210684 kB) free.
[1581808703] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20888992 kB) free!
[1581808822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37969336 kB) free.
[1581809478] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (210520 kB) free!
[1581809599] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181176 kB) free!
[1581809719] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165808 kB) free!
[1581810022] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22968456 kB) free!
[1581810142] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19732624 kB) free!
[1581810262] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37598804 kB) free.
[1581810467] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 62.5% (15418672 kB) free.
[1581810620] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.5% (18255520 kB) free.
[1581811462] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21932876 kB) free!
[1581811582] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39239468 kB) free.
[1581812782] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22836964 kB) free!
[1581812902] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39246200 kB) free.
[1581813502] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21259164 kB) free!
[1581813622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38487628 kB) free.
[1581813806] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 409 processes
[1581813926] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1581814101] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 32.4% (42726236 kB) free.
[1581814162] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39736364 kB) free.
[1581814207] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.3% (21878912 kB) free.
[1581814222] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19677772 kB) free!
[1581814342] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37554316 kB) free.
[1581814463] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 236920 MB (22% inode=99%):
[1581815267] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.3% (3037516 kB) free!
[1581815387] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (3028756 kB) free!
[1581815507] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (3019804 kB) free!
[1581815542] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18640744 kB) free!
[1581815662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38926092 kB) free.
[1581816262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20108872 kB) free!
[1581816382] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39234568 kB) free.
[1581816982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21341136 kB) free!
[1581817102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38799692 kB) free.
[1581817702] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21033800 kB) free!
[1581817822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37555104 kB) free.
[1581818507] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198984 kB) free!
[1581819022] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20916384 kB) free!
[1581819142] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38666052 kB) free.
[1581819742] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20730140 kB) free!
[1581819862] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37545024 kB) free.
[1581820520] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1581820520] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21949096 kB) free.
[1581820564] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(49) < 49.50.64.122 [16/Feb/2020:08:02:44 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1581820684] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581821062] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20562836 kB) free!
[1581821183] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38244300 kB) free.
[1581821631] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2020/02/16 08:19:52 [error] 19828#0: *54629382 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581821752] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581821884] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 172.105.59.173 [16/Feb/2020:08:22:22 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581822004] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581822383] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20990704 kB) free!
[1581822503] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38206540 kB) free.
[1581823103] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20307000 kB) free!
[1581823223] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37571460 kB) free.
[1581823527] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581823647] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1581823823] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19234608 kB) free!
[1581823943] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36896304 kB) free.
[1581824543] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15715976 kB) free!
[1581824663] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31985320 kB) free.
[1581825863] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16992628 kB) free!
[1581825983] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35630664 kB) free.
[1581826583] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16548516 kB) free!
[1581826703] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35267404 kB) free.
[1581827304] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15577948 kB) free!
[1581827424] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34280176 kB) free.
[1581827797] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.0% (4657836 kB) free!
[1581827917] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.5% (4535984 kB) free!
[1581828024] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11777064 kB) free!
[1581828037] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.9% (4391576 kB) free!
[1581828144] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33649344 kB) free.
[1581828744] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10450288 kB) free!
[1581828864] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32153688 kB) free.
[1581828927] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1581829237] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1916596 kB) free!
[1581829464] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9284504 kB) free!
[1581829584] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30767040 kB) free.
[1581830184] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8995396 kB) free!
[1581830304] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4884968 kB) free!
[1581830424] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23660268 kB) free!
[1581831024] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381064 kB) free!
[1581831427] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.5% (21683484 kB) free.
[1581831448] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 400 processes
[1581831490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 51.9% (68418440 kB) free.
[1581831568] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1581832949] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4213472 kB) free!
[1581833069] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (3939516 kB) free!
[1581833127] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.0% (3218768 kB) free!
[1581833189] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.4% (3684284 kB) free!
[1581833247] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.7% (3125896 kB) free!
[1581833368] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.3% (3032048 kB) free!
[1581833968] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2444780 kB) free!
[1581834389] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (866536 kB) free!
[1581834457] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 708 processes
[1581834577] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 709 processes
[1581834697] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 712 processes
[1581834951] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(52) < 2020/02/16 12:00:51 [error] 19838#0: *54722579 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581834989] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.8% (15888752 kB) free.
[1581835072] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581835205] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 172.105.59.173 [16/Feb/2020:12:01:52 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.000, 0.000"
[1581835297] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1581835325] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581835698] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15273540 kB) free!
[1581835765] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9232872 kB) free!
[1581835818] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41619112 kB) free.
[1581835867] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21656620 kB) free.
[1581835885] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9232180 kB) free!
[1581836005] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (9232240 kB) free!
[1581836419] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15699160 kB) free!
[1581836539] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (11490888 kB) free!
[1581836659] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38581732 kB) free.
[1581837259] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13361736 kB) free!
[1581837379] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9307040 kB) free!
[1581837499] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36855692 kB) free.
[1581837527] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4899404 kB) free!
[1581837647] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.9% (4881208 kB) free!
[1581837767] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.8% (4857456 kB) free!
[1581838099] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11549768 kB) free!
[1581838219] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38607428 kB) free.
[1581838818] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.7% (3368992 kB) free!
[1581838819] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13487368 kB) free!
[1581838938] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.8% (3136840 kB) free!
[1581838939] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9109076 kB) free!
[1581839005] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4064336 kB) free!
[1581839058] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (2891708 kB) free!
[1581839059] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27141456 kB) free.
[1581839658] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1724972 kB) free!
[1581839659] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12654260 kB) free!
[1581839779] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8227044 kB) free!
[1581839899] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14790544 kB) free!
[1581839918] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9342696 kB) free!
[1581840020] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4733664 kB) free!
[1581840038] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9142680 kB) free!
[1581840077] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.0% (3960576 kB) free!
[1581840140] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4680140 kB) free!
[1581840159] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8926308 kB) free!
[1581840197] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.7% (3133108 kB) free!
[1581840260] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4645584 kB) free!
[1581840317] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.3% (3032128 kB) free!
[1581840499] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1581840499] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (11574624 kB) free!
[1581841100] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19682600 kB) free!
[1581841517] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1581841517] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1852216 kB) free!
[1581841700] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28906292 kB) free.
[1581842300] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3447480 kB) free!
[1581842420] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32285716 kB) free.
[1581842559] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4797328 kB) free!
[1581842701] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.6% (21361232 kB) free.
[1581842743] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 55.3% (27238576 kB) free.
[1581842769] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 38.8% (19110504 kB) free.
[1581842800] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 74.0% (18129704 kB) free.
[1581842837] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4868184 kB) free!
[1581842958] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4747932 kB) free!
[1581843020] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7118948 kB) free!
[1581843078] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.0% (4657348 kB) free!
[1581843140] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2848652 kB) free!
[1581843260] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31895272 kB) free.
[1581843860] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7367176 kB) free!
[1581843860] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2438860 kB) free!
[1581843906] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (707612 kB) free!
[1581843980] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2931584 kB) free!
[1581844026] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (317888 kB) free!
[1581844100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31713168 kB) free.
[1581844146] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175580 kB) free!
[1581844589] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4296312 kB) free!
[1581844700] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7144260 kB) free!
[1581844709] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4159436 kB) free!
[1581844820] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2760904 kB) free!
[1581844829] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.9% (4011344 kB) free!
[1581844940] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31624488 kB) free.
[1581845429] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1581845429] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.7% (15870932 kB) free.
[1581845525] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(170) < 92.42.104.166 [16/Feb/2020:15:00:22 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581845540] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7008388 kB) free!
[1581845646] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581845660] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2702584 kB) free!
[1581845781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31622956 kB) free.
[1581846078] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2375292 kB) free!
[1581846380] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7175780 kB) free!
[1581846500] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2735528 kB) free!
[1581846620] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31554548 kB) free.
[1581847220] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8158800 kB) free!
[1581847340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3903772 kB) free!
[1581847460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31039904 kB) free.
[1581848060] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6797624 kB) free!
[1581848180] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2295100 kB) free!
[1581848300] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31148796 kB) free.
[1581848900] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6884692 kB) free!
[1581849020] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2700708 kB) free!
[1581849140] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30732500 kB) free.
[1581849245] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [16/Feb/2020:16:02:42 +0530] TCP 502 0 0 121.001 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "1.001, 60.000, 60.000, 0.000"
[1581849305] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21662880 kB) free.
[1581849338] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.7% (21889444 kB) free.
[1581849366] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581849370] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 67.3% (16491004 kB) free.
[1581849471] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2020/02/16 16:00:58 [error] 19827#0: *54845899 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581849591] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581849740] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4362080 kB) free!
[1581849860] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399992 kB) free!
[1581849980] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32563388 kB) free.
[1581850581] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10077032 kB) free!
[1581850701] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6009672 kB) free!
[1581850821] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32222408 kB) free.
[1581851421] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6075380 kB) free!
[1581851541] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1847588 kB) free!
[1581851661] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30230656 kB) free.
[1581852261] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3629708 kB) free!
[1581852381] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (377524 kB) free!
[1581852501] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37435580 kB) free.
[1581853101] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11823808 kB) free!
[1581853221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17462276 kB) free!
[1581853341] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37787532 kB) free.
[1581853941] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11212408 kB) free!
[1581854061] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (7036656 kB) free!
[1581854181] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37658088 kB) free.
[1581854744] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1581854746] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9084836 kB) free!
[1581854782] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11338952 kB) free!
[1581854866] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9080012 kB) free!
[1581854902] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7189628 kB) free!
[1581854986] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (9077100 kB) free!
[1581855022] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37967136 kB) free.
[1581855622] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4822692 kB) free!
[1581855625] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9659856 kB) free!
[1581855742] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378260 kB) free!
[1581855745] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9565572 kB) free!
[1581855862] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392136 kB) free!
[1581855865] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (9460636 kB) free!
[1581855968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 42.2% (55572852 kB) free.
[1581856830] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581857206] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4602732 kB) free!
[1581857326] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4490128 kB) free!
[1581857446] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4399024 kB) free!
[1581858377] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25878204 kB) free!
[1581858497] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.7% (21983152 kB) free!
[1581858617] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18033184 kB) free!
[1581859217] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28367336 kB) free.
[1581859611] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1581859731] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1581859851] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1581859967] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 10.0% (2448196 kB) free!
[1581860417] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18367784 kB) free!
[1581860446] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2385124 kB) free!
[1581860537] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 41.8% (55105072 kB) free.
[1581861137] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25391636 kB) free!
[1581861258] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20489072 kB) free!
[1581861344] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4907656 kB) free!
[1581861378] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16841780 kB) free!
[1581861464] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.3% (4278548 kB) free!
[1581861584] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4257736 kB) free!
[1581861586] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4170880 kB) free!
[1581861978] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27515028 kB) free.
[1581862251] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1581862371] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 329 processes
[1581862491] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1581863765] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [16/Feb/2020:20:01:53 +0530] TCP 502 0 0 60.001 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1581863885] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581863991] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/02/16 20:00:39 [error] 19830#0: *54846951 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:5515145/1588129, bytes from/to upstream:1588129/5515145
[1581864111] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581864584] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2408536 kB) free!
[1581864865] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4878344 kB) free!
[1581869779] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25616800 kB) free!
[1581869898] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.2% (50377916 kB) free.
[1581872513] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.3% (20311848 kB) free.
[1581872532] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 27.8% (13708728 kB) free.
[1581872555] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 50.7% (24996052 kB) free.
[1581872572] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 75.7% (18565964 kB) free.
[1581872599] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 71.5% (17517704 kB) free.
[1581873116] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (1823648 kB) free!
[1581873159] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1581873236] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (1782264 kB) free!
[1581873279] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 351 processes
[1581873356] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1312472 kB) free!
[1581873377] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198224 kB) free!
[1581873399] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1581873497] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175552 kB) free!
[1581873498] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581873618] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189888 kB) free!
[1581876592] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224012 kB) free!
[1581876712] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241260 kB) free!
[1581876832] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231400 kB) free!
[1581877101] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7444992 kB) free!
[1581877221] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (480916 kB) free!
[1581877341] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2341000 kB) free!
[1581877463] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 214632 MB (20% inode=99%):
[1581877583] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 214632 MB (20% inode=99%):
[1581877703] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 214632 MB (20% inode=99%):