[1573929085] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29524600 kB) free.
[1573929384] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (2462604 kB) free!
[1573929435] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 31.0% (6964884 kB) free.
[1573929494] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175924 kB) free!
[1573929502] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187792 kB) free!
[1573929507] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196692 kB) free!
[1573929533] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179368 kB) free!
[1573929596] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 94.4% (46521432 kB) free.
[1573929614] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180012 kB) free!
[1573929622] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185944 kB) free!
[1573929627] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185332 kB) free!
[1573929654] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186700 kB) free!
[1573929685] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10448372 kB) free!
[1573929735] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177976 kB) free!
[1573929742] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179056 kB) free!
[1573929774] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184500 kB) free!
[1573929805] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6850788 kB) free!
[1573929925] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3235420 kB) free!
[1573930043] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9732416 kB) free!
[1573930147] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11674352 kB) free!
[1573930163] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9731860 kB) free!
[1573930283] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (9732956 kB) free!
[1573930525] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13919100 kB) free!
[1573930747] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.2% (13469520 kB) free!
[1573931152] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 301 processes
[1573931272] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1573931535] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4268184 kB) free!
[1573931725] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6156560 kB) free!
[1573932135] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1573932135] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193636 kB) free!
[1573932325] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17248180 kB) free!
[1573932925] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28289392 kB) free.
[1573933526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9866616 kB) free!
[1573933646] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6307300 kB) free!
[1573933725] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4830932 kB) free!
[1573933766] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2819000 kB) free!
[1573933883] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3312136 kB) free!
[1573934366] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13624492 kB) free!
[1573935465] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(156) < 2019/11/17 01:47:39 [error] 23648#0: *4495168 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:56/17, bytes from/to upstream:17/56
[1573935566] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (5269116 kB) free!
[1573935585] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/11/17 01:48:03 [error] 23634#0: *4495448 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:56/17, bytes from/to upstream:17/56
[1573935705] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573936035] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4227520 kB) free!
[1573936147] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (26970704 kB) free.
[1573936155] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4141492 kB) free!
[1573936166] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15578776 kB) free!
[1573936275] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.1% (4057308 kB) free!
[1573936478] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1573936598] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1573936747] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.8% (23479612 kB) free!
[1573936766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26805044 kB) free.
[1573936867] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26479912 kB) free.
[1573937366] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7659508 kB) free!
[1573937468] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (22664196 kB) free!
[1573937486] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4115184 kB) free!
[1573937588] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.3% (25411064 kB) free!
[1573937606] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (698868 kB) free!
[1573937708] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28084520 kB) free.
[1573937872] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220672 kB) free!
[1573937992] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (397152 kB) free!
[1573938112] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224700 kB) free!
[1573938308] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (24879204 kB) free!
[1573938428] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (27739772 kB) free.
[1573938806] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22212400 kB) free!
[1573939028] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.3% (24097692 kB) free!
[1573939126] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.9% (43308352 kB) free.
[1573939148] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (26798088 kB) free.
[1573939275] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1992536 kB) free!
[1573939406] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3415536 kB) free!
[1573939748] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.7% (23282576 kB) free!
[1573939868] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.4% (22912488 kB) free!
[1573939875] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 29.9% (6706472 kB) free.
[1573939912] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.5% (45112468 kB) free.
[1573939988] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.4% (25528116 kB) free!
[1573940006] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13469016 kB) free!
[1573940506] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1315) < 2019/11/17 03:11:22 [error] 23641#0: *4530804 upstream timed out (110: Connection timed out) while connecting to upstream, client: 121.241.242.109, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573940535] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.2% (10651792 kB) free.
[1573940588] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.8% (28773772 kB) free.
[1573940626] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573940706] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1590) < 203.212.70.170 [17/Nov/2019:03:11:21 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1573940800] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1247) < 121.241.242.109 [17/Nov/2019:03:10:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1573940803] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1840) < 2019/11/17 03:12:01 [error] 24135#0: *7035073 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:83/0, bytes from/to upstream:0/83
[1573940826] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573940920] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573940923] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573941016] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573941137] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573941207] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5201624 kB) free!
[1573941788] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.5% (25750280 kB) free!
[1573941807] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15189412 kB) free!
[1573941908] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28065696 kB) free.
[1573943007] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6521236 kB) free!
[1573943108] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.9% (21021984 kB) free!
[1573943228] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (26609452 kB) free.
[1573943607] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16986140 kB) free!
[1573944428] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.0% (22452324 kB) free!
[1573944548] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.8% (24824936 kB) free!
[1573944668] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.9% (30245448 kB) free.
[1573944807] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (8600516 kB) free!
[1573945268] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.0% (22370888 kB) free!
[1573945388] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.3% (24144076 kB) free!
[1573945407] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19085188 kB) free!
[1573945509] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.6% (29776416 kB) free.
[1573946007] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (957284 kB) free!
[1573946475] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4467336 kB) free!
[1573946535] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573946595] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4370720 kB) free!
[1573946715] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 19.1% (4281640 kB) free!
[1573947207] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20191508 kB) free!
[1573947807] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2291616 kB) free!
[1573948935] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4696456 kB) free!
[1573949007] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21634552 kB) free!
[1573949055] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4659580 kB) free!
[1573949175] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4624688 kB) free!
[1573949607] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3853548 kB) free!
[1573949715] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2217672 kB) free!
[1573950208] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13271208 kB) free!
[1573950915] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 28.6% (6424232 kB) free.
[1573950976] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182348 kB) free!
[1573951408] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3950052 kB) free!
[1573951827] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 60.5% (14823540 kB) free.
[1573951942] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.2% (12299568 kB) free.
[1573951974] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.8% (20061976 kB) free.
[1573951980] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1573951980] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.7% (22863716 kB) free.
[1573952027] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.4% (57184312 kB) free.
[1573952070] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.9% (39857276 kB) free.
[1573952708] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573955032] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23708432 kB) free!
[1573955152] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.2% (20051832 kB) free!
[1573955272] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16366380 kB) free!
[1573956316] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4127672 kB) free!
[1573956436] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (3964176 kB) free!
[1573956472] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35553700 kB) free.
[1573956556] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.0% (3813400 kB) free!
[1573957072] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15763660 kB) free!
[1573957192] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40727848 kB) free.
[1573957756] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2192940 kB) free!
[1573957792] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21138980 kB) free!
[1573957912] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17159296 kB) free!
[1573957985] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4689752 kB) free!
[1573958032] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13304524 kB) free!
[1573958105] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4575380 kB) free!
[1573958225] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (4508760 kB) free!
[1573958469] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.14564.664> (10.147.212.158:35008 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573958589] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573959232] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33492732 kB) free.
[1573959425] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1546816 kB) free!
[1573959832] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13136956 kB) free!
[1573959952] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39021648 kB) free.
[1573960552] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18405152 kB) free!
[1573960672] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14357640 kB) free!
[1573960792] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (10374396 kB) free!
[1573961392] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20157196 kB) free!
[1573961910] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22515920 kB) free.
[1573961941] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 27.9% (6263148 kB) free.
[1573961992] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30831812 kB) free.
[1573962592] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8784516 kB) free!
[1573962712] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36377824 kB) free.
[1573963114] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176756 kB) free!
[1573963227] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4435880 kB) free!
[1573963234] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195352 kB) free!
[1573963312] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15027556 kB) free!
[1573963347] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (4223356 kB) free!
[1573963354] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173264 kB) free!
[1573963432] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10622240 kB) free!
[1573963467] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (3989360 kB) free!
[1573963552] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6412992 kB) free!
[1573964153] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16519860 kB) free!
[1573964626] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/17 09:49:25 [error] 23643#0: *4715684 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573964667] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1675544 kB) free!
[1573964746] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573964753] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29602572 kB) free.
[1573965142] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (4126116 kB) free!
[1573965262] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (3981888 kB) free!
[1573965353] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6468300 kB) free!
[1573965382] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (3866860 kB) free!
[1573965473] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2350212 kB) free!
[1573965545] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4421996 kB) free!
[1573965593] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32630756 kB) free.
[1573965665] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4126764 kB) free!
[1573965785] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.0% (3806812 kB) free!
[1573966193] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9587280 kB) free!
[1573966313] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5448776 kB) free!
[1573966433] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1465324 kB) free!
[1573966986] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1317876 kB) free!
[1573967074] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9077188 kB) free!
[1573967182] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2350472 kB) free!
[1573967194] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9077156 kB) free!
[1573967314] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (9077792 kB) free!
[1573967633] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15004336 kB) free!
[1573968233] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401396 kB) free!
[1573968346] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/17 10:52:00 [error] 23642#0: *4744860 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573968466] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573968524] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176196 kB) free!
[1573968644] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (192136 kB) free!
[1573968765] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (190008 kB) free!
[1573968867] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.9% (3413308 kB) free!
[1573969433] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25661316 kB) free!
[1573970033] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4738144 kB) free!
[1573970068] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1573970068] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2004836 kB) free!
[1573970634] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18246588 kB) free!
[1573971234] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8380936 kB) free!
[1573971467] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/17 11:38:16 [error] 23645#0: *4764529 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573971587] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573971834] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19616920 kB) free!
[1573971981] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22324756 kB) free.
[1573972014] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.0% (22546324 kB) free.
[1573972180] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 27.4% (6154760 kB) free.
[1573972434] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35167424 kB) free.
[1573972582] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 58.3% (14282904 kB) free.
[1573972585] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181972 kB) free!
[1573972705] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183332 kB) free!
[1573972825] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168840 kB) free!
[1573972915] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.4% (22305344 kB) free.
[1573972965] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 81.7% (13274124 kB) free.
[1573973034] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16032596 kB) free!
[1573973154] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12298772 kB) free!
[1573973274] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35227044 kB) free.
[1573973874] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15083020 kB) free!
[1573973914] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4229280 kB) free!
[1573973994] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12267944 kB) free!
[1573974114] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9458676 kB) free!
[1573974714] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23967768 kB) free!
[1573975314] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34667172 kB) free.
[1573975914] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15195008 kB) free!
[1573976034] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (11422236 kB) free!
[1573976154] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34697532 kB) free.
[1573976754] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13185356 kB) free!
[1573976875] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7549188 kB) free!
[1573976995] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2562708 kB) free!
[1573977125] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185924 kB) free!
[1573977245] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184064 kB) free!
[1573977365] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180996 kB) free!
[1573977382] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180280 kB) free!
[1573977423] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.6% (3567104 kB) free!
[1573977490] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1573977502] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176340 kB) free!
[1573977543] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.0% (3438952 kB) free!
[1573977590] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.3% (2545356 kB) free!
[1573977611] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1573977622] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197676 kB) free!
[1573977663] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (3259568 kB) free!
[1573977710] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (2091864 kB) free!
[1573977713] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8490428 kB) free!
[1573977830] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1637628 kB) free!
[1573977833] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.4% (8082360 kB) free!
[1573977953] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.6% (7678280 kB) free!
[1573978195] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31178028 kB) free.
[1573978263] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2378636 kB) free!
[1573978795] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6455004 kB) free!
[1573978915] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2553528 kB) free!
[1573979035] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407488 kB) free!
[1573979574] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4908192 kB) free!
[1573979635] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13657000 kB) free!
[1573979694] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4863316 kB) free!
[1573979753] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4244136 kB) free!
[1573979814] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (4825664 kB) free!
[1573979923] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(75) < 2019/11/17 14:08:36 [error] 24135#0: *7158722 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:86/0, bytes from/to upstream:0/86
[1573979988] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/17 14:06:11 [error] 23638#0: *4844025 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573980043] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573980108] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573980235] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27168548 kB) free.
[1573980414] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 21.2% (5197404 kB) free.
[1573980426] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 49.50.64.122 [17/Nov/2019:14:12:18 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1573980520] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(166) < 49.50.86.152 [17/Nov/2019:14:13:49 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1573980546] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573980641] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(40) < 49.50.67.175 [17/Nov/2019:14:20:28 +0530] TCP 502 0 0 60.001 "10.147.212.49:4000, 10.147.212.38:4000, 10.147.212.61:4000, 10.147.212.37:4000" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 0.000, 60.000"
[1573980644] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(51) < 2019/11/17 14:12:18 [error] 24134#0: *7160780 no live upstreams while connecting to upstream, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "backend_3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573980708] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(428) < 2019/11/17 14:20:38 [error] 23647#0: *4853657 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573980761] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573980764] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573980828] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573980835] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5104256 kB) free!
[1573980955] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (16994092 kB) free!
[1573981075] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30851924 kB) free.
[1573981428] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2019/11/17 14:33:10 [error] 23646#0: *4861618 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.67.175, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573981548] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(123) < 2019/11/17 14:35:14 [error] 23640#0: *4862672 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573981668] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2019/11/17 14:35:53 [error] 23636#0: *4862238 recv() failed (104: Connection reset by peer) while proxying connection, client: 124.153.112.226, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573981675] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5602680 kB) free!
[1573981795] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1648408 kB) free!
[1573981915] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18676812 kB) free!
[1573981964] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(193) < 2019/11/17 14:42:28 [error] 24128#0: *7166870 no live upstreams while connecting to upstream, client: 49.50.67.175, server: 0.0.0.0:13001, upstream: "backend_13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573982084] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573982268] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573982346] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(125) < 49.50.67.175 [17/Nov/2019:14:42:28 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1573982466] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573983230] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 26.8% (6024692 kB) free.
[1573983715] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2883764 kB) free!
[1573983884] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(67) < 2019/11/17 15:14:44 [error] 24130#0: *7173021 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573983953] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.1% (47345716 kB) free.
[1573984005] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(31) < 2019/11/17 15:15:42 [error] 24130#0: *7173194 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573984125] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573984315] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (19024376 kB) free!
[1573984565] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1573984572] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1573984692] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1573984812] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 17 users currently logged in
[1573984915] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29969656 kB) free.
[1573985366] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4908492 kB) free!
[1573985486] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4904396 kB) free!
[1573985515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4364892 kB) free!
[1573985606] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4890420 kB) free!
[1573985635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (710832 kB) free!
[1573985755] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33743788 kB) free.
[1573985925] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(33) < 2019/11/17 15:46:50 [error] 24130#0: *7179114 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573986045] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(49) < 2019/11/17 15:49:32 [error] 24131#0: *7179772 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573986165] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(24) < 2019/11/17 15:51:13 [error] 24131#0: *7180211 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573986355] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10613496 kB) free!
[1573986475] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6955432 kB) free!
[1573986595] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2929784 kB) free!
[1573986766] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573987195] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14209468 kB) free!
[1573987508] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5620416 kB) free!
[1573987628] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6293208 kB) free!
[1573987748] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6317768 kB) free!
[1573987966] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2019/11/17 16:21:50 [error] 24130#0: *7186398 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573988086] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(46) < 2019/11/17 16:23:41 [error] 24130#0: *7186918 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573988206] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(25) < 2019/11/17 16:25:45 [error] 24130#0: *7187402 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573988395] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5610704 kB) free!
[1573988412] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 21 users currently logged in
[1573988806] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1573988806] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573988814] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4826884 kB) free!
[1573988934] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4794892 kB) free!
[1573989054] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4755460 kB) free!
[1573989189] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.30921.675> (10.147.212.158:38244 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573989309] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573989468] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/17 16:46:03 [error] 23645#0: *4948132 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/uuid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/uuid"
[1573989588] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573989595] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24108160 kB) free!
[1573989664] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573989777] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (179832 kB) free!
[1573989830] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1573989897] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173940 kB) free!
[1573990007] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2019/11/17 16:56:39 [error] 24131#0: *7193211 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573990017] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164744 kB) free!
[1573990127] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573990188] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/17 16:53:33 [error] 23648#0: *4953028 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/DlrLoad/uuid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/dashboard"
[1573990196] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3723848 kB) free!
[1573990210] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1573990309] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573990330] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1573990340] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.2% (22750312 kB) free.
[1573990365] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.9% (23272156 kB) free.
[1573990403] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.8% (34021724 kB) free.
[1573990439] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46306376 kB) free.
[1573990464] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 68.1% (33557168 kB) free.
[1573990496] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 72.6% (17788144 kB) free.
[1573990535] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.3% (14676928 kB) free.
[1573990727] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(74) < 2019/11/17 17:00:48 [error] 24130#0: *7193969 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573990847] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573990943] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177464 kB) free!
[1573991064] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182112 kB) free!
[1573991184] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187684 kB) free!
[1573991579] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (2287068 kB) free!
[1573991630] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4378464 kB) free!
[1573991700] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (2318780 kB) free!
[1573991702] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1831232 kB) free!
[1573991750] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4356556 kB) free!
[1573991820] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2297660 kB) free!
[1573991822] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (1796824 kB) free!
[1573991871] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 19.3% (4340656 kB) free!
[1573991942] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1784308 kB) free!
[1573992108] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(249) < 2019/11/17 17:30:05 [error] 23639#0: *4976662 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573992228] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/11/17 17:32:52 [error] 23642#0: *4978373 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/uuid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/uuid"
[1573992252] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1573992342] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157932 kB) free!
[1573992348] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2019/11/17 17:34:46 [error] 23648#0: *4979542 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573992372] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1573992462] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (196756 kB) free!
[1573992582] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (196288 kB) free!
[1573992647] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(85) < 2019/11/17 17:36:25 [error] 24130#0: *7200677 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573992767] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573992842] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19140120 kB) free!
[1573992948] SERVICE FLAPPING ALERT: secondary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1573992948] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573992962] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43708548 kB) free.
[1573993196] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8360824 kB) free!
[1573993316] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240416 kB) free!
[1573993436] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222544 kB) free!
[1573993549] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/17 17:54:34 [error] 23645#0: *4991867 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/uuid HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/DlrLoad/uuid"
[1573993562] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22586592 kB) free!
[1573993669] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573993682] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (18978748 kB) free!
[1573993802] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43659940 kB) free.
[1573993812] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 19 users currently logged in
[1573994037] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 88.3% (43505248 kB) free.
[1573994220] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (3667496 kB) free!
[1573994271] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 25.5% (5732544 kB) free.
[1573994402] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23997712 kB) free!
[1573994522] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20331044 kB) free!
[1573994567] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(94) < 2019/11/17 18:12:14 [error] 24130#0: *7207367 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573994642] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17023068 kB) free!
[1573994687] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573994798] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 406 processes
[1573994918] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 407 processes
[1573995242] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26858336 kB) free.
[1573996487] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(45) < 2019/11/17 18:43:40 [error] 24131#0: *7213383 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573996608] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(20) < 2019/11/17 18:46:34 [error] 24130#0: *7213992 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573996728] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(30) < 2019/11/17 18:47:36 [error] 24130#0: *7214072 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573997043] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25220584 kB) free!
[1573997163] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.9% (22301616 kB) free!
[1573997220] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1573997220] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (938324 kB) free!
[1573997283] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19387876 kB) free!
[1573997328] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573997918] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1573998038] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1573998483] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35687216 kB) free.
[1573998528] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2019/11/17 19:17:57 [error] 24130#0: *7219899 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573998638] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1573998648] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(25) < 2019/11/17 19:19:04 [error] 24131#0: *7220122 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:38/0, bytes from/to upstream:0/38
[1573998759] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1573998768] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(26) < 2019/11/17 19:21:31 [error] 24130#0: *7220607 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573998777] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.1% (21732120 kB) free.
[1573998807] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.6% (22214716 kB) free.
[1573998837] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 72.5% (17773812 kB) free.
[1573998845] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1573998880] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.8% (14760772 kB) free.
[1573998965] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1573999381] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (204136 kB) free!
[1573999501] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202596 kB) free!
[1573999621] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1573999621] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201292 kB) free!
[1573999968] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574000153] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7380672 kB) free!
[1574000273] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.0% (7371988 kB) free!
[1574000393] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.9% (7364556 kB) free!
[1574000568] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(22) < 2019/11/17 19:51:28 [error] 24130#0: *7226243 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574000688] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(27) < 2019/11/17 19:53:46 [error] 24130#0: *7226654 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574000808] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(27) < 2019/11/17 19:55:00 [error] 24130#0: *7226896 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574001483] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26122104 kB) free!
[1574001603] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.3% (22853840 kB) free!
[1574001723] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43674444 kB) free.
[1574002008] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574002021] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4803600 kB) free!
[1574002323] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24016144 kB) free!
[1574002443] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20321700 kB) free!
[1574002563] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43611760 kB) free.
[1574002621] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (208112 kB) free!
[1574002930] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1574003050] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1574003164] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24690964 kB) free!
[1574003208] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(104) < 2019/11/17 20:32:49 [error] 24130#0: *7234042 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574003284] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.4% (21563416 kB) free!
[1574003328] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574003404] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43665516 kB) free.
[1574003558] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1574003606] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 6.7% (1632264 kB) free!
[1574003678] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1574004004] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26163124 kB) free!
[1574004124] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.0% (22471040 kB) free!
[1574004244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43644540 kB) free.
[1574004469] SERVICE FLAPPING ALERT: secondary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574004844] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25485460 kB) free!
[1574004964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.7% (21971972 kB) free!
[1574005084] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43672016 kB) free.
[1574005128] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(89) < 2019/11/17 21:08:34 [error] 24135#0: *7240770 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574005248] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574005478] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 406 processes
[1574005598] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1574005824] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.5% (23314192 kB) free.
[1574005871] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 92.0% (22540964 kB) free.
[1574007484] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22329964 kB) free!
[1574007604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43710020 kB) free.
[1574007632] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170904 kB) free!
[1574007648] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(84) < 2019/11/17 21:45:25 [error] 24132#0: *7247696 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574007752] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172296 kB) free!
[1574007768] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574007814] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189048 kB) free!
[1574007845] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (529988 kB) free!
[1574007872] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192016 kB) free!
[1574007934] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173368 kB) free!
[1574007965] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (527752 kB) free!
[1574008054] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163608 kB) free!
[1574008069] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/17 21:50:09 [error] 23646#0: *5136935 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/uuid HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/DlrLoad/uuid"
[1574008086] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (517620 kB) free!
[1574008189] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574009568] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(92) < 2019/11/17 22:20:49 [error] 24130#0: *7254322 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574009689] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574010604] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26375888 kB) free!
[1574010724] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (22675076 kB) free!
[1574010844] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1574010844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43460020 kB) free.
[1574011444] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25930084 kB) free!
[1574011488] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(61) < 2019/11/17 22:54:21 [error] 24130#0: *7260710 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574011564] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.3% (22863216 kB) free!
[1574011608] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574011684] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43804928 kB) free.
[1574012208] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/11/17 22:58:40 [error] 24130#0: *7261450 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574012284] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (26021592 kB) free!
[1574012328] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574012404] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.4% (22924928 kB) free!
[1574012412] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 11 users currently logged in
[1574012524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41905288 kB) free.
[1574012535] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 181345 MB (16% inode=99%):
[1574012637] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233848 kB) free!
[1574012655] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 170665 MB (15% inode=99%):
[1574012757] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 32.3% (15893900 kB) free.
[1574012775] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 166611 MB (15% inode=99%):
[1574013124] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25420168 kB) free!
[1574013244] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.7% (22002464 kB) free!
[1574013364] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42761024 kB) free.
[1574013528] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(48) < 2019/11/17 23:28:40 [error] 24128#0: *7267108 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574013648] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1574013648] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574013964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25584332 kB) free!
[1574014054] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1574014085] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41957476 kB) free.
[1574014248] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(46) < 2019/11/17 23:33:48 [error] 24131#0: *7268048 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574014368] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found