[1579545605] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14586552 kB) free!
[1579545725] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31773196 kB) free.
[1579546744] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579546864] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1579548409] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.0% (11855844 kB) free.
[1579548997] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(108) < 2020/01/21 01:06:28 [error] 21673#0: *25603172 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:472171/6846532, bytes from/to upstream:6846532/472171
[1579549117] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(120) < 2020/01/21 01:08:29 [error] 21680#0: *25645751 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:236/70, bytes from/to upstream:70/236
[1579549238] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(2) < 2020/01/21 01:09:17 [error] 21676#0: *25646018 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:51/17, bytes from/to upstream:17/51
[1579549609] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165068 kB) free!
[1579549730] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190580 kB) free!
[1579549838] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579549850] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178124 kB) free!
[1579553081] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.8% (10237580 kB) free.
[1579553525] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25500648 kB) free!
[1579553645] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21491116 kB) free!
[1579553765] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17684532 kB) free!
[1579554260] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.3% (42039032 kB) free.
[1579554365] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31905768 kB) free.
[1579555838] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1120) < 2020/01/21 03:00:30 [error] 21682#0: *25653170 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:39/0, bytes from/to upstream:0/39
[1579555879] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1145) < 203.212.70.170 [21/Jan/2020:03:01:14 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579555958] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(320) < 2020/01/21 03:02:11 [error] 21675#0: *25653428 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579555999] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(74) < 203.212.70.170 [21/Jan/2020:03:01:21 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579556079] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579556119] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579556678] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(38) < 2020/01/21 03:07:17 [error] 21680#0: *25654338 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:1055/2887, bytes from/to upstream:2887/1055
[1579556798] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(256) < 2020/01/21 03:16:26 [error] 21672#0: *25656110 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:39/0, bytes from/to upstream:0/39
[1579556918] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(427) < 2020/01/21 03:18:02 [error] 21682#0: *25656695 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579557115] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1579557320] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(237) < 172.105.59.173 [21/Jan/2020:03:16:52 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579557440] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579557518] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1579557518] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579559318] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(41) < 2020/01/21 03:54:32 [error] 21680#0: *25658880 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:135/49, bytes from/to upstream:49/135
[1579559438] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579562439] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(123) < 2020/01/21 04:45:55 [error] 21683#0: *25663412 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579562559] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579563281] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171932 kB) free!
[1579563401] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.9% (3657192 kB) free!
[1579563521] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192648 kB) free!
[1579563950] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.7% (18313964 kB) free.
[1579564250] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.2% (11884332 kB) free.
[1579568240] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 49.50.64.122 [21/Jan/2020:06:27:15 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579568360] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579568559] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(230) < 2020/01/21 06:27:29 [error] 21681#0: *25675206 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:71/17, bytes from/to upstream:17/71
[1579568679] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579569279] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(119) < 2020/01/21 06:39:22 [error] 21678#0: *25676948 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.144.246, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:122/0, bytes from/to upstream:0/122
[1579569401] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579569401] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579570601] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(131) < 2020/01/21 07:04:08 [error] 21680#0: *25680554 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:67/17, bytes from/to upstream:17/67
[1579570721] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579571322] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2020/01/21 07:15:36 [error] 21677#0: *25682174 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579571442] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579574464] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579574584] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579575050] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4540240 kB) free!
[1579575171] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4517132 kB) free!
[1579575199] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/21 08:17:41 [error] 18361#0: *42027082 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.1", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "webserver2.com:8091", referrer: "http://webserver2.com:8091/temtapp/index.php/consolidate/lrnreport"
[1579575291] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (4501132 kB) free!
[1579575319] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579576109] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.7% (4823104 kB) free!
[1579576229] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.6% (4802648 kB) free!
[1579576349] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.5% (4783508 kB) free!
[1579576842] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(352) < 2020/01/21 08:49:52 [error] 21681#0: *25693381 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579576962] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579577562] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(44) < 2020/01/21 08:57:57 [error] 21679#0: *25694602 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579577682] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579579491] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1444164 kB) free!
[1579580082] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(86) < 2020/01/21 09:40:47 [error] 21680#0: *25699736 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:1593167/2578765, bytes from/to upstream:2578765/1593167
[1579580202] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579580802] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(438) < 2020/01/21 09:56:42 [error] 21672#0: *25702414 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:41/0, bytes from/to upstream:0/41
[1579580922] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(88) < 2020/01/21 09:58:18 [error] 21679#0: *25701669 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:40/0, bytes from/to upstream:0/40
[1579580960] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(204) < 203.212.70.170 [21/Jan/2020:09:57:50 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579581043] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579581080] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579581749] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2397600 kB) free!
[1579582242] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/21 10:15:14 [error] 21681#0: *25704447 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:209/17, bytes from/to upstream:17/209
[1579582362] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579582962] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(68) < 2020/01/21 10:24:27 [error] 21680#0: *25704126 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:45638/692082, bytes from/to upstream:692082/45638
[1579583082] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579584494] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.1% (21833764 kB) free.
[1579584525] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.5% (22089928 kB) free.
[1579584578] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.9% (40366988 kB) free.
[1579584619] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.4% (22654192 kB) free.
[1579584732] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579584852] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579584972] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579586772] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579586928] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1899132 kB) free!
[1579587049] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (1724460 kB) free!
[1579587169] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1544996 kB) free!
[1579587336] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1579587336] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22021656 kB) free.
[1579587372] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579587492] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579587612] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579587950] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3972556 kB) free!
[1579588070] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193632 kB) free!
[1579588086] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3418256 kB) free!
[1579588190] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 22.2% (5433976 kB) free.
[1579588206] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16256324 kB) free.
[1579588483] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(118) < 2020/01/21 12:04:22 [error] 21681#0: *25710435 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:51/17, bytes from/to upstream:17/51
[1579588549] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1579588577] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 2.7% (660664 kB) free!
[1579588603] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579588669] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1579588697] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 2.6% (634420 kB) free!
[1579588790] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.7% (3369932 kB) free!
[1579588817] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 2.5% (605540 kB) free!
[1579588910] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192448 kB) free!
[1579588915] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1579589031] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193140 kB) free!
[1579589036] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579589156] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579589203] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/21 12:16:33 [error] 21682#0: *25711421 recv() failed (104: Connection reset by peer) while proxying connection, client: 139.59.21.208, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:57/17, bytes from/to upstream:17/57
[1579589269] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1579589323] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/21 12:16:43 [error] 21681#0: *25711356 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:51/17, bytes from/to upstream:17/51
[1579589378] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.6% (19763016 kB) free.
[1579589389] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1579589443] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579589476] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 91.0% (22311616 kB) free.
[1579589510] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1579589634] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 37.31, 51.47, 33.94
[1579589720] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(90) < 2020/01/21 12:17:49 [error] 18365#0: *42088326 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1579589754] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 19.90, 40.16, 31.92
[1579589756] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1579589840] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579589875] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 17.26, 32.18, 29.98
[1579590012] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1579590012] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579590643] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/21 12:35:29 [error] 21681#0: *25712685 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:49198/104, bytes from/to upstream:104/49198
[1579590764] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579590940] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4383860 kB) free!
[1579590965] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24447216 kB) free!
[1579591060] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4218520 kB) free!
[1579591085] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19350664 kB) free!
[1579591180] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (4071072 kB) free!
[1579591205] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12652704 kB) free!
[1579591584] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 45.6% (60142672 kB) free.
[1579592380] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2217088 kB) free!
[1579593788] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21655936 kB) free.
[1579594074] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 18.12, 45.15, 38.13
[1579594194] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 18.61, 36.41, 35.78
[1579594314] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 17.91, 30.39, 33.66
[1579594388] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.4% (3316676 kB) free!
[1579594508] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.7% (3124384 kB) free!
[1579594556] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1579594628] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.9% (2943760 kB) free!
[1579594676] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1579594787] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170048 kB) free!
[1579594796] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1579594891] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.6% (21963176 kB) free.
[1579594914] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 13.49, 17.65, 25.41
[1579595228] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166596 kB) free!
[1579595429] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4655004 kB) free!
[1579595549] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (4315480 kB) free!
[1579595660] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8987376 kB) free!
[1579595669] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (4173284 kB) free!
[1579595698] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.9% (21687348 kB) free.
[1579595780] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8757352 kB) free!
[1579595900] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (8519496 kB) free!
[1579596714] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.44, 88.38, 82.82
[1579596763] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/21 14:13:08 [error] 21677#0: *25718353 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:10054/87, bytes from/to upstream:87/10054
[1579596834] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.34, 62.84, 74.15
[1579596883] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579596954] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.81, 45.47, 66.40
[1579597040] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/21 14:21:41 [error] 18362#0: *42137571 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.3483429262699449 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1579597160] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579597469] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2072364 kB) free!
[1579597792] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9230984 kB) free!
[1579597806] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4236312 kB) free!
[1579597912] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9231300 kB) free!
[1579597926] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.7% (3981480 kB) free!
[1579598032] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (9231608 kB) free!
[1579598046] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.5% (3710920 kB) free!
[1579598247] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.3% (22379804 kB) free.
[1579598300] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4466660 kB) free!
[1579598412] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579598532] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579598652] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579598755] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1579598755] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 12.83, 13.82, 28.90
[1579598890] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9705084 kB) free!
[1579599010] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9682240 kB) free!
[1579599130] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (9572156 kB) free!
[1579599246] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.0% (16841116 kB) free.
[1579599596] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1579599604] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.0% (45346884 kB) free.
[1579600483] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(637) < 2020/01/21 15:23:15 [error] 21680#0: *25722884 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579600603] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579600880] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(626) < 103.255.146.164 [21/Jan/2020:15:23:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579601000] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579601108] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.8% (2909756 kB) free!
[1579601228] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.2% (2775904 kB) free!
[1579601349] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2654644 kB) free!
[1579601601] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(8) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1179_n"
[1579601721] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1579601789] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13743428 kB) free!
[1579601909] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30129308 kB) free.
[1579601949] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1910232 kB) free!
[1579602391] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.5% (21841884 kB) free.
[1579602494] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.7% (4573188 kB) free!
[1579602614] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.9% (4388788 kB) free!
[1579602734] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.2% (4220452 kB) free!
[1579603432] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4107020 kB) free!
[1579603930] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4653712 kB) free!
[1579604000] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(626) < 103.255.146.164 [21/Jan/2020:16:23:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579604075] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.4% (45020248 kB) free.
[1579604119] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.2% (40013260 kB) free.
[1579604121] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579604203] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1070) < 2020/01/21 16:26:41 [error] 21678#0: *25727629 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:40/0, bytes from/to upstream:0/40
[1579604309] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (21000276 kB) free!
[1579604323] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(297) < 2020/01/21 16:28:26 [error] 21682#0: *25727531 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:99/17, bytes from/to upstream:17/99
[1579604429] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42981328 kB) free.
[1579604443] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(309) < 2020/01/21 16:30:39 [error] 21681#0: *25728891 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:70/0, bytes from/to upstream:0/70
[1579604652] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1579604687] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 751 processes
[1579604720] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(73) < 172.105.59.173 [21/Jan/2020:16:34:36 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579604800] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170096 kB) free!
[1579604807] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 749 processes
[1579604840] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579604920] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (417028 kB) free!
[1579604927] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 340 processes
[1579605029] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18432708 kB) free!
[1579605040] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (289612 kB) free!
[1579605134] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2083004 kB) free!
[1579605149] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40838040 kB) free.
[1579605440] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1529) < 103.255.146.164 [21/Jan/2020:16:46:54 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579605560] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(4) < 172.105.59.173 [21/Jan/2020:16:49:05 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1579605596] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1579605681] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(78) < 182.18.144.246 [21/Jan/2020:16:50:20 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579605716] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1579605905] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.1% (21826556 kB) free.
[1579605941] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22512344 kB) free.
[1579605956] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 453.36, 250.60, 105.85
[1579606076] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 70.62, 172.12, 94.70
[1579606196] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16.04, 117.42, 84.03
[1579606349] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19237440 kB) free!
[1579606469] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40716788 kB) free.
[1579606844] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579606881] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579607069] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16155972 kB) free!
[1579607189] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37620112 kB) free.
[1579607396] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.37, 16.48, 33.98
[1579607996] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 17.29, 16.06, 25.07
[1579608389] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15998752 kB) free!
[1579608509] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37909780 kB) free.
[1579609109] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14336736 kB) free!
[1579609229] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36930540 kB) free.
[1579609601] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1579609722] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1579609830] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12800608 kB) free!
[1579609950] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35161060 kB) free.
[1579610126] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (2438216 kB) free!
[1579610246] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.0% (2432756 kB) free!
[1579610366] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (2430008 kB) free!
[1579610748] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.2% (3012952 kB) free!
[1579610868] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.5% (2832984 kB) free!
[1579610988] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2664160 kB) free!
[1579611589] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (500664 kB) free!
[1579612435] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.1% (22227688 kB) free.
[1579612844] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/21 18:41:00 [error] 21675#0: *25755505 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:2603/87, bytes from/to upstream:87/2603
[1579612964] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579613251] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (4297060 kB) free!
[1579613371] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (4116652 kB) free!
[1579613491] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (3749308 kB) free!
[1579613652] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579613749] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2951676 kB) free!
[1579613869] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.4% (4452812 kB) free.
[1579613996] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15.22, 100.34, 87.07
[1579614116] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.12, 71.83, 78.20
[1579614236] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 44.53, 68.45, 76.39
[1579614691] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2279720 kB) free!
[1579614766] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24960884 kB) free!
[1579614886] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39790760 kB) free.
[1579615436] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 16.10, 17.12, 32.66
[1579615486] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14941988 kB) free!
[1579615527] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9185368 kB) free!
[1579615606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39384708 kB) free.
[1579615648] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9181448 kB) free!
[1579615768] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (9182592 kB) free!
[1579616206] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14715004 kB) free!
[1579616326] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39533716 kB) free.
[1579616497] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4763648 kB) free!
[1579616617] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4736160 kB) free!
[1579616636] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 11.51, 14.21, 22.24
[1579616737] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (4692364 kB) free!
[1579616926] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14680536 kB) free!
[1579617046] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38620184 kB) free.
[1579617646] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14446860 kB) free!
[1579617766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37851388 kB) free.
[1579618366] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15273924 kB) free!
[1579618486] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10894508 kB) free!
[1579618606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34388716 kB) free.
[1579619039] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (1710996 kB) free!
[1579619159] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (1682192 kB) free!
[1579619206] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10647508 kB) free!
[1579619279] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1607940 kB) free!
[1579619326] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35477484 kB) free.
[1579619636] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15.93, 89.09, 59.62
[1579619756] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8.13, 61.91, 53.23
[1579619876] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 13.68, 45.79, 48.34
[1579619926] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11023192 kB) free!
[1579620046] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36089380 kB) free.
[1579620476] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 7.31, 13.19, 29.39
[1579620646] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11411844 kB) free!
[1579620766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35680848 kB) free.
[1579621366] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13702288 kB) free!
[1579621486] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29156576 kB) free.
[1579621768] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4824792 kB) free!
[1579622086] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14107364 kB) free!
[1579622206] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10122300 kB) free!
[1579622326] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32512656 kB) free.
[1579622921] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1579622927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11507520 kB) free!
[1579623041] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1579623047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33372912 kB) free.
[1579623100] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.2% (20283044 kB) free.
[1579623134] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 58.7% (28930300 kB) free.
[1579623169] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 59.2% (14500460 kB) free.
[1579623647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10693860 kB) free!
[1579623767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33722416 kB) free.
[1579623785] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579623905] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1579624367] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12837600 kB) free!
[1579624487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34761000 kB) free.
[1579624881] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(36) < 49.50.64.122 [21/Jan/2020:22:11:15 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579624965] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(819) < 2020/01/21 22:12:25 [error] 21682#0: *25778025 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1579625001] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(17) < 182.18.144.246 [21/Jan/2020:22:11:57 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579625085] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/21 22:12:53 [error] 21680#0: *25777587 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579625087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11396708 kB) free!
[1579625121] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579625205] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579625207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32785096 kB) free.
[1579625507] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4717036 kB) free!
[1579625627] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4705928 kB) free!
[1579625747] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4273944 kB) free!
[1579625807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11782140 kB) free!
[1579625927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32315584 kB) free.
[1579626527] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11967808 kB) free!
[1579626647] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32367776 kB) free.
[1579626947] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2197788 kB) free!
[1579627847] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15211292 kB) free!
[1579627967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34470288 kB) free.
[1579628567] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13983504 kB) free!
[1579628688] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34423216 kB) free.
[1579629287] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13469992 kB) free!
[1579629407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33021508 kB) free.
[1579629938] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2198964 kB) free!
[1579630005] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(473) < 2020/01/21 23:36:45 [error] 21677#0: *25787814 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:40/0, bytes from/to upstream:0/40
[1579630007] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22971164 kB) free!
[1579630125] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(369) < 2020/01/21 23:38:43 [error] 21680#0: *25787913 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:99/17, bytes from/to upstream:17/99
[1579630127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31056004 kB) free.
[1579630246] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(10) < 2020/01/21 23:38:52 [error] 21681#0: *25787455 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579630521] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(30) < 203.212.70.170 [21/Jan/2020:23:37:52 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1579630538] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (2820356 kB) free!
[1579630641] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579630846] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579631327] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13727912 kB) free!