[1559932230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559932350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559932402] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.7% (1732800 kB) free!
[1559932412] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437504 kB) free!
[1559932775] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2794504 kB) free!
[1559932895] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11677412 kB) free.
[1559933002] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171044 kB) free!
[1559933090] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 398 processes
[1559933603] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.4% (2018184 kB) free!
[1559933690] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 378 processes
[1559934151] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559934203] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167652 kB) free!
[1559934695] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3509160 kB) free!
[1559934751] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559934815] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12221364 kB) free.
[1559934871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559934991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559936015] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3085636 kB) free!
[1559936135] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11870316 kB) free.
[1559937335] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2658892 kB) free!
[1559937391] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559937455] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11382892 kB) free.
[1559937992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559938020] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1559938112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559938232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559938490] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 387 processes
[1559938655] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2929592 kB) free!
[1559938775] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11533764 kB) free.
[1559939690] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 392 processes
[1559940203] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (3902576 kB) free.
[1559940575] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3382424 kB) free!
[1559940632] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559940695] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11909980 kB) free.
[1559940803] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166708 kB) free!
[1559940923] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161956 kB) free!
[1559941043] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (3643716 kB) free.
[1559941232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559941353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559941473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559941643] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (2794732 kB) free!
[1559941763] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178612 kB) free!
[1559941883] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.8% (2407740 kB) free!
[1559941895] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3380396 kB) free!
[1559942015] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11902072 kB) free.
[1559943215] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2838256 kB) free!
[1559943273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559943336] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11363580 kB) free.
[1559944473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559944490] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 386 processes
[1559944535] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4986820 kB) free!
[1559944593] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559944655] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (2752452 kB) free!
[1559944714] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559944775] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12114388 kB) free.
[1559945090] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 375 processes
[1559945313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559945352] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 504.24, 293.91, 123.19
[1559945375] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5679204 kB) free!
[1559945433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559945439] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 513.21, 403.71, 201.66
[1559945473] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 510.35, 365.88, 170.26
[1559945495] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (3299868 kB) free!
[1559945553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559945559] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 338.52, 398.81, 225.82
[1559945593] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 88.69, 262.51, 156.96
[1559945615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12119112 kB) free.
[1559945679] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.44, 267.55, 198.70
[1559946061] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2896704 kB) free!
[1559946083] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (3294592 kB) free.
[1559946153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559946241] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.06, 10.03, 6.68
[1559946273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559946361] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.50, 6.82, 5.91
[1559946481] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 4.72, 5.25
[1559946815] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7139136 kB) free!
[1559946873] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559946935] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5201324 kB) free!
[1559946993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559947055] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2881340 kB) free!
[1559947081] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.90, 2.90
[1559947114] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559947490] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 363 processes
[1559947656] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7563512 kB) free!
[1559947713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559947833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559947953] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559947993] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.86, 0.57, 12.25
[1559948079] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.95, 2.08, 16.70
[1559948090] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 382 processes
[1559948256] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10429008 kB) free.
[1559948483] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.1% (2120912 kB) free!
[1559948553] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559948603] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.2% (3593644 kB) free.
[1559948673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559948793] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559949193] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 0.47, 3.71
[1559949393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559949514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559949634] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559949662] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461792 kB) free!
[1559949879] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.58, 1.70, 3.89
[1559950056] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3472468 kB) free!
[1559950176] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12106120 kB) free.
[1559950776] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2919236 kB) free!
[1559950834] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559950896] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11268668 kB) free.
[1559951434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559951554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559951674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559952096] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2825636 kB) free!
[1559952216] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11070188 kB) free.
[1559952803] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (3194676 kB) free!
[1559952890] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 400 processes
[1559952923] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.7% (3844016 kB) free.
[1559953490] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 419 processes
[1559954123] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.4% (2008252 kB) free!
[1559954243] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161884 kB) free!
[1559954363] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158356 kB) free!
[1559954494] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (997376 kB) free!
[1559954604] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.8% (5096500 kB) free.
[1559954614] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (926632 kB) free!
[1559954616] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (2981660 kB) free!
[1559954674] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559954734] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (834732 kB) free!
[1559954736] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11699700 kB) free.
[1559954904] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (309128 kB) free!
[1559954963] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.0% (4061412 kB) free.
[1559954964] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (310232 kB) free!
[1559955024] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 1.3% (311056 kB) free!
[1559955084] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 1.3% (310500 kB) free!
[1559955274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559955394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559955514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559955687] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 565 processes
[1559955747] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 563 processes
[1559955807] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 559 processes
[1559955936] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2376320 kB) free!
[1559956056] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11014752 kB) free.
[1559956284] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.2% (4207704 kB) free!
[1559956707] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 567 processes
[1559956767] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 570 processes
[1559957067] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 566 processes
[1559957127] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 567 processes
[1559957187] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 560 processes
[1559957363] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (3147408 kB) free!
[1559957483] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.0% (3890036 kB) free.
[1559957856] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2962992 kB) free!
[1559957914] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559957976] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11361200 kB) free.
[1559958083] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (3076876 kB) free!
[1559958204] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.9% (3875656 kB) free.
[1559958514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559958634] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559958754] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559958879] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 94.18, 63.91, 29.68
[1559958890] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 403 processes
[1559958999] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 94.30, 73.99, 37.52
[1559959119] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 94.58, 80.80, 44.44
[1559959393] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 86.36, 58.54, 26.51
[1559959490] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 384 processes
[1559959513] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 29.20, 50.85, 27.93
[1559959534] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (2911340 kB) free!
[1559959633] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.15, 34.12, 24.57
[1559959777] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3106492 kB) free!
[1559959862] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2733844 kB) free!
[1559959897] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11441668 kB) free.
[1559960233] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.81, 5.00, 13.08
[1559960304] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 366 processes
[1559960319] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.95, 3.36, 14.16
[1559960424] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559961433] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.53, 3.95
[1559961662] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2473144 kB) free!
[1559961697] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2669056 kB) free!
[1559961754] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559961817] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11087356 kB) free.
[1559962119] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.64, 1.76, 3.53
[1559962354] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559962403] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3071852 kB) free!
[1559962474] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559962523] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.7% (3675388 kB) free.
[1559962594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559963617] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (2337164 kB) free!
[1559963737] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10957516 kB) free.
[1559963833] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 38.69, 225.18, 152.27
[1559963919] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 513.15, 434.87, 230.47
[1559963953] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 5.81, 150.86, 133.87
[1559964039] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 125.36, 334.01, 219.96
[1559964062] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2744416 kB) free!
[1559964073] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1559964073] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.04, 101.41, 117.81
[1559964159] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1559964159] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 18.35, 224.05, 193.50
[1559964323] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3241948 kB) free!
[1559964334] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.4% (4997396 kB) free.
[1559964443] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (3355416 kB) free.
[1559964890] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 410 processes
[1559964934] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4825184 kB) free!
[1559965054] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4795316 kB) free!
[1559965174] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1559965174] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (4740156 kB) free!
[1559965262] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1559965262] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2347344 kB) free!
[1559965490] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 440 processes
[1559965538] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (1412428 kB) free!
[1559965594] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559965643] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.2% (2469632 kB) free!
[1559965658] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10514736 kB) free.
[1559965681] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.17, 91.38, 65.90
[1559965763] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.3% (2969448 kB) free!
[1559965801] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.72, 61.21, 57.93
[1559965873] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.96, 1.09, 17.57
[1559965884] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.9% (2907036 kB) free!
[1559965921] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.36, 41.02, 50.93
[1559966090] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 431 processes
[1559966559] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.88, 1.92, 16.37
[1559966690] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 427 processes
[1559966794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559966831] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47053 MB (20% inode=99%):
[1559966914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559966951] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46924 MB (20% inode=99%):
[1559967034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559967071] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46852 MB (20% inode=99%):
[1559967121] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.55, 1.37, 14.41
[1559967673] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.96, 1.24, 3.56
[1559968058] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (1211328 kB) free!
[1559968178] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11422800 kB) free.
[1559968284] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1525184 kB) free!
[1559968321] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.61, 0.58, 4.33
[1559968360] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.15, 2.14, 4.19
[1559968766] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.26, 5.66, 5.06
[1559968813] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 366 processes
[1559968884] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.8% (2404884 kB) free!
[1559968886] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.44, 5.30, 5.00
[1559968933] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 366 processes
[1559969053] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1559969378] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (1296856 kB) free!
[1559969434] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1559969486] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.61, 6.40, 5.70
[1559969498] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10907860 kB) free.
[1559969606] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.11, 5.57, 5.48
[1559969727] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.61, 5.26, 5.38
[1559969848] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (2091112 kB) free!
[1559969974] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2457412 kB) free!
[1559970034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559970154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559970253] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 5 crit, 0 warn out of 369 processes
[1559970274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559970373] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1559970448] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168204 kB) free!
[1559970973] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 366 processes
[1559970984] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.9% (2425432 kB) free!
[1559971093] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1559971298] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (1324360 kB) free!
[1559971419] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11173800 kB) free.
[1559971956] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.5% (2188796 kB) free!
[1559972018] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (1271904 kB) free!
[1559972074] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559972138] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.9% (9816680 kB) free!
[1559972258] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7307472 kB) free!
[1559972293] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1559972414] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1559972484] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (170084 kB) free!
[1559972556] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1050252 kB) free!
[1559973084] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.9% (1935160 kB) free!
[1559973274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559973394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559973514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559973684] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (336640 kB) free!
[1559974058] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10331844 kB) free.
[1559974213] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 1 warn out of 368 processes
[1559974333] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1559974571] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.2% (974308 kB) free!
[1559974691] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 10.3% (758680 kB) free!
[1559974811] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (523700 kB) free!
[1559974884] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.8% (3531712 kB) free.
[1559975484] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176764 kB) free!
[1559975533] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 365 processes
[1559975604] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (149736 kB) free!
[1559975653] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1559975724] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (154760 kB) free!
[1559975774] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1559975858] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232492 kB) free!
[1559975914] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559975978] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9906800 kB) free.
[1559976012] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.0% (6763260 kB) free.
[1559976062] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559976324] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.7% (2865980 kB) free!
[1559976973] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 367 processes
[1559977093] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 367 processes
[1559977115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559977174] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1559977213] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 367 processes
[1559977235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559977355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559977524] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (199596 kB) free!
[1559977778] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236052 kB) free!
[1559977873] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559977898] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10173720 kB) free.
[1559978090] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559978124] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.1% (1639656 kB) free!
[1559978413] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1559978471] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24410 MB (10% inode=99%):
[1559978560] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559978724] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159260 kB) free!
[1559978848] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (4218128 kB) free.
[1559979071] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27049 MB (12% inode=99%):
[1559979098] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (393104 kB) free!
[1559979155] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559979218] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10497456 kB) free.
[1559979448] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (199648 kB) free!
[1559979569] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161636 kB) free!
[1559979689] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1559979689] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181496 kB) free!
[1559979755] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559979875] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559979925] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.8% (2894448 kB) free!
[1559979995] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559980271] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24013 MB (10% inode=99%):
[1559980321] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 262.13, 372.40, 214.71
[1559980441] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.63, 249.28, 188.71
[1559980525] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152512 kB) free!
[1559980561] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.34, 167.00, 165.90
[1559980889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (4730996 kB) free.
[1559981018] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (1108640 kB) free!
[1559981138] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11360104 kB) free.
[1559981220] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1559981340] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1559981460] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1559981489] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166136 kB) free!
[1559981560] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 37.02, 33.58, 15.41
[1559981609] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (160544 kB) free!
[1559981680] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.59, 23.35, 13.86
[1559981725] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.4% (3150620 kB) free!
[1559981729] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165264 kB) free!
[1559981738] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (228904 kB) free!
[1559981795] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1559981800] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.03, 16.40, 12.46
[1559981858] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10407600 kB) free.
[1559982013] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.8% (795972 kB) free!
[1559982133] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (568828 kB) free!
[1559982253] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (362360 kB) free!
[1559982395] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559982400] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 85.42, 21.93, 13.35
[1559982515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559982613] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 375 processes
[1559982635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559982673] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.30, 13.40, 9.42
[1559982733] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1559982793] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.48, 10.10, 8.69
[1559982913] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.93, 7.76, 8.01
[1559982925] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (3287528 kB) free.
[1559983000] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.13, 8.99, 12.78
[1559983059] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (277152 kB) free!
[1559983179] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10707052 kB) free.
[1559983513] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.80, 1.63, 4.54
[1559983525] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (149900 kB) free!
[1559983561] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 1.05, 11.00
[1559983645] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.1% (2445084 kB) free!
[1559983765] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (148048 kB) free!
[1559983933] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1559984053] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1559984379] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (747664 kB) free!
[1559984435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559984499] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10582248 kB) free.
[1559984761] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.44, 3.35
[1559984800] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1559984800] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 121.98, 100.00, 54.03
[1559985036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559985156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559985253] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1559985276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559985313] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 96.13, 56.98, 24.98
[1559985374] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1559985433] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 97.51, 70.49, 33.81
[1559985494] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 371 processes
[1559985554] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 27.06, 56.96, 33.64
[1559985565] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.3% (1831416 kB) free!
[1559985699] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (1339112 kB) free!
[1559985819] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10871416 kB) free.
[1559986000] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.50, 4.21, 19.12
[1559986094] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1559986094] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1559986154] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1559986154] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.36, 8.14, 17.95
[1559986453] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.3% (6860836 kB) free.
[1559986694] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1559986814] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1559986934] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 0 warn out of 373 processes
[1559987129] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (4086612 kB) free.
[1559987365] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (555988 kB) free!
[1559987619] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (731012 kB) free!
[1559987676] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559987727] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 6.04, 4.80, 4.93
[1559987729] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168108 kB) free!
[1559987739] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10745836 kB) free.
[1559987800] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.17, 2.04, 4.39
[1559987850] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (151452 kB) free!
[1559987954] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.88, 1.15, 3.46
[1559987970] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (198324 kB) free!
[1559988276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559988396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559988927] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.45, 5.30, 5.22
[1559988939] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (1085932 kB) free!
[1559988996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559989047] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.77, 5.30, 5.24
[1559989059] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11203528 kB) free.
[1559989117] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559989167] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.59, 4.98, 5.12
[1559989237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559989344] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1559989561] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.63, 18.33, 9.89
[1559989681] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.04, 12.36, 8.72
[1559989765] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.7% (1735876 kB) free!
[1559989767] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.69, 3.46, 4.34
[1559989801] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1559989801] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 8.47, 7.74
[1559989944] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 1 warn out of 376 processes
[1559990064] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 5 crit, 1 warn out of 376 processes
[1559990184] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 2 warn out of 375 processes
[1559990259] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (493112 kB) free!
[1559990365] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155156 kB) free!
[1559990379] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10689652 kB) free.
[1559990401] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.39, 1.70, 4.37
[1559990965] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.2% (1661752 kB) free!
[1559991384] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1559991565] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166744 kB) free!
[1559991579] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (754260 kB) free!
[1559991637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559991699] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10316500 kB) free.
[1559991984] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1559992104] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1559992225] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 372 processes
[1559992237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559992357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559992477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559992765] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.6% (2854552 kB) free!
[1559992899] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (515352 kB) free!
[1559993019] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10589108 kB) free.
[1559993365] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163224 kB) free!
[1559993401] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 80.65, 26.63, 9.82
[1559993521] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 102.58, 52.90, 21.48
[1559993642] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 105.79, 70.57, 31.77
[1559994212] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2835164 kB) free!
[1559994219] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (1359808 kB) free!
[1559994253] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (1245892 kB) free!
[1559994277] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559994339] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10832568 kB) free.
[1559994373] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (1126564 kB) free!
[1559994493] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (1038996 kB) free!
[1559994625] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1559994842] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 3.27, 15.08
[1559994877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559994997] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559995093] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (563488 kB) free!
[1559995539] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (1088548 kB) free!
[1559995659] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10596848 kB) free.
[1559995765] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.1% (2615816 kB) free!
[1559995977] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3146560 kB) free!
[1559996013] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2328080 kB) free!
[1559996042] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.41, 4.41
[1559996097] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3102944 kB) free!
[1559996197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559996217] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.8% (3060808 kB) free!
[1559996317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559996365] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160540 kB) free!
[1559996438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559996859] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (925520 kB) free!
[1559996965] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (3448544 kB) free.
[1559996979] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10286260 kB) free.
[1559997417] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.0% (13660328 kB) free.
[1559997494] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1559997494] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.8% (6897368 kB) free.
[1559997565] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162052 kB) free!
[1559997686] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (3725784 kB) free.
[1559998000] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559998154] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559998180] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (737064 kB) free!
[1559998238] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559998285] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (158724 kB) free!
[1559998300] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10678260 kB) free.
[1559998405] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (3771232 kB) free.
[1559998838] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559998958] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559999005] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (165600 kB) free!
[1559999078] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559999125] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (663236 kB) free!
[1559999245] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (3287196 kB) free.
[1559999370] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559999500] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (1139988 kB) free!
[1559999620] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10610840 kB) free.
[1559999845] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167760 kB) free!
[1559999965] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (3279976 kB) free.
[1560000566] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (164936 kB) free!
[1560000686] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.4% (3151184 kB) free!
[1560000806] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161124 kB) free!
[1560000820] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (1488100 kB) free!
[1560000940] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10880332 kB) free.
[1560001406] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.5% (3168612 kB) free!
[1560001540] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (647132 kB) free!
[1560001660] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10179200 kB) free.
[1560002006] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164544 kB) free!
[1560002090] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1560002690] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 352 processes
[1560002860] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (1805744 kB) free!
[1560002981] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10474684 kB) free.
[1560003570] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (1966868 kB) free!
[1560003580] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1562376 kB) free!
[1560003700] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10464860 kB) free.
[1560003891] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 395 processes
[1560004170] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162660 kB) free!
[1560004226] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1560004300] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (2343992 kB) free!
[1560004406] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.0% (1782024 kB) free!
[1560004420] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10803516 kB) free.
[1560005020] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2717116 kB) free!
[1560005140] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10787860 kB) free.
[1560005606] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (638408 kB) free!
[1560005740] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (1539032 kB) free!
[1560005860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10137568 kB) free.
[1560006207] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.4% (2829076 kB) free!
[1560006400] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 189.21, 138.71, 66.17
[1560006520] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 189.67, 155.58, 81.15
[1560006640] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 106.95, 146.24, 87.41
[1560006842] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560007060] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1786576 kB) free!
[1560007094] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1560007180] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10610060 kB) free.
[1560007407] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152316 kB) free!
[1560007780] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (1827904 kB) free!
[1560007825] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 376 processes
[1560007860] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1560007900] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10432896 kB) free.
[1560007945] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1560008007] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.1% (1642748 kB) free!
[1560008355] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 124.19, 85.98, 39.33
[1560008440] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.54, 2.48, 14.30
[1560008463] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2731776 kB) free!
[1560008475] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 49.82, 77.83, 42.51
[1560008545] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 376 processes
[1560008556] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.8% (2890648 kB) free!
[1560008595] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7.05, 52.21, 37.40
[1560008666] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 376 processes
[1560008691] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1560008691] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 383 processes
[1560008786] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 374 processes
[1560008970] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (1692144 kB) free!
[1560009071] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 38136 MB (16% inode=99%):
[1560009100] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (833600 kB) free!
[1560009156] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (149360 kB) free!
[1560009195] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 7.63, 19.94
[1560009207] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (990072 kB) free!
[1560009220] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10627984 kB) free.
[1560009386] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1560009571] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1560009571] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167192 kB) free!
[1560009663] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2352372 kB) free!
[1560009671] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57265 MB (25% inode=99%):
[1560009756] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.0% (13491716 kB) free.
[1560009820] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230844 kB) free!
[1560009878] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560009940] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10186088 kB) free.
[1560010240] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.12, 1.89, 3.65
[1560010395] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 509.91, 340.10, 157.04
[1560010478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560010492] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1560010598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560010718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560010863] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2807528 kB) free!
[1560011092] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 425 processes
[1560011740] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (2101948 kB) free!
[1560011786] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 374 processes
[1560011860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10947668 kB) free.
[1560011906] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 374 processes
[1560012026] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 375 processes
[1560012156] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (182232 kB) free!
[1560012276] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (182632 kB) free!
[1560012396] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1560012396] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162748 kB) free!
[1560012626] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1560012626] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 373 processes
[1560012795] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1560012795] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.33, 0.62, 16.21
[1560013060] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (1970216 kB) free!
[1560013118] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560013180] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10911832 kB) free.
[1560013719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560013780] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (1229252 kB) free!
[1560013839] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1560013863] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1560013863] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2451496 kB) free!
[1560013900] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10274008 kB) free.
[1560013995] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.40, 4.76
[1560014092] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 423 processes
[1560014439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560014460] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1560014559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560014679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560014692] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 428 processes
[1560015060] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1560015100] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (1469080 kB) free!
[1560015220] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10427048 kB) free.
[1560015626] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 374 processes
[1560015746] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 373 processes
[1560015807] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.3% (1665696 kB) free!
[1560015866] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1560015892] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 418 processes
[1560016407] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164648 kB) free!
[1560016420] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (1303760 kB) free!
[1560016479] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560016493] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 428 processes
[1560016540] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10234320 kB) free.
[1560017679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560017796] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.9% (13959480 kB) free.
[1560017799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1560017919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560018060] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1560018060] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1560018293] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 382 processes
[1560018340] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (1890244 kB) free!
[1560018460] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10434480 kB) free.