[1590258943] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25182568 kB) free!
[1590259059] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 357 processes
[1590259095] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1590259543] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4033728 kB) free!
[1590260143] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26330500 kB) free!
[1590260575] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1590260743] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4587288 kB) free!
[1590261156] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152624 kB) free!
[1590261344] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28417944 kB) free.
[1590261459] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1590261944] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6066100 kB) free!
[1590262060] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1590262064] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2046216 kB) free!
[1590262184] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400344 kB) free!
[1590262660] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1590262784] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18642860 kB) free!
[1590262956] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (2140496 kB) free!
[1590263384] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403964 kB) free!
[1590263445] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 79.6% (19499828 kB) free.
[1590263795] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (1286760 kB) free!
[1590263915] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244320 kB) free!
[1590263984] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20568680 kB) free!
[1590264035] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233580 kB) free!
[1590264584] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403680 kB) free!
[1590264586] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (470036 kB) free!
[1590264706] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (476996 kB) free!
[1590264826] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1678052 kB) free!
[1590265060] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590265180] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1590265184] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19729000 kB) free!
[1590265356] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.2% (3933688 kB) free.
[1590265784] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387412 kB) free!
[1590265956] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (476512 kB) free!
[1590266076] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.7% (4009456 kB) free.
[1590266384] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20440572 kB) free!
[1590266435] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 21.8% (10724660 kB) free.
[1590266676] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.5% (2513272 kB) free!
[1590266796] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 37.3% (6063400 kB) free.
[1590266984] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391448 kB) free!
[1590267584] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20300428 kB) free!
[1590267621] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 13.9% (6860364 kB) free!
[1590268117] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (4164364 kB) free!
[1590268180] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590268184] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402220 kB) free!
[1590268221] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (1041800 kB) free!
[1590268301] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1590268717] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1404048 kB) free!
[1590268784] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19703852 kB) free!
[1590268821] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 92.3% (45477460 kB) free.
[1590269272] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.0% (45329360 kB) free.
[1590269385] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401952 kB) free!
[1590269482] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2508) < 103.13.97.241 [24/May/2020:03:00:37 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590269500] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1590269515] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(341) < 103.16.101.72 [24/May/2020:03:01:32 +0530] TCP 502 0 0 60.001 "10.147.212.49:4000, backend3" "0, 0" "0, 0" "60.001, 0.000"
[1590269602] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(175) < 172.105.59.173 [24/May/2020:03:02:32 +0530] TCP 502 0 0 120.002 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1590269620] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1590269635] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590269637] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(466) < 2020/05/24 03:01:32 [error] 22133#0: *90373765 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590269722] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590269757] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590269917] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2550048 kB) free!
[1590269985] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18199064 kB) free!
[1590270585] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398264 kB) free!
[1590271118] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1590271118] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2342036 kB) free!
[1590271185] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16329548 kB) free!
[1590271785] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418680 kB) free!
[1590272385] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14533652 kB) free!
[1590272985] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32554444 kB) free.
[1590273153] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.6% (5057252 kB) free.
[1590273585] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11904852 kB) free!
[1590273705] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8219372 kB) free!
[1590273753] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (1143268 kB) free!
[1590273825] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4607428 kB) free!
[1590273873] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189900 kB) free!
[1590273993] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (234336 kB) free!
[1590274420] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590274425] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20389292 kB) free!
[1590275025] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (2004104 kB) free!
[1590275197] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1590275625] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16762104 kB) free!
[1590276225] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32998368 kB) free.
[1590276825] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12454428 kB) free!
[1590276946] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8924812 kB) free!
[1590277066] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5443980 kB) free!
[1590278266] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36510656 kB) free.
[1590278579] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1631432 kB) free!
[1590278683] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 44.1% (10818116 kB) free.
[1590278745] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 14.6% (3581892 kB) free!
[1590278805] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3585576 kB) free!
[1590278865] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 14.6% (3585492 kB) free!
[1590278866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10623504 kB) free!
[1590278882] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.6% (10763788 kB) free.
[1590278922] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 55.0% (13484808 kB) free.
[1590278925] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 14.6% (3579092 kB) free!
[1590278986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6481868 kB) free!
[1590279106] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2456672 kB) free!
[1590279706] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15053360 kB) free!
[1590280306] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33634468 kB) free.
[1590280906] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10884176 kB) free!
[1590281026] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6881468 kB) free!
[1590281146] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2457420 kB) free!
[1590281746] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22248168 kB) free!
[1590282346] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35351560 kB) free.
[1590282482] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (211288 kB) free!
[1590282602] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193112 kB) free!
[1590282722] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189504 kB) free!
[1590282946] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17519480 kB) free!
[1590283067] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13854844 kB) free!
[1590283122] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (1623912 kB) free!
[1590283187] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (10251376 kB) free!
[1590283242] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (1603912 kB) free!
[1590283362] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1584520 kB) free!
[1590283379] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (1681272 kB) free!
[1590283787] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22133408 kB) free!
[1590284387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34994416 kB) free.
[1590284626] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 74.7% (98432268 kB) free.
[1590284987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14365660 kB) free!
[1590285107] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10485784 kB) free!
[1590285227] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6623172 kB) free!
[1590285827] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19923852 kB) free!
[1590286379] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (271772 kB) free!
[1590286427] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32793140 kB) free.
[1590287027] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12937496 kB) free!
[1590287147] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9164868 kB) free!
[1590287267] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5433336 kB) free!
[1590287587] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 372 processes
[1590287707] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1590287867] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16400856 kB) free!
[1590288467] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30241060 kB) free.
[1590289067] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8872708 kB) free!
[1590289187] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4894580 kB) free!
[1590289308] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (894400 kB) free!
[1590289908] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13340908 kB) free!
[1590290508] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28497912 kB) free.
[1590291108] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5888872 kB) free!
[1590291162] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1590291228] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1286284 kB) free!
[1590291348] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448868 kB) free!
[1590291517] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.60, 7.84, 5.99
[1590291594] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1590291637] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.52, 8.09, 6.31
[1590291715] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1590291757] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.36, 8.77, 6.76
[1590292483] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4755740 kB) free!
[1590292548] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31219324 kB) free.
[1590292603] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4642560 kB) free!
[1590292723] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (4537608 kB) free!
[1590293123] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(913) < 103.255.146.164 [24/May/2020:09:26:39 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1590293148] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (439812 kB) free!
[1590293243] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590293268] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (433336 kB) free!
[1590293388] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392836 kB) free!
[1590293988] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22382640 kB) free!
[1590294145] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1590294588] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (456660 kB) free!
[1590294745] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1590294865] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1590294985] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1590295123] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1685556 kB) free!
[1590295420] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590295520] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1590295540] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1590295585] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 694 processes
[1590295640] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1590295660] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590295760] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1590296260] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1590296360] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590296389] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13308832 kB) free!
[1590296860] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590296980] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590296989] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (444892 kB) free!
[1590297100] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1590297385] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 727 processes
[1590297495] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (2138592 kB) free!
[1590297506] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 725 processes
[1590297615] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (1539664 kB) free!
[1590297626] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1590297626] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 736 processes
[1590297735] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (858396 kB) free!
[1590297757] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.07, 32.35, 23.89
[1590298450] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 50.6% (12394316 kB) free.
[1590300239] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.1% (5997840 kB) free!
[1590300359] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 15.7% (5502000 kB) free!
[1590300479] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (4985396 kB) free!
[1590300850] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.9% (4641764 kB) free!
[1590300970] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.5% (4284448 kB) free!
[1590301079] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (2515280 kB) free!
[1590301090] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.0% (3931908 kB) free!
[1590301189] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13333968 kB) free!
[1590301690] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2290700 kB) free!
[1590301789] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388240 kB) free!
[1590301900] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1590301900] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 373 processes
[1590301936] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.7% (16101968 kB) free.
[1590302227] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 13.9% (3431480 kB) free!
[1590302280] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.8% (31794196 kB) free.
[1590302347] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 12.6% (3107392 kB) free!
[1590302389] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27879704 kB) free.
[1590302467] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 11.3% (2785480 kB) free!
[1590302500] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 356 processes
[1590302989] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (478528 kB) free!
[1590303021] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8752264 kB) free!
[1590303067] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.8% (22903936 kB) free.
[1590303109] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (429504 kB) free!
[1590303142] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (8438552 kB) free!
[1590303229] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381332 kB) free!
[1590303252] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 555 processes
[1590303262] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (8121996 kB) free!
[1590303372] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 555 processes
[1590303473] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9052324 kB) free!
[1590303492] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 605 processes
[1590303593] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8769576 kB) free!
[1590303713] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.2% (8477864 kB) free!
[1590303829] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18880212 kB) free!
[1590304314] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1590304430] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (489916 kB) free!
[1590304434] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590304596] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2959808 kB) free!
[1590304679] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.41, 6.08, 5.39
[1590304716] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (2685704 kB) free!
[1590304799] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.77, 6.40, 5.56
[1590304836] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (2416832 kB) free!
[1590304919] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.42, 6.72, 5.80
[1590305030] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33895808 kB) free.
[1590305062] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3801676 kB) free!
[1590305436] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3468648 kB) free.
[1590305513] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4669840 kB) free!
[1590305630] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (678836 kB) free!
[1590305750] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (383504 kB) free!
[1590305870] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389120 kB) free!
[1590306470] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22108072 kB) free!
[1590307070] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (427896 kB) free!
[1590307236] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2893832 kB) free!
[1590307356] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (2776620 kB) free!
[1590307467] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1726 processes
[1590307476] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (2654564 kB) free!
[1590307826] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1590307900] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 354 processes
[1590308390] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(13) < {writer,send_failed,{error,timeout}}
[1590308510] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(11) < {writer,send_failed,{error,timeout}}
[1590308536] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (2136828 kB) free!
[1590308630] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(29) < {writer,send_failed,{error,timeout}}
[1590308656] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (1619424 kB) free!
[1590308676] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1590308676] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3284528 kB) free.
[1590308776] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1094324 kB) free!
[1590308923] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 25.5% (6249736 kB) free.
[1590309276] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (2915856 kB) free!
[1590309396] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2864884 kB) free!
[1590309470] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (24052560 kB) free!
[1590309516] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (2825716 kB) free!
[1590310070] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378332 kB) free!
[1590310123] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4329772 kB) free!
[1590310243] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (4202740 kB) free!
[1590310363] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (4086244 kB) free!
[1590310451] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 958 processes
[1590310571] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 961 processes
[1590310691] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 960 processes
[1590311030] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590311880] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 15.3% (5349616 kB) free!
[1590312000] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 14.0% (4911180 kB) free!
[1590312120] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 12.6% (4420892 kB) free!
[1590312163] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2071132 kB) free!
[1590312470] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18165436 kB) free!
[1590312721] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1590312721] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.8% (31815292 kB) free.
[1590312977] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16266524 kB) free.
[1590313044] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(126) < 185.255.8.58 [24/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590313070] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393852 kB) free!
[1590313092] SERVICE ALERT: qservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 777 processes
[1590313136] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.29011.593> (10.147.212.158:53068 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590313164] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590313256] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590313341] SERVICE ALERT: sigqscluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1590313434] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1590313461] SERVICE ALERT: sigqscluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1590313554] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1590313670] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42414592 kB) free.
[1590314031] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1590314152] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590314154] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1590314270] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4625224 kB) free!
[1590314274] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1590314316] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1454572 kB) free!
[1590314390] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406880 kB) free!
[1590314500] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1590314510] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396868 kB) free!
[1590314892] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 885 processes
[1590314916] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (2129592 kB) free!
[1590315110] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40026044 kB) free.
[1590315710] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1228204 kB) free!
[1590315830] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (376852 kB) free!
[1590315950] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384104 kB) free!
[1590316116] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (408040 kB) free!
[1590316561] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1590316681] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(35) < {writer,send_failed,{error,timeout}}
[1590316716] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2702600 kB) free!
[1590316802] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(22) < {writer,send_failed,{error,timeout}}
[1590317293] SERVICE FLAPPING ALERT: qservercluster2;Total Processes;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1590317293] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 705 processes
[1590317456] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.21029.673> (10.147.212.158:34426 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590317576] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590317893] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 779 processes
[1590317916] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1134000 kB) free!
[1590318002] SERVICE FLAPPING ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1590318002] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590318013] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 779 processes
[1590318133] SERVICE ALERT: qservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 775 processes
[1590318350] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23868108 kB) free!
[1590318357] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/24 16:32:54 [error] 22132#0: *90501179 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590318478] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590318516] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.6% (2050536 kB) free!
[1590318539] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.26938.3351> (10.147.212.143:57822 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1590318659] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590318950] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424892 kB) free!
[1590319077] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/24 16:46:17 [error] 22135#0: *90504172 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net"
[1590319197] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590319450] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 42.1% (10330828 kB) free.
[1590319578] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.5% (3478656 kB) free!
[1590319698] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3062496 kB) free!
[1590319818] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (2875640 kB) free!
[1590320150] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13756036 kB) free!
[1590320397] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/24 17:00:40 [error] 22124#0: *90507382 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net"
[1590320418] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1590320418] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1736624 kB) free!
[1590320500] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1590320517] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590320576] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.29825.745> (10.147.212.158:42264 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590320697] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590320750] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424812 kB) free!
[1590320916] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3799636 kB) free.
[1590321117] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/24 17:19:24 [error] 22135#0: *90511277 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590321133] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 543 processes
[1590321237] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590321350] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18844712 kB) free!
[1590321689] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1590321812] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1590321950] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (435604 kB) free!
[1590322116] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2969068 kB) free!
[1590322236] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (179408 kB) free!
[1590322356] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (2729888 kB) free!
[1590322437] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/24 17:42:20 [error] 22135#0: *90516070 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590322550] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33650312 kB) free.
[1590322557] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/05/24 17:45:42 [error] 22132#0: *90516746 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net"
[1590322677] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590322956] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (820456 kB) free!
[1590323150] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2260056 kB) free!
[1590323270] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (375072 kB) free!
[1590323277] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/24 17:56:50 [error] 22135#0: *90519114 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590323390] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408136 kB) free!
[1590323397] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/24 17:58:43 [error] 22135#0: *90519473 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net"
[1590323419] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.9% (16372548 kB) free.
[1590323517] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2020/05/24 18:00:04 [error] 22135#0: *90519744 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net"
[1590323990] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32654656 kB) free.
[1590324121] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590324590] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (433052 kB) free!
[1590324711] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407468 kB) free!
[1590324717] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590324756] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.0% (4883860 kB) free.
[1590324831] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (430104 kB) free!
[1590325431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31186644 kB) free.
[1590325956] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.3% (2331604 kB) free!
[1590326031] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1228512 kB) free!
[1590326076] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3679776 kB) free.
[1590326151] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (450300 kB) free!
[1590326271] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424008 kB) free!
[1590326871] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26498500 kB) free.
[1590326874] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1590326994] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1590327219] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1590327341] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1590327471] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (441952 kB) free!
[1590327591] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391616 kB) free!
[1590327711] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (442908 kB) free!
[1590327876] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.7% (2066900 kB) free!
[1590327996] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3304100 kB) free.
[1590328311] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22161724 kB) free!
[1590328794] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1590328802] SERVICE FLAPPING ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1590328911] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378868 kB) free!
[1590328914] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 342 processes
[1590329034] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1590329511] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29627544 kB) free.
[1590329517] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/24 19:37:03 [error] 22135#0: *90536870 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590329637] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590329796] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (695724 kB) free!
[1590329916] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.2% (1979084 kB) free!
[1590330036] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3052664 kB) free!
[1590330111] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (405648 kB) free!
[1590330133] SERVICE FLAPPING ALERT: qservercluster2;Total Processes;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1590330231] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378404 kB) free!
[1590330351] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406096 kB) free!
[1590330636] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3570580 kB) free.
[1590330952] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24090096 kB) free!
[1590331164] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(516) < 103.255.146.164 [24/May/2020:20:07:58 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590331285] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590331552] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (431212 kB) free!
[1590332037] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/05/24 20:21:14 [error] 22133#0: *90543664 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net"
[1590332152] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32898772 kB) free.
[1590332157] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590332419] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1590332752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (738048 kB) free!
[1590332758] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/24 20:32:49 [error] 22128#0: *90545575 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590332872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406092 kB) free!
[1590332878] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590332992] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389008 kB) free!
[1590333036] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (615036 kB) free!
[1590333156] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.5% (1710020 kB) free!
[1590333179] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (1676316 kB) free!
[1590333276] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (2797488 kB) free!
[1590333478] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/24 20:43:04 [error] 22133#0: *90546860 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590333592] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28949660 kB) free.
[1590333598] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590334192] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (394388 kB) free!
[1590334198] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/24 20:56:09 [error] 22131#0: *90548331 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590334312] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (376488 kB) free!
[1590334318] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590334432] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386172 kB) free!
[1590335032] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21769196 kB) free!
[1590335632] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387916 kB) free!
[1590335676] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (699536 kB) free!
[1590336232] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27351648 kB) free.
[1590336832] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (397176 kB) free!
[1590336834] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1590336953] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (375720 kB) free!
[1590336955] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1590337073] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (452292 kB) free!
[1590337673] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14251456 kB) free!
[1590338076] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3007668 kB) free!
[1590338273] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393764 kB) free!
[1590339085] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1097) < 103.255.146.164 [24/May/2020:22:13:54 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590339205] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590339757] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.61, 10.12, 19.82
[1590339868] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1638 processes
[1590341273] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14300232 kB) free!
[1590341676] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1117636 kB) free!
[1590341873] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412940 kB) free!
[1590342276] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (2657060 kB) free!
[1590343673] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15147044 kB) free!
[1590344076] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (878752 kB) free!
[1590344273] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448180 kB) free!
[1590344676] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3280400 kB) free.