[1547923639] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1547930361] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547952438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547952559] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1547960958] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547961078] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547961198] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547962399] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547963804] HOST ALERT: sig2;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1547963855] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1547963861] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig2' was out of bounds)
[1547963919] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1547963931] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1547963981] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38727 MB (75% inode=98%):
[1547964363] SERVICE ALERT: sig4;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547964366] HOST ALERT: sig4;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.56)
[1547964366] SERVICE ALERT: sig4;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig4' was out of bounds)
[1547964369] SERVICE ALERT: sig4;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4' was out of bounds)
[1547964429] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547964431] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51214 MB (22% inode=99%):
[1547964484] SERVICE ALERT: sig4;Sensors Status;OK;HARD;1;SENSORS OK
[1547964486] SERVICE ALERT: sig4;Total Processes;OK;HARD;1;PROCS OK: 295 processes
[1547964489] SERVICE ALERT: sig4;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547964631] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547964734] SERVICE ALERT: sig1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547964744] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547964764] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1547964778] SERVICE ALERT: sig1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig1' was out of bounds)
[1547964824] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig1' was out of bounds)
[1547964824] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig1' was out of bounds)
[1547964847] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547964854] SERVICE ALERT: sig1;Sensors Status;OK;HARD;1;SENSORS OK
[1547964864] SERVICE ALERT: sig1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 194871 MB (86% inode=99%):
[1547964898] SERVICE ALERT: sig1;Total Processes;OK;HARD;1;PROCS OK: 313 processes
[1547964945] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547964945] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.6% (15375848 kB) free.
[1547965011] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (1122072 kB) free!
[1547965070] SERVICE ALERT: db1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 14242 MB (27% inode=99%):
[1547965231] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1547965351] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1547965471] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1547965791] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547965794] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1547965903] SERVICE ALERT: sig2;CPU_Procs;WARNING;HARD;1;CPU WARNING: 1 warn out of 313 processes
[1547965905] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1547965911] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38727 MB (75% inode=98%):
[1547966024] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1547966126] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547966246] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547966834] SERVICE ALERT: sig4;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547966837] SERVICE ALERT: sig4;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547966851] SERVICE ALERT: sig4;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'sig4' was out of bounds)
[1547966851] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4' was out of bounds)
[1547966854] SERVICE ALERT: sig4;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig4' was out of bounds)
[1547966855] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1547966887] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1547966894] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547966930] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 26.6% (4324768 kB) free.
[1547966954] SERVICE ALERT: sig4;Memory;OK;HARD;1;OK - 95.0% (15433008 kB) free.
[1547966958] SERVICE ALERT: sig4;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1547966971] SERVICE ALERT: sig4;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 220689 MB (98% inode=99%):
[1547966971] SERVICE ALERT: sig4;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37472 MB (73% inode=98%):
[1547966974] SERVICE ALERT: sig4;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.16, 0.06
[1547967014] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1547967354] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547967354] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547967384] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1547967474] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;HARD;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1547967474] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.7% (15399160 kB) free.
[1547967478] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547967799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547967919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547968039] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547968074] SERVICE FLAPPING ALERT: sig1;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1547968074] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1547968639] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1547968639] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547968761] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4446792 kB) free!
[1547971039] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547971120] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24685 MB (10% inode=99%):
[1547971160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547971280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547971761] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5138588 kB) free!
[1547972320] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24939 MB (11% inode=99%):
[1547972331] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (2950348 kB) free!
[1547972451] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.4% (2820244 kB) free!
[1547972460] SERVICE ALERT: sig1;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig1' was out of bounds)
[1547972460] SERVICE ALERT: sig1;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig1' was out of bounds)
[1547972463] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547972479] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547972509] HOST ALERT: webservercluster2;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1547972527] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1547972537] HOST ALERT: webservercluster2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1547972571] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.5% (2686548 kB) free!
[1547972580] SERVICE ALERT: sig1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1547972580] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.18, 0.06
[1547972599] SERVICE ALERT: webservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1547972831] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1547972920] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23403 MB (10% inode=99%):
[1547972934] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 454 processes
[1547972951] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1547972995] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 453 processes
[1547973071] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1547973771] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1347024 kB) free!
[1547974871] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1547975361] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3932208 kB) free!
[1547975471] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1547975591] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1547975711] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1547976447] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547976567] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1547978231] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47135 MB (20% inode=99%):
[1547978351] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47135 MB (20% inode=99%):
[1547978471] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47095 MB (20% inode=99%):
[1547978480] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547978967] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547979080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547979087] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1547979200] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547979201] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1547979321] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1547979441] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1547979474] SERVICE FLAPPING ALERT: sig1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547979800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547979921] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547980041] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547980042] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 319 processes
[1547980288] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547980408] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547980528] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547980642] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1547981128] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547981248] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547981368] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547981842] SERVICE FLAPPING ALERT: webservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1547981842] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1547981968] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547982088] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547982208] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547982809] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1547983409] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1547983529] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547983649] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547984249] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1547984249] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547984841] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547984849] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547984969] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547985089] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547985813] HOST ALERT: sig4;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1547985937] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547986041] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547986161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547986281] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547987832] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.31, 6.88, 6.09
[1547987952] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.47, 5.76, 5.77
[1547987961] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4998108 kB) free!
[1547988072] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.87, 4.77, 5.41
[1547988089] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1547988672] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.11, 3.30, 4.29
[1547989161] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2240208 kB) free!
[1547989889] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547990010] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547990130] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547991442] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 2 warn out of 320 processes
[1547991562] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 319 processes
[1547991683] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 320 processes
[1547991930] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547992283] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1547993130] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547993250] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547993371] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547993970] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547994090] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547994210] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547994682] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547995282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547995402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547995410] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547995511] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1547995522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547995824] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547995827] HOST ALERT: sig4;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.56)
[1547995891] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547995944] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 296 processes
[1547996010] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547996064] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547996111] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1547996130] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547996251] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547996471] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24066 MB (10% inode=99%):
[1547996850] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547996970] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547997090] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547997443] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 31.6% (7733680 kB) free.
[1547997690] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547997810] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547997930] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547998283] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 326 processes
[1547998403] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 326 processes
[1547998511] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1547998523] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1547998611] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.3% (5060500 kB) free!
[1547999122] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547999123] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 325 processes
[1547999130] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1547999243] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 325 processes
[1547999251] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547999363] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 325 processes
[1547999471] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 28695 MB (12% inode=99%):
[1547999722] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547999811] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4870196 kB) free!
[1547999842] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547999850] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547999962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547999963] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1547999970] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548000820] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46976 MB (20% inode=99%):
[1548000911] SERVICE FLAPPING ALERT: webservercluster2;Current Users;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1548000911] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1548000940] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46848 MB (20% inode=99%):
[1548001060] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46848 MB (20% inode=99%):
[1548001163] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 318 processes
[1548001284] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 319 processes
[1548001291] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1548001404] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 318 processes
[1548001411] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1548001531] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1548001643] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4222172 kB) free!
[1548001762] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1548001764] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (4122824 kB) free!
[1548001884] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (4013848 kB) free!
[1548002132] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1548002604] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 317 processes
[1548002970] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548003090] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548003204] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1548003211] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548003562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1548003682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1548003684] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 28.5% (6980052 kB) free.
[1548003802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1548003804] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 312 processes
[1548003811] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548003924] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 311 processes
[1548004044] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 311 processes
[1548005011] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548005131] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548005251] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548007051] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548007096] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 4 warn out of 332 processes
[1548007172] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548007292] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548007884] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4890864 kB) free!
[1548007892] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548008002] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1548008004] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4850392 kB) free!
[1548008124] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1548008124] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (4812056 kB) free!
[1548008296] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 328 processes