[1589049070] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4815492 kB) free!
[1589049096] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1589049182] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589049216] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1589049336] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1589049670] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19558424 kB) free!
[1589050270] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37538156 kB) free.
[1589050870] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14848184 kB) free!
[1589050990] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10733704 kB) free!
[1589051110] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (6836264 kB) free!
[1589051159] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196300 kB) free!
[1589051220] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178164 kB) free!
[1589051280] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (189268 kB) free!
[1589051340] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (195984 kB) free!
[1589051582] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (253248 kB) free!
[1589051702] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (1090152 kB) free!
[1589051710] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21894876 kB) free!
[1589051822] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (479132 kB) free!
[1589051940] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 76.2% (18667636 kB) free.
[1589052311] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27099828 kB) free.
[1589052911] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15247784 kB) free!
[1589053031] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11324788 kB) free!
[1589053151] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7491572 kB) free!
[1589053751] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (22010176 kB) free!
[1589053791] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589054011] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5954104 kB) free!
[1589054131] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (775700 kB) free!
[1589054251] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397332 kB) free!
[1589054351] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37196688 kB) free.
[1589054422] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220380 kB) free!
[1589054542] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217064 kB) free!
[1589054662] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213076 kB) free!
[1589054851] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.6% (13948888 kB) free!
[1589054951] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15221244 kB) free!
[1589055071] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11308112 kB) free!
[1589055191] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7549864 kB) free!
[1589055791] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19851372 kB) free!
[1589056391] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (36998496 kB) free.
[1589056462] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.9% (6845032 kB) free!
[1589056991] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13789268 kB) free!
[1589057111] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10002152 kB) free!
[1589057231] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6365064 kB) free!
[1589057831] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18568416 kB) free!
[1589058422] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (2801912 kB) free!
[1589058431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32600828 kB) free.
[1589058497] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 85.4% (42067048 kB) free.
[1589058932] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.9% (44795140 kB) free.
[1589059022] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4060268 kB) free.
[1589059032] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11241728 kB) free!
[1589059152] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7616588 kB) free!
[1589059272] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36868024 kB) free.
[1589059872] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15528572 kB) free!
[1589059886] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(340) < 49.50.86.152 [10/May/2020:03:00:23 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589059978] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1589059992] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11760944 kB) free!
[1589060007] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589060021] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(447) < 2020/05/10 03:01:32 [error] 22131#0: *87401145 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589060052] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (3871464 kB) free!
[1589060062] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.1% (44381192 kB) free.
[1589060074] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2512) < 103.13.97.241 [10/May/2020:03:02:22 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1589060112] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (8142372 kB) free!
[1589060141] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589060194] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589060652] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (1312880 kB) free!
[1589060712] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13597176 kB) free!
[1589060896] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.9% (5604064 kB) free.
[1589061312] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31830072 kB) free.
[1589061496] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (210428 kB) free!
[1589061616] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191684 kB) free!
[1589061736] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (225732 kB) free!
[1589061912] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5000508 kB) free!
[1589062032] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (912304 kB) free!
[1589062152] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37189600 kB) free.
[1589062752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10091092 kB) free!
[1589062872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6349324 kB) free!
[1589062992] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2710512 kB) free!
[1589063592] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17888800 kB) free!
[1589064192] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32970504 kB) free.
[1589064793] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13444568 kB) free!
[1589064913] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9827072 kB) free!
[1589065033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39251528 kB) free.
[1589065633] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19200552 kB) free!
[1589065753] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15511700 kB) free!
[1589065873] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (12020596 kB) free!
[1589066251] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589066473] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21498320 kB) free!
[1589067073] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31273496 kB) free.
[1589067673] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12720132 kB) free!
[1589067793] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37489636 kB) free.
[1589068393] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18360500 kB) free!
[1589068436] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 57.0% (14064320 kB) free.
[1589068452] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 57.7% (14153744 kB) free.
[1589068476] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 74.1% (18172228 kB) free.
[1589068513] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14882200 kB) free!
[1589068633] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41330744 kB) free.
[1589069233] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22626064 kB) free!
[1589069353] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18859412 kB) free!
[1589069474] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41315372 kB) free.
[1589070073] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22922336 kB) free!
[1589070193] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (18963516 kB) free!
[1589070313] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15048912 kB) free!
[1589071051] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.4% (29585324 kB) free.
[1589071513] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32308652 kB) free.
[1589071651] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.1% (15958136 kB) free!
[1589071771] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.6% (15289180 kB) free!
[1589071891] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.2% (14721232 kB) free!
[1589072491] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (679652 kB) free!
[1589072713] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21721764 kB) free!
[1589072833] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18103776 kB) free!
[1589072953] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14623624 kB) free!
[1589073091] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1589073091] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 40.0% (52685036 kB) free.
[1589073236] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170276 kB) free!
[1589073356] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171304 kB) free!
[1589073476] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182788 kB) free!
[1589074153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30774672 kB) free.
[1589075353] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19940780 kB) free!
[1589075473] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16231444 kB) free!
[1589075594] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39917308 kB) free.
[1589076193] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20386252 kB) free!
[1589076240] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4714372 kB) free!
[1589076300] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4715856 kB) free!
[1589076313] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (16550592 kB) free!
[1589076360] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.2% (4715204 kB) free!
[1589076420] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.2% (4715704 kB) free!
[1589076433] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39732476 kB) free.
[1589077033] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21098508 kB) free!
[1589077153] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17308364 kB) free!
[1589077273] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13543536 kB) free!
[1589078473] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29185588 kB) free.
[1589079673] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16502816 kB) free!
[1589079793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12373024 kB) free!
[1589079913] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37946348 kB) free.
[1589080513] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16671160 kB) free!
[1589080633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12441708 kB) free!
[1589080753] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18685044 kB) free!
[1589081885] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.60, 6.12, 5.09
[1589082005] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.80, 7.18, 5.61
[1589082125] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.77, 8.15, 6.19
[1589082553] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35119236 kB) free.
[1589083153] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12718980 kB) free!
[1589083273] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8521880 kB) free!
[1589083393] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4205584 kB) free!
[1589083452] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174544 kB) free!
[1589083572] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198080 kB) free!
[1589083692] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167092 kB) free!
[1589084491] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589084526] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.04, 3.31, 4.47
[1589084593] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24595392 kB) free!
[1589085126] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.08, 6.21, 5.46
[1589085179] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1589085193] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404424 kB) free!
[1589085246] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.11, 5.52, 5.30
[1589085299] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1589085366] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.35, 5.52, 5.33
[1589085419] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1589085793] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13313808 kB) free!
[1589086010] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1589086019] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1589086130] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589086250] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589086476] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183360 kB) free!
[1589086596] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166588 kB) free!
[1589086716] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193824 kB) free!
[1589086850] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1589086993] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (516512 kB) free!
[1589087146] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 55.8% (13681900 kB) free.
[1589087219] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1589087593] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14618036 kB) free!
[1589088193] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28164840 kB) free.
[1589088793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2050276 kB) free!
[1589088914] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384508 kB) free!
[1589088945] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 62.5% (15308596 kB) free.
[1589089019] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1589089019] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1589089022] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2712312 kB) free!
[1589089034] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15220616 kB) free!
[1589089143] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3420688 kB) free.
[1589089533] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4497128 kB) free!
[1589089634] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5466428 kB) free!
[1589089653] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4314316 kB) free!
[1589089742] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3226800 kB) free!
[1589089773] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.4% (4126460 kB) free!
[1589089862] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3181424 kB) free!
[1589089982] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2760536 kB) free!
[1589090234] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13321708 kB) free!
[1589090582] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3341576 kB) free.
[1589090973] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (477128 kB) free!
[1589091182] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (2363240 kB) free!
[1589091302] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2864860 kB) free!
[1589091422] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2812412 kB) free!
[1589091434] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405800 kB) free!
[1589092020] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1589092034] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18688924 kB) free!
[1589092489] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 33.6% (5460424 kB) free.
[1589092634] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412052 kB) free!
[1589092773] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.5% (16287436 kB) free.
[1589093220] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1589093234] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18049940 kB) free!
[1589093834] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (449768 kB) free!
[1589095034] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34673152 kB) free.
[1589095594] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(845) < 103.255.146.164 [10/May/2020:12:55:30 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589095634] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6719228 kB) free!
[1589095714] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589095754] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2439552 kB) free!
[1589095874] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397852 kB) free!
[1589096166] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 6.43, 5.01, 4.92
[1589096474] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14147148 kB) free!
[1589096766] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.19, 6.06, 5.47
[1589096820] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1589096886] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.63, 6.68, 5.77
[1589097006] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.51, 5.89, 5.60
[1589097074] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32565056 kB) free.
[1589097345] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.0% (4164084 kB) free!
[1589097466] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.5% (4037624 kB) free!
[1589097516] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (3851888 kB) free!
[1589097586] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.0% (3928100 kB) free!
[1589097674] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4460264 kB) free!
[1589097794] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407912 kB) free!
[1589097914] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40105812 kB) free.
[1589098114] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(137) < 103.255.146.164 [10/May/2020:13:31:22 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589098234] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589098514] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12081700 kB) free!
[1589098634] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7425232 kB) free!
[1589098754] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3181220 kB) free!
[1589098806] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.38, 2.61, 3.99
[1589099220] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 352 processes
[1589099316] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203908 kB) free!
[1589099354] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17519560 kB) free!
[1589099386] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 55.8% (13684516 kB) free.
[1589099820] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1589099823] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1591192 kB) free!
[1589099897] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9707976 kB) free!
[1589099954] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32744888 kB) free.
[1589100017] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9524460 kB) free!
[1589100137] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9379996 kB) free!
[1589100336] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 329 processes
[1589100421] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 354 processes
[1589100423] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1589100423] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3378692 kB) free.
[1589100456] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1589100554] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6574344 kB) free!
[1589100674] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2410928 kB) free!
[1589100794] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38379748 kB) free.
[1589101169] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;1;USERS CRITICAL - 11 users currently logged in
[1589101289] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;USERS CRITICAL - 12 users currently logged in
[1589101394] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10829504 kB) free!
[1589101409] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1589101514] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5763428 kB) free!
[1589101623] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.2% (1651964 kB) free!
[1589101635] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (500384 kB) free!
[1589101743] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (1500772 kB) free!
[1589101806] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.95, 11.07, 7.74
[1589101834] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(239) < 103.255.146.164 [10/May/2020:14:31:04 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589101863] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1521728 kB) free!
[1589101926] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.02, 9.41, 7.56
[1589101954] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589102046] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1589102046] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.43, 7.69, 7.15
[1589102235] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18745896 kB) free!
[1589102463] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3315064 kB) free.
[1589102835] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33871288 kB) free.
[1589102973] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4316312 kB) free!
[1589103063] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2993164 kB) free!
[1589103093] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.2% (16205568 kB) free.
[1589103154] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(837) < 185.255.8.58 [10/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1589103183] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2757580 kB) free!
[1589103275] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589103303] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2535160 kB) free!
[1589103435] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6935356 kB) free!
[1589103555] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2562376 kB) free!
[1589103675] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377468 kB) free!
[1589104275] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13395760 kB) free!
[1589104446] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.77, 3.47, 4.30
[1589104875] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29030240 kB) free.
[1589105074] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(634) < 103.255.146.164 [10/May/2020:15:31:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589105194] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589105475] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3673700 kB) free!
[1589105537] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4495916 kB) free!
[1589105595] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384188 kB) free!
[1589105703] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1294856 kB) free!
[1589105715] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34717048 kB) free.
[1589105821] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1589105916] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 32.5% (7959500 kB) free.
[1589106303] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2540072 kB) free!
[1589106315] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7931504 kB) free!
[1589106435] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3512448 kB) free!
[1589106555] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (12097736 kB) free!
[1589106995] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(345) < 103.255.146.164 [10/May/2020:16:01:19 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589107115] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589107755] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27211852 kB) free.
[1589108103] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1411396 kB) free!
[1589108356] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1703644 kB) free!
[1589108476] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36913760 kB) free.
[1589109076] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10979156 kB) free!
[1589109196] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6253972 kB) free!
[1589109303] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.5% (2033760 kB) free!
[1589109316] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1862460 kB) free!
[1589109421] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 365 processes
[1589109541] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 364 processes
[1589109586] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4771428 kB) free!
[1589109661] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 371 processes
[1589109707] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.2% (4694124 kB) free!
[1589109827] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.8% (4615072 kB) free!
[1589109916] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16074912 kB) free!
[1589110116] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(987) < 103.255.146.164 [10/May/2020:16:54:21 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589110236] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589110516] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30150640 kB) free.
[1589111009] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1589111116] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4165572 kB) free!
[1589111236] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (415368 kB) free!
[1589111356] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37008620 kB) free.
[1589111956] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7819784 kB) free!
[1589112076] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3183588 kB) free!
[1589112196] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383320 kB) free!
[1589112227] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2341072 kB) free!
[1589112246] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.60, 6.20, 5.46
[1589112303] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1584620 kB) free!
[1589112333] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9467716 kB) free!
[1589112366] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.23, 6.27, 5.58
[1589112453] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9381456 kB) free!
[1589112486] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.52, 5.69, 5.46
[1589112573] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9309872 kB) free!
[1589113086] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.42, 2.82, 3.99
[1589113116] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4837204 kB) free!
[1589113236] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4783756 kB) free!
[1589113356] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4746624 kB) free!
[1589113396] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28921580 kB) free.
[1589113997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2232312 kB) free!
[1589114103] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (2288560 kB) free!
[1589114117] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 20.0% (26365124 kB) free!
[1589114237] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30173192 kB) free.
[1589114827] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (7160832 kB) free!
[1589114837] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2937820 kB) free!
[1589114957] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (449964 kB) free!
[1589115077] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23074676 kB) free!
[1589115677] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374696 kB) free!
[1589116503] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1362136 kB) free!
[1589116627] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4838560 kB) free!
[1589116877] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36806864 kB) free.
[1589117103] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3123432 kB) free!
[1589117477] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6765284 kB) free!
[1589117597] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (2001292 kB) free!
[1589117717] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399100 kB) free!
[1589117827] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (5498816 kB) free!
[1589118061] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1589118156] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2139732 kB) free!
[1589118427] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1589118427] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4413280 kB) free!
[1589118917] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26373676 kB) free!
[1589119517] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1415688 kB) free!
[1589119686] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 19.63, 7.39, 4.86
[1589119774] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4919240 kB) free!
[1589119806] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.84, 5.82, 4.59
[1589120118] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14113892 kB) free!
[1589120718] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28574528 kB) free.
[1589121318] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3411448 kB) free!
[1589121438] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36647652 kB) free.
[1589122038] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10268380 kB) free!
[1589122158] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5793792 kB) free!
[1589122278] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1665532 kB) free!
[1589123103] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3493764 kB) free.
[1589123478] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25279148 kB) free!
[1589124078] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (1916456 kB) free!
[1589124303] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2217720 kB) free!
[1589124423] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3338536 kB) free.
[1589124606] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589125023] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2101752 kB) free!
[1589125143] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3155736 kB) free!
[1589125263] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3126488 kB) free!
[1589125278] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21665836 kB) free!
[1589126478] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11780236 kB) free!
[1589127078] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22480008 kB) free!
[1589127678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34655536 kB) free.
[1589128263] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3866748 kB) free.
[1589128279] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9839636 kB) free!
[1589128399] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5706288 kB) free!
[1589128519] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34440428 kB) free.
[1589128864] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (2652452 kB) free!
[1589128984] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3826700 kB) free.
[1589129119] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11053112 kB) free!
[1589129239] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6895684 kB) free!
[1589129359] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3089120 kB) free!
[1589129827] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589130559] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21688708 kB) free!
[1589131159] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32974944 kB) free.
[1589131759] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9088136 kB) free!
[1589131879] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5235452 kB) free!
[1589131999] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33217108 kB) free.
[1589132599] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9625596 kB) free!
[1589132719] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5845868 kB) free!
[1589132839] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33189420 kB) free.
[1589133184] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2894288 kB) free!
[1589133304] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3561984 kB) free.
[1589133439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10143612 kB) free!
[1589133560] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6416620 kB) free!
[1589133680] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32993264 kB) free.
[1589134280] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10113796 kB) free!
[1589134400] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6397228 kB) free!
[1589134504] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3084132 kB) free!
[1589134520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32163728 kB) free.
[1589134624] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3504788 kB) free.
[1589135120] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9958704 kB) free!
[1589135240] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6331812 kB) free!
[1589135360] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31248768 kB) free.