[1566934152] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.2% (9922420 kB) free.
[1566934541] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 293 processes
[1566934893] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (1050476 kB) free!
[1566935013] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1308184 kB) free!
[1566935133] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1306320 kB) free!
[1566935141] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 288 processes
[1566935255] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 16.47, 15.82, 7.05
[1566935352] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (4213196 kB) free!
[1566935375] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.41, 10.70, 6.24
[1566935472] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4212052 kB) free!
[1566935495] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.49, 7.24, 5.52
[1566935592] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (4210468 kB) free!
[1566936095] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.66, 1.43, 3.15
[1566936192] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1491752 kB) free!
[1566937175] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.28, 16.01, 9.18
[1566937295] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.03, 10.89, 8.13
[1566937415] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 7.38, 7.18
[1566938016] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.25, 1.95, 4.26
[1566939355] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238164 kB) free!
[1566939475] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240396 kB) free!
[1566939595] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231752 kB) free!
[1566940195] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.5% (44610440 kB) free.
[1566941960] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 24.6% (6031056 kB) free.
[1566941975] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 35.3% (8659280 kB) free.
[1566943293] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 69.83, 24.95, 9.47
[1566943413] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.89, 17.03, 8.46
[1566943533] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.48, 11.46, 7.46
[1566944133] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.11, 1.64, 3.97
[1566952533] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.2% (13451036 kB) free!
[1566952993] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.6% (14717836 kB) free.
[1566953733] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 37.9% (49984508 kB) free.
[1566953972] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 480 processes
[1566954032] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 476 processes
[1566957934] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566958054] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566959961] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4852628 kB) free!
[1566960081] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4829660 kB) free!
[1566960201] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4813700 kB) free!
[1566961393] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.3% (3279188 kB) free!
[1566961513] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (1581544 kB) free!
[1566961633] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1534712 kB) free!
[1566965375] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4830252 kB) free!
[1566965495] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4772136 kB) free!
[1566965615] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (4714488 kB) free!
[1566965863] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.7% (22635636 kB) free.
[1566965958] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40729420 kB) free.
[1566967055] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 301 processes
[1566967175] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1566967401] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2350136 kB) free!
[1566968615] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1589760 kB) free!
[1566971274] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (301136 kB) free!
[1566971394] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (256128 kB) free!
[1566971514] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202736 kB) free!
[1566972032] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 508 processes
[1566972092] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 508 processes
[1566973177] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 312 processes
[1566973297] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1566974733] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 56.76, 90.25, 52.63
[1566974853] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8.06, 60.53, 46.30
[1566974973] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.44, 40.63, 40.73
[1566975114] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2772844 kB) free!
[1566975454] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566975574] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1566975693] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566975714] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1566975714] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187408 kB) free!
[1566976162] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.1% (5986752 kB) free!
[1566976173] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.18, 2.51, 12.25
[1566976282] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.2% (5499688 kB) free!
[1566976403] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4831020 kB) free!
[1566977141] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566977261] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566977381] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566977973] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.68, 2.99
[1566981697] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1566981817] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1566982181] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566983373] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 23.06, 10.03, 4.72
[1566983374] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566983395] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8824396 kB) free!
[1566983493] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.32, 7.84, 4.55
[1566983494] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1566983516] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8587676 kB) free!
[1566983617] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 320 processes
[1566983636] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.0% (8374888 kB) free!
[1566983737] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1566984131] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40744852 kB) free.
[1566984168] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.2% (23344688 kB) free.
[1566984202] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.1% (22819868 kB) free.
[1566984245] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21995392 kB) free.
[1566984693] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566986036] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4518236 kB) free!
[1566986653] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191244 kB) free!
[1566986773] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190980 kB) free!
[1566986894] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191364 kB) free!
[1566987215] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194276 kB) free!
[1566987335] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (205180 kB) free!
[1566987455] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184364 kB) free!
[1566987772] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190440 kB) free!
[1566987892] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185652 kB) free!
[1566988013] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172004 kB) free!
[1566988331] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9649484 kB) free!
[1566988451] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8876064 kB) free!
[1566988571] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (8096856 kB) free!
[1566989137] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 319 processes
[1566989258] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1566989771] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (907852 kB) free!
[1566991197] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.6% (22360324 kB) free.
[1566991230] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.5% (22668520 kB) free.
[1566991260] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1566991260] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 59.1% (29130400 kB) free.
[1566991294] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 34.4% (16938016 kB) free.
[1566991337] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.6% (23180812 kB) free.
[1566991801] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190540 kB) free!
[1566991864] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.4% (5627228 kB) free!
[1566991921] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185556 kB) free!
[1566991984] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215208 kB) free!
[1566992041] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168900 kB) free!
[1566992104] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229772 kB) free!
[1566992257] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 318 processes
[1566992377] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1566994052] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.27, 5.48, 5.02
[1566994172] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.28, 5.05, 4.92
[1566995741] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.7% (4824868 kB) free!
[1566995861] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4763300 kB) free!
[1566995973] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.73, 5.33, 5.21
[1566995981] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.2% (4716296 kB) free!
[1566996093] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.76, 5.01, 5.10
[1566996093] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566996213] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.44, 5.58, 5.29
[1566996817] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46400 MB (20% inode=99%):
[1566996937] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45740 MB (20% inode=99%):
[1566997057] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44966 MB (20% inode=99%):
[1566998014] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.41, 5.68, 5.12
[1566998134] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 3.44, 4.78, 4.86
[1566998493] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566999934] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.02, 5.55, 5.19
[1567000054] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.70, 5.19, 5.10
[1567000174] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.68, 5.58, 5.26
[1567000233] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (203604 kB) free!
[1567000293] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.93, 5.75, 5.24
[1567000342] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194716 kB) free!
[1567000353] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202464 kB) free!
[1567000413] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.73, 5.53, 5.23
[1567000462] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194824 kB) free!
[1567000473] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183932 kB) free!
[1567000533] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.43, 4.52, 4.91
[1567000582] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164812 kB) free!
[1567000774] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 4.22, 3.92, 4.47
[1567001614] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.42, 3.58, 4.62
[1567002841] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567002904] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567003181] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2390040 kB) free!
[1567003416] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 156.42, 101.56, 45.41
[1567003536] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 156.76, 119.87, 58.93
[1567003657] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 158.01, 132.53, 70.95
[1567005457] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.49, 2.00, 15.83
[1567007257] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.00, 1.12, 3.35
[1567009781] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1567009894] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567009901] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1567009904] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9874568 kB) free!
[1567010014] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567010021] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1567010024] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9819708 kB) free!
[1567010144] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.8% (9763688 kB) free!
[1567012057] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24614 MB (10% inode=99%):
[1567012421] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in