[1562869876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562869944] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7225116 kB) free!
[1562869996] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562870064] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11828256 kB) free.
[1562870596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562870664] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7230608 kB) free!
[1562870695] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562870711] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181888 kB) free!
[1562870717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562870784] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11819164 kB) free.
[1562870785] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201532 kB) free!
[1562870815] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1562870831] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173796 kB) free!
[1562870905] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190572 kB) free!
[1562870951] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175832 kB) free!
[1562871025] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184452 kB) free!
[1562871316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562871384] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7421024 kB) free!
[1562871436] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562871504] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11827244 kB) free.
[1562872036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562872078] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 113.43, 45.05, 18.35
[1562872105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7459408 kB) free!
[1562872156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562872198] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 129.07, 73.60, 32.05
[1562872225] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11817768 kB) free.
[1562872318] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 131.02, 92.66, 44.07
[1562872756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562872825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7363512 kB) free!
[1562872876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562872945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11828560 kB) free.
[1562873015] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172244 kB) free!
[1562873135] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (183932 kB) free!
[1562873255] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (331152 kB) free!
[1562873414] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 420 processes
[1562873474] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 425 processes
[1562873476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562873545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7504736 kB) free!
[1562873597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562873665] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11826448 kB) free.
[1562873693] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.39, 0.41, 16.59
[1562874118] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.66, 3.05, 12.09
[1562874196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562874265] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7682224 kB) free!
[1562874316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562874385] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11819560 kB) free.
[1562874436] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562874893] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.45, 4.91
[1562874985] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7839820 kB) free!
[1562875036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562875105] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11827432 kB) free.
[1562875156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562875705] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8000912 kB) free!
[1562875756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562875825] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11820704 kB) free.
[1562875876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562875918] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.56, 2.60, 4.05
[1562876426] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8101596 kB) free!
[1562876477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562876546] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11824676 kB) free.
[1562876597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562877146] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8355212 kB) free!
[1562877197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562877266] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11827988 kB) free.
[1562877317] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562877866] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8437232 kB) free!
[1562877890] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14260.06, 8757.75, 3748.06
[1562877893] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 7.74, 281.40, 227.58
[1562877917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562877986] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11820960 kB) free.
[1562878010] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14376.02, 10622.69, 5038.89
[1562878013] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.29, 188.38, 200.02
[1562878037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1562878130] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14391.80, 11870.58, 6173.21
[1562878133] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562878133] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.47, 126.13, 175.79
[1562878586] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8476348 kB) free!
[1562878637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562878706] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11722048 kB) free.
[1562878751] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.8% (3141580 kB) free!
[1562878757] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562879306] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8538992 kB) free!
[1562879358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562879426] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6847980 kB) free!
[1562879478] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562879546] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11822668 kB) free.
[1562880078] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562880146] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7025708 kB) free!
[1562880198] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562880267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11829640 kB) free.
[1562880798] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562880866] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7121652 kB) free!
[1562880918] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562880986] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11816172 kB) free.
[1562881133] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.39, 2.66, 17.67
[1562881518] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562881638] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562882186] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8528604 kB) free!
[1562882239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562882306] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6826952 kB) free!
[1562882359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562882426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11817704 kB) free.
[1562882933] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.41, 0.45, 2.93
[1562882959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562883026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7026796 kB) free!
[1562883079] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562883146] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11827016 kB) free.
[1562883679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562883718] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 12101.46, 5110.00, 1923.01
[1562883746] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7152104 kB) free!
[1562883799] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562883838] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 13121.62, 7813.51, 3300.13
[1562883866] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11826400 kB) free.
[1562883899] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5605 MB (10% inode=97%):
[1562883958] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 13259.44, 9622.52, 4510.29
[1562884130] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.15, 0.27, 15.75
[1562884399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562884467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7294836 kB) free!
[1562884499] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6489 MB (12% inode=97%):
[1562884519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562884587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11818908 kB) free.
[1562884629] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (3134168 kB) free!
[1562885120] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562885187] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7422044 kB) free!
[1562885240] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562885307] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11826332 kB) free.
[1562885330] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 0.29, 4.59
[1562885840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562885907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7563980 kB) free!
[1562885933] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 271.31, 481.32, 227.40
[1562885951] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452196 kB) free!
[1562885960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562886027] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11819024 kB) free.
[1562886053] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 36.72, 322.14, 199.86
[1562886173] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.20, 215.66, 175.66
[1562886560] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562886627] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7649680 kB) free!
[1562886680] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562886747] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11823396 kB) free.
[1562887118] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (13172868 kB) free!
[1562887238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (21774752 kB) free.
[1562887280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562887347] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7764776 kB) free!
[1562887400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562887467] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11826676 kB) free.
[1562887521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562888067] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7832552 kB) free!
[1562888120] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562888187] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11819096 kB) free.
[1562888240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562888360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562888438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11846044 kB) free!
[1562888558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (21444844 kB) free.
[1562888573] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.02, 0.80, 13.82
[1562889387] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9192488 kB) free!
[1562889507] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7898828 kB) free!
[1562889530] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2548.93, 2385.67, 1054.48
[1562889551] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2639268 kB) free!
[1562889627] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11793744 kB) free.
[1562889650] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 344.02, 1596.38, 926.65
[1562889770] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1562889770] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 46.93, 1068.38, 814.37
[1562889773] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 0.35, 4.09
[1562889958] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.87, 3.12, 19.59
[1562890827] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9330440 kB) free!
[1562890947] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (8008156 kB) free!
[1562890958] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11584288 kB) free!
[1562891067] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11774388 kB) free.
[1562891079] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (21289404 kB) free.
[1562891158] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.92, 256.20, 198.56
[1562891351] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2424064 kB) free!
[1562891829] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2163804 kB) free!
[1562892267] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9545004 kB) free!
[1562892387] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8269964 kB) free!
[1562892455] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 52.6% (8553480 kB) free.
[1562892507] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11753924 kB) free.
[1562892624] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4909524 kB) free!
[1562892744] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4907844 kB) free!
[1562892864] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4904112 kB) free!
[1562893160] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562893370] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 0.19, 17.07
[1562893558] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562893558] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.10, 3.13, 17.72
[1562894307] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8496472 kB) free!
[1562894360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562894427] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6821512 kB) free!
[1562894480] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562894547] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11746000 kB) free.
[1562894570] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.28, 0.32, 4.92
[1562895080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562895147] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6864348 kB) free!
[1562895170] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.85, 137.62, 147.08
[1562895201] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562895267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11749468 kB) free.
[1562895290] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.29, 92.17, 129.27
[1562895358] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.67, 2.76, 4.94
[1562895410] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.45, 61.81, 113.65
[1562895800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562895867] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6925276 kB) free!
[1562895920] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562895987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11738272 kB) free.
[1562896520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562896640] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562897187] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8010560 kB) free!
[1562897210] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 0.41, 16.58
[1562897240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562897307] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11744772 kB) free.
[1562897360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562897480] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562897678] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12204444 kB) free!
[1562897798] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (21093632 kB) free.
[1562898237] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4257768 kB) free!
[1562898357] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (245988 kB) free!
[1562898410] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.10, 0.22, 4.73
[1562898477] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223732 kB) free!
[1562898680] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562898680] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562898974] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 422 processes
[1562898998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (12600196 kB) free!
[1562899034] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 421 processes
[1562899094] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 416 processes
[1562899107] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7883780 kB) free!
[1562899118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (21644700 kB) free.
[1562899227] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11735944 kB) free.
[1562899827] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7852992 kB) free!
[1562899880] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562899947] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11741144 kB) free.
[1562900000] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562900547] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7874724 kB) free!
[1562900573] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562900600] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562900667] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11742012 kB) free.
[1562900720] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562901267] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7909304 kB) free!
[1562901320] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562901387] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11732944 kB) free.
[1562901441] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562901518] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11576464 kB) free!
[1562901638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (20982324 kB) free.
[1562901958] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 28.08, 1155.02, 1020.37
[1562901987] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7945300 kB) free!
[1562902040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562902078] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6.13, 773.83, 897.02
[1562902107] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11739048 kB) free.
[1562902160] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562902198] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.85, 518.89, 788.65
[1562902373] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8209.94, 3972.16, 1554.48
[1562902494] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8548.67, 5503.87, 2408.92
[1562902614] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2452.86, 4618.96, 2492.33
[1562902708] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7965152 kB) free!
[1562902760] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562902828] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11727456 kB) free.
[1562902880] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562903000] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562903438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (10258304 kB) free!
[1562903558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (20445232 kB) free.
[1562903676] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3167688 kB) free!
[1562903796] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (3079720 kB) free!
[1562903916] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.5% (2999252 kB) free!
[1562904758] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (9780192 kB) free!
[1562904800] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562904878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (18756028 kB) free.
[1562905228] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8762820 kB) free!
[1562905348] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7667492 kB) free!
[1562905401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562905468] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11729064 kB) free.
[1562905521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562905716] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (917864 kB) free!
[1562905798] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.26, 3.24, 19.37
[1562906068] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8387792 kB) free!
[1562906078] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (8994768 kB) free!
[1562906121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562906188] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6791172 kB) free!
[1562906198] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (19753224 kB) free.
[1562906241] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562906308] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11732572 kB) free.
[1562906762] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.53, 7.42, 5.01
[1562906841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562906882] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.82, 8.95, 5.91
[1562906961] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562907002] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.97, 9.04, 6.32
[1562907398] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (9982456 kB) free!
[1562907414] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.15, 4.41, 18.28
[1562907508] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7962840 kB) free!
[1562907518] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (19884456 kB) free.
[1562907628] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11720568 kB) free.
[1562908198] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.16, 3.34, 4.55
[1562908718] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (11281040 kB) free!
[1562908761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562908838] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (7148060 kB) free!
[1562908881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562908958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (18810048 kB) free.
[1562909001] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562909210] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562909428] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9058884 kB) free!
[1562909548] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7674768 kB) free!
[1562909668] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11721264 kB) free.
[1562910158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (9025172 kB) free!
[1562910278] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (4864296 kB) free!
[1562910398] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (16681416 kB) free.
[1562910415] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1562910535] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1562910655] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562910868] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8879284 kB) free!
[1562910988] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7322060 kB) free!
[1562910998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (11793220 kB) free!
[1562911108] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11712600 kB) free.
[1562911118] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (7469416 kB) free!
[1562911238] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (3360456 kB) free!
[1562911838] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (15094876 kB) free.
[1562911910] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1562911990] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (1382952 kB) free!
[1562912030] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1562912110] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (1145872 kB) free!
[1562912150] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1562912210] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8788.06, 4762.49, 1937.35
[1562912230] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (957096 kB) free!
[1562912308] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9530088 kB) free!
[1562912332] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8972.47, 6211.15, 2826.85
[1562912428] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8130056 kB) free!
[1562912438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (10577016 kB) free!
[1562912452] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2503.20, 5097.71, 2862.45
[1562912548] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11710980 kB) free.
[1562912559] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (6356912 kB) free!
[1562912602] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562912679] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (2353296 kB) free!
[1562912830] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (120096 kB) free!
[1562913279] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1562913279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (13931976 kB) free.
[1562913431] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.0% (6767828 kB) free.
[1562913802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562913879] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (9741268 kB) free!
[1562913922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562913999] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (5588248 kB) free!
[1562914042] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562914119] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (1427796 kB) free!
[1562914642] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562914719] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (13490152 kB) free.
[1562914948] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9043664 kB) free!
[1562915068] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7638480 kB) free!
[1562915188] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11688908 kB) free.
[1562915242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562915319] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (9067300 kB) free!
[1562915363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562915439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (4864044 kB) free!
[1562915483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562915559] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (780992 kB) free!
[1562915697] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 443 processes
[1562915817] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 447 processes
[1562916083] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562916159] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13823680 kB) free.
[1562916388] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8473564 kB) free!
[1562916456] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175116 kB) free!
[1562916509] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6706300 kB) free!
[1562916576] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (158600 kB) free!
[1562916629] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11111736 kB) free.
[1562916683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562916696] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173976 kB) free!
[1562916759] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (9657288 kB) free!
[1562916803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562916879] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (5042104 kB) free!
[1562916923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562916999] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (977912 kB) free!
[1562917014] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.86, 4.09, 4.72
[1562917229] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9796264 kB) free!
[1562917252] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.67, 0.57, 16.96
[1562917349] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8472984 kB) free!
[1562917469] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6782204 kB) free!
[1562917523] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562917600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (14532892 kB) free.
[1562917716] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 34.8% (5656852 kB) free.
[1562917798] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562917896] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (2075408 kB) free!
[1562918123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562918200] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (6688564 kB) free!
[1562918243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562918320] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (1635488 kB) free!
[1562918440] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (244980 kB) free!
[1562918496] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (767424 kB) free!
[1562918916] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.5% (2688960 kB) free!
[1562918998] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 8146.80, 3881.96, 1513.65
[1562919036] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.5% (2513984 kB) free!
[1562919052] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 0.61, 3.04
[1562919118] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8513.75, 5433.43, 2369.35
[1562919156] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.5% (2349032 kB) free!
[1562919238] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8564.94, 6472.12, 3121.49
[1562919443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562919563] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562919640] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (15261432 kB) free.
[1562919756] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (851572 kB) free!
[1562920031] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (378884 kB) free!
[1562920151] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (176028 kB) free!
[1562920240] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (9570968 kB) free!
[1562920271] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1562920271] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (115188 kB) free!
[1562920360] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (5258688 kB) free!
[1562920480] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (1100212 kB) free!
[1562920763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562920883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562921003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562921080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (18991192 kB) free.
[1562921603] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562921669] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11673876 kB) free.
[1562922280] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (12869232 kB) free!
[1562922400] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (8026196 kB) free!
[1562922520] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (3425276 kB) free!
[1562922803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562922923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562923043] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562923120] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (18037568 kB) free.
[1562923469] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9793128 kB) free!
[1562923589] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.3% (8522640 kB) free!
[1562923709] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.1% (6968056 kB) free!
[1562923852] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.96, 76.11, 43.95
[1562923871] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6832704 kB) free.
[1562923972] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.83, 51.16, 38.71
[1562924092] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1562924092] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.61, 34.44, 34.09
[1562924175] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 41.3% (10115104 kB) free.
[1562924214] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.06, 6.92, 5.64
[1562924243] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562924309] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1562924309] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11252844 kB) free.
[1562924320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (8414308 kB) free!
[1562924334] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.38, 7.05, 5.84
[1562924440] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (3764528 kB) free!
[1562924454] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.70, 5.69, 5.51
[1562924560] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (17863940 kB) free.
[1562924692] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.27, 5.50, 18.35
[1562924843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562924909] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9829124 kB) free!
[1562924963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562925029] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.3% (8517692 kB) free!
[1562925054] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.86, 1.40, 3.24
[1562925084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562925149] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6829620 kB) free!
[1562925238] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.92, 4.44, 15.31
[1562925684] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562925696] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.1% (1970716 kB) free!
[1562925749] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11647084 kB) free.
[1562925760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10946488 kB) free!
[1562925880] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (5963280 kB) free!
[1562926000] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (1231904 kB) free!
[1562926284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562926296] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1562926296] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (153328 kB) free!
[1562926404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562926492] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.25, 0.85, 3.37
[1562926524] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562926574] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.99, 7.18, 5.01
[1562926600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (17578656 kB) free.
[1562926694] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.68, 8.03, 5.58
[1562926814] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.90, 8.28, 5.97
[1562927064] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.1% (2237076 kB) free!
[1562927200] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (13152984 kB) free!
[1562927320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (8730852 kB) free!
[1562927359] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1562927414] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 1.64, 3.62
[1562927425] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 51.3% (12655384 kB) free.
[1562927440] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (4354944 kB) free!
[1562927479] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1562927599] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1562927638] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.10, 3.26, 4.21
[1562928002] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.73, 3.79, 4.51
[1562928025] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189616 kB) free!
[1562928040] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (16939468 kB) free.
[1562928146] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181908 kB) free!
[1562928266] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175752 kB) free!
[1562929240] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (9576344 kB) free!
[1562929360] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (4897904 kB) free!
[1562929399] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1562929481] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (500372 kB) free!
[1562929949] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8570364 kB) free!
[1562930069] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6902732 kB) free!
[1562930081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (14426936 kB) free.
[1562930124] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562930189] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11450524 kB) free.
[1562930681] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (10381452 kB) free!
[1562930692] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.44, 14.65, 13.72
[1562930801] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (5663632 kB) free!
[1562930812] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.76, 10.38, 12.27
[1562930921] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (827348 kB) free!
[1562930932] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.68, 7.47, 10.97
[1562931199] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 362 processes
[1562931319] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 364 processes
[1562931440] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1562931521] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (16020996 kB) free.
[1562932121] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (11900420 kB) free!
[1562932132] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.39, 1.50, 4.00
[1562932241] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (7438968 kB) free!
[1562932271] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (1460652 kB) free!
[1562932361] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (3289716 kB) free!
[1562932391] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 18.6% (1367544 kB) free!
[1562932511] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 17.1% (1259524 kB) free!
[1562932524] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562932589] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9816612 kB) free!
[1562932645] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562932710] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8463588 kB) free!
[1562932732] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 42.34, 53.98, 25.98
[1562932765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562932830] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6746308 kB) free!
[1562932852] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.23, 36.69, 23.04
[1562932854] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 18.28, 93.05, 54.00
[1562932961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (15932404 kB) free.
[1562932972] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.28, 25.05, 20.43
[1562932974] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.77, 63.09, 47.75
[1562933038] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 6.30, 119.08, 98.94
[1562933094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.69, 43.54, 42.46
[1562933099] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5542 MB (10% inode=97%):
[1562933111] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (672508 kB) free!
[1562933158] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.76, 80.74, 87.33
[1562933278] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.15, 55.13, 77.15
[1562933365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562933430] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10882948 kB) free.
[1562933561] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12302248 kB) free!
[1562933681] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (8017148 kB) free!
[1562933801] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (3786684 kB) free!
[1562933829] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (4144880 kB) free!
[1562934014] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.20, 7.48, 5.83
[1562934135] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.82, 7.97, 6.21
[1562934255] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.69, 6.30, 5.84
[1562934294] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1562934294] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.20, 7.35, 18.33
[1562934299] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 8017 MB (15% inode=98%):
[1562934401] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (15545672 kB) free.
[1562934855] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562934855] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.23, 1.12, 3.22
[1562934912] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.6% (6811660 kB) free.
[1562935001] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12144540 kB) free!
[1562935078] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1562935078] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.55, 3.34, 16.23
[1562935122] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (8135060 kB) free!
[1562935165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562935242] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (3352236 kB) free!
[1562935285] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562935405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562935494] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7085.70, 3269.36, 1266.91
[1562935629] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188004 kB) free!
[1562935842] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (15191972 kB) free.
[1562936442] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (11492180 kB) free!
[1562936562] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (7524480 kB) free!
[1562936682] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (3445328 kB) free!
[1562936878] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.96, 19.28, 21.89
[1562937255] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.88, 7.44, 5.68
[1562937282] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (14882544 kB) free.
[1562937375] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.86, 6.05, 5.42
[1562937478] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.36, 5.54, 13.12
[1562937495] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 4.13, 4.79
[1562937655] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1562937696] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562937775] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1562937882] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10871004 kB) free!
[1562938002] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (6800112 kB) free!
[1562938122] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2732172 kB) free!
[1562938372] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 0.51, 18.67
[1562938722] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (14383096 kB) free.
[1562939278] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.96, 3.03, 4.49
[1562939322] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (10109132 kB) free!
[1562939430] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9170604 kB) free!
[1562939442] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (6174228 kB) free!
[1562939551] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7864192 kB) free!
[1562939562] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (2256468 kB) free!
[1562939671] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11548112 kB) free.
[1562940162] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (13636356 kB) free.
[1562940172] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.45, 0.39, 3.03
[1562940205] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562940495] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.74, 7.20, 5.21
[1562940615] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.88, 6.74, 5.31
[1562940735] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.18, 4.69, 4.73
[1562940763] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (9184164 kB) free!
[1562940883] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (5304424 kB) free!
[1562941003] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (1440340 kB) free!
[1562941406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562941471] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9684916 kB) free!
[1562941494] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.12, 4.16, 13.57
[1562941526] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562941591] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8396616 kB) free!
[1562941603] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (12181176 kB) free!
[1562941646] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562941711] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6869192 kB) free!
[1562942803] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (2658628 kB) free!
[1562942846] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562942911] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11087404 kB) free.
[1562943294] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.90, 2.77, 4.66
[1562943403] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (14272204 kB) free.
[1562943446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562943511] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9825284 kB) free!
[1562943566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562943631] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8563324 kB) free!
[1562943686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562943752] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7130824 kB) free!
[1562944003] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (10087880 kB) free!
[1562944123] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (6215768 kB) free!
[1562944243] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (2352132 kB) free!
[1562944287] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1562944352] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11382248 kB) free.
[1562944843] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (12482804 kB) free!
[1562944887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562944972] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.25, 55.07, 50.87
[1562945007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562945092] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.54, 37.00, 44.76
[1562945094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4.55, 55.10, 56.22
[1562945112] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562945127] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562945212] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.37, 24.89, 39.38
[1562945214] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.58, 38.33, 49.93
[1562945278] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4.32, 86.20, 74.76
[1562945334] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.13, 27.03, 44.38
[1562945398] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.00, 58.61, 66.04
[1562945518] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.08, 40.20, 58.39
[1562945552] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9022804 kB) free!
[1562945575] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1562945672] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7773640 kB) free!
[1562945695] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1562945792] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11511328 kB) free.
[1562946043] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (4159268 kB) free!
[1562946295] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562946392] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7833240 kB) free!
[1562946412] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 0.89, 11.14
[1562946415] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1562946512] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11518320 kB) free.
[1562946534] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.44, 3.53, 14.79
[1562946535] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1562946643] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13781560 kB) free.
[1562946718] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.61, 3.45, 18.11
[1562946735] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562947112] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7921000 kB) free!
[1562947135] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1562947232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11511748 kB) free.
[1562947243] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (9862908 kB) free!
[1562947255] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1562947363] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (6215680 kB) free!
[1562947484] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (2639068 kB) free!
[1562947612] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 0.46, 3.34
[1562947832] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7913776 kB) free!
[1562947952] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11517780 kB) free.
[1562948084] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (12863364 kB) free!
[1562948128] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562948334] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.76, 3.68, 4.77
[1562948518] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.53, 2.64, 4.94
[1562948552] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7957172 kB) free!
[1562948673] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11518436 kB) free.
[1562949272] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7949608 kB) free!
[1562949328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562949392] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11509416 kB) free.
[1562949448] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562949884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (18675052 kB) free.
[1562949992] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7918852 kB) free!
[1562950048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562950112] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11518392 kB) free.
[1562950168] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562950712] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7896912 kB) free!
[1562950768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562950832] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11506400 kB) free.
[1562950888] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562951432] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8031996 kB) free!
[1562951489] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562951552] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11514960 kB) free.
[1562951609] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562951684] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (10112836 kB) free!
[1562951804] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (19109176 kB) free.
[1562952152] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8048136 kB) free!
[1562952209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562952272] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11517304 kB) free.
[1562952329] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562952872] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8021148 kB) free!
[1562952929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562952993] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11510192 kB) free.
[1562953004] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12189264 kB) free!
[1562953049] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562953124] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (8749100 kB) free!
[1562953244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (18032256 kB) free.
[1562953592] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8017772 kB) free!
[1562953649] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562953712] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11514516 kB) free.
[1562953734] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 87.72, 26.15, 11.45
[1562953769] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562953844] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12707480 kB) free!
[1562953854] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 143.94, 68.02, 28.58
[1562953964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (9414704 kB) free!
[1562953974] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 152.46, 96.22, 43.69
[1562954084] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (18304872 kB) free.
[1562954312] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7992268 kB) free!
[1562954370] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562954432] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11507388 kB) free.
[1562954490] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562954685] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11846984 kB) free!
[1562954805] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (19141900 kB) free.
[1562955032] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7950520 kB) free!
[1562955090] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562955152] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11513592 kB) free.
[1562955210] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562955752] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7984332 kB) free!
[1562955774] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.61, 4.32, 13.00
[1562955810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562955872] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11513624 kB) free.
[1562955930] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes