[1576866818] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 22.3% (5455392 kB) free.
[1576867295] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (207564 kB) free!
[1576867319] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (233268 kB) free!
[1576867415] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196792 kB) free!
[1576867418] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172972 kB) free!
[1576867439] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (229584 kB) free!
[1576867467] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186232 kB) free!
[1576867535] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184292 kB) free!
[1576867538] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168240 kB) free!
[1576867559] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (224028 kB) free!
[1576867587] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178288 kB) free!
[1576867658] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164216 kB) free!
[1576867707] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184608 kB) free!
[1576870108] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.4% (10714612 kB) free.
[1576870708] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181460 kB) free!
[1576870828] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193664 kB) free!
[1576870845] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9854356 kB) free!
[1576870948] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172980 kB) free!
[1576870965] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9506608 kB) free!
[1576871085] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (9076340 kB) free!
[1576871685] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3297148 kB) free!
[1576875665] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232156 kB) free!
[1576875785] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.6% (43647532 kB) free.
[1576876348] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1576876348] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 31.1% (7686056 kB) free.
[1576877916] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1550) < 103.255.146.164 [21/Dec/2019:03:03:21 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576878036] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1675) < 103.209.99.7 [21/Dec/2019:03:10:14 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576878156] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(108) < 172.105.41.64 [21/Dec/2019:03:11:22 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1576878388] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/21 03:10:02 [error] 28867#0: *24065945 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:6221/6177, bytes from/to upstream:6177/6221
[1576878508] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576878757] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576882948] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1576883331] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1576883451] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576886309] SERVICE ALERT: webservercluster1;Total Processes;OK;HARD;1;PROCS OK: 469 processes
[1576888348] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (4345592 kB) free!
[1576888468] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (867536 kB) free!
[1576888588] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201256 kB) free!
[1576889160] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.4% (21669884 kB) free.
[1576889188] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.4% (18357800 kB) free.
[1576889258] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.2% (21128004 kB) free.
[1576892276] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.5% (38706780 kB) free.
[1576892308] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 07:00:01 [error] 28866#0: *24228912 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:113/17, bytes from/to upstream:17/113
[1576892323] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.4% (22895316 kB) free.
[1576892428] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576893051] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576893172] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576895326] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.2% (2754872 kB) free!
[1576895446] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.1% (2726300 kB) free!
[1576895566] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2608208 kB) free!
[1576897366] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2352816 kB) free!
[1576897588] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4944648 kB) free!
[1576897708] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4902612 kB) free!
[1576897828] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (4867464 kB) free!
[1576898571] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576898691] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576899028] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/21 08:57:45 [error] 28871#0: *24268889 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:125/17, bytes from/to upstream:17/125
[1576899148] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576899473] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1576899473] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.8% (22757420 kB) free.
[1576899748] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 09:11:03 [error] 28871#0: *24273520 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1576899868] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576900468] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/21 09:17:19 [error] 28874#0: *24275633 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1576900588] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576902388] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 09:55:57 [error] 28873#0: *24288799 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:372/17, bytes from/to upstream:17/372
[1576902508] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576903228] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1550436 kB) free!
[1576904308] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 10:22:46 [error] 28872#0: *24297989 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1576904428] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576905028] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 10:34:04 [error] 28866#0: *24301937 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1576905149] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576906657] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22320368 kB) free.
[1576906962] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 10.0% (2450280 kB) free!
[1576907548] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 11:14:57 [error] 28872#0: *24316189 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:242/17, bytes from/to upstream:17/242
[1576907668] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576908145] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1576908265] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 337 processes
[1576908268] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 11:31:16 [error] 28873#0: *24321879 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:225/17, bytes from/to upstream:17/225
[1576908385] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1576908389] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576908459] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3721732 kB) free!
[1576908477] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1576908579] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.7% (2629600 kB) free!
[1576908699] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (1085896 kB) free!
[1576909038] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4218140 kB) free!
[1576909158] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.4% (3910864 kB) free!
[1576909278] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.9% (3578408 kB) free!
[1576909550] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9114292 kB) free!
[1576909671] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.7% (8740480 kB) free!
[1576909791] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.0% (8397468 kB) free!
[1576909879] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 78.1% (17533656 kB) free.
[1576910391] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.7% (27955408 kB) free.
[1576910451] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.2% (22357748 kB) free.
[1576910483] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4823548 kB) free!
[1576910603] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.8% (22915224 kB) free.
[1576910788] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 12:16:26 [error] 28875#0: *24337124 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:248/17, bytes from/to upstream:17/248
[1576910908] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576911466] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4837076 kB) free!
[1576911508] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 12:22:42 [error] 28865#0: *24339128 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:201/17, bytes from/to upstream:17/201
[1576911586] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4254800 kB) free!
[1576911629] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/21 12:29:13 [error] 28872#0: *24341140 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:233/17, bytes from/to upstream:17/233
[1576911706] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2688232 kB) free!
[1576911749] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576912306] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1576912306] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1656748 kB) free!
[1576913549] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/21 12:58:36 [error] 28866#0: *24351150 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:248/17, bytes from/to upstream:17/248
[1576913669] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/21 13:03:37 [error] 28871#0: *24353262 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:253/17, bytes from/to upstream:17/253
[1576913789] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576913914] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22059056 kB) free.
[1576914553] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 91.5% (22419032 kB) free.
[1576915088] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8667172 kB) free!
[1576915208] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8606604 kB) free!
[1576915328] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (7524720 kB) free!
[1576917389] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 14:02:03 [error] 28873#0: *24377805 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:253/17, bytes from/to upstream:17/253
[1576917509] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576918230] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.5% (38704720 kB) free.
[1576918385] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9556156 kB) free!
[1576918506] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9410684 kB) free!
[1576918626] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9286240 kB) free!
[1576918677] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4911064 kB) free!
[1576918718] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (4610256 kB) free!
[1576918797] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (4750400 kB) free!
[1576918838] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.8% (2914364 kB) free!
[1576918917] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4597624 kB) free!
[1576918958] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181392 kB) free!
[1576920361] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4838228 kB) free!
[1576920481] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4810368 kB) free!
[1576920509] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 14:57:53 [error] 28875#0: *24399284 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:179/17, bytes from/to upstream:17/179
[1576920601] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4756896 kB) free!
[1576920629] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576920985] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 389 processes
[1576921105] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 394 processes
[1576921225] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 389 processes
[1576921317] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1947696 kB) free!
[1576922430] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 15:29:38 [error] 28875#0: *24412218 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1576922550] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576923426] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4486180 kB) free!
[1576923750] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 15:52:19 [error] 28875#0: *24421379 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1576923870] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576924834] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8769928 kB) free!
[1576924954] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.4% (5634196 kB) free!
[1576925074] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (5745528 kB) free!
[1576925458] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4406368 kB) free!
[1576925558] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1576925579] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (4292776 kB) free!
[1576925670] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 16:23:48 [error] 28863#0: *24434279 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:225/17, bytes from/to upstream:17/225
[1576925699] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (4160452 kB) free!
[1576925791] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/21 16:26:18 [error] 28865#0: *24435285 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:225/17, bytes from/to upstream:17/225
[1576925911] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576926080] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9490936 kB) free!
[1576926200] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9401408 kB) free!
[1576926275] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1576926275] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4530520 kB) free!
[1576926320] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (9312884 kB) free!
[1576927201] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2323496 kB) free!
[1576927502] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199104 kB) free!
[1576927711] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 16:53:56 [error] 28875#0: *24446749 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:251/17, bytes from/to upstream:17/251
[1576927717] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 567 processes
[1576927831] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576927837] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 573 processes
[1576927957] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 575 processes
[1576929032] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 17:17:45 [error] 28874#0: *24456478 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:250/17, bytes from/to upstream:17/250
[1576929152] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576929157] SERVICE ALERT: webservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 614 processes
[1576929207] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.1% (21490728 kB) free.
[1576929238] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.7% (21996492 kB) free.
[1576929291] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 93.5% (46080436 kB) free.
[1576929318] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.3% (38599272 kB) free.
[1576929365] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (45291180 kB) free.
[1576929392] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.5% (22181432 kB) free.
[1576929417] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.8% (22005908 kB) free.
[1576930418] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (1029748 kB) free!
[1576930538] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (893940 kB) free!
[1576930578] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 983.44, 1011.23, 556.59
[1576930658] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (799248 kB) free!
[1576930698] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 133.67, 677.01, 489.23
[1576930952] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/21 17:49:07 [error] 28871#0: *24469319 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:241/17, bytes from/to upstream:17/241
[1576931072] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576931672] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 18:00:23 [error] 28875#0: *24473956 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:253/17, bytes from/to upstream:17/253
[1576931792] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576932447] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.8% (22166564 kB) free.
[1576933317] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24859208 kB) free!
[1576933437] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19624292 kB) free!
[1576933557] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15508912 kB) free!
[1576933958] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1382) < 103.255.146.164 [21/Dec/2019:18:37:16 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576934078] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576934757] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27026792 kB) free.
[1576935928] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1576936557] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17208688 kB) free!
[1576936592] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 19:21:03 [error] 28873#0: *24507261 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:150/17, bytes from/to upstream:17/150
[1576936677] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.8% (45896780 kB) free.
[1576936713] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576937277] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16131848 kB) free!
[1576937397] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12817548 kB) free!
[1576937518] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42917336 kB) free.
[1576938117] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15797432 kB) free!
[1576938237] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (45007032 kB) free.
[1576938512] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/21 19:50:37 [error] 28875#0: *24519083 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:127/17, bytes from/to upstream:17/127
[1576938632] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576938837] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16853996 kB) free!
[1576938857] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1169_n"
[1576938957] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.1% (13274808 kB) free!
[1576938978] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;2;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1170_n"
[1576939077] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41324188 kB) free.
[1576939100] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1576939677] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13652016 kB) free!
[1576939797] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43261960 kB) free.
[1576940397] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14254088 kB) free!
[1576940517] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44885056 kB) free.
[1576941117] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19681492 kB) free!
[1576941237] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44516636 kB) free.
[1576941837] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20552060 kB) free!
[1576941957] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44198716 kB) free.
[1576942058] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (756008 kB) free!
[1576942178] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177684 kB) free!
[1576942298] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192232 kB) free!
[1576942447] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173020 kB) free!
[1576942557] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20234356 kB) free!
[1576942567] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169008 kB) free!
[1576942595] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (2325944 kB) free!
[1576942626] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (627028 kB) free!
[1576942678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43937044 kB) free.
[1576942687] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182032 kB) free!
[1576942715] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (2285376 kB) free!
[1576942747] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (624568 kB) free!
[1576942835] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2243180 kB) free!
[1576942867] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (616328 kB) free!
[1576943128] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9653536 kB) free!
[1576943248] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9653356 kB) free!
[1576943277] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20959084 kB) free!
[1576943368] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (9652568 kB) free!
[1576943397] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43742328 kB) free.
[1576943997] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20599752 kB) free!
[1576944117] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40996480 kB) free.
[1576944717] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18131764 kB) free!
[1576944837] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.8% (51202448 kB) free.