[1577903402] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12266764 kB) free!
[1577903522] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34515828 kB) free.
[1577903859] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186708 kB) free!
[1577903912] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191972 kB) free!
[1577903939] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.3% (27722868 kB) free.
[1577904034] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180768 kB) free!
[1577904122] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13254316 kB) free!
[1577904155] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186488 kB) free!
[1577904156] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184260 kB) free!
[1577904242] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36035596 kB) free.
[1577904842] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13381628 kB) free!
[1577904962] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35858972 kB) free.
[1577905543] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194060 kB) free!
[1577905562] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13179980 kB) free!
[1577905603] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 66.4% (16274472 kB) free.
[1577905682] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35586264 kB) free.
[1577905812] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.5% (22911884 kB) free.
[1577905854] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21946168 kB) free.
[1577905890] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.7% (21977772 kB) free.
[1577905931] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21875940 kB) free.
[1577905983] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577906282] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12449104 kB) free!
[1577906402] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34073436 kB) free.
[1577906586] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/02 00:47:57 [error] 20757#0: *21780059 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577906706] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577907306] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2020/01/02 01:02:57 [error] 20760#0: *21782503 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577907426] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2020/01/02 01:06:22 [error] 20751#0: *21783064 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577907546] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577907602] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14930188 kB) free!
[1577907657] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195668 kB) free!
[1577907722] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36695544 kB) free.
[1577907777] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189116 kB) free!
[1577907897] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1577907897] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172348 kB) free!
[1577908146] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/02 01:15:44 [error] 20756#0: *21784595 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577908166] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9174072 kB) free!
[1577908267] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577908286] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (8288932 kB) free!
[1577908322] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14532028 kB) free!
[1577908338] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (1367900 kB) free!
[1577908406] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (7101460 kB) free!
[1577908443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36614884 kB) free.
[1577908458] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (204492 kB) free!
[1577908578] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168060 kB) free!
[1577908866] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2020/01/02 01:30:44 [error] 20756#0: *21787046 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577908986] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577909006] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2248276 kB) free!
[1577909042] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13118920 kB) free!
[1577909162] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31078664 kB) free.
[1577909586] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2020/01/02 01:41:42 [error] 20750#0: *21788829 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577909706] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577910306] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(11) < 2020/01/02 01:46:42 [error] 20759#0: *21789673 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577910426] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2020/01/02 01:56:22 [error] 20760#0: *21791246 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577910547] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577911146] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2020/01/02 02:01:21 [error] 20759#0: *21792074 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577911266] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577911747] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.0% (21825740 kB) free.
[1577911802] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.9% (38873720 kB) free.
[1577911839] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 78.7% (19413676 kB) free.
[1577911866] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2020/01/02 02:16:41 [error] 20761#0: *21794580 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577911986] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577912277] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (212724 kB) free!
[1577912337] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8551172 kB) free!
[1577912397] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230496 kB) free!
[1577912446] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189956 kB) free!
[1577912457] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.6% (44659224 kB) free.
[1577912517] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241404 kB) free!
[1577912566] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (1099600 kB) free!
[1577912586] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2020/01/02 02:30:43 [error] 20750#0: *21796869 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577912686] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1577912686] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191736 kB) free!
[1577912706] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577913117] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.4% (42580948 kB) free.
[1577913307] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/02 02:36:21 [error] 20755#0: *21797777 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577913427] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(15) < 2020/01/02 02:45:45 [error] 20750#0: *21799314 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577913547] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577914147] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/02 02:51:42 [error] 20755#0: *21800263 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577914267] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2020/01/02 03:00:42 [error] 20755#0: *21801715 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577914387] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(6) < 2020/01/02 03:02:57 [error] 20752#0: *21802109 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577917223] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1292.65, 664.00, 265.59
[1577917343] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1330.05, 886.30, 395.36
[1577917463] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 879.76, 927.92, 473.95
[1577917486] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2606824 kB) free!
[1577917760] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1577918063] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.82, 125.17, 248.74
[1577918686] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2424104 kB) free!
[1577918760] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.1% (22487476 kB) free.
[1577918807] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577919414] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2020/01/02 04:26:42 [error] 20752#0: *21811808 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577919534] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577920134] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(11) < 2020/01/02 04:32:57 [error] 20759#0: *21812843 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577920254] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577920854] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2020/01/02 04:47:57 [error] 20756#0: *21815290 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577920974] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577921575] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/02 05:01:22 [error] 20754#0: *21817471 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577921695] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2020/01/02 05:02:57 [error] 20752#0: *21817734 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577921815] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577921815] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1577921935] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1577922415] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/02 05:15:38 [error] 20755#0: *21819782 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577922535] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2020/01/02 05:17:57 [error] 20757#0: *21820173 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577922655] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577923255] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/02 05:30:42 [error] 20757#0: *21822238 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577923375] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577923735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1577923855] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1577923976] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/02 05:36:42 [error] 20757#0: *21823233 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577924096] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577924696] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2020/01/02 05:52:57 [error] 20757#0: *21825906 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577924816] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577925416] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/02 06:00:46 [error] 20760#0: *21827183 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577925536] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1577925536] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577925559] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193688 kB) free!
[1577925615] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180932 kB) free!
[1577925679] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194852 kB) free!
[1577925698] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195804 kB) free!
[1577925735] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196348 kB) free!
[1577925799] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190312 kB) free!
[1577925818] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (183908 kB) free!
[1577925855] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.6% (21708360 kB) free.
[1577925938] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193428 kB) free!
[1577925968] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195608 kB) free!
[1577926088] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21983548 kB) free.
[1577926136] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2020/01/02 06:17:57 [error] 20758#0: *21829970 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577926257] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577926399] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.3% (22387432 kB) free.
[1577926448] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.3% (22619592 kB) free.
[1577926856] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/02 06:30:40 [error] 20757#0: *21832032 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577926976] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2020/01/02 06:31:42 [error] 20752#0: *21832218 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577927096] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577927696] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/01/02 06:41:42 [error] 20754#0: *21833844 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577927816] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(6) < 2020/01/02 06:45:38 [error] 20757#0: *21834501 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577927924] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) <                  {amqp_error,frame_error,
[1577927936] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577928045] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1577928536] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2020/01/02 06:56:42 [error] 20760#0: *21836314 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577928657] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2020/01/02 07:00:44 [error] 20757#0: *21836978 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577928777] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2020/01/02 07:01:42 [error] 20753#0: *21837143 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577929089] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1577931655] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1577931775] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1577932977] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577933577] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2020/01/02 08:21:21 [error] 20761#0: *21850098 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577933697] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577934297] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2020/01/02 08:32:57 [error] 20757#0: *21851979 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577934417] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2020/01/02 08:36:41 [error] 20755#0: *21852580 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577934537] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577935137] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2020/01/02 08:46:41 [error] 20756#0: *21854222 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577935258] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577935857] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2020/01/02 09:00:42 [error] 20756#0: *21856480 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577935977] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2020/01/02 09:01:42 [error] 20757#0: *21856641 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577936097] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577936697] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/01/02 09:11:42 [error] 20759#0: *21858299 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577936817] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(15) < 2020/01/02 09:15:46 [error] 20755#0: *21858969 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577936937] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577937538] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2020/01/02 09:27:57 [error] 20761#0: *21860934 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577937658] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2020/01/02 09:30:45 [error] 20759#0: *21861401 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577937778] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2020/01/02 09:31:42 [error] 20752#0: *21861568 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577942288] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (1128792 kB) free!
[1577942408] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (1030716 kB) free!
[1577942528] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (941044 kB) free!
[1577943459] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.9% (21692528 kB) free.
[1577944456] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184576 kB) free!
[1577944576] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190324 kB) free!
[1577944668] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187104 kB) free!
[1577944696] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186220 kB) free!
[1577944788] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187848 kB) free!
[1577944908] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196440 kB) free!
[1577945450] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1280) < 103.255.146.164 [02/Jan/2020:11:33:20 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1577945570] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577945607] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4347016 kB) free!
[1577945727] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (4126424 kB) free!
[1577945847] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (3882848 kB) free!
[1577947047] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1639140 kB) free!
[1577947052] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1577947052] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22834044 kB) free.
[1577947089] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.3% (22379952 kB) free.
[1577947126] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.9% (22765264 kB) free.
[1577951256] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (2184384 kB) free!
[1577951298] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (1377392 kB) free!
[1577951376] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (1990224 kB) free!
[1577951405] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8416476 kB) free!
[1577951418] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195104 kB) free!
[1577951496] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1761516 kB) free!
[1577951525] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8417192 kB) free!
[1577951538] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1577951538] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178044 kB) free!
[1577951645] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (8416504 kB) free!
[1577951762] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21755220 kB) free!
[1577951882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44311664 kB) free.
[1577951935] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179980 kB) free!
[1577952055] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177648 kB) free!
[1577952175] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173392 kB) free!
[1577952339] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.5% (4295524 kB) free!
[1577952459] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.7% (4099028 kB) free!
[1577952482] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19882032 kB) free!
[1577952579] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.0% (3920984 kB) free!
[1577952603] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42619308 kB) free.
[1577952845] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225016 kB) free!
[1577953202] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18435316 kB) free!
[1577953318] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8748536 kB) free!
[1577953322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42273764 kB) free.
[1577953438] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8492304 kB) free!
[1577953558] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.7% (8234904 kB) free!
[1577953779] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2368232 kB) free!
[1577953922] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16898692 kB) free!
[1577954042] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39687880 kB) free.
[1577954055] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4542864 kB) free!
[1577954175] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4382700 kB) free!
[1577954295] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4236884 kB) free!
[1577954642] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14920412 kB) free!
[1577954762] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35348048 kB) free.
[1577955324] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.9% (3113324 kB) free!
[1577955358] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3668980 kB) free!
[1577955362] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11507776 kB) free!
[1577955445] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.5% (2801932 kB) free!
[1577955482] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.5% (48073500 kB) free.
[1577955565] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.3% (2545760 kB) free!
[1577955770] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2020/01/02 14:31:35 [error] 28872#0: *31285157 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:913988/13708767, bytes from/to upstream:13708767/913988
[1577955890] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577956053] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 75.8% (18577316 kB) free.
[1577956082] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24686452 kB) free!
[1577956086] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 71.0% (17525804 kB) free.
[1577956095] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2186128 kB) free!
[1577956125] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.0% (22064740 kB) free.
[1577956140] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9250624 kB) free!
[1577956146] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 59.9% (29495380 kB) free.
[1577956165] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (967208 kB) free!
[1577956179] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 27.7% (13637924 kB) free.
[1577956202] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.7% (48439220 kB) free.
[1577956204] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 56.9% (12766880 kB) free.
[1577956224] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 57.9% (14192744 kB) free.
[1577956256] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.7% (22724260 kB) free.
[1577956260] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.0% (8866848 kB) free!
[1577956380] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.2% (8466032 kB) free!
[1577956802] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24194232 kB) free!
[1577956922] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.3% (49180400 kB) free.
[1577957522] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24313696 kB) free!
[1577957580] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.9% (28036796 kB) free.
[1577957605] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3839380 kB) free!
[1577957642] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.5% (48086880 kB) free.
[1577957725] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3665288 kB) free!
[1577957845] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3492692 kB) free!
[1577957896] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4806180 kB) free!
[1577958016] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (4142956 kB) free!
[1577958136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1400756 kB) free!
[1577958243] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24942120 kB) free!
[1577958363] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36927312 kB) free.
[1577958445] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.1% (22990300 kB) free.
[1577958513] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1577958627] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164132 kB) free!
[1577958633] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1577958747] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175884 kB) free!
[1577958867] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176100 kB) free!
[1577958963] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24735420 kB) free!
[1577959083] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20562856 kB) free!
[1577959203] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44469800 kB) free.
[1577959758] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9844744 kB) free!
[1577959803] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20678364 kB) free!
[1577959878] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9840896 kB) free!
[1577959923] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44561740 kB) free.
[1577959970] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(889) < 103.255.146.164 [02/Jan/2020:15:35:34 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1577959998] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9840748 kB) free!
[1577960091] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577960523] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19882416 kB) free!
[1577960543] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1577960643] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44839952 kB) free.
[1577960663] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1577960783] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1577961243] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20055976 kB) free!
[1577961363] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43915932 kB) free.
[1577961963] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20731300 kB) free!
[1577962083] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16826544 kB) free!
[1577962203] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41410248 kB) free.
[1577962490] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(757) < 103.255.146.164 [02/Jan/2020:16:15:56 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1577962610] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577962803] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17193572 kB) free!
[1577962924] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42410964 kB) free.
[1577963211] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(878) < 103.255.146.164 [02/Jan/2020:16:30:49 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1577963325] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1577963331] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577963523] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17223936 kB) free!
[1577963643] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41541056 kB) free.
[1577964199] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (253472 kB) free!
[1577964243] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17930828 kB) free!
[1577964363] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13846664 kB) free!
[1577964483] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39060512 kB) free.
[1577965083] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13865144 kB) free!
[1577965203] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39780792 kB) free.
[1577965375] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1577965495] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1577965803] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14342464 kB) free!
[1577965923] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10251952 kB) free!
[1577966043] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35912928 kB) free.
[1577966457] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9680936 kB) free!
[1577966536] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1577966577] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9539176 kB) free!
[1577966643] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10707404 kB) free!
[1577966697] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (9420804 kB) free!
[1577966763] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37695924 kB) free.
[1577967363] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12232268 kB) free!
[1577967483] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8035272 kB) free!
[1577967604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34193064 kB) free.
[1577968203] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9031932 kB) free!
[1577968323] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35005808 kB) free.
[1577968725] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4650892 kB) free!
[1577968833] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 337 processes
[1577968845] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4529240 kB) free!
[1577968923] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9305888 kB) free!
[1577968953] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 4 warn out of 342 processes
[1577968965] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4398292 kB) free!
[1577969043] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35288628 kB) free.
[1577969073] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1577969643] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12081520 kB) free!
[1577969673] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1577969763] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8079268 kB) free!
[1577969883] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33884348 kB) free.
[1577969931] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(653) < 103.255.146.164 [02/Jan/2020:18:27:26 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1577970051] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577970483] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7429972 kB) free!
[1577970603] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34046108 kB) free.
[1577971203] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10430644 kB) free!
[1577971323] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6444784 kB) free!
[1577971443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33379076 kB) free.
[1577971965] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2454972 kB) free!
[1577972043] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8110548 kB) free!
[1577972163] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (4013412 kB) free!
[1577972284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27613024 kB) free.
[1577972883] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (4019000 kB) free!
[1577973003] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (45969796 kB) free.
[1577973603] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22328424 kB) free!
[1577973723] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (47921120 kB) free.
[1577973897] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4855896 kB) free!
[1577974323] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24352512 kB) free!
[1577974443] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20339252 kB) free!
[1577974563] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46674404 kB) free.
[1577975063] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 776.29, 486.90, 210.96
[1577975163] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22501648 kB) free!
[1577975183] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 263.81, 431.74, 227.37
[1577975283] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.9% (47330316 kB) free.
[1577975304] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 36.27, 289.12, 199.89
[1577975883] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24619416 kB) free!
[1577976003] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20577380 kB) free!
[1577976051] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(819) < 103.255.146.164 [02/Jan/2020:20:06:51 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1577976123] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.7% (47035304 kB) free.
[1577976172] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577976723] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20685560 kB) free!
[1577976844] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16872252 kB) free!
[1577976964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44597008 kB) free.
[1577977564] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17869468 kB) free!
[1577977583] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1577977684] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.0% (46140692 kB) free.
[1577978284] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19450020 kB) free!
[1577978404] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15153248 kB) free!
[1577978524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45486596 kB) free.
[1577979124] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20681932 kB) free!
[1577979244] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16334116 kB) free!
[1577979364] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42322804 kB) free.
[1577979964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19031912 kB) free!
[1577980084] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44907716 kB) free.
[1577980256] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 61.6% (15204444 kB) free.
[1577980305] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 75.1% (18410792 kB) free.
[1577980323] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 33.5% (16492860 kB) free.
[1577980355] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 49.4% (24364872 kB) free.
[1577980377] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.4% (15780524 kB) free.
[1577980490] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(31) < 2020/01/02 21:21:42 [error] 28875#0: *31669776 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.67.175, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577980611] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577980684] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21649980 kB) free!
[1577980804] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17347380 kB) free!
[1577980924] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42206096 kB) free.
[1577981466] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187620 kB) free!
[1577981524] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17043828 kB) free!
[1577981557] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.7% (5767748 kB) free!
[1577981586] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186664 kB) free!
[1577981645] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44376740 kB) free.
[1577981678] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.7% (5768292 kB) free!
[1577981706] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176440 kB) free!
[1577981798] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (5770452 kB) free!
[1577982244] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21436880 kB) free!
[1577982364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17713092 kB) free!
[1577982484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42054908 kB) free.
[1577983084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18931512 kB) free!
[1577983204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (44058440 kB) free.
[1577983539] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22663356 kB) free.
[1577983598] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1385008 kB) free!
[1577983804] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20089500 kB) free!
[1577983924] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (16571656 kB) free!
[1577984044] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40229080 kB) free.
[1577984644] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17834600 kB) free!
[1577984764] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42894096 kB) free.
[1577985347] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201940 kB) free!
[1577985364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19315704 kB) free!
[1577985467] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188744 kB) free!
[1577985484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43640784 kB) free.
[1577985587] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1577985587] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180764 kB) free!
[1577986084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21223368 kB) free!
[1577986204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43624660 kB) free.
[1577986473] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1577986593] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1577986713] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1577986805] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22089152 kB) free!
[1577986925] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40298852 kB) free.
[1577987525] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22728320 kB) free!
[1577987580] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234012 kB) free!
[1577987645] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43609628 kB) free.
[1577987700] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235892 kB) free!
[1577987820] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217020 kB) free!
[1577988245] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22825064 kB) free!
[1577988365] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43572284 kB) free.
[1577988965] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22899296 kB) free!
[1577989085] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43567756 kB) free.
[1577989346] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.0% (27575512 kB) free.
[1577989395] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.3% (22526900 kB) free.
[1577989436] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.5% (38181924 kB) free.
[1577989493] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577989685] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21894772 kB) free!