[1576607766] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4584792 kB) free!
[1576607887] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187556 kB) free!
[1576607984] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1576608007] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185980 kB) free!
[1576608105] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1576608176] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200012 kB) free!
[1576608296] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200156 kB) free!
[1576608360] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174364 kB) free!
[1576608416] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188232 kB) free!
[1576608480] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192188 kB) free!
[1576608600] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166764 kB) free!
[1576610588] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576611000] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 46.3% (11428532 kB) free.
[1576611600] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172492 kB) free!
[1576611720] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191596 kB) free!
[1576611840] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170316 kB) free!
[1576612342] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2019/12/18 01:14:57 [error] 20756#0: *17862519 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576612463] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576613049] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 554 processes
[1576613649] SERVICE ALERT: webservercluster1;Total Processes;OK;HARD;1;PROCS OK: 492 processes
[1576616147] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238492 kB) free!
[1576616267] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231076 kB) free!
[1576616387] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (244372 kB) free!
[1576616987] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.4% (43067704 kB) free.
[1576617840] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.7% (10279708 kB) free.
[1576618462] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(108) < 2019/12/18 03:02:24 [error] 20756#0: *17879163 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1576618467] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/12/18 03:01:16 [error] 28866#0: *22643062 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1576618582] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576618587] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/18 03:06:25 [error] 28861#0: *22644266 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1576618707] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576623840] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1576624371] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1576624491] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1576624582] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2019/12/18 04:41:24 [error] 20756#0: *17895202 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1576624703] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576625091] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576625211] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1576626992] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(44) < 103.209.99.7 [18/Dec/2019:05:20:14 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1576627011] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576627102] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(60) < 2019/12/18 05:20:14 [error] 20758#0: *17901512 no live upstreams while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "backend_13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576627112] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576627131] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576627222] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(77) < 2019/12/18 05:29:00 [error] 20758#0: *17903108 no live upstreams while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "backend_13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576627342] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576627707] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2019/12/18 05:32:34 [error] 28872#0: *22682020 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1576627712] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(78) < 203.212.70.170 [18/Dec/2019:05:37:18 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1576627771] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(16) < 49.50.86.152 [18/Dec/2019:05:32:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1576627828] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576627832] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576627891] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576627942] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(127) < 2019/12/18 05:37:39 [error] 20755#0: *17905225 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:1611/0, bytes from/to upstream:0/1611
[1576628063] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576629240] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181616 kB) free!
[1576629262] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1264) < 2019/12/18 06:01:03 [error] 20759#0: *17910580 no live upstreams while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "backend_3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576629360] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178036 kB) free!
[1576629382] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576629480] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197092 kB) free!
[1576629632] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1197) < 103.255.146.164 [18/Dec/2019:06:01:03 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576629727] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 79.3% (19558252 kB) free.
[1576629752] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576629763] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.7% (22477156 kB) free.
[1576629787] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 22.3% (10981572 kB) free.
[1576629828] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.3% (23116424 kB) free.
[1576629871] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 95.7% (23456692 kB) free.
[1576630335] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201952 kB) free!
[1576630455] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193600 kB) free!
[1576630575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182224 kB) free!
[1576633038] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1576633038] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22937508 kB) free.
[1576634331] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576634451] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1576640248] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1576641623] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1576641743] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1576642851] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1576642971] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1576644171] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576644292] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1576645047] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4810836 kB) free!
[1576645167] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4689704 kB) free!
[1576645287] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.4% (3806088 kB) free!
[1576645583] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2019/12/18 10:27:22 [error] 20761#0: *17953590 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576645703] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576646303] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(35) < 2019/12/18 10:47:39 [error] 20757#0: *17956987 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576646423] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(8) < 2019/12/18 10:49:02 [error] 20755#0: *17957252 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576646543] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576647143] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(353) < 2019/12/18 11:02:02 [error] 20757#0: *17960374 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576647152] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(34) < 103.255.146.164 [18/Dec/2019:11:02:01 +0530] TCP 502 0 0 0.001 "10.147.212.143:3000, backend_3000" "0, 0" "0, 0" "0.001, 0.000"
[1576647263] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(36) < 2019/12/18 11:03:41 [error] 20757#0: *17960735 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576647272] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(4) < 103.255.146.164 [18/Dec/2019:11:03:41 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576647384] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576647392] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576647687] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (599396 kB) free!
[1576647983] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(382) < 2019/12/18 11:16:17 [error] 20758#0: *17963718 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576647992] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(73) < 103.255.146.164 [18/Dec/2019:11:15:59 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576648103] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(96) < 2019/12/18 11:18:14 [error] 20760#0: *17964340 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576648112] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(2) < 103.255.146.164 [18/Dec/2019:11:16:40 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576648223] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1321) < 2019/12/18 11:19:35 [error] 20760#0: *17965905 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:193631/53, bytes from/to upstream:53/193631
[1576648233] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(1249) < 103.255.146.164 [18/Dec/2019:11:19:25 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576650195] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3266164 kB) free!
[1576650315] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (3019732 kB) free!
[1576650435] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.2% (2729116 kB) free!
[1576650991] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22513176 kB) free.
[1576651036] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 78.5% (17617488 kB) free.
[1576651978] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25113256 kB) free!
[1576652098] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.5% (53405704 kB) free.
[1576652795] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (3587000 kB) free!
[1576652915] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201712 kB) free!
[1576653035] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1576653035] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165560 kB) free!
[1576653860] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1576654366] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (4044052 kB) free!
[1576654486] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3950736 kB) free!
[1576654585] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.1% (22491044 kB) free.
[1576654606] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (3818456 kB) free!
[1576654997] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8811284 kB) free!
[1576655035] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4711272 kB) free!
[1576655117] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9477056 kB) free!
[1576655155] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (4501540 kB) free!
[1576655237] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (9436960 kB) free!
[1576655275] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4327756 kB) free!
[1576656260] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1576656380] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1576656406] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1980356 kB) free!
[1576656500] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1576657075] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2313880 kB) free!
[1576659385] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (4220796 kB) free!
[1576659437] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4943236 kB) free!
[1576659505] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (4012376 kB) free!
[1576659625] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.1% (3720516 kB) free!
[1576660787] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9419436 kB) free!
[1576660907] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9219212 kB) free!
[1576660962] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 71.6% (17558844 kB) free.
[1576661008] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.7% (38804188 kB) free.
[1576661027] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9046760 kB) free!
[1576661044] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 61.8% (15144016 kB) free.
[1576661425] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2425536 kB) free!
[1576661650] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 559 processes
[1576661770] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 561 processes
[1576661890] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 570 processes
[1576662023] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576662033] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576663090] SERVICE ALERT: webservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 626 processes
[1576663477] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4823280 kB) free!
[1576663498] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26307904 kB) free!
[1576663597] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4757744 kB) free!
[1576663618] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.4% (22909464 kB) free!
[1576663717] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (4679140 kB) free!
[1576663738] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43587404 kB) free.
[1576664338] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23202040 kB) free!
[1576664458] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (43975276 kB) free.
[1576664617] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3803164 kB) free!
[1576664628] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4654628 kB) free!
[1576664737] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (3803588 kB) free!
[1576664857] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3802964 kB) free!
[1576665058] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23892968 kB) free!
[1576665178] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20484520 kB) free!
[1576665298] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41272588 kB) free.
[1576665898] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20790152 kB) free!
[1576666019] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42500360 kB) free.
[1576666343] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 765 processes
[1576666463] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 764 processes
[1576666618] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22764424 kB) free!
[1576666738] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19413436 kB) free!
[1576666858] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40313824 kB) free.
[1576667423] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/18 16:39:51 [error] 20755#0: *18093221 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576667458] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20084744 kB) free!
[1576667544] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2019/12/18 16:40:51 [error] 20760#0: *18093436 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576667578] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41387912 kB) free.
[1576667664] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2019/12/18 16:43:01 [error] 20758#0: *18093785 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576667917] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2407056 kB) free!
[1576668178] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21552400 kB) free!
[1576668298] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18211108 kB) free!
[1576668418] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39404764 kB) free.
[1576669018] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18834444 kB) free!
[1576669078] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.1% (18542460 kB) free.
[1576669107] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.7% (38802440 kB) free.
[1576669136] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 24.2% (11908832 kB) free.
[1576669138] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39069492 kB) free.
[1576669166] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.6% (18539168 kB) free.
[1576669234] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(24175) < 103.255.146.164 [18/Dec/2019:17:07:47 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1576669354] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576669491] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1576669612] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1576669685] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (389716 kB) free!
[1576669738] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18835280 kB) free!
[1576669805] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (353972 kB) free!
[1576669858] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (14962360 kB) free!
[1576669925] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (317868 kB) free!
[1576669978] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.2% (49089140 kB) free.
[1576670290] SERVICE ALERT: webservercluster1;Total Processes;OK;HARD;1;PROCS OK: 349 processes
[1576670665] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576670996] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1576672465] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2019/12/18 18:04:02 [error] 20757#0: *18130035 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576672586] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/18 18:05:46 [error] 20758#0: *18130318 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576672706] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576672978] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25932556 kB) free!
[1576673098] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.8% (49867308 kB) free.
[1576673306] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2019/12/18 18:15:11 [error] 20756#0: *18131809 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576673426] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1576674027] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2019/12/18 18:26:36 [error] 20758#0: *18133643 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576674147] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/18 18:31:36 [error] 20759#0: *18134443 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576674267] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1576674267] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(2) < 2019/12/18 18:34:02 [error] 20758#0: *18134836 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1576675091] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 580 processes
[1576675211] SERVICE ALERT: webservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 585 processes
[1576675331] SERVICE ALERT: webservercluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 594 processes
[1576675463] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 609 processes
[1576675583] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 610 processes
[1576675703] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 614 processes
[1576675931] SERVICE ALERT: webservercluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 622 processes
[1576677847] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4814144 kB) free!
[1576677968] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186784 kB) free!
[1576678088] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175832 kB) free!
[1576678176] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (209452 kB) free!
[1576678296] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186432 kB) free!
[1576678366] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188816 kB) free!
[1576678416] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186436 kB) free!
[1576678486] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168972 kB) free!
[1576678606] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188664 kB) free!
[1576679312] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 382 processes
[1576679432] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 387 processes
[1576679525] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1576679552] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 495 processes
[1576679616] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (2809712 kB) free!
[1576679806] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (2722804 kB) free!
[1576680012] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.0% (22697092 kB) free.
[1576680057] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.3% (23360424 kB) free.
[1576680152] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1576680299] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23406944 kB) free!
[1576680419] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.8% (47196892 kB) free.
[1576681006] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2318844 kB) free!
[1576681019] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25257832 kB) free!
[1576681139] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.4% (21645096 kB) free!
[1576681259] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.6% (46898244 kB) free.
[1576681416] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2457996 kB) free!
[1576681859] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24578296 kB) free!
[1576681979] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20942584 kB) free!
[1576682099] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46655780 kB) free.
[1576682699] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24539248 kB) free!
[1576682819] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.1% (21201088 kB) free!
[1576682939] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (45009360 kB) free.
[1576683267] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1576683539] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20788016 kB) free!
[1576683659] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39295284 kB) free.
[1576684259] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25411836 kB) free!
[1576684379] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.7% (22022596 kB) free!
[1576684500] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (45963480 kB) free.
[1576685099] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24547580 kB) free!
[1576685219] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.1% (21235236 kB) free!
[1576685339] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45439220 kB) free.
[1576685939] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23836316 kB) free!
[1576686059] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20543068 kB) free!
[1576686179] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (43034360 kB) free.
[1576686246] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1576686246] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.8% (23000720 kB) free.
[1576686344] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1576686344] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 95.1% (23308800 kB) free.
[1576686779] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22675032 kB) free!
[1576686899] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45570792 kB) free.
[1576687499] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25615804 kB) free!
[1576687619] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22525932 kB) free!
[1576687739] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45467228 kB) free.
[1576688339] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25664100 kB) free!
[1576688459] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22599884 kB) free!
[1576688579] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45028448 kB) free.
[1576689180] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24919432 kB) free!
[1576689300] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39344592 kB) free.
[1576691055] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1576692300] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26255844 kB) free!
[1576692346] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1576692420] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.6% (23181196 kB) free!
[1576692540] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44342848 kB) free.
[1576693140] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24946240 kB) free!
[1576693260] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45027292 kB) free.